Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.pred_set
% (82853)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 738
% Passive clauses: 5177
% Generated clauses: 19863
% Final active clauses: 727
% Final passive clauses: 4111
% Input formulas: 45
% Initial clauses: 189
% 
% Duplicate literals: 912
% Trivial inequalities: 45
% Fw subsumption resolutions: 2398
% Fw demodulations: 437
% Bw demodulations: 11
% 
% Simple tautologies: 266
% Equational tautologies: 77
% Forward subsumptions: 8112
% 
% Binary resolution: 6592
% Factoring: 1350
% Forward superposition: 2588
% Backward superposition: 4940
% Self superposition: 108
% Equality factoring: 108
% Equality resolution: 7
% 
% Split clauses: 1450
% Split components: 3085
% SAT solver clauses: 1834
% SAT solver unit clauses: 26
% SAT solver binary clauses: 194
% 
% Sat splits: 1449
% 
% TWLsolver clauses: 1828
% TWLsolver calls for satisfiability: 431
% 
% Memory used [KB]: 9210
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.pred_set
% (82863)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 440
% Passive clauses: 2362
% Generated clauses: 10401
% Final active clauses: 432
% Final passive clauses: 1922
% Input formulas: 45
% Initial clauses: 189
% 
% Duplicate literals: 233
% Trivial inequalities: 38
% Fw subsumption resolutions: 1039
% Fw demodulations: 289
% Bw demodulations: 6
% Condensations: 222
% Global subsumptions: 17
% 
% Simple tautologies: 267
% Equational tautologies: 26
% Forward subsumptions: 3094
% 
% Binary resolution: 3714
% Factoring: 325
% Forward superposition: 1145
% Backward superposition: 762
% Self superposition: 22
% Equality factoring: 26
% Equality resolution: 7
% 
% Split clauses: 68
% Split components: 141
% SAT solver clauses: 2933
% SAT solver unit clauses: 205
% SAT solver binary clauses: 91
% SAT solver learnt clauses: 34
% SAT solver learnt literals: 10
% 
% Sat splits: 68
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 2865
% TWLsolver calls for satisfiability: 1816
% 
% Memory used [KB]: 11769
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.pred_set
% (82882)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 2966
% Passive clauses: 11463
% Generated clauses: 164890
% Final active clauses: 1553
% Final passive clauses: 1797
% Input formulas: 45
% Initial clauses: 189
% 
% Duplicate literals: 9011
% Trivial inequalities: 358
% Fw subsumption resolutions: 7127
% Fw demodulations: 20243
% Bw demodulations: 1603
% Condensations: 11783
% Global subsumptions: 1048
% 
% Simple tautologies: 413
% Equational tautologies: 3002
% Forward subsumptions: 97969
% Bw demodulations to eq. taut.: 235
% 
% Unit resulting resolution: 53539
% Forward superposition: 30558
% Backward superposition: 22622
% Self superposition: 325
% Equality factoring: 264
% Equality resolution: 24
% 
% Split clauses: 404
% Split components: 831
% SAT solver clauses: 15161
% SAT solver unit clauses: 1870
% SAT solver binary clauses: 4193
% 
% Sat splits: 404
% Sat splitting refutations: 1839
% 
% Memory used [KB]: 80467
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.pred_set
% (83194)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 2534
% Passive clauses: 45183
% Generated clauses: 193994
% Final active clauses: 2500
% Final passive clauses: 40930
% Input formulas: 45
% Initial clauses: 188
% 
% Selected by SInE selection: 44
% SInE iterations: 1
% 
% Duplicate literals: 5315
% Trivial inequalities: 514
% Fw subsumption resolutions: 21777
% Fw demodulations: 4725
% Bw demodulations: 21
% 
% Simple tautologies: 5314
% Equational tautologies: 443
% Forward subsumptions: 100398
% 
% Binary resolution: 150005
% Factoring: 476
% Forward superposition: 7436
% Backward superposition: 2086
% Self superposition: 20
% Equality factoring: 2
% Equality resolution: 50
% 
% Split clauses: 11743
% Split components: 27567
% SAT solver clauses: 12696
% SAT solver unit clauses: 46
% SAT solver binary clauses: 596
% 
% Sat splits: 11743
% Sat splitting refutations: 63
% 
% Memory used [KB]: 57824
% Time elapsed: 4.0000 s
% ------------------------------
% ------------------------------
% remaining time: 142 next slice time: 6
lrs+1002_2:3_br=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:sas=minisat:stl=30:sac=on:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none:sp=reverse_arity:urr=on_4 on thm.pred_set
% (83295)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 587
% Passive clauses: 2517
% Generated clauses: 23418
% Final active clauses: 577
% Final passive clauses: 1821
% Input formulas: 45
% Initial clauses: 189
% 
% Duplicate literals: 598
% Trivial inequalities: 256
% Fw subsumption resolutions: 2541
% Fw demodulations: 564
% Bw demodulations: 30
% Global subsumptions: 22
% 
% Simple tautologies: 462
% Equational tautologies: 73
% Forward subsumptions: 15781
% 
% Unit resulting resolution: 6152
% Factoring: 27
% Forward superposition: 6235
% Backward superposition: 6569
% Self superposition: 128
% Equality factoring: 6
% Equality resolution: 26
% 
% Split clauses: 32
% Split components: 64
% SAT solver clauses: 2857
% SAT solver unit clauses: 404
% SAT solver binary clauses: 252
% 
% Sat splits: 32
% Sat splitting refutations: 284
% 
% Memory used [KB]: 9594
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 135 next slice time: 15
lrs+1011_12_bs=on:bsr=unit_only:cond=on:gs=on:gsaa=from_current:gsssp=full:nwc=1.1:sas=minisat:stl=60:sos=all:sac=on:sdd=off:sser=off:sfr=on:ssfp=100000:ssfq=1.2:ssnc=none:sp=reverse_arity:updr=off_14 on thm.pred_set
% (83329)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1063
% Passive clauses: 6619
% Generated clauses: 17410
% Final active clauses: 703
% Final passive clauses: 4459
% Input formulas: 45
% Initial clauses: 189
% 
% Duplicate literals: 531
% Trivial inequalities: 52
% Fw subsumption resolutions: 876
% Bw subsumption resolutions: 626
% Fw demodulations: 816
% Bw demodulations: 49
% Condensations: 1350
% Global subsumptions: 2
% 
% Simple tautologies: 1042
% Equational tautologies: 41
% Forward subsumptions: 5972
% Backward subsumptions: 902
% 
% Binary resolution: 8282
% Factoring: 42
% Forward superposition: 1765
% Backward superposition: 1665
% Self superposition: 26
% Equality factoring: 4
% Equality resolution: 12
% 
% Split clauses: 1
% Split components: 2
% SAT solver clauses: 6723
% SAT solver unit clauses: 352
% SAT solver binary clauses: 1473
% 
% Sat splits: 1
% Sat splitting refutations: 99
% 
% Memory used [KB]: 10490
% Time elapsed: 1.600 s
% ------------------------------
% ------------------------------
% remaining time: 119 next slice time: 9
lrs+10_8:1_bd=preordered:bs=on:ccuc=first:fsr=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:nicw=on:sas=minisat:stl=120:sos=on:sscc=on:sser=off:sfr=on:ssfp=4000:ssfq=1.0:ssnc=none:sp=reverse_arity:urr=on_7 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f20643,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f248,f443,f247,f445,f246,f447,f787,f784,f783,f782,f831,f781,f829,f778,f777,f776,f833,f762,f835,f760,f757,f756,f755,f839,f754,f837,f751,f750,f749,f841,f745,f843,f718,f845,f715,f712,f711,f710,f849,f709,f847,f706,f705,f704,f851,f701,f698,f697,f696,f683,f680,f679,f678,f677,f674,f673,f672,f671,f668,f667,f666,f859,f665,f662,f661,f660,f855,f647,f644,f643,f642,f641,f638,f637,f636,f635,f632,f631,f630,f863,f629,f853,f626,f625,f624,f861,f611,f608,f607,f606,f605,f602,f601,f600,f599,f857,f596,f595,f594,f865,f592,f589,f588,f587,f869,f586,f867,f583,f582,f581,f871,f574,f571,f570,f569,f875,f568,f873,f565,f564,f563,f877,f561,f558,f557,f556,f881,f555,f879,f552,f551,f550,f883,f549,f546,f545,f544,f887,f543,f885,f540,f539,f538,f889,f488,f485,f484,f483,f482,f479,f478,f477,f476,f475,f474,f899,f473,f897,f472,f469,f468,f467,f895,f466,f891,f463,f462,f461,f893,f451,f450,f901,f449,f448,f790,f791,f904,f903,f906,f908,f924,f923,f922,f930,f921,f932,f915,f914,f913,f934,f912,f936,f1280,f1277,f1276,f1275,f1326,f1274,f1324,f1271,f1270,f1269,f1328,f1255,f1330,f1253,f1250,f1249,f1248,f1334,f1247,f1332,f1244,f1243,f1242,f1336,f1238,f1338,f1211,f1340,f1208,f1205,f1204,f1203,f1344,f1202,f1342,f1199,f1198,f1197,f1346,f1194,f1191,f1190,f1189,f1176,f1173,f1172,f1171,f1170,f1167,f1166,f1165,f1164,f1161,f1160,f1159,f1354,f1158,f1155,f1154,f1153,f1350,f1140,f1137,f1136,f1135,f1134,f1131,f1130,f1129,f1128,f1125,f1124,f1123,f1358,f1122,f1348,f1119,f1118,f1117,f1356,f1104,f1101,f1100,f1099,f1098,f1095,f1094,f1093,f1092,f1352,f1089,f1088,f1087,f1360,f1085,f1082,f1081,f1080,f1364,f1079,f1362,f1076,f1075,f1074,f1366,f1071,f1368,f1070,f1065,f1062,f1061,f1060,f1374,f1059,f1372,f1056,f1055,f1054,f1376,f1052,f1049,f1048,f1047,f1380,f1046,f1378,f1043,f1042,f1041,f1382,f1040,f1037,f1036,f1035,f1386,f1034,f1384,f1031,f1030,f1029,f1370,f979,f976,f975,f974,f973,f970,f969,f968,f967,f966,f965,f1396,f964,f1394,f963,f960,f959,f958,f1392,f957,f1388,f954,f953,f952,f1390,f942,f941,f1398,f940,f939,f938,f937,f1283,f1404,f1284,f1402,f1285,f1286,f1405,f1400,f1519,f1518,f1517,f1524,f1540,f1539,f1538,f1543,f1552,f1544,f1550,f1545,f1548,f1849,f1848,f1847,f1851,f1861,f1852,f1859,f1853,f1857,f1897,f1903,f1905,f1973,f1979,f1984,f1990,f1996,f2001,f2026,f2030,f2036,f2040,f2046,f2050,f2056,f2060,f2120,f2119,f2118,f2199,f2114,f2113,f2112,f2201,f2411,f2417,f2424,f2444,f2448,f2456,f2460,f3366,f3365,f3364,f3374,f3390,f3389,f3388,f3395,f3410,f3409,f3408,f3415,f3717,f3716,f3721,f3729,f3722,f3727,f3789,f3795,f3797,f3821,f3825,f3829,f3872,f3878,f3883,f3889,f3895,f3900,f3927,f3931,f3937,f3941,f3947,f3951,f4472,f4469,f4468,f4467,f4465,f4462,f4461,f4460,f4457,f4524,f4456,f4526,f4445,f4442,f4441,f4440,f4532,f4435,f4530,f4432,f4431,f4430,f4538,f4414,f4542,f4410,f4407,f4406,f4405,f4548,f4402,f4546,f4399,f4398,f4397,f4552,f4345,f4568,f4310,f4574,f4285,f4282,f4281,f4280,f4588,f4279,f4586,f4276,f4275,f4274,f4590,f4271,f4268,f4267,f4266,f4253,f4250,f4249,f4248,f4247,f4244,f4243,f4242,f4241,f4238,f4237,f4236,f4598,f4235,f4232,f4231,f4230,f4594,f4217,f4214,f4213,f4212,f4211,f4208,f4207,f4206,f4205,f4202,f4201,f4200,f4602,f4199,f4592,f4196,f4195,f4194,f4600,f4181,f4178,f4177,f4176,f4175,f4172,f4171,f4170,f4169,f4596,f4166,f4165,f4164,f4604,f4160,f4157,f4156,f4155,f4610,f4152,f4608,f4149,f4148,f4147,f4614,f4141,f4138,f4137,f4136,f4622,f4134,f4131,f4130,f4129,f4520,f4127,f4620,f4124,f4123,f4122,f4624,f4120,f4518,f4117,f4116,f4115,f4522,f4112,f4626,f4111,f4110,f4630,f4109,f4632,f4102,f4099,f4098,f4097,f4638,f4094,f4636,f4091,f4090,f4089,f4642,f4085,f4082,f4081,f4080,f4648,f4077,f4646,f4074,f4073,f4072,f4652,f4071,f4654,f4070,f4656,f4069,f4658,f4068,f4660,f4067,f4064,f4063,f4062,f4664,f4061,f4662,f4058,f4057,f4056,f4628,f4006,f4003,f4002,f4001,f4000,f3997,f3996,f3995,f3994,f3993,f3992,f4674,f3991,f4672,f3990,f3987,f3986,f3985,f4670,f3984,f4666,f3981,f3980,f3979,f4668,f3969,f3968,f4676,f3967,f3966,f3965,f4478,f4680,f4479,f4480,f4682,f4678,f5193,f5190,f5189,f5188,f5186,f5183,f5182,f5181,f5178,f5244,f5177,f5246,f5166,f5163,f5162,f5161,f5252,f5156,f5250,f5153,f5152,f5151,f5258,f5135,f5262,f5131,f5128,f5127,f5126,f5268,f5123,f5266,f5120,f5119,f5118,f5272,f5066,f5288,f5031,f5294,f5006,f5003,f5002,f5001,f5308,f5000,f5306,f4997,f4996,f4995,f5310,f4992,f4989,f4988,f4987,f4974,f4971,f4970,f4969,f4968,f4965,f4964,f4963,f4962,f4959,f4958,f4957,f5318,f4956,f4953,f4952,f4951,f5314,f4938,f4935,f4934,f4933,f4932,f4929,f4928,f4927,f4926,f4923,f4922,f4921,f5322,f4920,f5312,f4917,f4916,f4915,f5320,f4902,f4899,f4898,f4897,f4896,f4893,f4892,f4891,f4890,f5316,f4887,f4886,f4885,f5324,f4881,f4878,f4877,f4876,f5330,f4873,f5328,f4870,f4869,f4868,f5334,f4862,f4859,f4858,f4857,f5342,f4855,f4852,f4851,f4850,f5240,f4848,f5340,f4845,f4844,f4843,f5344,f4841,f5238,f4838,f4837,f4836,f5242,f4833,f5346,f4832,f4831,f5350,f4830,f5352,f4823,f4820,f4819,f4818,f5358,f4815,f5356,f4812,f4811,f4810,f5362,f4806,f4803,f4802,f4801,f5368,f4798,f5366,f4795,f4794,f4793,f5372,f4792,f5374,f4791,f5376,f4790,f5378,f4789,f5380,f4788,f4785,f4784,f4783,f5384,f4782,f5382,f4779,f4778,f4777,f5348,f4727,f4724,f4723,f4722,f4721,f4718,f4717,f4716,f4715,f4714,f4713,f5394,f4712,f5392,f4711,f4708,f4707,f4706,f5390,f4705,f5386,f4702,f4701,f4700,f5388,f4690,f4689,f5396,f4688,f4687,f4686,f5400,f5199,f5200,f5401,f5398,f5461,f5460,f5459,f5542,f5455,f5454,f5453,f5544,f6858,f6864,f6871,f6891,f6895,f6903,f6907,f7389,f7390,f7454,f7450,f7392,f7453,f7393,f7455,f7394,f7395,f7396,f7457,f7397,f7456,f7398,f7458,f7400,f7459,f7401,f7460,f7402,f7403,f7462,f7404,f7463,f7407,f7464,f7410,f7465,f7411,f7466,f7412,f7467,f7418,f7468,f7425,f7469,f7426,f7470,f7427,f7471,f7432,f7472,f7433,f7473,f7435,f7474,f7437,f7475,f7438,f7476,f7439,f7477,f7440,f7461,f7478,f7626,f7625,f7624,f7632,f7623,f7622,f7621,f7634,f7649,f7648,f7647,f7652,f7666,f7665,f7664,f7669,f8097,f8103,f8105,f8210,f8216,f8221,f8227,f8233,f8238,f8309,f8313,f8319,f8323,f8329,f8333,f8405,f8404,f8403,f8411,f8399,f8398,f8397,f8413,f8396,f8395,f8394,f8415,f8436,f8435,f8434,f8439,f8430,f8429,f8428,f8441,f8458,f8457,f8456,f8464,f8859,f8865,f8867,f8967,f8973,f8978,f8984,f8990,f8995,f9039,f9043,f9049,f9053,f9059,f9063,f9665,f9664,f9663,f9662,f9660,f9659,f9658,f9657,f9654,f9720,f9653,f9722,f9648,f9647,f9646,f9645,f9643,f9642,f9641,f9640,f9637,f9730,f9636,f9732,f9614,f9613,f9612,f9611,f9740,f9606,f9738,f9605,f9604,f9603,f9746,f9589,f9752,f9579,f9578,f9577,f9576,f9764,f9567,f9762,f9566,f9565,f9564,f9774,f9472,f9812,f9435,f9824,f9386,f9385,f9384,f9383,f9856,f9382,f9854,f9381,f9380,f9379,f9858,f9376,f9375,f9374,f9373,f9372,f9371,f9370,f9369,f9368,f9367,f9366,f9365,f9364,f9363,f9362,f9361,f9866,f9360,f9359,f9358,f9357,f9862,f9356,f9355,f9354,f9353,f9352,f9351,f9350,f9349,f9348,f9347,f9346,f9345,f9870,f9344,f9860,f9343,f9342,f9341,f9868,f9340,f9339,f9338,f9337,f9336,f9335,f9334,f9333,f9332,f9864,f9331,f9330,f9329,f9872,f9319,f9318,f9317,f9316,f9884,f9307,f9882,f9306,f9305,f9304,f9894,f9294,f9293,f9292,f9291,f9906,f9289,f9288,f9287,f9286,f9910,f9280,f9279,f9278,f9277,f9716,f9275,f9274,f9273,f9272,f9726,f9270,f9904,f9269,f9268,f9267,f9912,f9265,f9908,f9264,f9263,f9262,f9914,f9256,f9714,f9255,f9254,f9253,f9718,f9251,f9724,f9250,f9249,f9248,f9728,f9245,f9916,f9244,f9242,f9920,f9241,f9922,f9239,f9924,f9238,f9926,f9225,f9224,f9223,f9222,f9938,f9213,f9936,f9212,f9211,f9210,f9948,f9200,f9199,f9198,f9197,f9960,f9188,f9958,f9187,f9186,f9185,f9970,f9184,f9972,f9183,f9974,f9180,f9976,f9179,f9978,f9178,f9980,f9177,f9982,f9174,f9984,f9173,f9986,f9162,f9988,f9161,f9990,f9158,f9992,f9157,f9994,f9156,f9996,f9155,f9998,f9152,f10000,f9151,f10002,f9150,f9149,f9148,f9147,f10006,f9146,f10004,f9145,f9144,f9143,f9918,f9109,f9108,f9107,f9106,f9105,f9104,f9103,f9102,f9101,f9100,f9099,f10016,f9098,f10014,f9097,f9096,f9095,f9094,f10012,f9093,f10008,f9092,f9091,f9090,f10010,f9080,f9079,f10018,f9078,f9077,f10022,f10023,f9675,f9676,f10024,f10020,f10047,f10046,f10051,f10059,f10052,f10057,f10064,f10068,f10072,f10078,f10077,f10076,f10096,f10075,f10074,f10073,f10098,f10104,f10103,f10102,f10122,f10124,f10136,f10138,f10150,f10528,f10527,f10526,f10615,f10524,f10523,f10522,f10617,f10519,f10518,f10517,f10621,f10516,f10515,f10514,f10623,f10513,f10625,f10512,f10627,f10717,f10716,f10715,f10803,f10713,f10712,f10711,f10805,f10708,f10707,f10706,f10809,f10705,f10704,f10703,f10811,f10702,f10813,f10701,f10815,f11213,f11218,f11221,f11270,f11274,f11278,f11283,f11288,f11291,f13887,f13888,f13960,f13958,f13890,f13959,f13891,f13961,f13892,f13893,f13894,f13963,f13895,f13962,f13896,f13964,f13898,f13965,f13899,f13966,f13900,f13901,f13968,f13902,f13969,f13905,f13970,f13908,f13971,f13909,f13972,f13910,f13973,f13916,f13974,f13923,f13975,f13924,f13976,f13925,f13977,f13930,f13978,f13931,f13979,f13933,f13980,f13935,f13981,f13936,f13982,f13937,f13983,f13938,f13967,f14007,f14006,f14005,f14027,f14004,f14003,f14002,f14029,f13995,f13994,f13993,f14031,f14049,f14048,f14047,f14069,f14046,f14045,f14044,f14071,f14077,f14094,f14076,f14096,f14102,f14119,f14101,f14121,f14298,f14304,f14311,f14351,f14355,f14363,f14367,f14617,f14623,f14630,f14670,f14674,f14682,f14686,f16488,f7452,f16489,f16925,f16924,f16923,f16931,f16922,f16921,f16920,f16933,f16907,f16906,f16905,f16935,f16904,f16903,f16902,f16937,f16970,f16969,f16968,f16983,f16967,f16966,f16965,f16985,f17014,f17013,f17012,f17020,f17362,f17368,f17370,f17438,f17444,f17449,f17455,f17461,f17466,f17508,f17512,f17518,f17522,f17528,f17532,f17582,f17581,f17586,f17594,f17587,f17592,f17599,f17603,f17607,f17612,f17611,f17616,f17624,f17617,f17622,f17629,f17633,f17637,f17753,f17752,f17751,f17849,f17749,f17748,f17747,f17851,f17746,f17853,f17745,f17855,f17744,f17743,f17742,f17857,f17741,f17740,f17739,f17859,f17738,f17861,f17737,f17863,f17736,f17735,f17734,f17865,f17733,f17732,f17731,f17867,f18342,f18347,f18350,f18359,f18363,f18367,f18377,f18381,f18385,f18395,f18391,f18398,f18418,f18414,f18421,f18427,f18431,f18435,f18444,f18448,f18452,f18460,f18463,f18469,f18466,f18472,f18478,f18482,f18486,f18495,f18499,f18503,f18511,f18514,f18551,f18548,f18554,f18712,f18775,f18866,f18870,f18874,f19407,f19406,f19405,f19430,f19404,f19403,f19402,f19432,f19401,f19400,f19399,f19434,f19452,f19451,f19450,f19478,f19449,f19448,f19447,f19480,f19486,f19515,f19485,f19517,f19523,f19544,f19522,f19546,f19564,f19563,f19562,f19590,f19561,f19560,f19559,f19592,f19610,f19609,f19608,f19636,f19652,f19681,f19687,f19708,f20006,f20012,f20019,f20106,f20111,f20114,f20128,f20132,f20140,f20144,f20159,f20163,f20167,f20177,f20181,f20185,f20195,f20199,f20203,f20208,f20212,f20216,f20223,f20227,f20231,f20253,f20256,f20262,f20265,f20272,f20268,f20275,f20282,f20278,f20285,f20642,f20640])).
fof(f20640,plain,(
  $false | ($spl1 | $spl708)),
  inference(unit_resulting_resolution,[],[f443,f10023])).
fof(f20642,plain,(
  $false | ($spl1 | $spl708)),
  inference(trivial_inequality_removal,[],[f20641])).
fof(f20641,plain,(
  s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun(sK13,'type.num.num'),sK16),s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun(sK13,'type.num.num'),sK16),s(fun(sK13,bool),sK14))) | ($spl1 | $spl708)),
  inference(superposition,[],[f443,f10023])).
fof(f20285,plain,(
  ~sP4('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl979),
  inference(cnf_transformation,[],[f20285_D])).
fof(f20285_D,plain,(
  ~sP4('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') <=> ~$spl979),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl979])])).
fof(f20278,plain,(
  ~sP4('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | ($spl2 | $spl172 | $spl823)),
  inference(unit_resulting_resolution,[],[f255,f4829,f14121,f3372])).
fof(f3372,plain,(
  ( ! [X6,X7] : (~sP4('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X7,X6) | sP6(X6,'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X7) | p(s(bool,X6)) | ~sP3(X7,'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X6)) ) | $spl2),
  inference(resolution,[],[f1053,f367])).
fof(f367,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X2,X1) | p(s(bool,X0)) | sP6(X0,X1,X2) | ~sP4(X1,X2,X0) | ~sP3(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f213])).
fof(f213,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (sP3(X2,X1,X0) & sP4(X1,X2,X0) & sP5(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP3(X2,X1,X0) | ~sP4(X1,X2,X0) | ~sP5(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(rectify,[],[f212])).
fof(f212,plain,(
  ! [X2,X1,X0] : ((~sP6(X2,X1,X0) | (sP3(X0,X1,X2) & sP4(X1,X0,X2) & sP5(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP3(X0,X1,X2) | ~sP4(X1,X0,X2) | ~sP5(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP6(X2,X1,X0)))),
  inference(flattening,[],[f211])).
fof(f211,plain,(
  ! [X2,X1,X0] : ((~sP6(X2,X1,X0) | (sP3(X0,X1,X2) & sP4(X1,X0,X2) & sP5(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP3(X0,X1,X2) | ~sP4(X1,X0,X2) | ~sP5(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP6(X2,X1,X0)))),
  inference(nnf_transformation,[],[f148])).
fof(f148,plain,(
  ! [X2,X1,X0] : (sP6(X2,X1,X0) <=> (sP3(X0,X1,X2) & sP4(X1,X0,X2) & sP5(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f1053,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f445,f369])).
fof(f369,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f216])).
fof(f216,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f215])).
fof(f215,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP5(X2,X0,X1)))),
  inference(flattening,[],[f214])).
fof(f214,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP5(X2,X0,X1)))),
  inference(nnf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X2,X0,X1] : (sP5(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f4829,plain,(
  ( ! [X0,X1] : (sP3('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X0,X1)) ) | $spl172),
  inference(unit_resulting_resolution,[],[f1404,f379])).
fof(f379,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f221])).
fof(f221,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(flattening,[],[f220])).
fof(f220,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(nnf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0,X1,X2] : (sP3(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f255,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f51])).
fof(f51,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.pred_set.SUM_IMAGE_PERMUTES',aHL_FALSITY)).
fof(f20282,plain,(
  ~sP4('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | ($spl2 | $spl172 | $spl823)),
  inference(unit_resulting_resolution,[],[f255,f4829,f1053,f14121,f367])).
fof(f20275,plain,(
  ~sP4('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl977),
  inference(cnf_transformation,[],[f20275_D])).
fof(f20275_D,plain,(
  ~sP4('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') <=> ~$spl977),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl977])])).
fof(f20268,plain,(
  ~sP4('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | ($spl2 | $spl170 | $spl821)),
  inference(unit_resulting_resolution,[],[f255,f4108,f14119,f3372])).
fof(f4108,plain,(
  ( ! [X0,X1] : (sP3('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X0,X1)) ) | $spl170),
  inference(unit_resulting_resolution,[],[f1402,f379])).
fof(f20272,plain,(
  ~sP4('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | ($spl2 | $spl170 | $spl821)),
  inference(unit_resulting_resolution,[],[f255,f4108,f1053,f14119,f367])).
fof(f20265,plain,(
  ~sP4('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl975),
  inference(cnf_transformation,[],[f20265_D])).
fof(f20265_D,plain,(
  ~sP4('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') <=> ~$spl975),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl975])])).
fof(f20262,plain,(
  ~sP4('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | ($spl2 | $spl172 | $spl819)),
  inference(unit_resulting_resolution,[],[f255,f1069,f4809,f14096,f367])).
fof(f4809,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f1404,f369])).
fof(f1069,plain,(
  ( ! [X0,X1] : (sP3('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X0,X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f445,f379])).
fof(f20256,plain,(
  ~sP4('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl973),
  inference(cnf_transformation,[],[f20256_D])).
fof(f20256_D,plain,(
  ~sP4('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') <=> ~$spl973),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl973])])).
fof(f20253,plain,(
  ~sP4('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | ($spl2 | $spl170 | $spl817)),
  inference(unit_resulting_resolution,[],[f255,f1069,f4088,f14094,f367])).
fof(f4088,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl170),
  inference(unit_resulting_resolution,[],[f1402,f369])).
fof(f20231,plain,(
  sP4('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl970),
  inference(cnf_transformation,[],[f20231_D])).
fof(f20231_D,plain,(
  sP4('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl970),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl970])])).
fof(f20227,plain,(
  sP4('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl814),
  inference(resolution,[],[f14071,f362])).
fof(f362,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | sP4(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f213])).
fof(f20223,plain,(
  sP4('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl814),
  inference(unit_resulting_resolution,[],[f14071,f362])).
fof(f20216,plain,(
  sP4('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl968),
  inference(cnf_transformation,[],[f20216_D])).
fof(f20216_D,plain,(
  sP4('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl968),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl968])])).
fof(f20212,plain,(
  sP4('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl812),
  inference(resolution,[],[f14069,f362])).
fof(f20208,plain,(
  sP4('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl812),
  inference(unit_resulting_resolution,[],[f14069,f362])).
fof(f20203,plain,(
  sP4('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl966),
  inference(cnf_transformation,[],[f20203_D])).
fof(f20203_D,plain,(
  sP4('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl966),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl966])])).
fof(f20199,plain,(
  sP4('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl810),
  inference(resolution,[],[f14031,f362])).
fof(f20195,plain,(
  sP4('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl810),
  inference(unit_resulting_resolution,[],[f14031,f362])).
fof(f20185,plain,(
  sP4('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl964),
  inference(cnf_transformation,[],[f20185_D])).
fof(f20185_D,plain,(
  sP4('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl964),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl964])])).
fof(f20181,plain,(
  sP4('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl808),
  inference(resolution,[],[f14029,f362])).
fof(f20177,plain,(
  sP4('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl808),
  inference(unit_resulting_resolution,[],[f14029,f362])).
fof(f20167,plain,(
  sP4('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl962),
  inference(cnf_transformation,[],[f20167_D])).
fof(f20167_D,plain,(
  sP4('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl962),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl962])])).
fof(f20163,plain,(
  sP4('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl806),
  inference(resolution,[],[f14027,f362])).
fof(f20159,plain,(
  sP4('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl806),
  inference(unit_resulting_resolution,[],[f14027,f362])).
fof(f20144,plain,(
  ~sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl961),
  inference(cnf_transformation,[],[f20144_D])).
fof(f20144_D,plain,(
  ~sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) <=> ~$spl961),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl961])])).
fof(f20140,plain,(
  ~sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl653),
  inference(unit_resulting_resolution,[],[f9958,f363])).
fof(f363,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | sP5(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f213])).
fof(f20132,plain,(
  ~sP6('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl959),
  inference(cnf_transformation,[],[f20132_D])).
fof(f20132_D,plain,(
  ~sP6('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) <=> ~$spl959),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl959])])).
fof(f20128,plain,(
  ~sP6('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl647),
  inference(unit_resulting_resolution,[],[f9936,f362])).
fof(f20114,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl956),
  inference(cnf_transformation,[],[f20114_D])).
fof(f20114_D,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl956),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl956])])).
fof(f20111,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl642),
  inference(resolution,[],[f9924,f363])).
fof(f20106,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl642),
  inference(unit_resulting_resolution,[],[f9924,f363])).
fof(f20019,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | $spl954),
  inference(cnf_transformation,[],[f20019_D])).
fof(f20019_D,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') <=> ~$spl954),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl954])])).
fof(f20012,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | $spl610),
  inference(resolution,[],[f9868,f384])).
fof(f384,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X1,X2) | sP7(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f225])).
fof(f225,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | (sP7(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP7(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f224])).
fof(f224,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (sP7(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP7(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f223])).
fof(f223,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (sP7(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP7(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f151])).
fof(f151,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (sP7(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f20006,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | $spl610),
  inference(unit_resulting_resolution,[],[f9868,f384])).
fof(f19708,plain,(
  ~sP6('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl953),
  inference(cnf_transformation,[],[f19708_D])).
fof(f19708_D,plain,(
  ~sP6('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) <=> ~$spl953),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl953])])).
fof(f19687,plain,(
  ~sP6('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ($spl448 | $spl706)),
  inference(unit_resulting_resolution,[],[f255,f10020,f10765])).
fof(f10765,plain,(
  ( ! [X48,X49] : (s(bool,'T_0') != s(bool,X48) | p(s(bool,X49)) | ~sP6(X49,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X48)) ) | $spl448),
  inference(superposition,[],[f383,f5398])).
fof(f383,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | p(s(bool,X2)) | ~sP6(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f222])).
fof(f222,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP6(X2,X1,X0)) & (~sP6(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f149])).
fof(f149,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP6(X2,X1,X0))),
  inference(definition_folding,[],[f107,f148,f147,f146,f145])).
fof(f146,plain,(
  ! [X1,X0,X2] : (sP4(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f107,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f106])).
fof(f106,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X20,X21,X22] : ((p(s(bool,X22)) <=> s(bool,X20) = s(bool,X21)) <=> ((p(s(bool,X22)) | p(s(bool,X21)) | p(s(bool,X20))) & (p(s(bool,X22)) | ~p(s(bool,X20)) | ~p(s(bool,X21))) & (p(s(bool,X21)) | ~p(s(bool,X20)) | ~p(s(bool,X22))) & (p(s(bool,X20)) | ~p(s(bool,X21)) | ~p(s(bool,X22)))))),
  file('thm.pred_set.SUM_IMAGE_PERMUTES','thm.sat.sat_2|49|_prove_ltb')).
fof(f19681,plain,(
  ~sP6('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl951),
  inference(cnf_transformation,[],[f19681_D])).
fof(f19681_D,plain,(
  ~sP6('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl951),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl951])])).
fof(f19652,plain,(
  ~sP6('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl448 | $spl706)),
  inference(unit_resulting_resolution,[],[f255,f10020,f10764])).
fof(f10764,plain,(
  ( ! [X47,X46] : (s(bool,'T_0') != s(bool,X46) | p(s(bool,X47)) | ~sP6(X47,X46,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl448),
  inference(superposition,[],[f383,f5398])).
fof(f19636,plain,(
  sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl948),
  inference(cnf_transformation,[],[f19636_D])).
fof(f19636_D,plain,(
  sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) <=> ~$spl948),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl948])])).
fof(f19608,plain,(
  sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ($spl448 | $spl706)),
  inference(unit_resulting_resolution,[],[f437,f10020,f10763])).
fof(f10763,plain,(
  ( ! [X45,X44] : (s(bool,'T_0') != s(bool,X44) | sP6(X45,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X44) | ~p(s(bool,X45))) ) | $spl448),
  inference(superposition,[],[f381,f5398])).
fof(f381,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | sP6(X2,X1,X0) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f222])).
fof(f437,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f306])).
fof(f306,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f188])).
fof(f188,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f81])).
fof(f81,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X6] : ! [X11] : (s(X6,X11) = s(X6,X11) <=> p(s(bool,'T_0')))),
  file('thm.pred_set.SUM_IMAGE_PERMUTES','thm.bool.REFL_CLAUSE')).
fof(f19609,plain,(
  sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ($spl448 | $spl706)),
  inference(unit_resulting_resolution,[],[f250,f10020,f10763])).
fof(f250,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.SUM_IMAGE_PERMUTES',aHL_TRUTH)).
fof(f19610,plain,(
  sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ($spl448 | $spl706)),
  inference(unit_resulting_resolution,[],[f249,f10020,f10763])).
fof(f249,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.SUM_IMAGE_PERMUTES','thm.bool.TRUTH')).
fof(f19592,plain,(
  sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl946),
  inference(cnf_transformation,[],[f19592_D])).
fof(f19592_D,plain,(
  sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl946),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl946])])).
fof(f19559,plain,(
  sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl448),
  inference(unit_resulting_resolution,[],[f437,f5398,f10762])).
fof(f10762,plain,(
  ( ! [X43,X42] : (s(bool,'T_0') != s(bool,X42) | sP6(X43,X42,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ~p(s(bool,X43))) ) | $spl448),
  inference(superposition,[],[f381,f5398])).
fof(f19560,plain,(
  sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl448),
  inference(unit_resulting_resolution,[],[f250,f5398,f10762])).
fof(f19561,plain,(
  sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl448),
  inference(unit_resulting_resolution,[],[f249,f5398,f10762])).
fof(f19590,plain,(
  sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl944),
  inference(cnf_transformation,[],[f19590_D])).
fof(f19590_D,plain,(
  sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl944),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl944])])).
fof(f19562,plain,(
  sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl448 | $spl706)),
  inference(unit_resulting_resolution,[],[f437,f10020,f10762])).
fof(f19563,plain,(
  sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl448 | $spl706)),
  inference(unit_resulting_resolution,[],[f250,f10020,f10762])).
fof(f19564,plain,(
  sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl448 | $spl706)),
  inference(unit_resulting_resolution,[],[f249,f10020,f10762])).
fof(f19546,plain,(
  ~sP6('F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl943),
  inference(cnf_transformation,[],[f19546_D])).
fof(f19546_D,plain,(
  ~sP6('F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl943),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl943])])).
fof(f19522,plain,(
  ~sP6('F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl340 | $spl448)),
  inference(unit_resulting_resolution,[],[f255,f5398,f10577])).
fof(f10577,plain,(
  ( ! [X47,X46] : (s(bool,'T_0') != s(bool,X46) | p(s(bool,X47)) | ~sP6(X47,'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X46)) ) | $spl340),
  inference(superposition,[],[f383,f4678])).
fof(f19544,plain,(
  ~sP6('F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl941),
  inference(cnf_transformation,[],[f19544_D])).
fof(f19544_D,plain,(
  ~sP6('F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) <=> ~$spl941),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl941])])).
fof(f19523,plain,(
  ~sP6('F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ($spl340 | $spl706)),
  inference(unit_resulting_resolution,[],[f255,f10020,f10577])).
fof(f19517,plain,(
  ~sP6('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl939),
  inference(cnf_transformation,[],[f19517_D])).
fof(f19517_D,plain,(
  ~sP6('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl939),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl939])])).
fof(f19485,plain,(
  ~sP6('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl340 | $spl448)),
  inference(unit_resulting_resolution,[],[f255,f5398,f10576])).
fof(f10576,plain,(
  ( ! [X45,X44] : (s(bool,'T_0') != s(bool,X44) | p(s(bool,X45)) | ~sP6(X45,X44,'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl340),
  inference(superposition,[],[f383,f4678])).
fof(f19515,plain,(
  ~sP6('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl937),
  inference(cnf_transformation,[],[f19515_D])).
fof(f19515_D,plain,(
  ~sP6('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl937),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl937])])).
fof(f19486,plain,(
  ~sP6('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl340 | $spl706)),
  inference(unit_resulting_resolution,[],[f255,f10020,f10576])).
fof(f19480,plain,(
  sP6('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl934),
  inference(cnf_transformation,[],[f19480_D])).
fof(f19480_D,plain,(
  sP6('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl934),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl934])])).
fof(f19447,plain,(
  sP6('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl340 | $spl448)),
  inference(unit_resulting_resolution,[],[f437,f5398,f10575])).
fof(f10575,plain,(
  ( ! [X43,X42] : (s(bool,'T_0') != s(bool,X42) | sP6(X43,'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X42) | ~p(s(bool,X43))) ) | $spl340),
  inference(superposition,[],[f381,f4678])).
fof(f19448,plain,(
  sP6('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl340 | $spl448)),
  inference(unit_resulting_resolution,[],[f250,f5398,f10575])).
fof(f19449,plain,(
  sP6('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl340 | $spl448)),
  inference(unit_resulting_resolution,[],[f249,f5398,f10575])).
fof(f19478,plain,(
  sP6('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl932),
  inference(cnf_transformation,[],[f19478_D])).
fof(f19478_D,plain,(
  sP6('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) <=> ~$spl932),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl932])])).
fof(f19450,plain,(
  sP6('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ($spl340 | $spl706)),
  inference(unit_resulting_resolution,[],[f437,f10020,f10575])).
fof(f19451,plain,(
  sP6('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ($spl340 | $spl706)),
  inference(unit_resulting_resolution,[],[f250,f10020,f10575])).
fof(f19452,plain,(
  sP6('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ($spl340 | $spl706)),
  inference(unit_resulting_resolution,[],[f249,f10020,f10575])).
fof(f19434,plain,(
  sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl930),
  inference(cnf_transformation,[],[f19434_D])).
fof(f19434_D,plain,(
  sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl930),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl930])])).
fof(f19399,plain,(
  sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl340 | $spl448)),
  inference(unit_resulting_resolution,[],[f437,f5398,f10574])).
fof(f10574,plain,(
  ( ! [X41,X40] : (s(bool,'T_0') != s(bool,X40) | sP6(X41,X40,'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ~p(s(bool,X41))) ) | $spl340),
  inference(superposition,[],[f381,f4678])).
fof(f19400,plain,(
  sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl340 | $spl448)),
  inference(unit_resulting_resolution,[],[f250,f5398,f10574])).
fof(f19401,plain,(
  sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl340 | $spl448)),
  inference(unit_resulting_resolution,[],[f249,f5398,f10574])).
fof(f19432,plain,(
  sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl928),
  inference(cnf_transformation,[],[f19432_D])).
fof(f19432_D,plain,(
  sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl928),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl928])])).
fof(f19402,plain,(
  sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl340 | $spl706)),
  inference(unit_resulting_resolution,[],[f437,f10020,f10574])).
fof(f19403,plain,(
  sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl340 | $spl706)),
  inference(unit_resulting_resolution,[],[f250,f10020,f10574])).
fof(f19404,plain,(
  sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl340 | $spl706)),
  inference(unit_resulting_resolution,[],[f249,f10020,f10574])).
fof(f19430,plain,(
  sP6('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl926),
  inference(cnf_transformation,[],[f19430_D])).
fof(f19430_D,plain,(
  sP6('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl926),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl926])])).
fof(f19405,plain,(
  sP6('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl340),
  inference(unit_resulting_resolution,[],[f437,f4678,f10574])).
fof(f19406,plain,(
  sP6('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl340),
  inference(unit_resulting_resolution,[],[f250,f4678,f10574])).
fof(f19407,plain,(
  sP6('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl340),
  inference(unit_resulting_resolution,[],[f249,f4678,f10574])).
fof(f18874,plain,(
  sP4('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0','F_0') | $spl924),
  inference(cnf_transformation,[],[f18874_D])).
fof(f18874_D,plain,(
  sP4('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0','F_0') <=> ~$spl924),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl924])])).
fof(f18870,plain,(
  sP4('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0','F_0') | $spl922),
  inference(resolution,[],[f18775,f362])).
fof(f18866,plain,(
  sP4('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0','F_0') | $spl922),
  inference(unit_resulting_resolution,[],[f18775,f362])).
fof(f18775,plain,(
  sP6('F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0') | $spl922),
  inference(cnf_transformation,[],[f18775_D])).
fof(f18775_D,plain,(
  sP6('F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0') <=> ~$spl922),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl922])])).
fof(f18712,plain,(
  sP6('F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f255,f255,f577,f14135])).
fof(f14135,plain,(
  ( ! [X0,X1] : (~sP3(X1,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X0) | p(s(bool,X0)) | sP6(X0,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X1) | p(s(bool,X1))) ) | $spl4),
  inference(resolution,[],[f1495,f374])).
fof(f374,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f219])).
fof(f219,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f218])).
fof(f218,plain,(
  ! [X1,X0,X2] : ((~sP4(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP4(X1,X0,X2)))),
  inference(flattening,[],[f217])).
fof(f217,plain,(
  ! [X1,X0,X2] : ((~sP4(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP4(X1,X0,X2)))),
  inference(nnf_transformation,[],[f146])).
fof(f1495,plain,(
  ( ! [X6,X7] : (~sP4('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X7,X6) | sP6(X6,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X7) | p(s(bool,X6)) | ~sP3(X7,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X6)) ) | $spl4),
  inference(resolution,[],[f562,f367])).
fof(f562,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f447,f369])).
fof(f577,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f447,f378])).
fof(f378,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f221])).
fof(f18554,plain,(
  ~sP4('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl921),
  inference(cnf_transformation,[],[f18554_D])).
fof(f18554_D,plain,(
  ~sP4('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') <=> ~$spl921),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl921])])).
fof(f18548,plain,(
  ~sP4('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | ($spl4 | $spl172 | $spl749)),
  inference(unit_resulting_resolution,[],[f255,f4829,f10815,f1495])).
fof(f18551,plain,(
  ~sP4('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | ($spl4 | $spl172 | $spl749)),
  inference(unit_resulting_resolution,[],[f255,f4829,f562,f10815,f367])).
fof(f18514,plain,(
  ~sP4('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0') | $spl919),
  inference(cnf_transformation,[],[f18514_D])).
fof(f18514_D,plain,(
  ~sP4('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0') <=> ~$spl919),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl919])])).
fof(f18511,plain,(
  ~sP4('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0') | ($spl4 | $spl172 | $spl747)),
  inference(unit_resulting_resolution,[],[f255,f578,f4809,f10813,f367])).
fof(f578,plain,(
  ( ! [X0,X1] : (sP3('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X0,X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f447,f379])).
fof(f18503,plain,(
  sP4('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl916),
  inference(cnf_transformation,[],[f18503_D])).
fof(f18503_D,plain,(
  sP4('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl916),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl916])])).
fof(f18499,plain,(
  sP4('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl744),
  inference(resolution,[],[f10811,f362])).
fof(f18495,plain,(
  sP4('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl744),
  inference(unit_resulting_resolution,[],[f10811,f362])).
fof(f18486,plain,(
  sP4('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl914),
  inference(cnf_transformation,[],[f18486_D])).
fof(f18486_D,plain,(
  sP4('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl914),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl914])])).
fof(f18482,plain,(
  sP4('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl742),
  inference(resolution,[],[f10809,f362])).
fof(f18478,plain,(
  sP4('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl742),
  inference(unit_resulting_resolution,[],[f10809,f362])).
fof(f18472,plain,(
  ~sP4('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl913),
  inference(cnf_transformation,[],[f18472_D])).
fof(f18472_D,plain,(
  ~sP4('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') <=> ~$spl913),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl913])])).
fof(f18466,plain,(
  ~sP4('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | ($spl4 | $spl170 | $spl737)),
  inference(unit_resulting_resolution,[],[f255,f4108,f10627,f1495])).
fof(f18469,plain,(
  ~sP4('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | ($spl4 | $spl170 | $spl737)),
  inference(unit_resulting_resolution,[],[f255,f4108,f562,f10627,f367])).
fof(f18463,plain,(
  ~sP4('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0') | $spl911),
  inference(cnf_transformation,[],[f18463_D])).
fof(f18463_D,plain,(
  ~sP4('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0') <=> ~$spl911),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl911])])).
fof(f18460,plain,(
  ~sP4('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0') | ($spl4 | $spl170 | $spl735)),
  inference(unit_resulting_resolution,[],[f255,f578,f4088,f10625,f367])).
fof(f18452,plain,(
  sP4('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl908),
  inference(cnf_transformation,[],[f18452_D])).
fof(f18452_D,plain,(
  sP4('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl908),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl908])])).
fof(f18448,plain,(
  sP4('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl732),
  inference(resolution,[],[f10623,f362])).
fof(f18444,plain,(
  sP4('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl732),
  inference(unit_resulting_resolution,[],[f10623,f362])).
fof(f18435,plain,(
  sP4('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl906),
  inference(cnf_transformation,[],[f18435_D])).
fof(f18435_D,plain,(
  sP4('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl906),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl906])])).
fof(f18431,plain,(
  sP4('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl730),
  inference(resolution,[],[f10621,f362])).
fof(f18427,plain,(
  sP4('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl730),
  inference(unit_resulting_resolution,[],[f10621,f362])).
fof(f18421,plain,(
  ~sP4('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl905),
  inference(cnf_transformation,[],[f18421_D])).
fof(f18421_D,plain,(
  ~sP4('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') <=> ~$spl905),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl905])])).
fof(f18414,plain,(
  ~sP4('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | ($spl2 | $spl4 | $spl725)),
  inference(unit_resulting_resolution,[],[f255,f1069,f10150,f1495])).
fof(f18418,plain,(
  ~sP4('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | ($spl2 | $spl4 | $spl725)),
  inference(unit_resulting_resolution,[],[f255,f1069,f562,f10150,f367])).
fof(f18398,plain,(
  ~sP4('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0') | $spl903),
  inference(cnf_transformation,[],[f18398_D])).
fof(f18398_D,plain,(
  ~sP4('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0') <=> ~$spl903),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl903])])).
fof(f18391,plain,(
  ~sP4('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0') | ($spl2 | $spl4 | $spl723)),
  inference(unit_resulting_resolution,[],[f255,f578,f10136,f3372])).
fof(f18395,plain,(
  ~sP4('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0') | ($spl2 | $spl4 | $spl723)),
  inference(unit_resulting_resolution,[],[f255,f578,f1053,f10136,f367])).
fof(f18385,plain,(
  sP4('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl900),
  inference(cnf_transformation,[],[f18385_D])).
fof(f18385_D,plain,(
  sP4('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl900),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl900])])).
fof(f18381,plain,(
  sP4('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl720),
  inference(resolution,[],[f10122,f362])).
fof(f18377,plain,(
  sP4('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl720),
  inference(unit_resulting_resolution,[],[f10122,f362])).
fof(f18367,plain,(
  sP4('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl898),
  inference(cnf_transformation,[],[f18367_D])).
fof(f18367_D,plain,(
  sP4('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl898),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl898])])).
fof(f18363,plain,(
  sP4('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl716),
  inference(resolution,[],[f10096,f362])).
fof(f18359,plain,(
  sP4('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl716),
  inference(unit_resulting_resolution,[],[f10096,f362])).
fof(f18350,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl896),
  inference(cnf_transformation,[],[f18350_D])).
fof(f18350_D,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl896),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl896])])).
fof(f18347,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl644),
  inference(resolution,[],[f9926,f363])).
fof(f18342,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl644),
  inference(unit_resulting_resolution,[],[f9926,f363])).
fof(f17867,plain,(
  sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl894),
  inference(cnf_transformation,[],[f17867_D])).
fof(f17867_D,plain,(
  sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) <=> ~$spl894),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl894])])).
fof(f17731,plain,(
  sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | ($spl80 | $spl706)),
  inference(unit_resulting_resolution,[],[f437,f10020,f2166])).
fof(f2166,plain,(
  ( ! [X41,X40] : (s(bool,'T_0') != s(bool,X40) | sP6(X41,X40,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | ~p(s(bool,X41))) ) | $spl80),
  inference(superposition,[],[f381,f903])).
fof(f17732,plain,(
  sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | ($spl80 | $spl706)),
  inference(unit_resulting_resolution,[],[f250,f10020,f2166])).
fof(f17733,plain,(
  sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | ($spl80 | $spl706)),
  inference(unit_resulting_resolution,[],[f249,f10020,f2166])).
fof(f17865,plain,(
  sP6('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl892),
  inference(cnf_transformation,[],[f17865_D])).
fof(f17865_D,plain,(
  sP6('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) <=> ~$spl892),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl892])])).
fof(f17734,plain,(
  sP6('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ($spl80 | $spl706)),
  inference(unit_resulting_resolution,[],[f437,f10020,f2167])).
fof(f2167,plain,(
  ( ! [X43,X42] : (s(bool,'T_0') != s(bool,X42) | sP6(X43,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X42) | ~p(s(bool,X43))) ) | $spl80),
  inference(superposition,[],[f381,f903])).
fof(f17735,plain,(
  sP6('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ($spl80 | $spl706)),
  inference(unit_resulting_resolution,[],[f250,f10020,f2167])).
fof(f17736,plain,(
  sP6('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ($spl80 | $spl706)),
  inference(unit_resulting_resolution,[],[f249,f10020,f2167])).
fof(f17863,plain,(
  ~sP6('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl891),
  inference(cnf_transformation,[],[f17863_D])).
fof(f17863_D,plain,(
  ~sP6('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) <=> ~$spl891),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl891])])).
fof(f17737,plain,(
  ~sP6('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | ($spl80 | $spl706)),
  inference(unit_resulting_resolution,[],[f255,f10020,f2168])).
fof(f2168,plain,(
  ( ! [X45,X44] : (s(bool,'T_0') != s(bool,X44) | p(s(bool,X45)) | ~sP6(X45,X44,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl80),
  inference(superposition,[],[f383,f903])).
fof(f17861,plain,(
  ~sP6('F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl889),
  inference(cnf_transformation,[],[f17861_D])).
fof(f17861_D,plain,(
  ~sP6('F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) <=> ~$spl889),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl889])])).
fof(f17738,plain,(
  ~sP6('F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ($spl80 | $spl706)),
  inference(unit_resulting_resolution,[],[f255,f10020,f2169])).
fof(f2169,plain,(
  ( ! [X47,X46] : (s(bool,'T_0') != s(bool,X46) | p(s(bool,X47)) | ~sP6(X47,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X46)) ) | $spl80),
  inference(superposition,[],[f383,f903])).
fof(f17859,plain,(
  sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl886),
  inference(cnf_transformation,[],[f17859_D])).
fof(f17859_D,plain,(
  sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl886),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl886])])).
fof(f17739,plain,(
  sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl168 | $spl706)),
  inference(unit_resulting_resolution,[],[f437,f10020,f5509])).
fof(f5509,plain,(
  ( ! [X41,X40] : (s(bool,'T_0') != s(bool,X40) | sP6(X41,X40,'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ~p(s(bool,X41))) ) | $spl168),
  inference(superposition,[],[f381,f1400])).
fof(f17740,plain,(
  sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl168 | $spl706)),
  inference(unit_resulting_resolution,[],[f250,f10020,f5509])).
fof(f17741,plain,(
  sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl168 | $spl706)),
  inference(unit_resulting_resolution,[],[f249,f10020,f5509])).
fof(f17857,plain,(
  sP6('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl884),
  inference(cnf_transformation,[],[f17857_D])).
fof(f17857_D,plain,(
  sP6('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) <=> ~$spl884),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl884])])).
fof(f17742,plain,(
  sP6('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ($spl168 | $spl706)),
  inference(unit_resulting_resolution,[],[f437,f10020,f5510])).
fof(f5510,plain,(
  ( ! [X43,X42] : (s(bool,'T_0') != s(bool,X42) | sP6(X43,'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X42) | ~p(s(bool,X43))) ) | $spl168),
  inference(superposition,[],[f381,f1400])).
fof(f17743,plain,(
  sP6('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ($spl168 | $spl706)),
  inference(unit_resulting_resolution,[],[f250,f10020,f5510])).
fof(f17744,plain,(
  sP6('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ($spl168 | $spl706)),
  inference(unit_resulting_resolution,[],[f249,f10020,f5510])).
fof(f17855,plain,(
  ~sP6('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl883),
  inference(cnf_transformation,[],[f17855_D])).
fof(f17855_D,plain,(
  ~sP6('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl883),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl883])])).
fof(f17745,plain,(
  ~sP6('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl168 | $spl706)),
  inference(unit_resulting_resolution,[],[f255,f10020,f5511])).
fof(f5511,plain,(
  ( ! [X45,X44] : (s(bool,'T_0') != s(bool,X44) | p(s(bool,X45)) | ~sP6(X45,X44,'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl168),
  inference(superposition,[],[f383,f1400])).
fof(f17853,plain,(
  ~sP6('F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl881),
  inference(cnf_transformation,[],[f17853_D])).
fof(f17853_D,plain,(
  ~sP6('F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) <=> ~$spl881),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl881])])).
fof(f17746,plain,(
  ~sP6('F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ($spl168 | $spl706)),
  inference(unit_resulting_resolution,[],[f255,f10020,f5512])).
fof(f5512,plain,(
  ( ! [X47,X46] : (s(bool,'T_0') != s(bool,X46) | p(s(bool,X47)) | ~sP6(X47,'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X46)) ) | $spl168),
  inference(superposition,[],[f383,f1400])).
fof(f17851,plain,(
  sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | $spl878),
  inference(cnf_transformation,[],[f17851_D])).
fof(f17851_D,plain,(
  sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') <=> ~$spl878),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl878])])).
fof(f17747,plain,(
  sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | $spl706),
  inference(unit_resulting_resolution,[],[f437,f10020,f381])).
fof(f17748,plain,(
  sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | $spl706),
  inference(unit_resulting_resolution,[],[f250,f10020,f381])).
fof(f17749,plain,(
  sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | $spl706),
  inference(unit_resulting_resolution,[],[f249,f10020,f381])).
fof(f17849,plain,(
  sP6('T_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl876),
  inference(cnf_transformation,[],[f17849_D])).
fof(f17849_D,plain,(
  sP6('T_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) <=> ~$spl876),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl876])])).
fof(f17751,plain,(
  sP6('T_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl706),
  inference(unit_resulting_resolution,[],[f437,f10020,f381])).
fof(f17752,plain,(
  sP6('T_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl706),
  inference(unit_resulting_resolution,[],[f250,f10020,f381])).
fof(f17753,plain,(
  sP6('T_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl706),
  inference(unit_resulting_resolution,[],[f249,f10020,f381])).
fof(f17637,plain,(
  sP4('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','T_0') | $spl874),
  inference(cnf_transformation,[],[f17637_D])).
fof(f17637_D,plain,(
  sP4('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','T_0') <=> ~$spl874),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl874])])).
fof(f17633,plain,(
  sP4('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','T_0') | $spl740),
  inference(resolution,[],[f10805,f362])).
fof(f17629,plain,(
  sP4('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','T_0') | $spl740),
  inference(unit_resulting_resolution,[],[f10805,f362])).
fof(f17622,plain,(
  sP5('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl870),
  inference(cnf_transformation,[],[f17622_D])).
fof(f17622_D,plain,(
  sP5('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl870),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl870])])).
fof(f17617,plain,(
  sP5('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl738),
  inference(resolution,[],[f10803,f363])).
fof(f17624,plain,(
  sP4('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl872),
  inference(cnf_transformation,[],[f17624_D])).
fof(f17624_D,plain,(
  sP4('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl872),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl872])])).
fof(f17616,plain,(
  sP4('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl738),
  inference(resolution,[],[f10803,f362])).
fof(f17611,plain,(
  sP4('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl738),
  inference(unit_resulting_resolution,[],[f10803,f362])).
fof(f17612,plain,(
  sP5('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl738),
  inference(unit_resulting_resolution,[],[f10803,f363])).
fof(f17607,plain,(
  sP4('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','T_0') | $spl868),
  inference(cnf_transformation,[],[f17607_D])).
fof(f17607_D,plain,(
  sP4('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','T_0') <=> ~$spl868),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl868])])).
fof(f17603,plain,(
  sP4('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','T_0') | $spl728),
  inference(resolution,[],[f10617,f362])).
fof(f17599,plain,(
  sP4('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','T_0') | $spl728),
  inference(unit_resulting_resolution,[],[f10617,f362])).
fof(f17592,plain,(
  sP5('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl864),
  inference(cnf_transformation,[],[f17592_D])).
fof(f17592_D,plain,(
  sP5('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl864),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl864])])).
fof(f17587,plain,(
  sP5('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl726),
  inference(resolution,[],[f10615,f363])).
fof(f17594,plain,(
  sP4('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl866),
  inference(cnf_transformation,[],[f17594_D])).
fof(f17594_D,plain,(
  sP4('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl866),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl866])])).
fof(f17586,plain,(
  sP4('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl726),
  inference(resolution,[],[f10615,f362])).
fof(f17581,plain,(
  sP4('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl726),
  inference(unit_resulting_resolution,[],[f10615,f362])).
fof(f17582,plain,(
  sP5('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl726),
  inference(unit_resulting_resolution,[],[f10615,f363])).
fof(f17532,plain,(
  ~sP6('T_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl863),
  inference(cnf_transformation,[],[f17532_D])).
fof(f17532_D,plain,(
  ~sP6('T_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) <=> ~$spl863),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl863])])).
fof(f17528,plain,(
  ~sP6('T_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl655),
  inference(unit_resulting_resolution,[],[f9960,f363])).
fof(f17522,plain,(
  ~sP6('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | $spl861),
  inference(cnf_transformation,[],[f17522_D])).
fof(f17522_D,plain,(
  ~sP6('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') <=> ~$spl861),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl861])])).
fof(f17518,plain,(
  ~sP6('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | $spl651),
  inference(unit_resulting_resolution,[],[f9948,f362])).
fof(f17512,plain,(
  ~sP6('F_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl859),
  inference(cnf_transformation,[],[f17512_D])).
fof(f17512_D,plain,(
  ~sP6('F_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) <=> ~$spl859),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl859])])).
fof(f17508,plain,(
  ~sP6('F_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl649),
  inference(unit_resulting_resolution,[],[f9938,f362])).
fof(f17466,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0','T_0') | $spl856),
  inference(cnf_transformation,[],[f17466_D])).
fof(f17466_D,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0','T_0') <=> ~$spl856),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl856])])).
fof(f17461,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0','T_0') | $spl614),
  inference(resolution,[],[f9872,f384])).
fof(f17455,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0','T_0') | $spl614),
  inference(unit_resulting_resolution,[],[f9872,f384])).
fof(f17449,plain,(
  sP7('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | $spl854),
  inference(cnf_transformation,[],[f17449_D])).
fof(f17449_D,plain,(
  sP7('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') <=> ~$spl854),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl854])])).
fof(f17444,plain,(
  sP7('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | $spl612),
  inference(resolution,[],[f9870,f384])).
fof(f17438,plain,(
  sP7('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | $spl612),
  inference(unit_resulting_resolution,[],[f9870,f384])).
fof(f17370,plain,(
  sP11('F_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl852),
  inference(cnf_transformation,[],[f17370_D])).
fof(f17370_D,plain,(
  sP11('F_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) <=> ~$spl852),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl852])])).
fof(f17368,plain,(
  sP11('F_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl584),
  inference(resolution,[],[f9752,f420])).
fof(f420,plain,(
  ( ! [X2,X0,X1] : (~sP12(X0,X1,X2) | sP11(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f240])).
fof(f240,plain,(
  ! [X0,X1,X2] : ((~sP12(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP11(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP11(X0,X1,X2) | sP12(X0,X1,X2)))),
  inference(rectify,[],[f239])).
fof(f239,plain,(
  ! [X2,X0,X1] : ((~sP12(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP11(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP11(X2,X0,X1) | sP12(X2,X0,X1)))),
  inference(flattening,[],[f238])).
fof(f238,plain,(
  ! [X2,X0,X1] : ((~sP12(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP11(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP11(X2,X0,X1)) | sP12(X2,X0,X1)))),
  inference(nnf_transformation,[],[f157])).
fof(f157,plain,(
  ! [X2,X0,X1] : (sP12(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP11(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP12])])).
fof(f17362,plain,(
  sP11('F_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl584),
  inference(unit_resulting_resolution,[],[f9752,f420])).
fof(f17020,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | $spl850),
  inference(cnf_transformation,[],[f17020_D])).
fof(f17020_D,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') <=> ~$spl850),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl850])])).
fof(f17012,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | $spl450),
  inference(unit_resulting_resolution,[],[f437,f9209,f9234,f9235,f365])).
fof(f365,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X2,X1) | ~p(s(bool,X2)) | sP6(X0,X1,X2) | ~sP4(X1,X2,X0) | ~sP3(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f213])).
fof(f9235,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')))) ) | $spl450),
  inference(unit_resulting_resolution,[],[f5400,f377])).
fof(f377,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f221])).
fof(f9234,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')))) ) | $spl450),
  inference(unit_resulting_resolution,[],[f5400,f373])).
fof(f373,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f219])).
fof(f9209,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')))) ) | $spl450),
  inference(unit_resulting_resolution,[],[f5400,f369])).
fof(f17013,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | $spl450),
  inference(unit_resulting_resolution,[],[f250,f9209,f9234,f9235,f365])).
fof(f17014,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | $spl450),
  inference(unit_resulting_resolution,[],[f249,f9209,f9234,f9235,f365])).
fof(f16985,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl848),
  inference(cnf_transformation,[],[f16985_D])).
fof(f16985_D,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl848),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl848])])).
fof(f16965,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | ($spl172 | $spl450)),
  inference(unit_resulting_resolution,[],[f437,f4828,f4809,f9234,f365])).
fof(f4828,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X1)) ) | $spl172),
  inference(unit_resulting_resolution,[],[f1404,f378])).
fof(f16966,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | ($spl172 | $spl450)),
  inference(unit_resulting_resolution,[],[f250,f4828,f4809,f9234,f365])).
fof(f16967,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | ($spl172 | $spl450)),
  inference(unit_resulting_resolution,[],[f249,f4828,f4809,f9234,f365])).
fof(f16983,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl846),
  inference(cnf_transformation,[],[f16983_D])).
fof(f16983_D,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl846),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl846])])).
fof(f16968,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | ($spl170 | $spl450)),
  inference(unit_resulting_resolution,[],[f437,f4107,f4088,f9234,f365])).
fof(f4107,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X1)) ) | $spl170),
  inference(unit_resulting_resolution,[],[f1402,f378])).
fof(f16969,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | ($spl170 | $spl450)),
  inference(unit_resulting_resolution,[],[f250,f4107,f4088,f9234,f365])).
fof(f16970,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | ($spl170 | $spl450)),
  inference(unit_resulting_resolution,[],[f249,f4107,f4088,f9234,f365])).
fof(f16937,plain,(
  sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | $spl844),
  inference(cnf_transformation,[],[f16937_D])).
fof(f16937_D,plain,(
  sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') <=> ~$spl844),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl844])])).
fof(f16902,plain,(
  sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | ($spl4 | $spl450)),
  inference(unit_resulting_resolution,[],[f437,f576,f575,f9209,f365])).
fof(f575,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f447,f373])).
fof(f576,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f447,f377])).
fof(f16903,plain,(
  sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | ($spl4 | $spl450)),
  inference(unit_resulting_resolution,[],[f250,f576,f575,f9209,f365])).
fof(f16904,plain,(
  sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | ($spl4 | $spl450)),
  inference(unit_resulting_resolution,[],[f249,f576,f575,f9209,f365])).
fof(f16935,plain,(
  sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | $spl842),
  inference(cnf_transformation,[],[f16935_D])).
fof(f16935_D,plain,(
  sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') <=> ~$spl842),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl842])])).
fof(f16905,plain,(
  sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | ($spl2 | $spl450)),
  inference(unit_resulting_resolution,[],[f437,f1067,f1066,f9209,f365])).
fof(f1066,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f445,f373])).
fof(f1067,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f445,f377])).
fof(f16906,plain,(
  sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | ($spl2 | $spl450)),
  inference(unit_resulting_resolution,[],[f250,f1067,f1066,f9209,f365])).
fof(f16907,plain,(
  sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | ($spl2 | $spl450)),
  inference(unit_resulting_resolution,[],[f249,f1067,f1066,f9209,f365])).
fof(f16933,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | $spl840),
  inference(cnf_transformation,[],[f16933_D])).
fof(f16933_D,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') <=> ~$spl840),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl840])])).
fof(f16920,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | ($spl172 | $spl450)),
  inference(unit_resulting_resolution,[],[f437,f4827,f4826,f9209,f365])).
fof(f4826,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f1404,f373])).
fof(f4827,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f1404,f377])).
fof(f16921,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | ($spl172 | $spl450)),
  inference(unit_resulting_resolution,[],[f250,f4827,f4826,f9209,f365])).
fof(f16922,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | ($spl172 | $spl450)),
  inference(unit_resulting_resolution,[],[f249,f4827,f4826,f9209,f365])).
fof(f16931,plain,(
  sP6('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | $spl838),
  inference(cnf_transformation,[],[f16931_D])).
fof(f16931_D,plain,(
  sP6('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') <=> ~$spl838),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl838])])).
fof(f16923,plain,(
  sP6('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | ($spl170 | $spl450)),
  inference(unit_resulting_resolution,[],[f437,f4106,f4105,f9209,f365])).
fof(f4105,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl170),
  inference(unit_resulting_resolution,[],[f1402,f373])).
fof(f4106,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl170),
  inference(unit_resulting_resolution,[],[f1402,f377])).
fof(f16924,plain,(
  sP6('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | ($spl170 | $spl450)),
  inference(unit_resulting_resolution,[],[f250,f4106,f4105,f9209,f365])).
fof(f16925,plain,(
  sP6('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | ($spl170 | $spl450)),
  inference(unit_resulting_resolution,[],[f249,f4106,f4105,f9209,f365])).
fof(f16489,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X0),s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X0)))) | s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun(sK13,'type.num.num'),X1),s(fun(sK13,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),X0))))) = s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK13,'type.num.num'),X1),s(fun(sK13,sK13),sK15))),s(fun(sK13,bool),X0)))) ) | $spl836),
  inference(cnf_transformation,[],[f16489_D])).
fof(f16489_D,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X0),s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X0)))) | s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun(sK13,'type.num.num'),X1),s(fun(sK13,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),X0))))) = s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK13,'type.num.num'),X1),s(fun(sK13,sK13),sK15))),s(fun(sK13,bool),X0)))) ) <=> ~$spl836),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl836])])).
fof(f7452,plain,(
  ~p(s(bool,'T_0')) | $spl465),
  inference(cnf_transformation,[],[f7452_D])).
fof(f7452_D,plain,(
  ~p(s(bool,'T_0')) <=> ~$spl465),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl465])])).
fof(f16488,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X0),s(fun(sK13,bool),sK14)))) | s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun(sK13,'type.num.num'),X1),s(fun(sK13,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),X0))))) = s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK13,'type.num.num'),X1),s(fun(sK13,sK13),sK15))),s(fun(sK13,bool),X0))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X0))))) ) | ($spl4 | $spl80 | $spl172)),
  inference(forward_demodulation,[],[f16444,f7334])).
fof(f7334,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) ) | ($spl4 | $spl80)),
  inference(forward_demodulation,[],[f7330,f903])).
fof(f7330,plain,(
  ( ! [X0,X1] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f255,f2742,f380])).
fof(f380,plain,(
  ( ! [X2,X0,X1] : (sP6(X2,X1,X0) | p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f222])).
fof(f2742,plain,(
  ( ! [X0,X1] : (~sP6('F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f567,f362])).
fof(f567,plain,(
  ( ! [X0,X1] : (~sP4('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'F_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f305,f255,f447,f372])).
fof(f372,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f219])).
fof(f305,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0'))))) )),
  inference(cnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0,X1] : p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0'))))),
  inference(flattening,[],[f79])).
fof(f79,plain,(
  ! [X0] : ! [X1] : p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0'))))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X6] : ! [X7] : p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X7),s(fun(X6,bool),'const.pred_set.UNIV_0'))))),
  file('thm.pred_set.SUM_IMAGE_PERMUTES','thm.pred_set.SUBSET_UNIV')).
fof(f16444,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X0),s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')))) | s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun(sK13,'type.num.num'),X1),s(fun(sK13,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),X0))))) = s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK13,'type.num.num'),X1),s(fun(sK13,sK13),sK15))),s(fun(sK13,bool),X0))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),X0))))) ) | $spl172),
  inference(resolution,[],[f5198,f337])).
fof(f337,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X3),s(fun(X1,bool),X2),s(fun(X0,bool),'const.pred_set.UNIV_0')))) | s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun(X0,'type.num.num'),X4),s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X1,X0),X3),s(fun(X1,bool),X2))))) = s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun(X1,'type.num.num'),'const.combin.o_2'(s(fun(X0,'type.num.num'),X4),s(fun(X1,X0),X3))),s(fun(X1,bool),X2))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X1,bool),X2))))) )),
  inference(cnf_transformation,[],[f132])).
fof(f132,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X1,bool),X2)))) | ! [X3] : (~p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X3),s(fun(X1,bool),X2),s(fun(X0,bool),'const.pred_set.UNIV_0')))) | ! [X4] : s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun(X0,'type.num.num'),X4),s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X1,X0),X3),s(fun(X1,bool),X2))))) = s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun(X1,'type.num.num'),'const.combin.o_2'(s(fun(X0,'type.num.num'),X4),s(fun(X1,X0),X3))),s(fun(X1,bool),X2)))))),
  inference(ennf_transformation,[],[f95])).
fof(f95,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X1,bool),X2)))) => ! [X3] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X3),s(fun(X1,bool),X2),s(fun(X0,bool),'const.pred_set.UNIV_0')))) => ! [X4] : s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun(X0,'type.num.num'),X4),s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X1,X0),X3),s(fun(X1,bool),X2))))) = s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun(X1,'type.num.num'),'const.combin.o_2'(s(fun(X0,'type.num.num'),X4),s(fun(X1,X0),X3))),s(fun(X1,bool),X2)))))),
  inference(flattening,[],[f94])).
fof(f94,plain,(
  ! [X0,X1] : ! [X2] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X1,bool),X2)))) => ! [X3] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X3),s(fun(X1,bool),X2),s(fun(X0,bool),'const.pred_set.UNIV_0')))) => ! [X4] : s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun(X0,'type.num.num'),X4),s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X1,X0),X3),s(fun(X1,bool),X2))))) = s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun(X1,'type.num.num'),'const.combin.o_2'(s(fun(X0,'type.num.num'),X4),s(fun(X1,X0),X3))),s(fun(X1,bool),X2)))))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X6,X13] : ! [X7] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X13,bool),X7)))) => ! [X8] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X13,X6),X8),s(fun(X13,bool),X7),s(fun(X6,bool),'const.pred_set.UNIV_0')))) => ! [X9] : s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun(X6,'type.num.num'),X9),s(fun(X6,bool),'const.pred_set.IMAGE_2'(s(fun(X13,X6),X8),s(fun(X13,bool),X7))))) = s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun(X13,'type.num.num'),'const.combin.o_2'(s(fun(X6,'type.num.num'),X9),s(fun(X13,X6),X8))),s(fun(X13,bool),X7)))))),
  file('thm.pred_set.SUM_IMAGE_PERMUTES','thm.pred_set.SUM_IMAGE_INJ_o')).
fof(f5198,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),X1),s(fun(sK13,bool),X0)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X1),s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),sK14),s(fun(sK13,bool),X0))))) ) | $spl172),
  inference(resolution,[],[f1404,f344])).
fof(f344,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X1,bool),X4),s(fun(X1,bool),X6)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X5),s(fun(X0,bool),X3)))) | p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X5),s(fun(X1,bool),X6))))) )),
  inference(cnf_transformation,[],[f135])).
fof(f135,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : (~p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X5),s(fun(X0,bool),X3)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X1,bool),X4),s(fun(X1,bool),X6)))) | p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X5),s(fun(X1,bool),X6)))))),
  inference(flattening,[],[f134])).
fof(f134,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ((~p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X5),s(fun(X0,bool),X3)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X1,bool),X4),s(fun(X1,bool),X6))))) | p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X5),s(fun(X1,bool),X6)))))),
  inference(ennf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ((p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X5),s(fun(X0,bool),X3)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X1,bool),X4),s(fun(X1,bool),X6))))) => p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X5),s(fun(X1,bool),X6)))))),
  inference(flattening,[],[f102])).
fof(f102,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6] : ((p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X5),s(fun(X0,bool),X3)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X1,bool),X4),s(fun(X1,bool),X6))))) => p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X5),s(fun(X1,bool),X6)))))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X6,X13] : ! [X9,X7,X10,X18,X19] : ((p(s(bool,'const.pred_set.INJ_3'(s(fun(X6,X13),X9),s(fun(X6,bool),X7),s(fun(X13,bool),X10)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X18),s(fun(X6,bool),X7)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X13,bool),X10),s(fun(X13,bool),X19))))) => p(s(bool,'const.pred_set.INJ_3'(s(fun(X6,X13),X9),s(fun(X6,bool),X18),s(fun(X13,bool),X19)))))),
  file('thm.pred_set.SUM_IMAGE_PERMUTES','thm.pred_set.INJ_SUBSET')).
fof(f14686,plain,(
  ~sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl835),
  inference(cnf_transformation,[],[f14686_D])).
fof(f14686_D,plain,(
  ~sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl835),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl835])])).
fof(f14682,plain,(
  ~sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl419),
  inference(unit_resulting_resolution,[],[f5366,f363])).
fof(f14674,plain,(
  ~sP6('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl833),
  inference(cnf_transformation,[],[f14674_D])).
fof(f14674_D,plain,(
  ~sP6('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl833),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl833])])).
fof(f14670,plain,(
  ~sP6('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl413),
  inference(unit_resulting_resolution,[],[f5356,f362])).
fof(f14630,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl830),
  inference(cnf_transformation,[],[f14630_D])).
fof(f14630_D,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl830),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl830])])).
fof(f14623,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl386),
  inference(resolution,[],[f5320,f384])).
fof(f14617,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl386),
  inference(unit_resulting_resolution,[],[f5320,f384])).
fof(f14367,plain,(
  ~sP6('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl829),
  inference(cnf_transformation,[],[f14367_D])).
fof(f14367_D,plain,(
  ~sP6('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl829),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl829])])).
fof(f14363,plain,(
  ~sP6('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl311),
  inference(unit_resulting_resolution,[],[f4646,f363])).
fof(f14355,plain,(
  ~sP6('F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl827),
  inference(cnf_transformation,[],[f14355_D])).
fof(f14355_D,plain,(
  ~sP6('F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl827),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl827])])).
fof(f14351,plain,(
  ~sP6('F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl305),
  inference(unit_resulting_resolution,[],[f4636,f362])).
fof(f14311,plain,(
  sP7('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl824),
  inference(cnf_transformation,[],[f14311_D])).
fof(f14311_D,plain,(
  sP7('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl824),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl824])])).
fof(f14304,plain,(
  sP7('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl278),
  inference(resolution,[],[f4600,f384])).
fof(f14298,plain,(
  sP7('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl278),
  inference(unit_resulting_resolution,[],[f4600,f384])).
fof(f14121,plain,(
  ~sP6('F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl823),
  inference(cnf_transformation,[],[f14121_D])).
fof(f14121_D,plain,(
  ~sP6('F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl823),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl823])])).
fof(f14101,plain,(
  ~sP6('F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl168 | $spl448)),
  inference(unit_resulting_resolution,[],[f255,f5398,f5512])).
fof(f14119,plain,(
  ~sP6('F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl821),
  inference(cnf_transformation,[],[f14119_D])).
fof(f14119_D,plain,(
  ~sP6('F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl821),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl821])])).
fof(f14102,plain,(
  ~sP6('F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl168 | $spl340)),
  inference(unit_resulting_resolution,[],[f255,f4678,f5512])).
fof(f14096,plain,(
  ~sP6('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl819),
  inference(cnf_transformation,[],[f14096_D])).
fof(f14096_D,plain,(
  ~sP6('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl819),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl819])])).
fof(f14076,plain,(
  ~sP6('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl168 | $spl448)),
  inference(unit_resulting_resolution,[],[f255,f5398,f5511])).
fof(f14094,plain,(
  ~sP6('F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl817),
  inference(cnf_transformation,[],[f14094_D])).
fof(f14094_D,plain,(
  ~sP6('F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl817),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl817])])).
fof(f14077,plain,(
  ~sP6('F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl168 | $spl340)),
  inference(unit_resulting_resolution,[],[f255,f4678,f5511])).
fof(f14071,plain,(
  sP6('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl814),
  inference(cnf_transformation,[],[f14071_D])).
fof(f14071_D,plain,(
  sP6('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl814),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl814])])).
fof(f14044,plain,(
  sP6('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl168 | $spl448)),
  inference(unit_resulting_resolution,[],[f437,f5398,f5510])).
fof(f14045,plain,(
  sP6('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl168 | $spl448)),
  inference(unit_resulting_resolution,[],[f250,f5398,f5510])).
fof(f14046,plain,(
  sP6('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl168 | $spl448)),
  inference(unit_resulting_resolution,[],[f249,f5398,f5510])).
fof(f14069,plain,(
  sP6('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl812),
  inference(cnf_transformation,[],[f14069_D])).
fof(f14069_D,plain,(
  sP6('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl812),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl812])])).
fof(f14047,plain,(
  sP6('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl168 | $spl340)),
  inference(unit_resulting_resolution,[],[f437,f4678,f5510])).
fof(f14048,plain,(
  sP6('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl168 | $spl340)),
  inference(unit_resulting_resolution,[],[f250,f4678,f5510])).
fof(f14049,plain,(
  sP6('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl168 | $spl340)),
  inference(unit_resulting_resolution,[],[f249,f4678,f5510])).
fof(f14031,plain,(
  sP6('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl810),
  inference(cnf_transformation,[],[f14031_D])).
fof(f14031_D,plain,(
  sP6('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl810),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl810])])).
fof(f13993,plain,(
  sP6('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl168),
  inference(unit_resulting_resolution,[],[f437,f1400,f5509])).
fof(f13994,plain,(
  sP6('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl168),
  inference(unit_resulting_resolution,[],[f250,f1400,f5509])).
fof(f13995,plain,(
  sP6('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl168),
  inference(unit_resulting_resolution,[],[f249,f1400,f5509])).
fof(f14029,plain,(
  sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl808),
  inference(cnf_transformation,[],[f14029_D])).
fof(f14029_D,plain,(
  sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl808),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl808])])).
fof(f14002,plain,(
  sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl168 | $spl448)),
  inference(unit_resulting_resolution,[],[f437,f5398,f5509])).
fof(f14003,plain,(
  sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl168 | $spl448)),
  inference(unit_resulting_resolution,[],[f250,f5398,f5509])).
fof(f14004,plain,(
  sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl168 | $spl448)),
  inference(unit_resulting_resolution,[],[f249,f5398,f5509])).
fof(f14027,plain,(
  sP6('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl806),
  inference(cnf_transformation,[],[f14027_D])).
fof(f14027_D,plain,(
  sP6('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl806),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl806])])).
fof(f14005,plain,(
  sP6('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl168 | $spl340)),
  inference(unit_resulting_resolution,[],[f437,f4678,f5509])).
fof(f14006,plain,(
  sP6('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl168 | $spl340)),
  inference(unit_resulting_resolution,[],[f250,f4678,f5509])).
fof(f14007,plain,(
  sP6('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl168 | $spl340)),
  inference(unit_resulting_resolution,[],[f249,f4678,f5509])).
fof(f13967,plain,(
  ( ! [X78,X76,X77,X75] : (sP2(X77,'const.pred_set.SUBSET_2'(s(fun(X75,bool),X76),s(fun(X75,bool),'const.pred_set.UNIV_0')),X78)) ) | $spl772),
  inference(cnf_transformation,[],[f13967_D])).
fof(f13967_D,plain,(
  ( ! [X78,X76,X77,X75] : (sP2(X77,'const.pred_set.SUBSET_2'(s(fun(X75,bool),X76),s(fun(X75,bool),'const.pred_set.UNIV_0')),X78)) ) <=> ~$spl772),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl772])])).
fof(f13938,plain,(
  ( ! [X229,X227,X230,X228] : (~p(s(bool,'T_0')) | sP2(X229,'const.pred_set.SUBSET_2'(s(fun(X227,bool),X228),s(fun(X227,bool),'const.pred_set.UNIV_0')),X230)) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f438,f7334])).
fof(f438,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | sP2(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f347])).
fof(f347,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f206])).
fof(f206,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | ((p(s(bool,X1)) | p(s(bool,X2))) & (p(s(bool,X1)) | p(s(bool,X0))))) & ((~p(s(bool,X1)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(rectify,[],[f205])).
fof(f205,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(flattening,[],[f204])).
fof(f204,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2))))) & (((~p(s(bool,X0)) & ~p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X2)))) | sP2(X2,X0,X1)))),
  inference(nnf_transformation,[],[f143])).
fof(f143,plain,(
  ! [X2,X0,X1] : (sP2(X2,X0,X1) <=> ((p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f13983,plain,(
  ( ! [X225,X223,X226,X224] : (~p(s(bool,X225)) | sP12('const.pred_set.SUBSET_2'(s(fun(X223,bool),X224),s(fun(X223,bool),'const.pred_set.UNIV_0')),X225,X226)) ) | $spl804),
  inference(cnf_transformation,[],[f13983_D])).
fof(f13983_D,plain,(
  ( ! [X225,X223,X226,X224] : (~p(s(bool,X225)) | sP12('const.pred_set.SUBSET_2'(s(fun(X223,bool),X224),s(fun(X223,bool),'const.pred_set.UNIV_0')),X225,X226)) ) <=> ~$spl804),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl804])])).
fof(f13937,plain,(
  ( ! [X225,X223,X226,X224] : (~p(s(bool,'T_0')) | ~p(s(bool,X225)) | sP12('const.pred_set.SUBSET_2'(s(fun(X223,bool),X224),s(fun(X223,bool),'const.pred_set.UNIV_0')),X225,X226)) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f431,f7334])).
fof(f431,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP12(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f245])).
fof(f245,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP12(X2,X0,X1)) & (~sP12(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f244])).
fof(f244,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP12(X2,X0,X1)) & (~sP12(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f158])).
fof(f158,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP12(X2,X0,X1))),
  inference(definition_folding,[],[f138,f157,f156])).
fof(f156,plain,(
  ! [X2,X0,X1] : (sP11(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f138,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f113])).
fof(f113,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f112])).
fof(f112,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X20,X21,X22] : ((p(s(bool,X22)) <=> (p(s(bool,X21)) => p(s(bool,X20)))) <=> ((p(s(bool,X22)) | p(s(bool,X21))) & (p(s(bool,X22)) | ~p(s(bool,X20))) & (~p(s(bool,X21)) | p(s(bool,X20)) | ~p(s(bool,X22)))))),
  file('thm.pred_set.SUM_IMAGE_PERMUTES','thm.sat.sat_27_prove_ltb')).
fof(f13982,plain,(
  ( ! [X222,X220,X219,X221] : (p(s(bool,X221)) | sP12('const.pred_set.SUBSET_2'(s(fun(X219,bool),X220),s(fun(X219,bool),'const.pred_set.UNIV_0')),X222,X221)) ) | $spl802),
  inference(cnf_transformation,[],[f13982_D])).
fof(f13982_D,plain,(
  ( ! [X222,X220,X219,X221] : (p(s(bool,X221)) | sP12('const.pred_set.SUBSET_2'(s(fun(X219,bool),X220),s(fun(X219,bool),'const.pred_set.UNIV_0')),X222,X221)) ) <=> ~$spl802),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl802])])).
fof(f13936,plain,(
  ( ! [X222,X220,X219,X221] : (~p(s(bool,'T_0')) | p(s(bool,X221)) | sP12('const.pred_set.SUBSET_2'(s(fun(X219,bool),X220),s(fun(X219,bool),'const.pred_set.UNIV_0')),X222,X221)) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f430,f7334])).
fof(f430,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | sP12(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f245])).
fof(f13981,plain,(
  ( ! [X218,X216,X215,X217] : (p(s(bool,X217)) | p(s(bool,X218)) | sP12(X218,X217,'const.pred_set.SUBSET_2'(s(fun(X215,bool),X216),s(fun(X215,bool),'const.pred_set.UNIV_0')))) ) | $spl800),
  inference(cnf_transformation,[],[f13981_D])).
fof(f13981_D,plain,(
  ( ! [X218,X216,X215,X217] : (p(s(bool,X217)) | p(s(bool,X218)) | sP12(X218,X217,'const.pred_set.SUBSET_2'(s(fun(X215,bool),X216),s(fun(X215,bool),'const.pred_set.UNIV_0')))) ) <=> ~$spl800),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl800])])).
fof(f13935,plain,(
  ( ! [X218,X216,X215,X217] : (~p(s(bool,'T_0')) | p(s(bool,X217)) | sP12(X218,X217,'const.pred_set.SUBSET_2'(s(fun(X215,bool),X216),s(fun(X215,bool),'const.pred_set.UNIV_0'))) | p(s(bool,X218))) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f429,f7334])).
fof(f429,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP12(X2,X0,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f245])).
fof(f13980,plain,(
  ( ! [X210,X208,X209,X207] : (sP11(X209,'const.pred_set.SUBSET_2'(s(fun(X207,bool),X208),s(fun(X207,bool),'const.pred_set.UNIV_0')),X210)) ) | $spl798),
  inference(cnf_transformation,[],[f13980_D])).
fof(f13980_D,plain,(
  ( ! [X210,X208,X209,X207] : (sP11(X209,'const.pred_set.SUBSET_2'(s(fun(X207,bool),X208),s(fun(X207,bool),'const.pred_set.UNIV_0')),X210)) ) <=> ~$spl798),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl798])])).
fof(f13933,plain,(
  ( ! [X210,X208,X209,X207] : (~p(s(bool,'T_0')) | sP11(X209,'const.pred_set.SUBSET_2'(s(fun(X207,bool),X208),s(fun(X207,bool),'const.pred_set.UNIV_0')),X210)) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f427,f7334])).
fof(f427,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | sP11(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f243])).
fof(f243,plain,(
  ! [X0,X1,X2] : ((~sP11(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP11(X0,X1,X2)))),
  inference(rectify,[],[f242])).
fof(f242,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP11(X2,X0,X1)))),
  inference(flattening,[],[f241])).
fof(f241,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP11(X2,X0,X1)))),
  inference(nnf_transformation,[],[f156])).
fof(f13979,plain,(
  ( ! [X202,X200,X199,X201] : (~p(s(bool,X201)) | sP10('const.pred_set.SUBSET_2'(s(fun(X199,bool),X200),s(fun(X199,bool),'const.pred_set.UNIV_0')),X201,X202)) ) | $spl796),
  inference(cnf_transformation,[],[f13979_D])).
fof(f13979_D,plain,(
  ( ! [X202,X200,X199,X201] : (~p(s(bool,X201)) | sP10('const.pred_set.SUBSET_2'(s(fun(X199,bool),X200),s(fun(X199,bool),'const.pred_set.UNIV_0')),X201,X202)) ) <=> ~$spl796),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl796])])).
fof(f13931,plain,(
  ( ! [X202,X200,X199,X201] : (~p(s(bool,'T_0')) | ~p(s(bool,X201)) | sP10('const.pred_set.SUBSET_2'(s(fun(X199,bool),X200),s(fun(X199,bool),'const.pred_set.UNIV_0')),X201,X202)) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f414,f7334])).
fof(f414,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP10(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f237])).
fof(f237,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP10(X2,X0,X1)) & (~sP10(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f236])).
fof(f236,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP10(X2,X0,X1)) & (~sP10(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f155])).
fof(f155,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP10(X2,X0,X1))),
  inference(definition_folding,[],[f111,f154,f153])).
fof(f153,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f154,plain,(
  ! [X2,X0,X1] : (sP10(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP9(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f111,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f110])).
fof(f110,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X20,X21,X22] : ((p(s(bool,X22)) <=> (p(s(bool,X21)) | p(s(bool,X20)))) <=> ((p(s(bool,X22)) | ~p(s(bool,X21))) & (p(s(bool,X22)) | ~p(s(bool,X20))) & (p(s(bool,X21)) | p(s(bool,X20)) | ~p(s(bool,X22)))))),
  file('thm.pred_set.SUM_IMAGE_PERMUTES','thm.sat.sat_25_prove_ltb')).
fof(f13978,plain,(
  ( ! [X198,X196,X197,X195] : (~p(s(bool,X197)) | sP10('const.pred_set.SUBSET_2'(s(fun(X195,bool),X196),s(fun(X195,bool),'const.pred_set.UNIV_0')),X198,X197)) ) | $spl794),
  inference(cnf_transformation,[],[f13978_D])).
fof(f13978_D,plain,(
  ( ! [X198,X196,X197,X195] : (~p(s(bool,X197)) | sP10('const.pred_set.SUBSET_2'(s(fun(X195,bool),X196),s(fun(X195,bool),'const.pred_set.UNIV_0')),X198,X197)) ) <=> ~$spl794),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl794])])).
fof(f13930,plain,(
  ( ! [X198,X196,X197,X195] : (~p(s(bool,'T_0')) | ~p(s(bool,X197)) | sP10('const.pred_set.SUBSET_2'(s(fun(X195,bool),X196),s(fun(X195,bool),'const.pred_set.UNIV_0')),X198,X197)) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f413,f7334])).
fof(f413,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | sP10(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f237])).
fof(f13977,plain,(
  ( ! [X177,X175,X178,X176] : (sP9(X177,'const.pred_set.SUBSET_2'(s(fun(X175,bool),X176),s(fun(X175,bool),'const.pred_set.UNIV_0')),X178)) ) | $spl792),
  inference(cnf_transformation,[],[f13977_D])).
fof(f13977_D,plain,(
  ( ! [X177,X175,X178,X176] : (sP9(X177,'const.pred_set.SUBSET_2'(s(fun(X175,bool),X176),s(fun(X175,bool),'const.pred_set.UNIV_0')),X178)) ) <=> ~$spl792),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl792])])).
fof(f13925,plain,(
  ( ! [X177,X175,X178,X176] : (~p(s(bool,'T_0')) | sP9(X177,'const.pred_set.SUBSET_2'(s(fun(X175,bool),X176),s(fun(X175,bool),'const.pred_set.UNIV_0')),X178)) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f410,f7334])).
fof(f410,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f235])).
fof(f235,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f234])).
fof(f234,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP9(X2,X0,X1)))),
  inference(flattening,[],[f233])).
fof(f233,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP9(X2,X0,X1)))),
  inference(nnf_transformation,[],[f153])).
fof(f13976,plain,(
  ( ! [X173,X171,X174,X172] : (sP9(X173,X174,'const.pred_set.SUBSET_2'(s(fun(X171,bool),X172),s(fun(X171,bool),'const.pred_set.UNIV_0')))) ) | $spl790),
  inference(cnf_transformation,[],[f13976_D])).
fof(f13976_D,plain,(
  ( ! [X173,X171,X174,X172] : (sP9(X173,X174,'const.pred_set.SUBSET_2'(s(fun(X171,bool),X172),s(fun(X171,bool),'const.pred_set.UNIV_0')))) ) <=> ~$spl790),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl790])])).
fof(f13924,plain,(
  ( ! [X173,X171,X174,X172] : (~p(s(bool,'T_0')) | sP9(X173,X174,'const.pred_set.SUBSET_2'(s(fun(X171,bool),X172),s(fun(X171,bool),'const.pred_set.UNIV_0')))) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f409,f7334])).
fof(f409,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f235])).
fof(f13975,plain,(
  ( ! [X167,X169,X170,X168] : (~p(s(bool,X169)) | sP8('const.pred_set.SUBSET_2'(s(fun(X167,bool),X168),s(fun(X167,bool),'const.pred_set.UNIV_0')),X169,X170) | ~p(s(bool,X170))) ) | $spl788),
  inference(cnf_transformation,[],[f13975_D])).
fof(f13975_D,plain,(
  ( ! [X167,X169,X170,X168] : (~p(s(bool,X169)) | sP8('const.pred_set.SUBSET_2'(s(fun(X167,bool),X168),s(fun(X167,bool),'const.pred_set.UNIV_0')),X169,X170) | ~p(s(bool,X170))) ) <=> ~$spl788),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl788])])).
fof(f13923,plain,(
  ( ! [X167,X169,X170,X168] : (~p(s(bool,'T_0')) | ~p(s(bool,X169)) | ~p(s(bool,X170)) | sP8('const.pred_set.SUBSET_2'(s(fun(X167,bool),X168),s(fun(X167,bool),'const.pred_set.UNIV_0')),X169,X170)) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f397,f7334])).
fof(f397,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | sP8(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f229])).
fof(f229,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP8(X2,X0,X1)) & (~sP8(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f228])).
fof(f228,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP8(X2,X0,X1)) & (~sP8(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f152])).
fof(f152,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP8(X2,X0,X1))),
  inference(definition_folding,[],[f109,f151,f150])).
fof(f150,plain,(
  ! [X0,X1,X2] : (sP7(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f109,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f108])).
fof(f108,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X20,X21,X22] : ((p(s(bool,X22)) <=> (p(s(bool,X21)) & p(s(bool,X20)))) <=> ((p(s(bool,X22)) | ~p(s(bool,X21)) | ~p(s(bool,X20))) & (p(s(bool,X21)) | ~p(s(bool,X22))) & (p(s(bool,X20)) | ~p(s(bool,X22)))))),
  file('thm.pred_set.SUM_IMAGE_PERMUTES','thm.sat.sat_23_prove_ltb')).
fof(f13974,plain,(
  ( ! [X142,X140,X141,X139] : (sP7(X141,X142,'const.pred_set.SUBSET_2'(s(fun(X139,bool),X140),s(fun(X139,bool),'const.pred_set.UNIV_0')))) ) | $spl786),
  inference(cnf_transformation,[],[f13974_D])).
fof(f13974_D,plain,(
  ( ! [X142,X140,X141,X139] : (sP7(X141,X142,'const.pred_set.SUBSET_2'(s(fun(X139,bool),X140),s(fun(X139,bool),'const.pred_set.UNIV_0')))) ) <=> ~$spl786),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl786])])).
fof(f13916,plain,(
  ( ! [X142,X140,X141,X139] : (~p(s(bool,'T_0')) | sP7(X141,X142,'const.pred_set.SUBSET_2'(s(fun(X139,bool),X140),s(fun(X139,bool),'const.pred_set.UNIV_0')))) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f392,f7334])).
fof(f392,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP7(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f227])).
fof(f227,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(flattening,[],[f226])).
fof(f226,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(nnf_transformation,[],[f150])).
fof(f13973,plain,(
  ( ! [X118,X116,X117,X115] : (sP3('const.pred_set.SUBSET_2'(s(fun(X115,bool),X116),s(fun(X115,bool),'const.pred_set.UNIV_0')),X117,X118)) ) | $spl784),
  inference(cnf_transformation,[],[f13973_D])).
fof(f13973_D,plain,(
  ( ! [X118,X116,X117,X115] : (sP3('const.pred_set.SUBSET_2'(s(fun(X115,bool),X116),s(fun(X115,bool),'const.pred_set.UNIV_0')),X117,X118)) ) <=> ~$spl784),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl784])])).
fof(f13910,plain,(
  ( ! [X118,X116,X117,X115] : (~p(s(bool,'T_0')) | sP3('const.pred_set.SUBSET_2'(s(fun(X115,bool),X116),s(fun(X115,bool),'const.pred_set.UNIV_0')),X117,X118)) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f379,f7334])).
fof(f13972,plain,(
  ( ! [X111,X114,X112,X113] : (sP3(X113,'const.pred_set.SUBSET_2'(s(fun(X111,bool),X112),s(fun(X111,bool),'const.pred_set.UNIV_0')),X114)) ) | $spl782),
  inference(cnf_transformation,[],[f13972_D])).
fof(f13972_D,plain,(
  ( ! [X111,X114,X112,X113] : (sP3(X113,'const.pred_set.SUBSET_2'(s(fun(X111,bool),X112),s(fun(X111,bool),'const.pred_set.UNIV_0')),X114)) ) <=> ~$spl782),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl782])])).
fof(f13909,plain,(
  ( ! [X111,X114,X112,X113] : (~p(s(bool,'T_0')) | sP3(X113,'const.pred_set.SUBSET_2'(s(fun(X111,bool),X112),s(fun(X111,bool),'const.pred_set.UNIV_0')),X114)) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f378,f7334])).
fof(f13971,plain,(
  ( ! [X109,X107,X110,X108] : (sP3(X109,X110,'const.pred_set.SUBSET_2'(s(fun(X107,bool),X108),s(fun(X107,bool),'const.pred_set.UNIV_0')))) ) | $spl780),
  inference(cnf_transformation,[],[f13971_D])).
fof(f13971_D,plain,(
  ( ! [X109,X107,X110,X108] : (sP3(X109,X110,'const.pred_set.SUBSET_2'(s(fun(X107,bool),X108),s(fun(X107,bool),'const.pred_set.UNIV_0')))) ) <=> ~$spl780),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl780])])).
fof(f13908,plain,(
  ( ! [X109,X107,X110,X108] : (~p(s(bool,'T_0')) | sP3(X109,X110,'const.pred_set.SUBSET_2'(s(fun(X107,bool),X108),s(fun(X107,bool),'const.pred_set.UNIV_0')))) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f377,f7334])).
fof(f13970,plain,(
  ( ! [X97,X95,X98,X96] : (sP4(X97,X98,'const.pred_set.SUBSET_2'(s(fun(X95,bool),X96),s(fun(X95,bool),'const.pred_set.UNIV_0')))) ) | $spl778),
  inference(cnf_transformation,[],[f13970_D])).
fof(f13970_D,plain,(
  ( ! [X97,X95,X98,X96] : (sP4(X97,X98,'const.pred_set.SUBSET_2'(s(fun(X95,bool),X96),s(fun(X95,bool),'const.pred_set.UNIV_0')))) ) <=> ~$spl778),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl778])])).
fof(f13905,plain,(
  ( ! [X97,X95,X98,X96] : (~p(s(bool,'T_0')) | sP4(X97,X98,'const.pred_set.SUBSET_2'(s(fun(X95,bool),X96),s(fun(X95,bool),'const.pred_set.UNIV_0')))) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f373,f7334])).
fof(f13969,plain,(
  ( ! [X85,X83,X86,X84] : (sP5(X85,X86,'const.pred_set.SUBSET_2'(s(fun(X83,bool),X84),s(fun(X83,bool),'const.pred_set.UNIV_0')))) ) | $spl776),
  inference(cnf_transformation,[],[f13969_D])).
fof(f13969_D,plain,(
  ( ! [X85,X83,X86,X84] : (sP5(X85,X86,'const.pred_set.SUBSET_2'(s(fun(X83,bool),X84),s(fun(X83,bool),'const.pred_set.UNIV_0')))) ) <=> ~$spl776),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl776])])).
fof(f13902,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,'T_0')) | sP5(X85,X86,'const.pred_set.SUBSET_2'(s(fun(X83,bool),X84),s(fun(X83,bool),'const.pred_set.UNIV_0')))) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f369,f7334])).
fof(f13968,plain,(
  ( ! [X80,X81,X79,X82] : (sP2('const.pred_set.SUBSET_2'(s(fun(X79,bool),X80),s(fun(X79,bool),'const.pred_set.UNIV_0')),X81,X82) | ~p(s(bool,X82))) ) | $spl774),
  inference(cnf_transformation,[],[f13968_D])).
fof(f13968_D,plain,(
  ( ! [X80,X81,X79,X82] : (sP2('const.pred_set.SUBSET_2'(s(fun(X79,bool),X80),s(fun(X79,bool),'const.pred_set.UNIV_0')),X81,X82) | ~p(s(bool,X82))) ) <=> ~$spl774),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl774])])).
fof(f13901,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,'T_0')) | sP2('const.pred_set.SUBSET_2'(s(fun(X79,bool),X80),s(fun(X79,bool),'const.pred_set.UNIV_0')),X81,X82) | ~p(s(bool,X82))) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f352,f7334])).
fof(f352,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP2(X2,X0,X1) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f208])).
fof(f208,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X0)) & (~p(s(bool,X1)) | ~p(s(bool,X2)))) | sP2(X2,X0,X1)) & (~sP2(X2,X0,X1) | p(s(bool,X0)) | (p(s(bool,X1)) & p(s(bool,X2)))))),
  inference(flattening,[],[f207])).
fof(f207,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X0)) & (~p(s(bool,X1)) | ~p(s(bool,X2)))) | sP2(X2,X0,X1)) & (~sP2(X2,X0,X1) | (p(s(bool,X0)) | (p(s(bool,X1)) & p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f144])).
fof(f144,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X0)) | (p(s(bool,X1)) & p(s(bool,X2)))) <=> sP2(X2,X0,X1))),
  inference(definition_folding,[],[f104,f143])).
fof(f104,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X0)) | (p(s(bool,X1)) & p(s(bool,X2)))) <=> ((p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X6,X13,X14] : ((p(s(bool,X6)) | (p(s(bool,X13)) & p(s(bool,X14)))) <=> ((p(s(bool,X6)) | p(s(bool,X13))) & (p(s(bool,X6)) | p(s(bool,X14)))))),
  file('thm.pred_set.SUM_IMAGE_PERMUTES','thm.bool.LEFT_OR_OVER_AND')).
fof(f13900,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,'T_0')) | sP2(X77,'const.pred_set.SUBSET_2'(s(fun(X75,bool),X76),s(fun(X75,bool),'const.pred_set.UNIV_0')),X78)) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f351,f7334])).
fof(f351,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | sP2(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f208])).
fof(f13966,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,X73)) | sP2(X73,X74,'const.pred_set.SUBSET_2'(s(fun(X71,bool),X72),s(fun(X71,bool),'const.pred_set.UNIV_0')))) ) | $spl770),
  inference(cnf_transformation,[],[f13966_D])).
fof(f13966_D,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,X73)) | sP2(X73,X74,'const.pred_set.SUBSET_2'(s(fun(X71,bool),X72),s(fun(X71,bool),'const.pred_set.UNIV_0')))) ) <=> ~$spl770),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl770])])).
fof(f13899,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,'T_0')) | ~p(s(bool,X73)) | sP2(X73,X74,'const.pred_set.SUBSET_2'(s(fun(X71,bool),X72),s(fun(X71,bool),'const.pred_set.UNIV_0')))) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f350,f7334])).
fof(f350,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f206])).
fof(f13965,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,X69)) | sP2(X69,'const.pred_set.SUBSET_2'(s(fun(X67,bool),X68),s(fun(X67,bool),'const.pred_set.UNIV_0')),X70)) ) | $spl768),
  inference(cnf_transformation,[],[f13965_D])).
fof(f13965_D,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,X69)) | sP2(X69,'const.pred_set.SUBSET_2'(s(fun(X67,bool),X68),s(fun(X67,bool),'const.pred_set.UNIV_0')),X70)) ) <=> ~$spl768),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl768])])).
fof(f13898,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,'T_0')) | ~p(s(bool,X69)) | sP2(X69,'const.pred_set.SUBSET_2'(s(fun(X67,bool),X68),s(fun(X67,bool),'const.pred_set.UNIV_0')),X70)) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f349,f7334])).
fof(f349,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f206])).
fof(f13964,plain,(
  ( ! [X61,X62,X60] : (~sP1('const.pred_set.SUBSET_2'(s(fun(X60,bool),X61),s(fun(X60,bool),'const.pred_set.UNIV_0')),X62) | ~p(s(bool,X62))) ) | $spl766),
  inference(cnf_transformation,[],[f13964_D])).
fof(f13964_D,plain,(
  ( ! [X61,X62,X60] : (~sP1('const.pred_set.SUBSET_2'(s(fun(X60,bool),X61),s(fun(X60,bool),'const.pred_set.UNIV_0')),X62) | ~p(s(bool,X62))) ) <=> ~$spl766),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl766])])).
fof(f13896,plain,(
  ( ! [X61,X62,X60] : (~p(s(bool,'T_0')) | ~p(s(bool,X62)) | ~sP1('const.pred_set.SUBSET_2'(s(fun(X60,bool),X61),s(fun(X60,bool),'const.pred_set.UNIV_0')),X62)) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f335,f7334])).
fof(f335,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | ~sP1(X1,X0)) )),
  inference(cnf_transformation,[],[f199])).
fof(f199,plain,(
  ! [X0,X1] : ((((p(s(bool,X1)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)))) | sP1(X1,X0)) & (~sP1(X1,X0) | ((~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X0)) | p(s(bool,X1))))))),
  inference(nnf_transformation,[],[f142])).
fof(f142,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> sP1(X1,X0))),
  inference(definition_folding,[],[f93,f141])).
fof(f141,plain,(
  ! [X1,X0] : (sP1(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f93,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  inference(flattening,[],[f92])).
fof(f92,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X21,X22] : ((p(s(bool,X22)) <=> ~p(s(bool,X21))) <=> ((p(s(bool,X22)) | p(s(bool,X21))) & (~p(s(bool,X21)) | ~p(s(bool,X22)))))),
  file('thm.pred_set.SUM_IMAGE_PERMUTES','thm.sat.sat_29_prove_ltb')).
fof(f13962,plain,(
  ( ! [X50,X48,X49] : (sP1('const.pred_set.SUBSET_2'(s(fun(X48,bool),X49),s(fun(X48,bool),'const.pred_set.UNIV_0')),X50) | p(s(bool,X50))) ) | $spl762),
  inference(cnf_transformation,[],[f13962_D])).
fof(f13962_D,plain,(
  ( ! [X50,X48,X49] : (sP1('const.pred_set.SUBSET_2'(s(fun(X48,bool),X49),s(fun(X48,bool),'const.pred_set.UNIV_0')),X50) | p(s(bool,X50))) ) <=> ~$spl762),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl762])])).
fof(f13895,plain,(
  ( ! [X59,X57,X58] : (~p(s(bool,'T_0')) | p(s(bool,X59)) | sP1('const.pred_set.SUBSET_2'(s(fun(X57,bool),X58),s(fun(X57,bool),'const.pred_set.UNIV_0')),X59)) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f334,f7334])).
fof(f334,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP1(X1,X0)) )),
  inference(cnf_transformation,[],[f199])).
fof(f13963,plain,(
  ( ! [X52,X53,X51] : (sP1(X53,'const.pred_set.SUBSET_2'(s(fun(X51,bool),X52),s(fun(X51,bool),'const.pred_set.UNIV_0'))) | p(s(bool,X53))) ) | $spl764),
  inference(cnf_transformation,[],[f13963_D])).
fof(f13963_D,plain,(
  ( ! [X52,X53,X51] : (sP1(X53,'const.pred_set.SUBSET_2'(s(fun(X51,bool),X52),s(fun(X51,bool),'const.pred_set.UNIV_0'))) | p(s(bool,X53))) ) <=> ~$spl764),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl764])])).
fof(f13894,plain,(
  ( ! [X54,X56,X55] : (~p(s(bool,'T_0')) | sP1(X56,'const.pred_set.SUBSET_2'(s(fun(X54,bool),X55),s(fun(X54,bool),'const.pred_set.UNIV_0'))) | p(s(bool,X56))) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f333,f7334])).
fof(f333,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP1(X1,X0) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f199])).
fof(f13893,plain,(
  ( ! [X52,X53,X51] : (~p(s(bool,'T_0')) | p(s(bool,X53)) | sP1(X53,'const.pred_set.SUBSET_2'(s(fun(X51,bool),X52),s(fun(X51,bool),'const.pred_set.UNIV_0')))) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f332,f7334])).
fof(f332,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f198])).
fof(f198,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1)))),
  inference(rectify,[],[f197])).
fof(f197,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0)))),
  inference(flattening,[],[f196])).
fof(f196,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP1(X1,X0)))),
  inference(nnf_transformation,[],[f141])).
fof(f13892,plain,(
  ( ! [X50,X48,X49] : (~p(s(bool,'T_0')) | p(s(bool,X50)) | sP1('const.pred_set.SUBSET_2'(s(fun(X48,bool),X49),s(fun(X48,bool),'const.pred_set.UNIV_0')),X50)) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f329,f7334])).
fof(f329,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f198])).
fof(f13961,plain,(
  ( ! [X47,X45,X46] : (~sP1(X47,'const.pred_set.SUBSET_2'(s(fun(X45,bool),X46),s(fun(X45,bool),'const.pred_set.UNIV_0'))) | ~p(s(bool,X47))) ) | $spl760),
  inference(cnf_transformation,[],[f13961_D])).
fof(f13961_D,plain,(
  ( ! [X47,X45,X46] : (~sP1(X47,'const.pred_set.SUBSET_2'(s(fun(X45,bool),X46),s(fun(X45,bool),'const.pred_set.UNIV_0'))) | ~p(s(bool,X47))) ) <=> ~$spl760),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl760])])).
fof(f13891,plain,(
  ( ! [X47,X45,X46] : (~p(s(bool,'T_0')) | ~p(s(bool,X47)) | ~sP1(X47,'const.pred_set.SUBSET_2'(s(fun(X45,bool),X46),s(fun(X45,bool),'const.pred_set.UNIV_0')))) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f328,f7334])).
fof(f328,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | ~sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f198])).
fof(f13959,plain,(
  ( ! [X37,X38,X36] : (sP0('const.pred_set.SUBSET_2'(s(fun(X36,bool),X37),s(fun(X36,bool),'const.pred_set.UNIV_0')),X38)) ) | $spl758),
  inference(cnf_transformation,[],[f13959_D])).
fof(f13959_D,plain,(
  ( ! [X37,X38,X36] : (sP0('const.pred_set.SUBSET_2'(s(fun(X36,bool),X37),s(fun(X36,bool),'const.pred_set.UNIV_0')),X38)) ) <=> ~$spl758),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl758])])).
fof(f13890,plain,(
  ( ! [X43,X44,X42] : (~p(s(bool,'T_0')) | sP0('const.pred_set.SUBSET_2'(s(fun(X42,bool),X43),s(fun(X42,bool),'const.pred_set.UNIV_0')),X44)) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f319,f7334])).
fof(f319,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f193])).
fof(f193,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f192])).
fof(f192,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0')))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f140])).
fof(f140,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(definition_folding,[],[f129,f139])).
fof(f139,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f129,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f128])).
fof(f128,plain,(
  ! [X0,X1] : (((p(s(bool,X1)) | p(s(bool,X0))) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | (p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(ennf_transformation,[],[f89])).
fof(f89,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(flattening,[],[f88])).
fof(f88,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X13,X6] : ((~(p(s(bool,X6)) | p(s(bool,X13))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X6)) => p(s(bool,'F_0'))) => (~p(s(bool,X13)) => p(s(bool,'F_0')))))),
  file('thm.pred_set.SUM_IMAGE_PERMUTES','thm.sat.sat_9_prove_ltb')).
fof(f13958,plain,(
  ( ! [X35,X33,X34] : (sP0(X35,'const.pred_set.SUBSET_2'(s(fun(X33,bool),X34),s(fun(X33,bool),'const.pred_set.UNIV_0')))) ) | $spl756),
  inference(cnf_transformation,[],[f13958_D])).
fof(f13958_D,plain,(
  ( ! [X35,X33,X34] : (sP0(X35,'const.pred_set.SUBSET_2'(s(fun(X33,bool),X34),s(fun(X33,bool),'const.pred_set.UNIV_0')))) ) <=> ~$spl756),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl756])])).
fof(f13960,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,'T_0')) | sP0(X41,'const.pred_set.SUBSET_2'(s(fun(X39,bool),X40),s(fun(X39,bool),'const.pred_set.UNIV_0')))) ) | ($spl4 | $spl80)),
  inference(global_subsumption,[],[f248,f247,f246,f447,f789,f788,f787,f786,f785,f784,f783,f782,f781,f780,f779,f778,f777,f776,f775,f774,f773,f772,f768,f767,f766,f763,f762,f761,f760,f759,f758,f757,f756,f755,f754,f753,f752,f751,f750,f749,f747,f746,f745,f744,f743,f742,f741,f737,f736,f735,f732,f731,f730,f729,f725,f724,f723,f720,f719,f718,f715,f714,f713,f712,f711,f710,f709,f708,f707,f706,f705,f704,f703,f702,f701,f700,f699,f698,f697,f696,f695,f694,f693,f692,f689,f688,f687,f686,f683,f682,f681,f680,f679,f678,f677,f674,f673,f672,f671,f668,f667,f666,f665,f662,f661,f660,f659,f658,f657,f656,f653,f652,f651,f650,f647,f646,f645,f644,f643,f642,f641,f638,f637,f636,f635,f632,f631,f630,f629,f626,f625,f624,f622,f621,f620,f616,f615,f614,f611,f610,f609,f608,f607,f606,f605,f602,f601,f600,f599,f596,f595,f594,f593,f592,f591,f590,f589,f588,f587,f586,f585,f584,f583,f582,f581,f578,f577,f576,f575,f574,f573,f572,f571,f570,f569,f568,f567,f566,f565,f564,f563,f562,f561,f560,f559,f558,f557,f556,f555,f554,f553,f552,f551,f550,f549,f548,f547,f546,f545,f544,f543,f542,f541,f540,f539,f538,f537,f536,f535,f534,f530,f529,f528,f505,f504,f503,f499,f498,f497,f488,f487,f486,f485,f484,f483,f482,f481,f480,f479,f478,f477,f476,f475,f474,f473,f472,f469,f468,f467,f466,f463,f462,f461,f460,f459,f451,f450,f449,f448,f790,f791,f800,f801,f802,f805,f808,f810,f817,f820,f821,f823,f824,f825,f827,f904,f903,f1493,f1502,f1495,f1521,f1520,f1519,f1518,f1517,f1639,f1658,f1681,f1715,f1784,f1799,f1830,f2122,f2121,f2120,f2119,f2118,f2116,f2115,f2114,f2113,f2112,f2166,f2167,f2168,f2169,f2702,f2710,f2718,f2726,f2734,f2742,f2750,f2758,f2802,f2817,f2832,f2847,f2867,f2882,f2897,f2912,f3054,f3073,f3141,f3160,f3242,f3261,f5649,f5669,f5822,f5858,f5875,f5892,f5909,f5926,f5943,f5960,f7263,f7273,f7284,f7283,f7297,f7296,f7327,f7334,f7371,f7370,f7367,f7366,f7389,f7390,f7391,f7454,f7392,f7393,f7394,f7395,f7396,f7397,f7398,f7400,f7401,f7402,f7403,f7404,f7407,f7410,f7411,f7412,f7414,f7415,f7416,f7417,f7418,f7425,f7426,f7427,f7432,f7433,f7435,f7437,f7438,f7439,f7440,f10081,f10075,f10074,f10073,f10094,f10092,f10091,f10107,f10118,f10117,f10134,f10132,f10131,f10146,f10145,f10445,f10458,f11806,f12020,f12019,f13867,f13865,f13861,f13859,f13852,f13847,f13887,f13888,f13889])).
fof(f13889,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,'T_0')) | p(s(bool,'F_0')) | sP0(X41,'const.pred_set.SUBSET_2'(s(fun(X39,bool),X40),s(fun(X39,bool),'const.pred_set.UNIV_0')))) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f318,f7334])).
fof(f318,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,'F_0')) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f193])).
fof(f13847,plain,(
  ( ! [X0,X1] : (sP6('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) ) | ($spl4 | $spl80)),
  inference(unit_resulting_resolution,[],[f249,f7334,f2167])).
fof(f13852,plain,(
  ( ! [X0,X1] : (sP6('T_0','const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | ($spl4 | $spl80)),
  inference(unit_resulting_resolution,[],[f249,f7334,f2166])).
fof(f13859,plain,(
  ( ! [X0,X1] : (sP6('T_0','const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'T_0')) ) | ($spl4 | $spl80)),
  inference(unit_resulting_resolution,[],[f249,f7334,f381])).
fof(f13861,plain,(
  ( ! [X0,X1] : (~sP6('F_0','const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'T_0')) ) | ($spl4 | $spl80)),
  inference(unit_resulting_resolution,[],[f255,f7334,f383])).
fof(f13865,plain,(
  ( ! [X0,X1] : (sP6('T_0','T_0','const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) ) | ($spl4 | $spl80)),
  inference(unit_resulting_resolution,[],[f249,f7334,f381])).
fof(f13867,plain,(
  ( ! [X0,X1] : (~sP6('F_0','T_0','const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) ) | ($spl4 | $spl80)),
  inference(unit_resulting_resolution,[],[f255,f7334,f383])).
fof(f12019,plain,(
  ( ! [X0,X1] : (sP4('T_0','const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'T_0')) ) | ($spl4 | $spl80)),
  inference(unit_resulting_resolution,[],[f7370,f362])).
fof(f12020,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'T_0')) ) | ($spl4 | $spl80)),
  inference(unit_resulting_resolution,[],[f7370,f363])).
fof(f11806,plain,(
  ( ! [X0,X1] : (sP4('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'T_0','T_0')) ) | ($spl4 | $spl80)),
  inference(unit_resulting_resolution,[],[f7366,f362])).
fof(f10458,plain,(
  ( ! [X3] : (sP5(X3,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | ~p(s(bool,X3))) ) | $spl80),
  inference(resolution,[],[f10091,f363])).
fof(f10445,plain,(
  ( ! [X4] : (sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X4) | ~p(s(bool,X4))) ) | $spl4),
  inference(resolution,[],[f1502,f575])).
fof(f10145,plain,(
  ( ! [X0] : (~sP6(X0,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | p(s(bool,X0))) ) | $spl80),
  inference(equality_resolution,[],[f2169])).
fof(f10146,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X8)) | ~sP6(X8,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X7)))) ) | ($spl4 | $spl80)),
  inference(trivial_inequality_removal,[],[f10144])).
fof(f10144,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,X8)) | ~sP6(X8,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X7)))) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f2169,f7327])).
fof(f10131,plain,(
  ( ! [X0] : (~sP6(X0,'T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | p(s(bool,X0))) ) | $spl80),
  inference(equality_resolution,[],[f2168])).
fof(f10132,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X8)) | ~sP6(X8,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X7)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | ($spl4 | $spl80)),
  inference(trivial_inequality_removal,[],[f10130])).
fof(f10130,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,X8)) | ~sP6(X8,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X7)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f2168,f7327])).
fof(f10134,plain,(
  ( ! [X4] : (~sP6(X4,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | p(s(bool,X4))) ) | $spl80),
  inference(trivial_inequality_removal,[],[f10128])).
fof(f10128,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,X4)) | ~sP6(X4,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl80),
  inference(superposition,[],[f2168,f903])).
fof(f10117,plain,(
  ( ! [X0] : (sP6(X0,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | ~p(s(bool,X0))) ) | $spl80),
  inference(equality_resolution,[],[f2167])).
fof(f10118,plain,(
  ( ! [X6,X8,X7] : (sP6(X8,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X7))) | ~p(s(bool,X8))) ) | ($spl4 | $spl80)),
  inference(trivial_inequality_removal,[],[f10116])).
fof(f10116,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,'T_0') | sP6(X8,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X7))) | ~p(s(bool,X8))) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f2167,f7327])).
fof(f10107,plain,(
  ( ! [X0,X1] : (sP6('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)))) ) | ($spl4 | $spl80)),
  inference(unit_resulting_resolution,[],[f249,f7327,f2167])).
fof(f10091,plain,(
  ( ! [X0] : (sP6(X0,'T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | ~p(s(bool,X0))) ) | $spl80),
  inference(equality_resolution,[],[f2166])).
fof(f10092,plain,(
  ( ! [X6,X8,X7] : (sP6(X8,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X7)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | ~p(s(bool,X8))) ) | ($spl4 | $spl80)),
  inference(trivial_inequality_removal,[],[f10090])).
fof(f10090,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,'T_0') | sP6(X8,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X7)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | ~p(s(bool,X8))) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f2166,f7327])).
fof(f10094,plain,(
  ( ! [X4] : (sP6(X4,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | ~p(s(bool,X4))) ) | $spl80),
  inference(trivial_inequality_removal,[],[f10088])).
fof(f10088,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,'T_0') | sP6(X4,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | ~p(s(bool,X4))) ) | $spl80),
  inference(superposition,[],[f2166,f903])).
fof(f10081,plain,(
  ( ! [X0,X1] : (sP6('T_0','const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | ($spl4 | $spl80)),
  inference(unit_resulting_resolution,[],[f249,f7327,f2166])).
fof(f7417,plain,(
  ( ! [X134,X136,X135,X137] : (s(bool,'T_0') != s(bool,X136) | p(s(bool,X137)) | ~sP6(X137,'const.pred_set.SUBSET_2'(s(fun(X134,bool),X135),s(fun(X134,bool),X135)),X136)) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f383,f7327])).
fof(f7416,plain,(
  ( ! [X132,X130,X133,X131] : (s(bool,'T_0') != s(bool,X132) | p(s(bool,X133)) | ~sP6(X133,X132,'const.pred_set.SUBSET_2'(s(fun(X130,bool),X131),s(fun(X130,bool),X131)))) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f383,f7327])).
fof(f7415,plain,(
  ( ! [X127,X128,X126,X129] : (s(bool,'T_0') != s(bool,X128) | sP6(X129,'const.pred_set.SUBSET_2'(s(fun(X126,bool),X127),s(fun(X126,bool),X127)),X128) | ~p(s(bool,X129))) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f381,f7327])).
fof(f7414,plain,(
  ( ! [X125,X123,X124,X122] : (s(bool,'T_0') != s(bool,X124) | sP6(X125,X124,'const.pred_set.SUBSET_2'(s(fun(X122,bool),X123),s(fun(X122,bool),X123))) | ~p(s(bool,X125))) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f381,f7327])).
fof(f7391,plain,(
  ( ! [X39,X38,X40] : (~p(s(bool,'T_0')) | p(s(bool,'F_0')) | sP0(X40,'const.pred_set.SUBSET_2'(s(fun(X38,bool),X39),s(fun(X38,bool),X39)))) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f318,f7327])).
fof(f7366,plain,(
  ( ! [X0,X1] : (sP6('T_0','const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'T_0')) ) | ($spl4 | $spl80)),
  inference(unit_resulting_resolution,[],[f249,f7327,f381])).
fof(f7367,plain,(
  ( ! [X0,X1] : (~sP6('F_0','const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'T_0')) ) | ($spl4 | $spl80)),
  inference(unit_resulting_resolution,[],[f255,f7327,f383])).
fof(f7370,plain,(
  ( ! [X0,X1] : (sP6('T_0','T_0','const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)))) ) | ($spl4 | $spl80)),
  inference(unit_resulting_resolution,[],[f249,f7327,f381])).
fof(f7371,plain,(
  ( ! [X0,X1] : (~sP6('F_0','T_0','const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)))) ) | ($spl4 | $spl80)),
  inference(unit_resulting_resolution,[],[f255,f7327,f383])).
fof(f7327,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)))) ) | ($spl4 | $spl80)),
  inference(forward_demodulation,[],[f7322,f903])).
fof(f7322,plain,(
  ( ! [X0,X1] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f255,f2734,f380])).
fof(f7296,plain,(
  ( ! [X0,X1] : (sP4('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) ) | $spl80),
  inference(unit_resulting_resolution,[],[f2122,f362])).
fof(f7297,plain,(
  ( ! [X0,X1] : (sP5('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0')) ) | $spl80),
  inference(unit_resulting_resolution,[],[f2122,f363])).
fof(f7283,plain,(
  ( ! [X0,X1] : (sP4('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)))) ) | $spl80),
  inference(unit_resulting_resolution,[],[f2121,f362])).
fof(f7284,plain,(
  ( ! [X0,X1] : (sP5('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0')) ) | $spl80),
  inference(unit_resulting_resolution,[],[f2121,f363])).
fof(f7273,plain,(
  ( ! [X0,X1] : (sP4('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) ) | $spl80),
  inference(unit_resulting_resolution,[],[f2116,f362])).
fof(f7263,plain,(
  ( ! [X0,X1] : (sP4('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)))) ) | $spl80),
  inference(unit_resulting_resolution,[],[f2115,f362])).
fof(f5960,plain,(
  ( ! [X2,X0,X3,X1] : (sP7('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X2,bool),X3),s(fun(X2,bool),'const.pred_set.UNIV_0')))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f658,f384])).
fof(f5943,plain,(
  ( ! [X2,X0,X3,X1] : (sP7('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X2,bool),X3),s(fun(X2,bool),'const.pred_set.UNIV_0')))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f657,f384])).
fof(f5926,plain,(
  ( ! [X2,X0,X3,X1] : (sP7('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X2,bool),X3),s(fun(X2,bool),X3)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f652,f384])).
fof(f5909,plain,(
  ( ! [X2,X0,X3,X1] : (sP7('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X2,bool),X3),s(fun(X2,bool),X3)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f651,f384])).
fof(f5892,plain,(
  ( ! [X2,X0,X3,X1] : (sP7('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.SUBSET_2'(s(fun(X2,bool),X3),s(fun(X2,bool),'const.pred_set.UNIV_0')))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f622,f384])).
fof(f5875,plain,(
  ( ! [X2,X0,X3,X1] : (sP7('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.SUBSET_2'(s(fun(X2,bool),X3),s(fun(X2,bool),X3)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f621,f384])).
fof(f5858,plain,(
  ( ! [X2,X0,X3,X1] : (sP7('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.SUBSET_2'(s(fun(X2,bool),X3),s(fun(X2,bool),'const.pred_set.UNIV_0')))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f616,f384])).
fof(f5822,plain,(
  ( ! [X2,X0,X3,X1] : (sP7('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.SUBSET_2'(s(fun(X2,bool),X3),s(fun(X2,bool),X3)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f615,f384])).
fof(f5669,plain,(
  ( ! [X0,X1] : (sP7('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f659,f384])).
fof(f5649,plain,(
  ( ! [X0,X1] : (sP7('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f653,f384])).
fof(f3261,plain,(
  ( ! [X2,X0,X1] : (sP11('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),X2)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f774,f420])).
fof(f3242,plain,(
  ( ! [X2,X0,X1] : (sP11('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),X2)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f773,f420])).
fof(f3160,plain,(
  ( ! [X2,X0,X1] : (sP9('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),X2)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f743,f403])).
fof(f403,plain,(
  ( ! [X2,X0,X1] : (~sP10(X0,X1,X2) | sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f232])).
fof(f232,plain,(
  ! [X0,X1,X2] : ((~sP10(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP9(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP9(X0,X1,X2) | sP10(X0,X1,X2)))),
  inference(rectify,[],[f231])).
fof(f231,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP9(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP9(X2,X0,X1) | sP10(X2,X0,X1)))),
  inference(flattening,[],[f230])).
fof(f230,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP9(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP9(X2,X0,X1)) | sP10(X2,X0,X1)))),
  inference(nnf_transformation,[],[f154])).
fof(f3141,plain,(
  ( ! [X2,X0,X1] : (sP9('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),X2)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f742,f403])).
fof(f3073,plain,(
  ( ! [X2,X0,X1] : (sP9('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X0,'const.pred_set.SUBSET_2'(s(fun(X1,bool),X2),s(fun(X1,bool),'const.pred_set.UNIV_0')))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f731,f403])).
fof(f3054,plain,(
  ( ! [X2,X0,X1] : (sP9('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X0,'const.pred_set.SUBSET_2'(s(fun(X1,bool),X2),s(fun(X1,bool),X2)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f730,f403])).
fof(f2912,plain,(
  ( ! [X0,X1] : (sP7('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f656,f384])).
fof(f2897,plain,(
  ( ! [X0,X1] : (sP7('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f650,f384])).
fof(f2882,plain,(
  ( ! [X0,X1] : (sP7('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f646,f384])).
fof(f2867,plain,(
  ( ! [X0,X1] : (sP7('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f645,f384])).
fof(f2847,plain,(
  ( ! [X0,X1] : (sP7('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'T_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f620,f384])).
fof(f2832,plain,(
  ( ! [X0,X1] : (sP7('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'T_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f614,f384])).
fof(f2817,plain,(
  ( ! [X0,X1] : (sP7('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f610,f384])).
fof(f2802,plain,(
  ( ! [X0,X1] : (sP7('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f609,f384])).
fof(f2758,plain,(
  ( ! [X0,X1] : (~sP6('F_0','const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f573,f362])).
fof(f2750,plain,(
  ( ! [X0,X1] : (~sP6('F_0','const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f572,f362])).
fof(f2734,plain,(
  ( ! [X0,X1] : (~sP6('F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f566,f362])).
fof(f2726,plain,(
  ( ! [X0,X1] : (~sP6('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f560,f363])).
fof(f2718,plain,(
  ( ! [X0,X1] : (~sP6('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f559,f363])).
fof(f2710,plain,(
  ( ! [X0,X1] : (~sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0','const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f554,f363])).
fof(f2702,plain,(
  ( ! [X0,X1] : (~sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0','const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f553,f363])).
fof(f2115,plain,(
  ( ! [X0,X1] : (sP6('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0')) ) | $spl80),
  inference(unit_resulting_resolution,[],[f304,f903,f381])).
fof(f304,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0,X1] : p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1))))),
  inference(flattening,[],[f77])).
fof(f77,plain,(
  ! [X0] : ! [X1] : p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1))))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X6] : ! [X7] : p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X7))))),
  file('thm.pred_set.SUM_IMAGE_PERMUTES','thm.pred_set.SUBSET_REFL')).
fof(f2116,plain,(
  ( ! [X0,X1] : (sP6('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0')) ) | $spl80),
  inference(unit_resulting_resolution,[],[f305,f903,f381])).
fof(f2121,plain,(
  ( ! [X0,X1] : (sP6('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl80),
  inference(unit_resulting_resolution,[],[f304,f903,f381])).
fof(f2122,plain,(
  ( ! [X0,X1] : (sP6('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl80),
  inference(unit_resulting_resolution,[],[f305,f903,f381])).
fof(f1830,plain,(
  ( ! [X0] : (sP11('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0',X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f772,f420])).
fof(f1799,plain,(
  ( ! [X0] : (sP11('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X0,'F_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f763,f420])).
fof(f1784,plain,(
  ( ! [X2,X3] : (sP12(X2,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X3) | ~p(s(bool,X2))) ) | $spl4),
  inference(resolution,[],[f761,f441])).
fof(f441,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | sP12(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(duplicate_literal_removal,[],[f421])).
fof(f421,plain,(
  ( ! [X2,X0,X1] : (sP12(X0,X1,X2) | ~sP11(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f240])).
fof(f1715,plain,(
  ( ! [X0] : (sP9('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0',X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f741,f403])).
fof(f1681,plain,(
  ( ! [X0] : (sP9('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X0,'T_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f729,f403])).
fof(f1658,plain,(
  ( ! [X2,X3] : (sP10(X2,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X3) | ~p(s(bool,X2))) ) | $spl4),
  inference(resolution,[],[f720,f440])).
fof(f440,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | sP10(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(duplicate_literal_removal,[],[f404])).
fof(f404,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~sP9(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f232])).
fof(f1639,plain,(
  ( ! [X2,X3] : (sP10(X2,X3,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | ~p(s(bool,X2))) ) | $spl4),
  inference(resolution,[],[f719,f440])).
fof(f1520,plain,(
  ( ! [X0,X1] : (sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f304,f562,f575,f576,f365])).
fof(f1521,plain,(
  ( ! [X0,X1] : (sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f305,f562,f575,f576,f365])).
fof(f1502,plain,(
  ( ! [X2,X3] : (~sP4('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X2,X3) | sP6(X3,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X2) | ~p(s(bool,X2))) ) | $spl4),
  inference(global_subsumption,[],[f248,f247,f246,f447,f789,f788,f787,f786,f785,f784,f783,f782,f781,f780,f779,f778,f777,f776,f775,f774,f773,f772,f768,f767,f766,f763,f762,f761,f760,f759,f758,f757,f756,f755,f754,f753,f752,f751,f750,f749,f747,f746,f745,f744,f743,f742,f741,f737,f736,f735,f732,f731,f730,f729,f725,f724,f723,f720,f719,f718,f715,f714,f713,f712,f711,f710,f709,f708,f707,f706,f705,f704,f703,f702,f701,f700,f699,f698,f697,f696,f695,f694,f693,f692,f689,f688,f687,f686,f683,f682,f681,f680,f679,f678,f677,f674,f673,f672,f671,f668,f667,f666,f665,f662,f661,f660,f659,f658,f657,f656,f653,f652,f651,f650,f647,f646,f645,f644,f643,f642,f641,f638,f637,f636,f635,f632,f631,f630,f629,f626,f625,f624,f622,f621,f620,f616,f615,f614,f611,f610,f609,f608,f607,f606,f605,f602,f601,f600,f599,f596,f595,f594,f593,f592,f591,f590,f589,f588,f587,f586,f585,f584,f583,f582,f581,f578,f577,f576,f575,f574,f573,f572,f571,f570,f569,f568,f567,f566,f565,f564,f563,f562,f561,f560,f559,f558,f557,f556,f555,f554,f553,f552,f551,f550,f549,f548,f547,f546,f545,f544,f543,f542,f541,f540,f539,f538,f537,f536,f535,f534,f530,f529,f528,f505,f504,f503,f499,f498,f497,f488,f487,f486,f485,f484,f483,f482,f481,f480,f479,f478,f477,f476,f475,f474,f473,f472,f469,f468,f467,f466,f463,f462,f461,f460,f459,f451,f450,f449,f448,f790,f791,f800,f801,f802,f805,f808,f810,f817,f820,f821,f823,f824,f825,f827,f904,f1493])).
fof(f1493,plain,(
  ( ! [X2,X3] : (~p(s(bool,X2)) | sP6(X3,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X2) | ~sP4('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X2,X3) | ~sP3(X2,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X3)) ) | $spl4),
  inference(resolution,[],[f562,f365])).
fof(f827,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl4),
  inference(superposition,[],[f447,f256])).
fof(f256,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.pred_set.SUM_IMAGE_PERMUTES',aHL_BOOL_CASES)).
fof(f825,plain,(
  ( ! [X48,X49] : (sP12('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X48,X49) | ~p(s(bool,X48))) ) | $spl4),
  inference(resolution,[],[f447,f431])).
fof(f824,plain,(
  ( ! [X47,X46] : (sP12('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X47,X46) | p(s(bool,X46))) ) | $spl4),
  inference(resolution,[],[f447,f430])).
fof(f823,plain,(
  ( ! [X45,X44] : (sP12(X45,X44,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | p(s(bool,X44)) | p(s(bool,X45))) ) | $spl4),
  inference(resolution,[],[f447,f429])).
fof(f821,plain,(
  ( ! [X41,X40] : (sP10('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X40,X41) | ~p(s(bool,X40))) ) | $spl4),
  inference(resolution,[],[f447,f414])).
fof(f820,plain,(
  ( ! [X39,X38] : (sP10('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X39,X38) | ~p(s(bool,X38))) ) | $spl4),
  inference(resolution,[],[f447,f413])).
fof(f817,plain,(
  ( ! [X33,X32] : (sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X32,X33) | ~p(s(bool,X33)) | ~p(s(bool,X32))) ) | $spl4),
  inference(resolution,[],[f447,f397])).
fof(f810,plain,(
  ( ! [X19,X18] : (sP2('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X18,X19) | ~p(s(bool,X19))) ) | $spl4),
  inference(resolution,[],[f447,f352])).
fof(f808,plain,(
  ( ! [X14,X15] : (sP2(X14,X15,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | ~p(s(bool,X14))) ) | $spl4),
  inference(resolution,[],[f447,f350])).
fof(f805,plain,(
  ( ! [X9] : (~sP1('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X9) | ~p(s(bool,X9))) ) | $spl4),
  inference(resolution,[],[f447,f335])).
fof(f802,plain,(
  ( ! [X6] : (sP1(X6,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | p(s(bool,X6))) ) | $spl4),
  inference(resolution,[],[f447,f332])).
fof(f801,plain,(
  ( ! [X5] : (sP1('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X5) | p(s(bool,X5))) ) | $spl4),
  inference(resolution,[],[f447,f329])).
fof(f800,plain,(
  ( ! [X4] : (~sP1(X4,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | ~p(s(bool,X4))) ) | $spl4),
  inference(resolution,[],[f447,f328])).
fof(f459,plain,(
  ( ! [X0] : (sP0(X0,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f255,f447,f318])).
fof(f460,plain,(
  ( ! [X0] : (sP0('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f447,f319])).
fof(f480,plain,(
  ( ! [X0,X1] : (~sP1('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f304,f447,f335])).
fof(f481,plain,(
  ( ! [X0,X1] : (~sP1('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f305,f447,f335])).
fof(f486,plain,(
  ( ! [X0,X1] : (~sP1('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f304,f447,f335])).
fof(f487,plain,(
  ( ! [X0,X1] : (~sP1('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f305,f447,f335])).
fof(f497,plain,(
  ( ! [X0] : (sP2(X0,'T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f249,f447,f348])).
fof(f348,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f206])).
fof(f498,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,'const.pred_set.SUBSET_2'(s(fun(X1,bool),X2),s(fun(X1,bool),X2)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f304,f447,f348])).
fof(f499,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,'const.pred_set.SUBSET_2'(s(fun(X1,bool),X2),s(fun(X1,bool),'const.pred_set.UNIV_0')),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f305,f447,f348])).
fof(f503,plain,(
  ( ! [X0] : (sP2('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0',X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f249,f447,f349])).
fof(f504,plain,(
  ( ! [X2,X0,X1] : (sP2('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),X2)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f304,f447,f349])).
fof(f505,plain,(
  ( ! [X2,X0,X1] : (sP2('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),X2)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f305,f447,f349])).
fof(f528,plain,(
  ( ! [X0] : (sP2('T_0',X0,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f249,f447,f352])).
fof(f529,plain,(
  ( ! [X2,X0,X1] : (sP2('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),X2,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f304,f447,f352])).
fof(f530,plain,(
  ( ! [X2,X0,X1] : (sP2('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),X2,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f305,f447,f352])).
fof(f534,plain,(
  ( ! [X0] : (sP2('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X0,'T_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f249,f447,f352])).
fof(f535,plain,(
  ( ! [X2,X0,X1] : (sP2('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X0,'const.pred_set.SUBSET_2'(s(fun(X1,bool),X2),s(fun(X1,bool),X2)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f304,f447,f352])).
fof(f536,plain,(
  ( ! [X2,X0,X1] : (sP2('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X0,'const.pred_set.SUBSET_2'(s(fun(X1,bool),X2),s(fun(X1,bool),'const.pred_set.UNIV_0')))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f305,f447,f352])).
fof(f537,plain,(
  ( ! [X0] : (sP2('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X0,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f447,f447,f352])).
fof(f541,plain,(
  ( ! [X0,X1] : (~sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'F_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f304,f255,f447,f364])).
fof(f364,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f213])).
fof(f542,plain,(
  ( ! [X0,X1] : (~sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'F_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f305,f255,f447,f364])).
fof(f547,plain,(
  ( ! [X0,X1] : (~sP6('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f304,f255,f447,f364])).
fof(f548,plain,(
  ( ! [X0,X1] : (~sP6('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f305,f255,f447,f364])).
fof(f553,plain,(
  ( ! [X0,X1] : (~sP5('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'F_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f304,f255,f447,f368])).
fof(f368,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f216])).
fof(f554,plain,(
  ( ! [X0,X1] : (~sP5('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'F_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f305,f255,f447,f368])).
fof(f559,plain,(
  ( ! [X0,X1] : (~sP5('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f304,f255,f447,f368])).
fof(f560,plain,(
  ( ! [X0,X1] : (~sP5('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f305,f255,f447,f368])).
fof(f566,plain,(
  ( ! [X0,X1] : (~sP4('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'F_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f304,f255,f447,f372])).
fof(f572,plain,(
  ( ! [X0,X1] : (~sP4('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f304,f255,f447,f372])).
fof(f573,plain,(
  ( ! [X0,X1] : (~sP4('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f305,f255,f447,f372])).
fof(f584,plain,(
  ( ! [X0,X1] : (~sP7('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'F_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f304,f255,f447,f391])).
fof(f391,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f227])).
fof(f585,plain,(
  ( ! [X0,X1] : (~sP7('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'F_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f305,f255,f447,f391])).
fof(f590,plain,(
  ( ! [X0,X1] : (~sP7('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f304,f255,f447,f391])).
fof(f591,plain,(
  ( ! [X0,X1] : (~sP7('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f305,f255,f447,f391])).
fof(f593,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f447,f392])).
fof(f609,plain,(
  ( ! [X0,X1] : (sP8('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f304,f249,f447,f397])).
fof(f610,plain,(
  ( ! [X0,X1] : (sP8('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f305,f249,f447,f397])).
fof(f614,plain,(
  ( ! [X0,X1] : (sP8('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f249,f304,f447,f397])).
fof(f615,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X2,bool),X3),s(fun(X2,bool),X3)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f304,f304,f447,f397])).
fof(f616,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X2,bool),X3),s(fun(X2,bool),X3)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f305,f304,f447,f397])).
fof(f620,plain,(
  ( ! [X0,X1] : (sP8('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f249,f305,f447,f397])).
fof(f621,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X2,bool),X3),s(fun(X2,bool),'const.pred_set.UNIV_0')))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f304,f305,f447,f397])).
fof(f622,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X2,bool),X3),s(fun(X2,bool),'const.pred_set.UNIV_0')))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f305,f305,f447,f397])).
fof(f645,plain,(
  ( ! [X0,X1] : (sP8('T_0','const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f304,f249,f447,f397])).
fof(f646,plain,(
  ( ! [X0,X1] : (sP8('T_0','const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f305,f249,f447,f397])).
fof(f650,plain,(
  ( ! [X0,X1] : (sP8('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f249,f304,f447,f397])).
fof(f651,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.SUBSET_2'(s(fun(X2,bool),X3),s(fun(X2,bool),X3)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f304,f304,f447,f397])).
fof(f652,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.SUBSET_2'(s(fun(X2,bool),X3),s(fun(X2,bool),'const.pred_set.UNIV_0')),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f305,f304,f447,f397])).
fof(f653,plain,(
  ( ! [X0,X1] : (sP8('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f447,f304,f447,f397])).
fof(f656,plain,(
  ( ! [X0,X1] : (sP8('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f249,f305,f447,f397])).
fof(f657,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.SUBSET_2'(s(fun(X2,bool),X3),s(fun(X2,bool),X3)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f304,f305,f447,f397])).
fof(f658,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.SUBSET_2'(s(fun(X2,bool),X3),s(fun(X2,bool),'const.pred_set.UNIV_0')),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f305,f305,f447,f397])).
fof(f659,plain,(
  ( ! [X0,X1] : (sP8('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f447,f305,f447,f397])).
fof(f681,plain,(
  ( ! [X0,X1] : (sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'T_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f304,f249,f447,f397])).
fof(f682,plain,(
  ( ! [X0,X1] : (sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'T_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f305,f249,f447,f397])).
fof(f686,plain,(
  ( ! [X0,X1] : (sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f249,f304,f447,f397])).
fof(f687,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.SUBSET_2'(s(fun(X2,bool),X3),s(fun(X2,bool),X3)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f304,f304,f447,f397])).
fof(f688,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.SUBSET_2'(s(fun(X2,bool),X3),s(fun(X2,bool),X3)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f305,f304,f447,f397])).
fof(f689,plain,(
  ( ! [X0,X1] : (sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f447,f304,f447,f397])).
fof(f692,plain,(
  ( ! [X0,X1] : (sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f249,f305,f447,f397])).
fof(f693,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.SUBSET_2'(s(fun(X2,bool),X3),s(fun(X2,bool),'const.pred_set.UNIV_0')))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f304,f305,f447,f397])).
fof(f694,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.SUBSET_2'(s(fun(X2,bool),X3),s(fun(X2,bool),'const.pred_set.UNIV_0')))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f305,f305,f447,f397])).
fof(f695,plain,(
  ( ! [X0,X1] : (sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f447,f305,f447,f397])).
fof(f699,plain,(
  ( ! [X0,X1] : (sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f304,f447,f447,f397])).
fof(f700,plain,(
  ( ! [X0,X1] : (sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f305,f447,f447,f397])).
fof(f702,plain,(
  ( ! [X0] : (~sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X0,'F_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f255,f447,f398])).
fof(f398,plain,(
  ( ! [X2,X0,X1] : (~sP8(X2,X0,X1) | ~p(s(bool,X2)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f229])).
fof(f703,plain,(
  ( ! [X0] : (~sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0',X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f255,f447,f399])).
fof(f399,plain,(
  ( ! [X2,X0,X1] : (~sP8(X2,X0,X1) | ~p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f229])).
fof(f707,plain,(
  ( ! [X0,X1] : (~sP8('F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f255,f304,f447,f400])).
fof(f400,plain,(
  ( ! [X2,X0,X1] : (~sP8(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f229])).
fof(f708,plain,(
  ( ! [X0,X1] : (~sP8('F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f255,f305,f447,f400])).
fof(f713,plain,(
  ( ! [X0,X1] : (~sP8('F_0','const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f255,f304,f447,f400])).
fof(f714,plain,(
  ( ! [X0,X1] : (~sP8('F_0','const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f255,f305,f447,f400])).
fof(f719,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f447,f409])).
fof(f720,plain,(
  ( ! [X0,X1] : (sP9(X0,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f447,f410])).
fof(f723,plain,(
  ( ! [X0] : (sP10('T_0',X0,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f249,f447,f413])).
fof(f724,plain,(
  ( ! [X2,X0,X1] : (sP10('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),X2,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f304,f447,f413])).
fof(f725,plain,(
  ( ! [X2,X0,X1] : (sP10('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),X2,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f305,f447,f413])).
fof(f729,plain,(
  ( ! [X0] : (sP10('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X0,'T_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f249,f447,f413])).
fof(f730,plain,(
  ( ! [X2,X0,X1] : (sP10('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X0,'const.pred_set.SUBSET_2'(s(fun(X1,bool),X2),s(fun(X1,bool),X2)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f304,f447,f413])).
fof(f731,plain,(
  ( ! [X2,X0,X1] : (sP10('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X0,'const.pred_set.SUBSET_2'(s(fun(X1,bool),X2),s(fun(X1,bool),'const.pred_set.UNIV_0')))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f305,f447,f413])).
fof(f732,plain,(
  ( ! [X0] : (sP10('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X0,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f447,f447,f413])).
fof(f735,plain,(
  ( ! [X0] : (sP10('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f249,f447,f414])).
fof(f736,plain,(
  ( ! [X2,X0,X1] : (sP10('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X2)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f304,f447,f414])).
fof(f737,plain,(
  ( ! [X2,X0,X1] : (sP10('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X2)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f305,f447,f414])).
fof(f741,plain,(
  ( ! [X0] : (sP10('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0',X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f249,f447,f414])).
fof(f742,plain,(
  ( ! [X2,X0,X1] : (sP10('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),X2)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f304,f447,f414])).
fof(f743,plain,(
  ( ! [X2,X0,X1] : (sP10('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),X2)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f305,f447,f414])).
fof(f744,plain,(
  ( ! [X0] : (sP10('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f447,f447,f414])).
fof(f746,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f255,f447,f416])).
fof(f416,plain,(
  ( ! [X2,X0,X1] : (~sP10(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f237])).
fof(f747,plain,(
  ( ! [X0] : (~sP10('F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f255,f447,f417])).
fof(f417,plain,(
  ( ! [X2,X0,X1] : (~sP10(X2,X0,X1) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f237])).
fof(f752,plain,(
  ( ! [X0,X1] : (~sP11('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0','const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f255,f304,f447,f425])).
fof(f425,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f243])).
fof(f753,plain,(
  ( ! [X0,X1] : (~sP11('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0','const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f255,f305,f447,f425])).
fof(f758,plain,(
  ( ! [X0,X1] : (~sP11('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f304,f255,f447,f425])).
fof(f759,plain,(
  ( ! [X0,X1] : (~sP11('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f305,f255,f447,f425])).
fof(f761,plain,(
  ( ! [X0,X1] : (sP11(X0,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f447,f427])).
fof(f763,plain,(
  ( ! [X0] : (sP12('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X0,'F_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f255,f447,f430])).
fof(f766,plain,(
  ( ! [X0] : (sP12('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f249,f447,f431])).
fof(f767,plain,(
  ( ! [X2,X0,X1] : (sP12('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X2)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f304,f447,f431])).
fof(f768,plain,(
  ( ! [X2,X0,X1] : (sP12('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X2)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f305,f447,f431])).
fof(f772,plain,(
  ( ! [X0] : (sP12('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0',X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f249,f447,f431])).
fof(f773,plain,(
  ( ! [X2,X0,X1] : (sP12('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),X2)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f304,f447,f431])).
fof(f774,plain,(
  ( ! [X2,X0,X1] : (sP12('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),X2)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f305,f447,f431])).
fof(f775,plain,(
  ( ! [X0] : (sP12('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f447,f447,f431])).
fof(f779,plain,(
  ( ! [X0,X1] : (~sP12('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f255,f304,f447,f432])).
fof(f432,plain,(
  ( ! [X2,X0,X1] : (~sP12(X2,X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f245])).
fof(f780,plain,(
  ( ! [X0,X1] : (~sP12('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f255,f305,f447,f432])).
fof(f785,plain,(
  ( ! [X0,X1] : (~sP12('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0','const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f255,f304,f447,f432])).
fof(f786,plain,(
  ( ! [X0,X1] : (~sP12('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0','const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f255,f305,f447,f432])).
fof(f788,plain,(
  ( ! [X0] : (~sP12('F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f255,f447,f434])).
fof(f434,plain,(
  ( ! [X2,X0,X1] : (~sP12(X2,X0,X1) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f245])).
fof(f789,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f447,f438])).
fof(f13888,plain,(
  ( ! [X37,X38,X36] : (~p(s(bool,'T_0')) | sP0('const.pred_set.SUBSET_2'(s(fun(X36,bool),X37),s(fun(X36,bool),'const.pred_set.UNIV_0')),X38)) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f314,f7334])).
fof(f314,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f191])).
fof(f191,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f190])).
fof(f190,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f139])).
fof(f13887,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,'T_0')) | sP0(X35,'const.pred_set.SUBSET_2'(s(fun(X33,bool),X34),s(fun(X33,bool),'const.pred_set.UNIV_0')))) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f313,f7334])).
fof(f313,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f191])).
fof(f11291,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl754),
  inference(cnf_transformation,[],[f11291_D])).
fof(f11291_D,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl754),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl754])])).
fof(f11288,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl410),
  inference(resolution,[],[f5352,f363])).
fof(f11283,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl410),
  inference(unit_resulting_resolution,[],[f5352,f363])).
fof(f11278,plain,(
  sP4('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl752),
  inference(cnf_transformation,[],[f11278_D])).
fof(f11278_D,plain,(
  sP4('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl752),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl752])])).
fof(f11274,plain,(
  sP4('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl718),
  inference(resolution,[],[f10098,f362])).
fof(f11270,plain,(
  sP4('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl718),
  inference(unit_resulting_resolution,[],[f10098,f362])).
fof(f11221,plain,(
  sP5('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl750),
  inference(cnf_transformation,[],[f11221_D])).
fof(f11221_D,plain,(
  sP5('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl750),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl750])])).
fof(f11218,plain,(
  sP5('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl302),
  inference(resolution,[],[f4632,f363])).
fof(f11213,plain,(
  sP5('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl302),
  inference(unit_resulting_resolution,[],[f4632,f363])).
fof(f10815,plain,(
  ~sP6('F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl749),
  inference(cnf_transformation,[],[f10815_D])).
fof(f10815_D,plain,(
  ~sP6('F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl749),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl749])])).
fof(f10701,plain,(
  ~sP6('F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl80 | $spl448)),
  inference(unit_resulting_resolution,[],[f255,f5398,f2169])).
fof(f10813,plain,(
  ~sP6('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl747),
  inference(cnf_transformation,[],[f10813_D])).
fof(f10813_D,plain,(
  ~sP6('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) <=> ~$spl747),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl747])])).
fof(f10702,plain,(
  ~sP6('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | ($spl80 | $spl448)),
  inference(unit_resulting_resolution,[],[f255,f5398,f2168])).
fof(f10811,plain,(
  sP6('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl744),
  inference(cnf_transformation,[],[f10811_D])).
fof(f10811_D,plain,(
  sP6('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl744),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl744])])).
fof(f10703,plain,(
  sP6('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl80 | $spl448)),
  inference(unit_resulting_resolution,[],[f437,f5398,f2167])).
fof(f10704,plain,(
  sP6('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl80 | $spl448)),
  inference(unit_resulting_resolution,[],[f250,f5398,f2167])).
fof(f10705,plain,(
  sP6('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl80 | $spl448)),
  inference(unit_resulting_resolution,[],[f249,f5398,f2167])).
fof(f10809,plain,(
  sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl742),
  inference(cnf_transformation,[],[f10809_D])).
fof(f10809_D,plain,(
  sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) <=> ~$spl742),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl742])])).
fof(f10706,plain,(
  sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | ($spl80 | $spl448)),
  inference(unit_resulting_resolution,[],[f437,f5398,f2166])).
fof(f10707,plain,(
  sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | ($spl80 | $spl448)),
  inference(unit_resulting_resolution,[],[f250,f5398,f2166])).
fof(f10708,plain,(
  sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | ($spl80 | $spl448)),
  inference(unit_resulting_resolution,[],[f249,f5398,f2166])).
fof(f10805,plain,(
  sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl740),
  inference(cnf_transformation,[],[f10805_D])).
fof(f10805_D,plain,(
  sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl740),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl740])])).
fof(f10711,plain,(
  sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl448),
  inference(unit_resulting_resolution,[],[f437,f5398,f381])).
fof(f10712,plain,(
  sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl448),
  inference(unit_resulting_resolution,[],[f250,f5398,f381])).
fof(f10713,plain,(
  sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl448),
  inference(unit_resulting_resolution,[],[f249,f5398,f381])).
fof(f10803,plain,(
  sP6('T_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl738),
  inference(cnf_transformation,[],[f10803_D])).
fof(f10803_D,plain,(
  sP6('T_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl738),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl738])])).
fof(f10715,plain,(
  sP6('T_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl448),
  inference(unit_resulting_resolution,[],[f437,f5398,f381])).
fof(f10716,plain,(
  sP6('T_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl448),
  inference(unit_resulting_resolution,[],[f250,f5398,f381])).
fof(f10717,plain,(
  sP6('T_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl448),
  inference(unit_resulting_resolution,[],[f249,f5398,f381])).
fof(f10627,plain,(
  ~sP6('F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl737),
  inference(cnf_transformation,[],[f10627_D])).
fof(f10627_D,plain,(
  ~sP6('F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl737),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl737])])).
fof(f10512,plain,(
  ~sP6('F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl80 | $spl340)),
  inference(unit_resulting_resolution,[],[f255,f4678,f2169])).
fof(f10625,plain,(
  ~sP6('F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl735),
  inference(cnf_transformation,[],[f10625_D])).
fof(f10625_D,plain,(
  ~sP6('F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) <=> ~$spl735),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl735])])).
fof(f10513,plain,(
  ~sP6('F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | ($spl80 | $spl340)),
  inference(unit_resulting_resolution,[],[f255,f4678,f2168])).
fof(f10623,plain,(
  sP6('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl732),
  inference(cnf_transformation,[],[f10623_D])).
fof(f10623_D,plain,(
  sP6('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl732),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl732])])).
fof(f10514,plain,(
  sP6('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl80 | $spl340)),
  inference(unit_resulting_resolution,[],[f437,f4678,f2167])).
fof(f10515,plain,(
  sP6('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl80 | $spl340)),
  inference(unit_resulting_resolution,[],[f250,f4678,f2167])).
fof(f10516,plain,(
  sP6('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl80 | $spl340)),
  inference(unit_resulting_resolution,[],[f249,f4678,f2167])).
fof(f10621,plain,(
  sP6('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl730),
  inference(cnf_transformation,[],[f10621_D])).
fof(f10621_D,plain,(
  sP6('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) <=> ~$spl730),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl730])])).
fof(f10517,plain,(
  sP6('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | ($spl80 | $spl340)),
  inference(unit_resulting_resolution,[],[f437,f4678,f2166])).
fof(f10518,plain,(
  sP6('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | ($spl80 | $spl340)),
  inference(unit_resulting_resolution,[],[f250,f4678,f2166])).
fof(f10519,plain,(
  sP6('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | ($spl80 | $spl340)),
  inference(unit_resulting_resolution,[],[f249,f4678,f2166])).
fof(f10617,plain,(
  sP6('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl728),
  inference(cnf_transformation,[],[f10617_D])).
fof(f10617_D,plain,(
  sP6('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl728),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl728])])).
fof(f10522,plain,(
  sP6('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl340),
  inference(unit_resulting_resolution,[],[f437,f4678,f381])).
fof(f10523,plain,(
  sP6('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl340),
  inference(unit_resulting_resolution,[],[f250,f4678,f381])).
fof(f10524,plain,(
  sP6('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl340),
  inference(unit_resulting_resolution,[],[f249,f4678,f381])).
fof(f10615,plain,(
  sP6('T_0','T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl726),
  inference(cnf_transformation,[],[f10615_D])).
fof(f10615_D,plain,(
  sP6('T_0','T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl726),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl726])])).
fof(f10526,plain,(
  sP6('T_0','T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl340),
  inference(unit_resulting_resolution,[],[f437,f4678,f381])).
fof(f10527,plain,(
  sP6('T_0','T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl340),
  inference(unit_resulting_resolution,[],[f250,f4678,f381])).
fof(f10528,plain,(
  sP6('T_0','T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl340),
  inference(unit_resulting_resolution,[],[f249,f4678,f381])).
fof(f10150,plain,(
  ~sP6('F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl725),
  inference(cnf_transformation,[],[f10150_D])).
fof(f10150_D,plain,(
  ~sP6('F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl725),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl725])])).
fof(f10138,plain,(
  ~sP6('F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl80 | $spl168)),
  inference(unit_resulting_resolution,[],[f255,f1400,f2169])).
fof(f10136,plain,(
  ~sP6('F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl723),
  inference(cnf_transformation,[],[f10136_D])).
fof(f10136_D,plain,(
  ~sP6('F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) <=> ~$spl723),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl723])])).
fof(f10124,plain,(
  ~sP6('F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | ($spl80 | $spl168)),
  inference(unit_resulting_resolution,[],[f255,f1400,f2168])).
fof(f10122,plain,(
  sP6('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl720),
  inference(cnf_transformation,[],[f10122_D])).
fof(f10122_D,plain,(
  sP6('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl720),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl720])])).
fof(f10102,plain,(
  sP6('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl80 | $spl168)),
  inference(unit_resulting_resolution,[],[f437,f1400,f2167])).
fof(f10103,plain,(
  sP6('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl80 | $spl168)),
  inference(unit_resulting_resolution,[],[f250,f1400,f2167])).
fof(f10104,plain,(
  sP6('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl80 | $spl168)),
  inference(unit_resulting_resolution,[],[f249,f1400,f2167])).
fof(f10098,plain,(
  sP6('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl718),
  inference(cnf_transformation,[],[f10098_D])).
fof(f10098_D,plain,(
  sP6('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) <=> ~$spl718),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl718])])).
fof(f10073,plain,(
  sP6('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl80),
  inference(unit_resulting_resolution,[],[f437,f903,f2166])).
fof(f10074,plain,(
  sP6('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl80),
  inference(unit_resulting_resolution,[],[f250,f903,f2166])).
fof(f10075,plain,(
  sP6('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl80),
  inference(unit_resulting_resolution,[],[f249,f903,f2166])).
fof(f10096,plain,(
  sP6('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl716),
  inference(cnf_transformation,[],[f10096_D])).
fof(f10096_D,plain,(
  sP6('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) <=> ~$spl716),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl716])])).
fof(f10076,plain,(
  sP6('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | ($spl80 | $spl168)),
  inference(unit_resulting_resolution,[],[f437,f1400,f2166])).
fof(f10077,plain,(
  sP6('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | ($spl80 | $spl168)),
  inference(unit_resulting_resolution,[],[f250,f1400,f2166])).
fof(f10078,plain,(
  sP6('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | ($spl80 | $spl168)),
  inference(unit_resulting_resolution,[],[f249,f1400,f2166])).
fof(f10072,plain,(
  sP4('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','T_0') | $spl714),
  inference(cnf_transformation,[],[f10072_D])).
fof(f10072_D,plain,(
  sP4('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','T_0') <=> ~$spl714),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl714])])).
fof(f10068,plain,(
  sP4('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','T_0') | $spl454),
  inference(resolution,[],[f5544,f362])).
fof(f10064,plain,(
  sP4('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','T_0') | $spl454),
  inference(unit_resulting_resolution,[],[f5544,f362])).
fof(f10057,plain,(
  sP5('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl710),
  inference(cnf_transformation,[],[f10057_D])).
fof(f10057_D,plain,(
  sP5('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl710),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl710])])).
fof(f10052,plain,(
  sP5('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl452),
  inference(resolution,[],[f5542,f363])).
fof(f10059,plain,(
  sP4('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl712),
  inference(cnf_transformation,[],[f10059_D])).
fof(f10059_D,plain,(
  sP4('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl712),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl712])])).
fof(f10051,plain,(
  sP4('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl452),
  inference(resolution,[],[f5542,f362])).
fof(f10046,plain,(
  sP4('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl452),
  inference(unit_resulting_resolution,[],[f5542,f362])).
fof(f10047,plain,(
  sP5('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl452),
  inference(unit_resulting_resolution,[],[f5542,f363])).
fof(f10020,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl706),
  inference(cnf_transformation,[],[f10020_D])).
fof(f10020_D,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) <=> ~$spl706),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl706])])).
fof(f10024,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl450),
  inference(global_subsumption,[],[f248,f247,f246,f5400,f9620,f9615,f9614,f9613,f9612,f9611,f9606,f9605,f9604,f9603,f9602,f9601,f9597,f9594,f9589,f9588,f9579,f9578,f9577,f9576,f9567,f9566,f9565,f9564,f9492,f9473,f9472,f9471,f9470,f9462,f9459,f9458,f9450,f9447,f9446,f9435,f9386,f9385,f9384,f9383,f9382,f9381,f9380,f9379,f9378,f9377,f9376,f9375,f9374,f9373,f9372,f9371,f9370,f9369,f9368,f9367,f9366,f9365,f9364,f9363,f9362,f9361,f9360,f9359,f9358,f9357,f9356,f9355,f9354,f9353,f9352,f9351,f9350,f9349,f9348,f9347,f9346,f9345,f9344,f9343,f9342,f9341,f9340,f9339,f9338,f9337,f9336,f9335,f9334,f9333,f9332,f9331,f9330,f9329,f9328,f9319,f9318,f9317,f9316,f9307,f9306,f9305,f9304,f9237,f9236,f9235,f9234,f9225,f9224,f9223,f9222,f9213,f9212,f9211,f9210,f9209,f9200,f9199,f9198,f9197,f9188,f9187,f9186,f9185,f9150,f9149,f9148,f9147,f9146,f9145,f9144,f9143,f9142,f9141,f9137,f9120,f9116,f9109,f9108,f9107,f9106,f9105,f9104,f9103,f9102,f9101,f9100,f9099,f9098,f9097,f9096,f9095,f9094,f9093,f9092,f9091,f9090,f9089,f9088,f9080,f9079,f9078,f9077,f9674,f9675,f9676,f9685,f9686,f9687,f9690,f9693,f9695,f9702,f9705,f9706,f9708,f9709,f9710,f9712])).
fof(f9712,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl450),
  inference(superposition,[],[f5400,f256])).
fof(f9710,plain,(
  ( ! [X52,X51] : (sP12('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),X51,X52) | ~p(s(bool,X51))) ) | $spl450),
  inference(resolution,[],[f5400,f431])).
fof(f9709,plain,(
  ( ! [X50,X49] : (sP12('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),X50,X49) | p(s(bool,X49))) ) | $spl450),
  inference(resolution,[],[f5400,f430])).
fof(f9708,plain,(
  ( ! [X47,X48] : (sP12(X48,X47,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | p(s(bool,X47)) | p(s(bool,X48))) ) | $spl450),
  inference(resolution,[],[f5400,f429])).
fof(f9706,plain,(
  ( ! [X43,X44] : (sP10('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),X43,X44) | ~p(s(bool,X43))) ) | $spl450),
  inference(resolution,[],[f5400,f414])).
fof(f9705,plain,(
  ( ! [X41,X42] : (sP10('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),X42,X41) | ~p(s(bool,X41))) ) | $spl450),
  inference(resolution,[],[f5400,f413])).
fof(f9702,plain,(
  ( ! [X35,X36] : (sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),X35,X36) | ~p(s(bool,X36)) | ~p(s(bool,X35))) ) | $spl450),
  inference(resolution,[],[f5400,f397])).
fof(f9695,plain,(
  ( ! [X21,X22] : (sP2('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),X21,X22) | ~p(s(bool,X22))) ) | $spl450),
  inference(resolution,[],[f5400,f352])).
fof(f9693,plain,(
  ( ! [X17,X18] : (sP2(X17,X18,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ~p(s(bool,X17))) ) | $spl450),
  inference(resolution,[],[f5400,f350])).
fof(f9690,plain,(
  ( ! [X12] : (~sP1('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),X12) | ~p(s(bool,X12))) ) | $spl450),
  inference(resolution,[],[f5400,f335])).
fof(f9687,plain,(
  ( ! [X9] : (sP1(X9,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | p(s(bool,X9))) ) | $spl450),
  inference(resolution,[],[f5400,f332])).
fof(f9686,plain,(
  ( ! [X8] : (sP1('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),X8) | p(s(bool,X8))) ) | $spl450),
  inference(resolution,[],[f5400,f329])).
fof(f9685,plain,(
  ( ! [X7] : (~sP1(X7,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ~p(s(bool,X7))) ) | $spl450),
  inference(resolution,[],[f5400,f328])).
fof(f9674,plain,(
  ( ! [X2,X1] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),'const.pred_set.UNIV_0'),s(fun(sK13,bool),X1)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK13,bool),X2),s(fun(sK13,bool),sK14)))) | p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),X2),s(fun(sK13,bool),X1))))) ) | $spl450),
  inference(resolution,[],[f5400,f344])).
fof(f9088,plain,(
  ( ! [X0] : (sP0(X0,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')))) ) | $spl450),
  inference(unit_resulting_resolution,[],[f255,f5400,f318])).
fof(f9089,plain,(
  ( ! [X0] : (sP0('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),X0)) ) | $spl450),
  inference(unit_resulting_resolution,[],[f5400,f319])).
fof(f9116,plain,(
  ( ! [X0] : (sP2(X0,'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')))) ) | $spl450),
  inference(unit_resulting_resolution,[],[f249,f5400,f348])).
fof(f9120,plain,(
  ( ! [X0] : (sP2('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0',X0)) ) | $spl450),
  inference(unit_resulting_resolution,[],[f249,f5400,f349])).
fof(f9137,plain,(
  ( ! [X0] : (sP2('T_0',X0,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')))) ) | $spl450),
  inference(unit_resulting_resolution,[],[f249,f5400,f352])).
fof(f9141,plain,(
  ( ! [X0] : (sP2('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),X0,'T_0')) ) | $spl450),
  inference(unit_resulting_resolution,[],[f249,f5400,f352])).
fof(f9142,plain,(
  ( ! [X0] : (sP2('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),X0,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')))) ) | $spl450),
  inference(unit_resulting_resolution,[],[f5400,f5400,f352])).
fof(f9236,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),X1)) ) | $spl450),
  inference(unit_resulting_resolution,[],[f5400,f378])).
fof(f9237,plain,(
  ( ! [X0,X1] : (sP3('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),X0,X1)) ) | $spl450),
  inference(unit_resulting_resolution,[],[f5400,f379])).
fof(f9328,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')))) ) | $spl450),
  inference(unit_resulting_resolution,[],[f5400,f392])).
fof(f9377,plain,(
  ( ! [X0] : (~sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),X0,'F_0')) ) | $spl450),
  inference(unit_resulting_resolution,[],[f255,f5400,f398])).
fof(f9378,plain,(
  ( ! [X0] : (~sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0',X0)) ) | $spl450),
  inference(unit_resulting_resolution,[],[f255,f5400,f399])).
fof(f9446,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')))) ) | $spl450),
  inference(unit_resulting_resolution,[],[f5400,f409])).
fof(f9447,plain,(
  ( ! [X0,X1] : (sP9(X0,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),X1)) ) | $spl450),
  inference(unit_resulting_resolution,[],[f5400,f410])).
fof(f9450,plain,(
  ( ! [X0] : (sP10('T_0',X0,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')))) ) | $spl450),
  inference(unit_resulting_resolution,[],[f249,f5400,f413])).
fof(f9458,plain,(
  ( ! [X0] : (sP10('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),X0,'T_0')) ) | $spl450),
  inference(unit_resulting_resolution,[],[f249,f5400,f413])).
fof(f9459,plain,(
  ( ! [X0] : (sP10('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),X0,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')))) ) | $spl450),
  inference(unit_resulting_resolution,[],[f5400,f5400,f413])).
fof(f9462,plain,(
  ( ! [X0] : (sP10('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),X0)) ) | $spl450),
  inference(unit_resulting_resolution,[],[f249,f5400,f414])).
fof(f9470,plain,(
  ( ! [X0] : (sP10('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0',X0)) ) | $spl450),
  inference(unit_resulting_resolution,[],[f249,f5400,f414])).
fof(f9471,plain,(
  ( ! [X0] : (sP10('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),X0)) ) | $spl450),
  inference(unit_resulting_resolution,[],[f5400,f5400,f414])).
fof(f9473,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')))) ) | $spl450),
  inference(unit_resulting_resolution,[],[f255,f5400,f416])).
fof(f9492,plain,(
  ( ! [X0] : (~sP10('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),X0)) ) | $spl450),
  inference(unit_resulting_resolution,[],[f255,f5400,f417])).
fof(f9588,plain,(
  ( ! [X0,X1] : (sP11(X0,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),X1)) ) | $spl450),
  inference(unit_resulting_resolution,[],[f5400,f427])).
fof(f9594,plain,(
  ( ! [X0] : (sP12('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),X0,'F_0')) ) | $spl450),
  inference(unit_resulting_resolution,[],[f255,f5400,f430])).
fof(f9597,plain,(
  ( ! [X0] : (sP12('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),X0)) ) | $spl450),
  inference(unit_resulting_resolution,[],[f249,f5400,f431])).
fof(f9601,plain,(
  ( ! [X0] : (sP12('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0',X0)) ) | $spl450),
  inference(unit_resulting_resolution,[],[f249,f5400,f431])).
fof(f9602,plain,(
  ( ! [X0] : (sP12('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),X0)) ) | $spl450),
  inference(unit_resulting_resolution,[],[f5400,f5400,f431])).
fof(f9615,plain,(
  ( ! [X0] : (~sP12('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),X0)) ) | $spl450),
  inference(unit_resulting_resolution,[],[f255,f5400,f434])).
fof(f9620,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),X1)) ) | $spl450),
  inference(unit_resulting_resolution,[],[f5400,f438])).
fof(f9676,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl450),
  inference(resolution,[],[f5400,f265])).
fof(f265,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f163])).
fof(f163,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X10] : (s(bool,'T_0') = s(bool,X10) <=> p(s(bool,X10)))),
  file('thm.pred_set.SUM_IMAGE_PERMUTES','thm.bool.EQ_CLAUSES|split|1')).
fof(f9675,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl450),
  inference(resolution,[],[f5400,f263])).
fof(f263,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f162])).
fof(f162,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f56])).
fof(f56,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X10] : (s(bool,'T_0') = s(bool,X10) <=> p(s(bool,X10)))),
  file('thm.pred_set.SUM_IMAGE_PERMUTES','thm.bool.EQ_CLAUSES|split|0')).
fof(f10023,plain,(
  ( ! [X0] : (s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun(sK13,'type.num.num'),X0),s(fun(sK13,bool),sK14))) = s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK13,'type.num.num'),X0),s(fun(sK13,sK13),sK15))),s(fun(sK13,bool),sK14)))) ) | $spl708),
  inference(cnf_transformation,[],[f10023_D])).
fof(f10023_D,plain,(
  ( ! [X0] : (s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun(sK13,'type.num.num'),X0),s(fun(sK13,bool),sK14))) = s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK13,'type.num.num'),X0),s(fun(sK13,sK13),sK15))),s(fun(sK13,bool),sK14)))) ) <=> ~$spl708),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl708])])).
fof(f10022,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun(sK13,'type.num.num'),X0),s(fun(sK13,bool),sK14))) = s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK13,'type.num.num'),X0),s(fun(sK13,sK13),sK15))),s(fun(sK13,bool),sK14)))) ) | ($spl80 | $spl342 | $spl450)),
  inference(forward_demodulation,[],[f10021,f903])).
fof(f10021,plain,(
  ( ! [X0] : (s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun(sK13,'type.num.num'),X0),s(fun(sK13,bool),sK14))) = s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK13,'type.num.num'),X0),s(fun(sK13,sK13),sK15))),s(fun(sK13,bool),sK14))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | ($spl342 | $spl450)),
  inference(forward_demodulation,[],[f9673,f4680])).
fof(f9673,plain,(
  ( ! [X0] : (s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun(sK13,'type.num.num'),X0),s(fun(sK13,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14))))) = s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK13,'type.num.num'),X0),s(fun(sK13,sK13),sK15))),s(fun(sK13,bool),sK14))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))) ) | $spl450),
  inference(resolution,[],[f5400,f337])).
fof(f9077,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl450),
  inference(unit_resulting_resolution,[],[f5400,f263])).
fof(f9078,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl450),
  inference(unit_resulting_resolution,[],[f5400,f265])).
fof(f10018,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl705),
  inference(cnf_transformation,[],[f10018_D])).
fof(f10018_D,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) <=> ~$spl705),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl705])])).
fof(f9079,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl450),
  inference(unit_resulting_resolution,[],[f5400,f266])).
fof(f266,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f164])).
fof(f164,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f58])).
fof(f58,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X10] : (s(bool,'F_0') = s(bool,X10) <=> ~p(s(bool,X10)))),
  file('thm.pred_set.SUM_IMAGE_PERMUTES','thm.bool.EQ_CLAUSES|split|3')).
fof(f9080,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl450),
  inference(unit_resulting_resolution,[],[f5400,f268])).
fof(f268,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f165])).
fof(f165,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f61])).
fof(f61,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f60])).
fof(f60,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X10] : (s(bool,'F_0') = s(bool,X10) <=> ~p(s(bool,X10)))),
  file('thm.pred_set.SUM_IMAGE_PERMUTES','thm.bool.EQ_CLAUSES|split|2')).
fof(f10010,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | $spl697),
  inference(cnf_transformation,[],[f10010_D])).
fof(f10010_D,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') <=> ~$spl697),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl697])])).
fof(f9090,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | $spl450),
  inference(unit_resulting_resolution,[],[f437,f5400,f328])).
fof(f9091,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | $spl450),
  inference(unit_resulting_resolution,[],[f250,f5400,f328])).
fof(f9092,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | $spl450),
  inference(unit_resulting_resolution,[],[f249,f5400,f328])).
fof(f10008,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl695),
  inference(cnf_transformation,[],[f10008_D])).
fof(f10008_D,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) <=> ~$spl695),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl695])])).
fof(f9093,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl450),
  inference(unit_resulting_resolution,[],[f5400,f5400,f328])).
fof(f10012,plain,(
  ~sP1('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl699),
  inference(cnf_transformation,[],[f10012_D])).
fof(f10012_D,plain,(
  ~sP1('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) <=> ~$spl699),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl699])])).
fof(f9094,plain,(
  ~sP1('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl450),
  inference(unit_resulting_resolution,[],[f437,f5400,f328])).
fof(f9095,plain,(
  ~sP1('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl450),
  inference(unit_resulting_resolution,[],[f250,f5400,f328])).
fof(f9096,plain,(
  ~sP1('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl450),
  inference(unit_resulting_resolution,[],[f249,f5400,f328])).
fof(f9097,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl450),
  inference(unit_resulting_resolution,[],[f5400,f5400,f328])).
fof(f10014,plain,(
  sP1('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0') | $spl700),
  inference(cnf_transformation,[],[f10014_D])).
fof(f10014_D,plain,(
  sP1('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0') <=> ~$spl700),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl700])])).
fof(f9098,plain,(
  sP1('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0') | $spl450),
  inference(unit_resulting_resolution,[],[f255,f5400,f329])).
fof(f10016,plain,(
  sP1('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl702),
  inference(cnf_transformation,[],[f10016_D])).
fof(f10016_D,plain,(
  sP1('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) <=> ~$spl702),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl702])])).
fof(f9099,plain,(
  sP1('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl450),
  inference(unit_resulting_resolution,[],[f255,f5400,f332])).
fof(f9100,plain,(
  sP1('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl450),
  inference(unit_resulting_resolution,[],[f255,f5400,f333])).
fof(f9101,plain,(
  sP1('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0') | $spl450),
  inference(unit_resulting_resolution,[],[f255,f5400,f334])).
fof(f9102,plain,(
  ~sP1('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl450),
  inference(unit_resulting_resolution,[],[f437,f5400,f335])).
fof(f9103,plain,(
  ~sP1('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl450),
  inference(unit_resulting_resolution,[],[f250,f5400,f335])).
fof(f9104,plain,(
  ~sP1('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl450),
  inference(unit_resulting_resolution,[],[f249,f5400,f335])).
fof(f9105,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl450),
  inference(unit_resulting_resolution,[],[f5400,f5400,f335])).
fof(f9106,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | $spl450),
  inference(unit_resulting_resolution,[],[f437,f5400,f335])).
fof(f9107,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | $spl450),
  inference(unit_resulting_resolution,[],[f250,f5400,f335])).
fof(f9108,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | $spl450),
  inference(unit_resulting_resolution,[],[f249,f5400,f335])).
fof(f9109,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl450),
  inference(unit_resulting_resolution,[],[f5400,f5400,f335])).
fof(f9918,plain,(
  ~sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0','F_0') | $spl637),
  inference(cnf_transformation,[],[f9918_D])).
fof(f9918_D,plain,(
  ~sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0','F_0') <=> ~$spl637),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl637])])).
fof(f9143,plain,(
  ~sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0','F_0') | $spl450),
  inference(unit_resulting_resolution,[],[f437,f255,f5400,f364])).
fof(f9144,plain,(
  ~sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0','F_0') | $spl450),
  inference(unit_resulting_resolution,[],[f250,f255,f5400,f364])).
fof(f9145,plain,(
  ~sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0','F_0') | $spl450),
  inference(unit_resulting_resolution,[],[f249,f255,f5400,f364])).
fof(f10004,plain,(
  ~sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0') | $spl691),
  inference(cnf_transformation,[],[f10004_D])).
fof(f10004_D,plain,(
  ~sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0') <=> ~$spl691),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl691])])).
fof(f9146,plain,(
  ~sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0') | $spl450),
  inference(unit_resulting_resolution,[],[f5400,f255,f5400,f364])).
fof(f10006,plain,(
  ~sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0') | $spl693),
  inference(cnf_transformation,[],[f10006_D])).
fof(f10006_D,plain,(
  ~sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0') <=> ~$spl693),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl693])])).
fof(f9147,plain,(
  ~sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0') | $spl450),
  inference(unit_resulting_resolution,[],[f437,f255,f5400,f364])).
fof(f9148,plain,(
  ~sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0') | $spl450),
  inference(unit_resulting_resolution,[],[f250,f255,f5400,f364])).
fof(f9149,plain,(
  ~sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0') | $spl450),
  inference(unit_resulting_resolution,[],[f249,f255,f5400,f364])).
fof(f9150,plain,(
  ~sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0') | $spl450),
  inference(unit_resulting_resolution,[],[f5400,f255,f5400,f364])).
fof(f10002,plain,(
  sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl688),
  inference(cnf_transformation,[],[f10002_D])).
fof(f10002_D,plain,(
  sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) <=> ~$spl688),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl688])])).
fof(f9151,plain,(
  sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ($spl4 | $spl450)),
  inference(unit_resulting_resolution,[],[f577,f575,f562,f5400,f365])).
fof(f10000,plain,(
  sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl686),
  inference(cnf_transformation,[],[f10000_D])).
fof(f10000_D,plain,(
  sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) <=> ~$spl686),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl686])])).
fof(f9152,plain,(
  sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ($spl2 | $spl4 | $spl450)),
  inference(unit_resulting_resolution,[],[f577,f1066,f562,f5400,f365])).
fof(f9998,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl684),
  inference(cnf_transformation,[],[f9998_D])).
fof(f9998_D,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) <=> ~$spl684),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl684])])).
fof(f9155,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ($spl4 | $spl172 | $spl450)),
  inference(unit_resulting_resolution,[],[f577,f4826,f562,f5400,f365])).
fof(f9996,plain,(
  sP6('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl682),
  inference(cnf_transformation,[],[f9996_D])).
fof(f9996_D,plain,(
  sP6('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) <=> ~$spl682),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl682])])).
fof(f9156,plain,(
  sP6('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ($spl4 | $spl170 | $spl450)),
  inference(unit_resulting_resolution,[],[f577,f4105,f562,f5400,f365])).
fof(f9994,plain,(
  sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl680),
  inference(cnf_transformation,[],[f9994_D])).
fof(f9994_D,plain,(
  sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) <=> ~$spl680),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl680])])).
fof(f9157,plain,(
  sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ($spl2 | $spl4 | $spl450)),
  inference(unit_resulting_resolution,[],[f1068,f575,f1053,f5400,f365])).
fof(f1068,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f445,f378])).
fof(f9992,plain,(
  sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl678),
  inference(cnf_transformation,[],[f9992_D])).
fof(f9992_D,plain,(
  sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) <=> ~$spl678),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl678])])).
fof(f9158,plain,(
  sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ($spl2 | $spl450)),
  inference(unit_resulting_resolution,[],[f1068,f1066,f1053,f5400,f365])).
fof(f9990,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl676),
  inference(cnf_transformation,[],[f9990_D])).
fof(f9990_D,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) <=> ~$spl676),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl676])])).
fof(f9161,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ($spl2 | $spl172 | $spl450)),
  inference(unit_resulting_resolution,[],[f1068,f4826,f1053,f5400,f365])).
fof(f9988,plain,(
  sP6('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl674),
  inference(cnf_transformation,[],[f9988_D])).
fof(f9988_D,plain,(
  sP6('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) <=> ~$spl674),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl674])])).
fof(f9162,plain,(
  sP6('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ($spl2 | $spl170 | $spl450)),
  inference(unit_resulting_resolution,[],[f1068,f4105,f1053,f5400,f365])).
fof(f9986,plain,(
  sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl672),
  inference(cnf_transformation,[],[f9986_D])).
fof(f9986_D,plain,(
  sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) <=> ~$spl672),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl672])])).
fof(f9173,plain,(
  sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ($spl4 | $spl172 | $spl450)),
  inference(unit_resulting_resolution,[],[f4828,f575,f4809,f5400,f365])).
fof(f9984,plain,(
  sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl670),
  inference(cnf_transformation,[],[f9984_D])).
fof(f9984_D,plain,(
  sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) <=> ~$spl670),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl670])])).
fof(f9174,plain,(
  sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ($spl2 | $spl172 | $spl450)),
  inference(unit_resulting_resolution,[],[f4828,f1066,f4809,f5400,f365])).
fof(f9982,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl668),
  inference(cnf_transformation,[],[f9982_D])).
fof(f9982_D,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) <=> ~$spl668),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl668])])).
fof(f9177,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ($spl172 | $spl450)),
  inference(unit_resulting_resolution,[],[f4828,f4826,f4809,f5400,f365])).
fof(f9980,plain,(
  sP6('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl666),
  inference(cnf_transformation,[],[f9980_D])).
fof(f9980_D,plain,(
  sP6('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) <=> ~$spl666),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl666])])).
fof(f9178,plain,(
  sP6('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ($spl170 | $spl172 | $spl450)),
  inference(unit_resulting_resolution,[],[f4828,f4105,f4809,f5400,f365])).
fof(f9978,plain,(
  sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl664),
  inference(cnf_transformation,[],[f9978_D])).
fof(f9978_D,plain,(
  sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) <=> ~$spl664),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl664])])).
fof(f9179,plain,(
  sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ($spl4 | $spl170 | $spl450)),
  inference(unit_resulting_resolution,[],[f4107,f575,f4088,f5400,f365])).
fof(f9976,plain,(
  sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl662),
  inference(cnf_transformation,[],[f9976_D])).
fof(f9976_D,plain,(
  sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) <=> ~$spl662),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl662])])).
fof(f9180,plain,(
  sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ($spl2 | $spl170 | $spl450)),
  inference(unit_resulting_resolution,[],[f4107,f1066,f4088,f5400,f365])).
fof(f9974,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl660),
  inference(cnf_transformation,[],[f9974_D])).
fof(f9974_D,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) <=> ~$spl660),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl660])])).
fof(f9183,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ($spl170 | $spl172 | $spl450)),
  inference(unit_resulting_resolution,[],[f4107,f4826,f4088,f5400,f365])).
fof(f9972,plain,(
  sP6('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl658),
  inference(cnf_transformation,[],[f9972_D])).
fof(f9972_D,plain,(
  sP6('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) <=> ~$spl658),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl658])])).
fof(f9184,plain,(
  sP6('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ($spl170 | $spl450)),
  inference(unit_resulting_resolution,[],[f4107,f4105,f4088,f5400,f365])).
fof(f9970,plain,(
  ~sP5('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0','F_0') | $spl657),
  inference(cnf_transformation,[],[f9970_D])).
fof(f9970_D,plain,(
  ~sP5('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0','F_0') <=> ~$spl657),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl657])])).
fof(f9185,plain,(
  ~sP5('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0','F_0') | $spl450),
  inference(unit_resulting_resolution,[],[f437,f255,f5400,f368])).
fof(f9186,plain,(
  ~sP5('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0','F_0') | $spl450),
  inference(unit_resulting_resolution,[],[f250,f255,f5400,f368])).
fof(f9187,plain,(
  ~sP5('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0','F_0') | $spl450),
  inference(unit_resulting_resolution,[],[f249,f255,f5400,f368])).
fof(f9958,plain,(
  ~sP5('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0') | $spl653),
  inference(cnf_transformation,[],[f9958_D])).
fof(f9958_D,plain,(
  ~sP5('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0') <=> ~$spl653),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl653])])).
fof(f9188,plain,(
  ~sP5('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0') | $spl450),
  inference(unit_resulting_resolution,[],[f5400,f255,f5400,f368])).
fof(f9960,plain,(
  ~sP5('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0') | $spl655),
  inference(cnf_transformation,[],[f9960_D])).
fof(f9960_D,plain,(
  ~sP5('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0') <=> ~$spl655),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl655])])).
fof(f9197,plain,(
  ~sP5('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0') | $spl450),
  inference(unit_resulting_resolution,[],[f437,f255,f5400,f368])).
fof(f9198,plain,(
  ~sP5('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0') | $spl450),
  inference(unit_resulting_resolution,[],[f250,f255,f5400,f368])).
fof(f9199,plain,(
  ~sP5('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0') | $spl450),
  inference(unit_resulting_resolution,[],[f249,f255,f5400,f368])).
fof(f9200,plain,(
  ~sP5('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0') | $spl450),
  inference(unit_resulting_resolution,[],[f5400,f255,f5400,f368])).
fof(f9948,plain,(
  ~sP4('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0','F_0') | $spl651),
  inference(cnf_transformation,[],[f9948_D])).
fof(f9948_D,plain,(
  ~sP4('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0','F_0') <=> ~$spl651),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl651])])).
fof(f9210,plain,(
  ~sP4('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0','F_0') | $spl450),
  inference(unit_resulting_resolution,[],[f437,f255,f5400,f372])).
fof(f9211,plain,(
  ~sP4('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0','F_0') | $spl450),
  inference(unit_resulting_resolution,[],[f250,f255,f5400,f372])).
fof(f9212,plain,(
  ~sP4('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0','F_0') | $spl450),
  inference(unit_resulting_resolution,[],[f249,f255,f5400,f372])).
fof(f9936,plain,(
  ~sP4('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0') | $spl647),
  inference(cnf_transformation,[],[f9936_D])).
fof(f9936_D,plain,(
  ~sP4('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0') <=> ~$spl647),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl647])])).
fof(f9213,plain,(
  ~sP4('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0') | $spl450),
  inference(unit_resulting_resolution,[],[f5400,f255,f5400,f372])).
fof(f9938,plain,(
  ~sP4('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0') | $spl649),
  inference(cnf_transformation,[],[f9938_D])).
fof(f9938_D,plain,(
  ~sP4('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0') <=> ~$spl649),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl649])])).
fof(f9222,plain,(
  ~sP4('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0') | $spl450),
  inference(unit_resulting_resolution,[],[f437,f255,f5400,f372])).
fof(f9223,plain,(
  ~sP4('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0') | $spl450),
  inference(unit_resulting_resolution,[],[f250,f255,f5400,f372])).
fof(f9224,plain,(
  ~sP4('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0') | $spl450),
  inference(unit_resulting_resolution,[],[f249,f255,f5400,f372])).
fof(f9225,plain,(
  ~sP4('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0') | $spl450),
  inference(unit_resulting_resolution,[],[f5400,f255,f5400,f372])).
fof(f9926,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl644),
  inference(cnf_transformation,[],[f9926_D])).
fof(f9926_D,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) <=> ~$spl644),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl644])])).
fof(f9238,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | ($spl80 | $spl450)),
  inference(unit_resulting_resolution,[],[f903,f5400,f381])).
fof(f9924,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl642),
  inference(cnf_transformation,[],[f9924_D])).
fof(f9924_D,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl642),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl642])])).
fof(f9239,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl168 | $spl450)),
  inference(unit_resulting_resolution,[],[f1400,f5400,f381])).
fof(f9922,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl640),
  inference(cnf_transformation,[],[f9922_D])).
fof(f9922_D,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl640),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl640])])).
fof(f9241,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | ($spl80 | $spl450)),
  inference(unit_resulting_resolution,[],[f903,f5400,f381])).
fof(f9920,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl638),
  inference(cnf_transformation,[],[f9920_D])).
fof(f9920_D,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl638),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl638])])).
fof(f9242,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | ($spl168 | $spl450)),
  inference(unit_resulting_resolution,[],[f1400,f5400,f381])).
fof(f9244,plain,(
  ~sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0','F_0') | ($spl83 | $spl450)),
  inference(unit_resulting_resolution,[],[f908,f5400,f382])).
fof(f382,plain,(
  ( ! [X2,X0,X1] : (~sP6(X2,X1,X0) | ~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f222])).
fof(f9916,plain,(
  ~sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0','T_0') | $spl635),
  inference(cnf_transformation,[],[f9916_D])).
fof(f9916_D,plain,(
  ~sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0','T_0') <=> ~$spl635),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl635])])).
fof(f9245,plain,(
  ~sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0','T_0') | ($spl83 | $spl450)),
  inference(unit_resulting_resolution,[],[f908,f5400,f382])).
fof(f9728,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | $spl572),
  inference(cnf_transformation,[],[f9728_D])).
fof(f9728_D,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') <=> ~$spl572),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl572])])).
fof(f9248,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | ($spl4 | $spl450)),
  inference(unit_resulting_resolution,[],[f437,f593,f5400,f387])).
fof(f387,plain,(
  ( ! [X2,X0,X1] : (~sP7(X1,X2,X0) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f225])).
fof(f9249,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | ($spl4 | $spl450)),
  inference(unit_resulting_resolution,[],[f250,f593,f5400,f387])).
fof(f9250,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | ($spl4 | $spl450)),
  inference(unit_resulting_resolution,[],[f249,f593,f5400,f387])).
fof(f9724,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl568),
  inference(cnf_transformation,[],[f9724_D])).
fof(f9724_D,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) <=> ~$spl568),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl568])])).
fof(f9251,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ($spl4 | $spl450)),
  inference(unit_resulting_resolution,[],[f5400,f593,f5400,f387])).
fof(f9718,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | $spl562),
  inference(cnf_transformation,[],[f9718_D])).
fof(f9718_D,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') <=> ~$spl562),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl562])])).
fof(f9253,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | ($spl2 | $spl450)),
  inference(unit_resulting_resolution,[],[f437,f1086,f5400,f387])).
fof(f1086,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f445,f392])).
fof(f9254,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | ($spl2 | $spl450)),
  inference(unit_resulting_resolution,[],[f250,f1086,f5400,f387])).
fof(f9255,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | ($spl2 | $spl450)),
  inference(unit_resulting_resolution,[],[f249,f1086,f5400,f387])).
fof(f9714,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl558),
  inference(cnf_transformation,[],[f9714_D])).
fof(f9714_D,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) <=> ~$spl558),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl558])])).
fof(f9256,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ($spl2 | $spl450)),
  inference(unit_resulting_resolution,[],[f5400,f1086,f5400,f387])).
fof(f9914,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | $spl632),
  inference(cnf_transformation,[],[f9914_D])).
fof(f9914_D,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') <=> ~$spl632),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl632])])).
fof(f9262,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | ($spl172 | $spl450)),
  inference(unit_resulting_resolution,[],[f437,f4884,f5400,f387])).
fof(f4884,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f1404,f392])).
fof(f9263,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | ($spl172 | $spl450)),
  inference(unit_resulting_resolution,[],[f250,f4884,f5400,f387])).
fof(f9264,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | ($spl172 | $spl450)),
  inference(unit_resulting_resolution,[],[f249,f4884,f5400,f387])).
fof(f9908,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl626),
  inference(cnf_transformation,[],[f9908_D])).
fof(f9908_D,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) <=> ~$spl626),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl626])])).
fof(f9265,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ($spl172 | $spl450)),
  inference(unit_resulting_resolution,[],[f5400,f4884,f5400,f387])).
fof(f9912,plain,(
  sP8('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | $spl630),
  inference(cnf_transformation,[],[f9912_D])).
fof(f9912_D,plain,(
  sP8('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') <=> ~$spl630),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl630])])).
fof(f9267,plain,(
  sP8('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | ($spl170 | $spl450)),
  inference(unit_resulting_resolution,[],[f437,f4163,f5400,f387])).
fof(f4163,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl170),
  inference(unit_resulting_resolution,[],[f1402,f392])).
fof(f9268,plain,(
  sP8('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | ($spl170 | $spl450)),
  inference(unit_resulting_resolution,[],[f250,f4163,f5400,f387])).
fof(f9269,plain,(
  sP8('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | ($spl170 | $spl450)),
  inference(unit_resulting_resolution,[],[f249,f4163,f5400,f387])).
fof(f9904,plain,(
  sP8('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl622),
  inference(cnf_transformation,[],[f9904_D])).
fof(f9904_D,plain,(
  sP8('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) <=> ~$spl622),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl622])])).
fof(f9270,plain,(
  sP8('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ($spl170 | $spl450)),
  inference(unit_resulting_resolution,[],[f5400,f4163,f5400,f387])).
fof(f9726,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl570),
  inference(cnf_transformation,[],[f9726_D])).
fof(f9726_D,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) <=> ~$spl570),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl570])])).
fof(f9272,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ($spl4 | $spl450)),
  inference(unit_resulting_resolution,[],[f437,f593,f5400,f387])).
fof(f9273,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ($spl4 | $spl450)),
  inference(unit_resulting_resolution,[],[f250,f593,f5400,f387])).
fof(f9274,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ($spl4 | $spl450)),
  inference(unit_resulting_resolution,[],[f249,f593,f5400,f387])).
fof(f9275,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ($spl4 | $spl450)),
  inference(unit_resulting_resolution,[],[f5400,f593,f5400,f387])).
fof(f9716,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl560),
  inference(cnf_transformation,[],[f9716_D])).
fof(f9716_D,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) <=> ~$spl560),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl560])])).
fof(f9277,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ($spl2 | $spl450)),
  inference(unit_resulting_resolution,[],[f437,f1086,f5400,f387])).
fof(f9278,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ($spl2 | $spl450)),
  inference(unit_resulting_resolution,[],[f250,f1086,f5400,f387])).
fof(f9279,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ($spl2 | $spl450)),
  inference(unit_resulting_resolution,[],[f249,f1086,f5400,f387])).
fof(f9280,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ($spl2 | $spl450)),
  inference(unit_resulting_resolution,[],[f5400,f1086,f5400,f387])).
fof(f9910,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl628),
  inference(cnf_transformation,[],[f9910_D])).
fof(f9910_D,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) <=> ~$spl628),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl628])])).
fof(f9286,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ($spl172 | $spl450)),
  inference(unit_resulting_resolution,[],[f437,f4884,f5400,f387])).
fof(f9287,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ($spl172 | $spl450)),
  inference(unit_resulting_resolution,[],[f250,f4884,f5400,f387])).
fof(f9288,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ($spl172 | $spl450)),
  inference(unit_resulting_resolution,[],[f249,f4884,f5400,f387])).
fof(f9289,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ($spl172 | $spl450)),
  inference(unit_resulting_resolution,[],[f5400,f4884,f5400,f387])).
fof(f9906,plain,(
  sP8('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl624),
  inference(cnf_transformation,[],[f9906_D])).
fof(f9906_D,plain,(
  sP8('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) <=> ~$spl624),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl624])])).
fof(f9291,plain,(
  sP8('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ($spl170 | $spl450)),
  inference(unit_resulting_resolution,[],[f437,f4163,f5400,f387])).
fof(f9292,plain,(
  sP8('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ($spl170 | $spl450)),
  inference(unit_resulting_resolution,[],[f250,f4163,f5400,f387])).
fof(f9293,plain,(
  sP8('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ($spl170 | $spl450)),
  inference(unit_resulting_resolution,[],[f249,f4163,f5400,f387])).
fof(f9294,plain,(
  sP8('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ($spl170 | $spl450)),
  inference(unit_resulting_resolution,[],[f5400,f4163,f5400,f387])).
fof(f9894,plain,(
  ~sP7('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0','F_0') | $spl621),
  inference(cnf_transformation,[],[f9894_D])).
fof(f9894_D,plain,(
  ~sP7('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0','F_0') <=> ~$spl621),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl621])])).
fof(f9304,plain,(
  ~sP7('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0','F_0') | $spl450),
  inference(unit_resulting_resolution,[],[f437,f255,f5400,f391])).
fof(f9305,plain,(
  ~sP7('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0','F_0') | $spl450),
  inference(unit_resulting_resolution,[],[f250,f255,f5400,f391])).
fof(f9306,plain,(
  ~sP7('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0','F_0') | $spl450),
  inference(unit_resulting_resolution,[],[f249,f255,f5400,f391])).
fof(f9882,plain,(
  ~sP7('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0') | $spl617),
  inference(cnf_transformation,[],[f9882_D])).
fof(f9882_D,plain,(
  ~sP7('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0') <=> ~$spl617),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl617])])).
fof(f9307,plain,(
  ~sP7('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0') | $spl450),
  inference(unit_resulting_resolution,[],[f5400,f255,f5400,f391])).
fof(f9884,plain,(
  ~sP7('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0') | $spl619),
  inference(cnf_transformation,[],[f9884_D])).
fof(f9884_D,plain,(
  ~sP7('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0') <=> ~$spl619),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl619])])).
fof(f9316,plain,(
  ~sP7('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0') | $spl450),
  inference(unit_resulting_resolution,[],[f437,f255,f5400,f391])).
fof(f9317,plain,(
  ~sP7('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0') | $spl450),
  inference(unit_resulting_resolution,[],[f250,f255,f5400,f391])).
fof(f9318,plain,(
  ~sP7('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0') | $spl450),
  inference(unit_resulting_resolution,[],[f249,f255,f5400,f391])).
fof(f9319,plain,(
  ~sP7('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0') | $spl450),
  inference(unit_resulting_resolution,[],[f5400,f255,f5400,f391])).
fof(f9872,plain,(
  sP8('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | $spl614),
  inference(cnf_transformation,[],[f9872_D])).
fof(f9872_D,plain,(
  sP8('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') <=> ~$spl614),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl614])])).
fof(f9329,plain,(
  sP8('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | $spl450),
  inference(unit_resulting_resolution,[],[f437,f437,f5400,f397])).
fof(f9330,plain,(
  sP8('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | $spl450),
  inference(unit_resulting_resolution,[],[f250,f437,f5400,f397])).
fof(f9331,plain,(
  sP8('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | $spl450),
  inference(unit_resulting_resolution,[],[f249,f437,f5400,f397])).
fof(f9864,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | $spl606),
  inference(cnf_transformation,[],[f9864_D])).
fof(f9864_D,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') <=> ~$spl606),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl606])])).
fof(f9332,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | $spl450),
  inference(unit_resulting_resolution,[],[f5400,f437,f5400,f397])).
fof(f9333,plain,(
  sP8('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | $spl450),
  inference(unit_resulting_resolution,[],[f437,f250,f5400,f397])).
fof(f9334,plain,(
  sP8('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | $spl450),
  inference(unit_resulting_resolution,[],[f250,f250,f5400,f397])).
fof(f9335,plain,(
  sP8('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | $spl450),
  inference(unit_resulting_resolution,[],[f249,f250,f5400,f397])).
fof(f9336,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | $spl450),
  inference(unit_resulting_resolution,[],[f5400,f250,f5400,f397])).
fof(f9337,plain,(
  sP8('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | $spl450),
  inference(unit_resulting_resolution,[],[f437,f249,f5400,f397])).
fof(f9338,plain,(
  sP8('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | $spl450),
  inference(unit_resulting_resolution,[],[f250,f249,f5400,f397])).
fof(f9339,plain,(
  sP8('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | $spl450),
  inference(unit_resulting_resolution,[],[f249,f249,f5400,f397])).
fof(f9340,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | $spl450),
  inference(unit_resulting_resolution,[],[f5400,f249,f5400,f397])).
fof(f9868,plain,(
  sP8('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl610),
  inference(cnf_transformation,[],[f9868_D])).
fof(f9868_D,plain,(
  sP8('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) <=> ~$spl610),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl610])])).
fof(f9341,plain,(
  sP8('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl450),
  inference(unit_resulting_resolution,[],[f437,f5400,f5400,f397])).
fof(f9342,plain,(
  sP8('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl450),
  inference(unit_resulting_resolution,[],[f250,f5400,f5400,f397])).
fof(f9343,plain,(
  sP8('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl450),
  inference(unit_resulting_resolution,[],[f249,f5400,f5400,f397])).
fof(f9860,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl602),
  inference(cnf_transformation,[],[f9860_D])).
fof(f9860_D,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) <=> ~$spl602),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl602])])).
fof(f9344,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl450),
  inference(unit_resulting_resolution,[],[f5400,f5400,f5400,f397])).
fof(f9870,plain,(
  sP8('T_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl612),
  inference(cnf_transformation,[],[f9870_D])).
fof(f9870_D,plain,(
  sP8('T_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) <=> ~$spl612),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl612])])).
fof(f9345,plain,(
  sP8('T_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl450),
  inference(unit_resulting_resolution,[],[f437,f437,f5400,f397])).
fof(f9346,plain,(
  sP8('T_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl450),
  inference(unit_resulting_resolution,[],[f250,f437,f5400,f397])).
fof(f9347,plain,(
  sP8('T_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl450),
  inference(unit_resulting_resolution,[],[f249,f437,f5400,f397])).
fof(f9348,plain,(
  sP8('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl450),
  inference(unit_resulting_resolution,[],[f5400,f437,f5400,f397])).
fof(f9349,plain,(
  sP8('T_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl450),
  inference(unit_resulting_resolution,[],[f437,f250,f5400,f397])).
fof(f9350,plain,(
  sP8('T_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl450),
  inference(unit_resulting_resolution,[],[f250,f250,f5400,f397])).
fof(f9351,plain,(
  sP8('T_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl450),
  inference(unit_resulting_resolution,[],[f249,f250,f5400,f397])).
fof(f9352,plain,(
  sP8('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl450),
  inference(unit_resulting_resolution,[],[f5400,f250,f5400,f397])).
fof(f9353,plain,(
  sP8('T_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl450),
  inference(unit_resulting_resolution,[],[f437,f249,f5400,f397])).
fof(f9354,plain,(
  sP8('T_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl450),
  inference(unit_resulting_resolution,[],[f250,f249,f5400,f397])).
fof(f9355,plain,(
  sP8('T_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl450),
  inference(unit_resulting_resolution,[],[f249,f249,f5400,f397])).
fof(f9356,plain,(
  sP8('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl450),
  inference(unit_resulting_resolution,[],[f5400,f249,f5400,f397])).
fof(f9862,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl604),
  inference(cnf_transformation,[],[f9862_D])).
fof(f9862_D,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) <=> ~$spl604),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl604])])).
fof(f9357,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl450),
  inference(unit_resulting_resolution,[],[f437,f5400,f5400,f397])).
fof(f9358,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl450),
  inference(unit_resulting_resolution,[],[f250,f5400,f5400,f397])).
fof(f9359,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl450),
  inference(unit_resulting_resolution,[],[f249,f5400,f5400,f397])).
fof(f9360,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl450),
  inference(unit_resulting_resolution,[],[f5400,f5400,f5400,f397])).
fof(f9866,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0','T_0') | $spl608),
  inference(cnf_transformation,[],[f9866_D])).
fof(f9866_D,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0','T_0') <=> ~$spl608),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl608])])).
fof(f9361,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0','T_0') | $spl450),
  inference(unit_resulting_resolution,[],[f437,f437,f5400,f397])).
fof(f9362,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0','T_0') | $spl450),
  inference(unit_resulting_resolution,[],[f250,f437,f5400,f397])).
fof(f9363,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0','T_0') | $spl450),
  inference(unit_resulting_resolution,[],[f249,f437,f5400,f397])).
fof(f9364,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | $spl450),
  inference(unit_resulting_resolution,[],[f5400,f437,f5400,f397])).
fof(f9365,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0','T_0') | $spl450),
  inference(unit_resulting_resolution,[],[f437,f250,f5400,f397])).
fof(f9366,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0','T_0') | $spl450),
  inference(unit_resulting_resolution,[],[f250,f250,f5400,f397])).
fof(f9367,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0','T_0') | $spl450),
  inference(unit_resulting_resolution,[],[f249,f250,f5400,f397])).
fof(f9368,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | $spl450),
  inference(unit_resulting_resolution,[],[f5400,f250,f5400,f397])).
fof(f9369,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0','T_0') | $spl450),
  inference(unit_resulting_resolution,[],[f437,f249,f5400,f397])).
fof(f9370,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0','T_0') | $spl450),
  inference(unit_resulting_resolution,[],[f250,f249,f5400,f397])).
fof(f9371,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0','T_0') | $spl450),
  inference(unit_resulting_resolution,[],[f249,f249,f5400,f397])).
fof(f9372,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | $spl450),
  inference(unit_resulting_resolution,[],[f5400,f249,f5400,f397])).
fof(f9373,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl450),
  inference(unit_resulting_resolution,[],[f437,f5400,f5400,f397])).
fof(f9374,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl450),
  inference(unit_resulting_resolution,[],[f250,f5400,f5400,f397])).
fof(f9375,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl450),
  inference(unit_resulting_resolution,[],[f249,f5400,f5400,f397])).
fof(f9376,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl450),
  inference(unit_resulting_resolution,[],[f5400,f5400,f5400,f397])).
fof(f9858,plain,(
  ~sP8('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | $spl601),
  inference(cnf_transformation,[],[f9858_D])).
fof(f9858_D,plain,(
  ~sP8('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') <=> ~$spl601),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl601])])).
fof(f9379,plain,(
  ~sP8('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | $spl450),
  inference(unit_resulting_resolution,[],[f255,f437,f5400,f400])).
fof(f9380,plain,(
  ~sP8('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | $spl450),
  inference(unit_resulting_resolution,[],[f255,f250,f5400,f400])).
fof(f9381,plain,(
  ~sP8('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | $spl450),
  inference(unit_resulting_resolution,[],[f255,f249,f5400,f400])).
fof(f9854,plain,(
  ~sP8('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl597),
  inference(cnf_transformation,[],[f9854_D])).
fof(f9854_D,plain,(
  ~sP8('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) <=> ~$spl597),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl597])])).
fof(f9382,plain,(
  ~sP8('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl450),
  inference(unit_resulting_resolution,[],[f255,f5400,f5400,f400])).
fof(f9856,plain,(
  ~sP8('F_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl599),
  inference(cnf_transformation,[],[f9856_D])).
fof(f9856_D,plain,(
  ~sP8('F_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) <=> ~$spl599),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl599])])).
fof(f9383,plain,(
  ~sP8('F_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl450),
  inference(unit_resulting_resolution,[],[f255,f437,f5400,f400])).
fof(f9384,plain,(
  ~sP8('F_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl450),
  inference(unit_resulting_resolution,[],[f255,f250,f5400,f400])).
fof(f9385,plain,(
  ~sP8('F_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl450),
  inference(unit_resulting_resolution,[],[f255,f249,f5400,f400])).
fof(f9386,plain,(
  ~sP8('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl450),
  inference(unit_resulting_resolution,[],[f255,f5400,f5400,f400])).
fof(f9824,plain,(
  ~sP9('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0','F_0') | $spl595),
  inference(cnf_transformation,[],[f9824_D])).
fof(f9824_D,plain,(
  ~sP9('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0','F_0') <=> ~$spl595),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl595])])).
fof(f9435,plain,(
  ~sP9('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0','F_0') | $spl450),
  inference(unit_resulting_resolution,[],[f255,f255,f5400,f408])).
fof(f408,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f235])).
fof(f9812,plain,(
  ~sP10('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0','F_0') | $spl593),
  inference(cnf_transformation,[],[f9812_D])).
fof(f9812_D,plain,(
  ~sP10('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0','F_0') <=> ~$spl593),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl593])])).
fof(f9472,plain,(
  ~sP10('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0','F_0') | $spl450),
  inference(unit_resulting_resolution,[],[f255,f255,f5400,f415])).
fof(f415,plain,(
  ( ! [X2,X0,X1] : (~sP10(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f237])).
fof(f9774,plain,(
  ~sP11('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0','T_0') | $spl591),
  inference(cnf_transformation,[],[f9774_D])).
fof(f9774_D,plain,(
  ~sP11('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0','T_0') <=> ~$spl591),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl591])])).
fof(f9564,plain,(
  ~sP11('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0','T_0') | $spl450),
  inference(unit_resulting_resolution,[],[f255,f437,f5400,f425])).
fof(f9565,plain,(
  ~sP11('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0','T_0') | $spl450),
  inference(unit_resulting_resolution,[],[f255,f250,f5400,f425])).
fof(f9566,plain,(
  ~sP11('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0','T_0') | $spl450),
  inference(unit_resulting_resolution,[],[f255,f249,f5400,f425])).
fof(f9762,plain,(
  ~sP11('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl587),
  inference(cnf_transformation,[],[f9762_D])).
fof(f9762_D,plain,(
  ~sP11('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) <=> ~$spl587),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl587])])).
fof(f9567,plain,(
  ~sP11('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl450),
  inference(unit_resulting_resolution,[],[f255,f5400,f5400,f425])).
fof(f9764,plain,(
  ~sP11('T_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl589),
  inference(cnf_transformation,[],[f9764_D])).
fof(f9764_D,plain,(
  ~sP11('T_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) <=> ~$spl589),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl589])])).
fof(f9576,plain,(
  ~sP11('T_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl450),
  inference(unit_resulting_resolution,[],[f437,f255,f5400,f425])).
fof(f9577,plain,(
  ~sP11('T_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl450),
  inference(unit_resulting_resolution,[],[f250,f255,f5400,f425])).
fof(f9578,plain,(
  ~sP11('T_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl450),
  inference(unit_resulting_resolution,[],[f249,f255,f5400,f425])).
fof(f9579,plain,(
  ~sP11('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl450),
  inference(unit_resulting_resolution,[],[f5400,f255,f5400,f425])).
fof(f9752,plain,(
  sP12('F_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl584),
  inference(cnf_transformation,[],[f9752_D])).
fof(f9752_D,plain,(
  sP12('F_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) <=> ~$spl584),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl584])])).
fof(f9589,plain,(
  sP12('F_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl450),
  inference(unit_resulting_resolution,[],[f255,f255,f5400,f429])).
fof(f9746,plain,(
  ~sP12('T_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl583),
  inference(cnf_transformation,[],[f9746_D])).
fof(f9746_D,plain,(
  ~sP12('T_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) <=> ~$spl583),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl583])])).
fof(f9603,plain,(
  ~sP12('T_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl450),
  inference(unit_resulting_resolution,[],[f255,f437,f5400,f432])).
fof(f9604,plain,(
  ~sP12('T_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl450),
  inference(unit_resulting_resolution,[],[f255,f250,f5400,f432])).
fof(f9605,plain,(
  ~sP12('T_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl450),
  inference(unit_resulting_resolution,[],[f255,f249,f5400,f432])).
fof(f9738,plain,(
  ~sP12('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl579),
  inference(cnf_transformation,[],[f9738_D])).
fof(f9738_D,plain,(
  ~sP12('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) <=> ~$spl579),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl579])])).
fof(f9606,plain,(
  ~sP12('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl450),
  inference(unit_resulting_resolution,[],[f255,f5400,f5400,f432])).
fof(f9740,plain,(
  ~sP12('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0','T_0') | $spl581),
  inference(cnf_transformation,[],[f9740_D])).
fof(f9740_D,plain,(
  ~sP12('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0','T_0') <=> ~$spl581),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl581])])).
fof(f9611,plain,(
  ~sP12('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0','T_0') | $spl450),
  inference(unit_resulting_resolution,[],[f255,f437,f5400,f432])).
fof(f9612,plain,(
  ~sP12('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0','T_0') | $spl450),
  inference(unit_resulting_resolution,[],[f255,f250,f5400,f432])).
fof(f9613,plain,(
  ~sP12('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0','T_0') | $spl450),
  inference(unit_resulting_resolution,[],[f255,f249,f5400,f432])).
fof(f9614,plain,(
  ~sP12('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl450),
  inference(unit_resulting_resolution,[],[f255,f5400,f5400,f432])).
fof(f9732,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl577),
  inference(cnf_transformation,[],[f9732_D])).
fof(f9732_D,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) <=> ~$spl577),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl577])])).
fof(f9636,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | ($spl4 | $spl450)),
  inference(unit_resulting_resolution,[],[f5400,f800])).
fof(f9730,plain,(
  ~sP1('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl575),
  inference(cnf_transformation,[],[f9730_D])).
fof(f9730_D,plain,(
  ~sP1('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) <=> ~$spl575),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl575])])).
fof(f9637,plain,(
  ~sP1('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ($spl4 | $spl450)),
  inference(unit_resulting_resolution,[],[f5400,f805])).
fof(f9640,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | ($spl4 | $spl450)),
  inference(unit_resulting_resolution,[],[f437,f5400,f817])).
fof(f9641,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | ($spl4 | $spl450)),
  inference(unit_resulting_resolution,[],[f250,f5400,f817])).
fof(f9642,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | ($spl4 | $spl450)),
  inference(unit_resulting_resolution,[],[f249,f5400,f817])).
fof(f9643,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ($spl4 | $spl450)),
  inference(unit_resulting_resolution,[],[f5400,f5400,f817])).
fof(f9645,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ($spl4 | $spl450)),
  inference(unit_resulting_resolution,[],[f437,f5400,f817])).
fof(f9646,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ($spl4 | $spl450)),
  inference(unit_resulting_resolution,[],[f250,f5400,f817])).
fof(f9647,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ($spl4 | $spl450)),
  inference(unit_resulting_resolution,[],[f249,f5400,f817])).
fof(f9648,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ($spl4 | $spl450)),
  inference(unit_resulting_resolution,[],[f5400,f5400,f817])).
fof(f9722,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl567),
  inference(cnf_transformation,[],[f9722_D])).
fof(f9722_D,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl567),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl567])])).
fof(f9653,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl2 | $spl450)),
  inference(unit_resulting_resolution,[],[f5400,f1295])).
fof(f1295,plain,(
  ( ! [X4] : (~sP1(X4,'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ~p(s(bool,X4))) ) | $spl2),
  inference(resolution,[],[f445,f328])).
fof(f9720,plain,(
  ~sP1('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | $spl565),
  inference(cnf_transformation,[],[f9720_D])).
fof(f9720_D,plain,(
  ~sP1('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) <=> ~$spl565),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl565])])).
fof(f9654,plain,(
  ~sP1('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ($spl2 | $spl450)),
  inference(unit_resulting_resolution,[],[f5400,f1300])).
fof(f1300,plain,(
  ( ! [X9] : (~sP1('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X9) | ~p(s(bool,X9))) ) | $spl2),
  inference(resolution,[],[f445,f335])).
fof(f9657,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | ($spl2 | $spl450)),
  inference(unit_resulting_resolution,[],[f437,f5400,f1312])).
fof(f1312,plain,(
  ( ! [X33,X32] : (sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X32,X33) | ~p(s(bool,X33)) | ~p(s(bool,X32))) ) | $spl2),
  inference(resolution,[],[f445,f397])).
fof(f9658,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | ($spl2 | $spl450)),
  inference(unit_resulting_resolution,[],[f250,f5400,f1312])).
fof(f9659,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'T_0') | ($spl2 | $spl450)),
  inference(unit_resulting_resolution,[],[f249,f5400,f1312])).
fof(f9660,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ($spl2 | $spl450)),
  inference(unit_resulting_resolution,[],[f5400,f5400,f1312])).
fof(f9662,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ($spl2 | $spl450)),
  inference(unit_resulting_resolution,[],[f437,f5400,f1312])).
fof(f9663,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ($spl2 | $spl450)),
  inference(unit_resulting_resolution,[],[f250,f5400,f1312])).
fof(f9664,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ($spl2 | $spl450)),
  inference(unit_resulting_resolution,[],[f249,f5400,f1312])).
fof(f9665,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0'))) | ($spl2 | $spl450)),
  inference(unit_resulting_resolution,[],[f5400,f5400,f1312])).
fof(f9063,plain,(
  ~sP6('T_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl557),
  inference(cnf_transformation,[],[f9063_D])).
fof(f9063_D,plain,(
  ~sP6('T_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl557),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl557])])).
fof(f9059,plain,(
  ~sP6('T_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl421),
  inference(unit_resulting_resolution,[],[f5368,f363])).
fof(f9053,plain,(
  ~sP6('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl555),
  inference(cnf_transformation,[],[f9053_D])).
fof(f9053_D,plain,(
  ~sP6('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl555),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl555])])).
fof(f9049,plain,(
  ~sP6('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl417),
  inference(unit_resulting_resolution,[],[f5362,f362])).
fof(f9043,plain,(
  ~sP6('F_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl553),
  inference(cnf_transformation,[],[f9043_D])).
fof(f9043_D,plain,(
  ~sP6('F_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl553),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl553])])).
fof(f9039,plain,(
  ~sP6('F_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl415),
  inference(unit_resulting_resolution,[],[f5358,f362])).
fof(f8995,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','T_0') | $spl550),
  inference(cnf_transformation,[],[f8995_D])).
fof(f8995_D,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','T_0') <=> ~$spl550),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl550])])).
fof(f8990,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','T_0') | $spl390),
  inference(resolution,[],[f5324,f384])).
fof(f8984,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','T_0') | $spl390),
  inference(unit_resulting_resolution,[],[f5324,f384])).
fof(f8978,plain,(
  sP7('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl548),
  inference(cnf_transformation,[],[f8978_D])).
fof(f8978_D,plain,(
  sP7('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl548),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl548])])).
fof(f8973,plain,(
  sP7('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl388),
  inference(resolution,[],[f5322,f384])).
fof(f8967,plain,(
  sP7('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl388),
  inference(unit_resulting_resolution,[],[f5322,f384])).
fof(f8867,plain,(
  sP11('F_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl546),
  inference(cnf_transformation,[],[f8867_D])).
fof(f8867_D,plain,(
  sP11('F_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl546),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl546])])).
fof(f8865,plain,(
  sP11('F_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl360),
  inference(resolution,[],[f5262,f420])).
fof(f8859,plain,(
  sP11('F_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl360),
  inference(unit_resulting_resolution,[],[f5262,f420])).
fof(f8464,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl544),
  inference(cnf_transformation,[],[f8464_D])).
fof(f8464_D,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl544),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl544])])).
fof(f8456,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f437,f4809,f4826,f4827,f365])).
fof(f8457,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f250,f4809,f4826,f4827,f365])).
fof(f8458,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f249,f4809,f4826,f4827,f365])).
fof(f8441,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl542),
  inference(cnf_transformation,[],[f8441_D])).
fof(f8441_D,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl542),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl542])])).
fof(f8428,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | ($spl2 | $spl172)),
  inference(unit_resulting_resolution,[],[f437,f1068,f1053,f4826,f365])).
fof(f8429,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | ($spl2 | $spl172)),
  inference(unit_resulting_resolution,[],[f250,f1068,f1053,f4826,f365])).
fof(f8430,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | ($spl2 | $spl172)),
  inference(unit_resulting_resolution,[],[f249,f1068,f1053,f4826,f365])).
fof(f8439,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl540),
  inference(cnf_transformation,[],[f8439_D])).
fof(f8439_D,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl540),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl540])])).
fof(f8434,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | ($spl170 | $spl172)),
  inference(unit_resulting_resolution,[],[f437,f4107,f4088,f4826,f365])).
fof(f8435,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | ($spl170 | $spl172)),
  inference(unit_resulting_resolution,[],[f250,f4107,f4088,f4826,f365])).
fof(f8436,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | ($spl170 | $spl172)),
  inference(unit_resulting_resolution,[],[f249,f4107,f4088,f4826,f365])).
fof(f8415,plain,(
  sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl538),
  inference(cnf_transformation,[],[f8415_D])).
fof(f8415_D,plain,(
  sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl538),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl538])])).
fof(f8394,plain,(
  sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | ($spl4 | $spl172)),
  inference(unit_resulting_resolution,[],[f437,f576,f575,f4809,f365])).
fof(f8395,plain,(
  sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | ($spl4 | $spl172)),
  inference(unit_resulting_resolution,[],[f250,f576,f575,f4809,f365])).
fof(f8396,plain,(
  sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | ($spl4 | $spl172)),
  inference(unit_resulting_resolution,[],[f249,f576,f575,f4809,f365])).
fof(f8413,plain,(
  sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl536),
  inference(cnf_transformation,[],[f8413_D])).
fof(f8413_D,plain,(
  sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl536),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl536])])).
fof(f8397,plain,(
  sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | ($spl2 | $spl172)),
  inference(unit_resulting_resolution,[],[f437,f1067,f1066,f4809,f365])).
fof(f8398,plain,(
  sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | ($spl2 | $spl172)),
  inference(unit_resulting_resolution,[],[f250,f1067,f1066,f4809,f365])).
fof(f8399,plain,(
  sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | ($spl2 | $spl172)),
  inference(unit_resulting_resolution,[],[f249,f1067,f1066,f4809,f365])).
fof(f8411,plain,(
  sP6('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl534),
  inference(cnf_transformation,[],[f8411_D])).
fof(f8411_D,plain,(
  sP6('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl534),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl534])])).
fof(f8403,plain,(
  sP6('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | ($spl170 | $spl172)),
  inference(unit_resulting_resolution,[],[f437,f4106,f4105,f4809,f365])).
fof(f8404,plain,(
  sP6('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | ($spl170 | $spl172)),
  inference(unit_resulting_resolution,[],[f250,f4106,f4105,f4809,f365])).
fof(f8405,plain,(
  sP6('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | ($spl170 | $spl172)),
  inference(unit_resulting_resolution,[],[f249,f4106,f4105,f4809,f365])).
fof(f8333,plain,(
  ~sP6('T_0','F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl533),
  inference(cnf_transformation,[],[f8333_D])).
fof(f8333_D,plain,(
  ~sP6('T_0','F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl533),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl533])])).
fof(f8329,plain,(
  ~sP6('T_0','F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl313),
  inference(unit_resulting_resolution,[],[f4648,f363])).
fof(f8323,plain,(
  ~sP6('F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl531),
  inference(cnf_transformation,[],[f8323_D])).
fof(f8323_D,plain,(
  ~sP6('F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl531),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl531])])).
fof(f8319,plain,(
  ~sP6('F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl309),
  inference(unit_resulting_resolution,[],[f4642,f362])).
fof(f8313,plain,(
  ~sP6('F_0','T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl529),
  inference(cnf_transformation,[],[f8313_D])).
fof(f8313_D,plain,(
  ~sP6('F_0','T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl529),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl529])])).
fof(f8309,plain,(
  ~sP6('F_0','T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl307),
  inference(unit_resulting_resolution,[],[f4638,f362])).
fof(f8238,plain,(
  sP7('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','T_0') | $spl526),
  inference(cnf_transformation,[],[f8238_D])).
fof(f8238_D,plain,(
  sP7('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','T_0') <=> ~$spl526),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl526])])).
fof(f8233,plain,(
  sP7('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','T_0') | $spl282),
  inference(resolution,[],[f4604,f384])).
fof(f8227,plain,(
  sP7('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','T_0') | $spl282),
  inference(unit_resulting_resolution,[],[f4604,f384])).
fof(f8221,plain,(
  sP7('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl524),
  inference(cnf_transformation,[],[f8221_D])).
fof(f8221_D,plain,(
  sP7('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl524),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl524])])).
fof(f8216,plain,(
  sP7('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl280),
  inference(resolution,[],[f4602,f384])).
fof(f8210,plain,(
  sP7('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl280),
  inference(unit_resulting_resolution,[],[f4602,f384])).
fof(f8105,plain,(
  sP11('F_0','F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl522),
  inference(cnf_transformation,[],[f8105_D])).
fof(f8105_D,plain,(
  sP11('F_0','F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl522),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl522])])).
fof(f8103,plain,(
  sP11('F_0','F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl252),
  inference(resolution,[],[f4542,f420])).
fof(f8097,plain,(
  sP11('F_0','F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl252),
  inference(unit_resulting_resolution,[],[f4542,f420])).
fof(f7669,plain,(
  sP6('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl520),
  inference(cnf_transformation,[],[f7669_D])).
fof(f7669_D,plain,(
  sP6('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl520),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl520])])).
fof(f7664,plain,(
  sP6('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl170),
  inference(unit_resulting_resolution,[],[f437,f4088,f4105,f4106,f365])).
fof(f7665,plain,(
  sP6('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl170),
  inference(unit_resulting_resolution,[],[f250,f4088,f4105,f4106,f365])).
fof(f7666,plain,(
  sP6('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl170),
  inference(unit_resulting_resolution,[],[f249,f4088,f4105,f4106,f365])).
fof(f7652,plain,(
  sP6('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl518),
  inference(cnf_transformation,[],[f7652_D])).
fof(f7652_D,plain,(
  sP6('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl518),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl518])])).
fof(f7647,plain,(
  sP6('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | ($spl2 | $spl170)),
  inference(unit_resulting_resolution,[],[f437,f1068,f1053,f4105,f365])).
fof(f7648,plain,(
  sP6('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | ($spl2 | $spl170)),
  inference(unit_resulting_resolution,[],[f250,f1068,f1053,f4105,f365])).
fof(f7649,plain,(
  sP6('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | ($spl2 | $spl170)),
  inference(unit_resulting_resolution,[],[f249,f1068,f1053,f4105,f365])).
fof(f7634,plain,(
  sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl516),
  inference(cnf_transformation,[],[f7634_D])).
fof(f7634_D,plain,(
  sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl516),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl516])])).
fof(f7621,plain,(
  sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | ($spl4 | $spl170)),
  inference(unit_resulting_resolution,[],[f437,f576,f575,f4088,f365])).
fof(f7622,plain,(
  sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | ($spl4 | $spl170)),
  inference(unit_resulting_resolution,[],[f250,f576,f575,f4088,f365])).
fof(f7623,plain,(
  sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | ($spl4 | $spl170)),
  inference(unit_resulting_resolution,[],[f249,f576,f575,f4088,f365])).
fof(f7632,plain,(
  sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl514),
  inference(cnf_transformation,[],[f7632_D])).
fof(f7632_D,plain,(
  sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl514),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl514])])).
fof(f7624,plain,(
  sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | ($spl2 | $spl170)),
  inference(unit_resulting_resolution,[],[f437,f1067,f1066,f4088,f365])).
fof(f7625,plain,(
  sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | ($spl2 | $spl170)),
  inference(unit_resulting_resolution,[],[f250,f1067,f1066,f4088,f365])).
fof(f7626,plain,(
  sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | ($spl2 | $spl170)),
  inference(unit_resulting_resolution,[],[f249,f1067,f1066,f4088,f365])).
fof(f7478,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | p(s(bool,'T_0'))),
  introduced(theory_axiom,[])).
fof(f7461,plain,(
  ( ! [X76,X74,X77,X75] : (sP2(X76,'const.pred_set.SUBSET_2'(s(fun(X74,bool),X75),s(fun(X74,bool),X75)),X77)) ) | $spl480),
  inference(cnf_transformation,[],[f7461_D])).
fof(f7461_D,plain,(
  ( ! [X76,X74,X77,X75] : (sP2(X76,'const.pred_set.SUBSET_2'(s(fun(X74,bool),X75),s(fun(X74,bool),X75)),X77)) ) <=> ~$spl480),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl480])])).
fof(f7440,plain,(
  ( ! [X229,X227,X228,X226] : (~p(s(bool,'T_0')) | sP2(X228,'const.pred_set.SUBSET_2'(s(fun(X226,bool),X227),s(fun(X226,bool),X227)),X229)) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f438,f7327])).
fof(f7477,plain,(
  ( ! [X222,X225,X223,X224] : (~p(s(bool,X224)) | sP12('const.pred_set.SUBSET_2'(s(fun(X222,bool),X223),s(fun(X222,bool),X223)),X224,X225)) ) | $spl512),
  inference(cnf_transformation,[],[f7477_D])).
fof(f7477_D,plain,(
  ( ! [X222,X225,X223,X224] : (~p(s(bool,X224)) | sP12('const.pred_set.SUBSET_2'(s(fun(X222,bool),X223),s(fun(X222,bool),X223)),X224,X225)) ) <=> ~$spl512),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl512])])).
fof(f7439,plain,(
  ( ! [X222,X225,X223,X224] : (~p(s(bool,'T_0')) | ~p(s(bool,X224)) | sP12('const.pred_set.SUBSET_2'(s(fun(X222,bool),X223),s(fun(X222,bool),X223)),X224,X225)) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f431,f7327])).
fof(f7476,plain,(
  ( ! [X218,X220,X219,X221] : (sP12('const.pred_set.SUBSET_2'(s(fun(X218,bool),X219),s(fun(X218,bool),X219)),X221,X220) | p(s(bool,X220))) ) | $spl510),
  inference(cnf_transformation,[],[f7476_D])).
fof(f7476_D,plain,(
  ( ! [X218,X220,X219,X221] : (sP12('const.pred_set.SUBSET_2'(s(fun(X218,bool),X219),s(fun(X218,bool),X219)),X221,X220) | p(s(bool,X220))) ) <=> ~$spl510),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl510])])).
fof(f7438,plain,(
  ( ! [X218,X220,X219,X221] : (~p(s(bool,'T_0')) | p(s(bool,X220)) | sP12('const.pred_set.SUBSET_2'(s(fun(X218,bool),X219),s(fun(X218,bool),X219)),X221,X220)) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f430,f7327])).
fof(f7475,plain,(
  ( ! [X216,X215,X217,X214] : (p(s(bool,X216)) | p(s(bool,X217)) | sP12(X217,X216,'const.pred_set.SUBSET_2'(s(fun(X214,bool),X215),s(fun(X214,bool),X215)))) ) | $spl508),
  inference(cnf_transformation,[],[f7475_D])).
fof(f7475_D,plain,(
  ( ! [X216,X215,X217,X214] : (p(s(bool,X216)) | p(s(bool,X217)) | sP12(X217,X216,'const.pred_set.SUBSET_2'(s(fun(X214,bool),X215),s(fun(X214,bool),X215)))) ) <=> ~$spl508),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl508])])).
fof(f7437,plain,(
  ( ! [X216,X215,X217,X214] : (~p(s(bool,'T_0')) | p(s(bool,X216)) | sP12(X217,X216,'const.pred_set.SUBSET_2'(s(fun(X214,bool),X215),s(fun(X214,bool),X215))) | p(s(bool,X217))) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f429,f7327])).
fof(f7474,plain,(
  ( ! [X208,X206,X209,X207] : (sP11(X208,'const.pred_set.SUBSET_2'(s(fun(X206,bool),X207),s(fun(X206,bool),X207)),X209)) ) | $spl506),
  inference(cnf_transformation,[],[f7474_D])).
fof(f7474_D,plain,(
  ( ! [X208,X206,X209,X207] : (sP11(X208,'const.pred_set.SUBSET_2'(s(fun(X206,bool),X207),s(fun(X206,bool),X207)),X209)) ) <=> ~$spl506),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl506])])).
fof(f7435,plain,(
  ( ! [X208,X206,X209,X207] : (~p(s(bool,'T_0')) | sP11(X208,'const.pred_set.SUBSET_2'(s(fun(X206,bool),X207),s(fun(X206,bool),X207)),X209)) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f427,f7327])).
fof(f7473,plain,(
  ( ! [X198,X200,X199,X201] : (sP10('const.pred_set.SUBSET_2'(s(fun(X198,bool),X199),s(fun(X198,bool),X199)),X200,X201) | ~p(s(bool,X200))) ) | $spl504),
  inference(cnf_transformation,[],[f7473_D])).
fof(f7473_D,plain,(
  ( ! [X198,X200,X199,X201] : (sP10('const.pred_set.SUBSET_2'(s(fun(X198,bool),X199),s(fun(X198,bool),X199)),X200,X201) | ~p(s(bool,X200))) ) <=> ~$spl504),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl504])])).
fof(f7433,plain,(
  ( ! [X198,X200,X199,X201] : (~p(s(bool,'T_0')) | ~p(s(bool,X200)) | sP10('const.pred_set.SUBSET_2'(s(fun(X198,bool),X199),s(fun(X198,bool),X199)),X200,X201)) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f414,f7327])).
fof(f7472,plain,(
  ( ! [X196,X194,X197,X195] : (sP10('const.pred_set.SUBSET_2'(s(fun(X194,bool),X195),s(fun(X194,bool),X195)),X197,X196) | ~p(s(bool,X196))) ) | $spl502),
  inference(cnf_transformation,[],[f7472_D])).
fof(f7472_D,plain,(
  ( ! [X196,X194,X197,X195] : (sP10('const.pred_set.SUBSET_2'(s(fun(X194,bool),X195),s(fun(X194,bool),X195)),X197,X196) | ~p(s(bool,X196))) ) <=> ~$spl502),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl502])])).
fof(f7432,plain,(
  ( ! [X196,X194,X197,X195] : (~p(s(bool,'T_0')) | ~p(s(bool,X196)) | sP10('const.pred_set.SUBSET_2'(s(fun(X194,bool),X195),s(fun(X194,bool),X195)),X197,X196)) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f413,f7327])).
fof(f7471,plain,(
  ( ! [X177,X175,X176,X174] : (sP9(X176,'const.pred_set.SUBSET_2'(s(fun(X174,bool),X175),s(fun(X174,bool),X175)),X177)) ) | $spl500),
  inference(cnf_transformation,[],[f7471_D])).
fof(f7471_D,plain,(
  ( ! [X177,X175,X176,X174] : (sP9(X176,'const.pred_set.SUBSET_2'(s(fun(X174,bool),X175),s(fun(X174,bool),X175)),X177)) ) <=> ~$spl500),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl500])])).
fof(f7427,plain,(
  ( ! [X177,X175,X176,X174] : (~p(s(bool,'T_0')) | sP9(X176,'const.pred_set.SUBSET_2'(s(fun(X174,bool),X175),s(fun(X174,bool),X175)),X177)) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f410,f7327])).
fof(f7470,plain,(
  ( ! [X173,X171,X172,X170] : (sP9(X172,X173,'const.pred_set.SUBSET_2'(s(fun(X170,bool),X171),s(fun(X170,bool),X171)))) ) | $spl498),
  inference(cnf_transformation,[],[f7470_D])).
fof(f7470_D,plain,(
  ( ! [X173,X171,X172,X170] : (sP9(X172,X173,'const.pred_set.SUBSET_2'(s(fun(X170,bool),X171),s(fun(X170,bool),X171)))) ) <=> ~$spl498),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl498])])).
fof(f7426,plain,(
  ( ! [X173,X171,X172,X170] : (~p(s(bool,'T_0')) | sP9(X172,X173,'const.pred_set.SUBSET_2'(s(fun(X170,bool),X171),s(fun(X170,bool),X171)))) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f409,f7327])).
fof(f7469,plain,(
  ( ! [X167,X169,X166,X168] : (~p(s(bool,X168)) | sP8('const.pred_set.SUBSET_2'(s(fun(X166,bool),X167),s(fun(X166,bool),X167)),X168,X169) | ~p(s(bool,X169))) ) | $spl496),
  inference(cnf_transformation,[],[f7469_D])).
fof(f7469_D,plain,(
  ( ! [X167,X169,X166,X168] : (~p(s(bool,X168)) | sP8('const.pred_set.SUBSET_2'(s(fun(X166,bool),X167),s(fun(X166,bool),X167)),X168,X169) | ~p(s(bool,X169))) ) <=> ~$spl496),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl496])])).
fof(f7425,plain,(
  ( ! [X167,X169,X166,X168] : (~p(s(bool,'T_0')) | ~p(s(bool,X168)) | ~p(s(bool,X169)) | sP8('const.pred_set.SUBSET_2'(s(fun(X166,bool),X167),s(fun(X166,bool),X167)),X168,X169)) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f397,f7327])).
fof(f7468,plain,(
  ( ! [X140,X138,X141,X139] : (sP7(X140,X141,'const.pred_set.SUBSET_2'(s(fun(X138,bool),X139),s(fun(X138,bool),X139)))) ) | $spl494),
  inference(cnf_transformation,[],[f7468_D])).
fof(f7468_D,plain,(
  ( ! [X140,X138,X141,X139] : (sP7(X140,X141,'const.pred_set.SUBSET_2'(s(fun(X138,bool),X139),s(fun(X138,bool),X139)))) ) <=> ~$spl494),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl494])])).
fof(f7418,plain,(
  ( ! [X140,X138,X141,X139] : (~p(s(bool,'T_0')) | sP7(X140,X141,'const.pred_set.SUBSET_2'(s(fun(X138,bool),X139),s(fun(X138,bool),X139)))) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f392,f7327])).
fof(f7467,plain,(
  ( ! [X116,X114,X117,X115] : (sP3('const.pred_set.SUBSET_2'(s(fun(X114,bool),X115),s(fun(X114,bool),X115)),X116,X117)) ) | $spl492),
  inference(cnf_transformation,[],[f7467_D])).
fof(f7467_D,plain,(
  ( ! [X116,X114,X117,X115] : (sP3('const.pred_set.SUBSET_2'(s(fun(X114,bool),X115),s(fun(X114,bool),X115)),X116,X117)) ) <=> ~$spl492),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl492])])).
fof(f7412,plain,(
  ( ! [X116,X114,X117,X115] : (~p(s(bool,'T_0')) | sP3('const.pred_set.SUBSET_2'(s(fun(X114,bool),X115),s(fun(X114,bool),X115)),X116,X117)) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f379,f7327])).
fof(f7466,plain,(
  ( ! [X111,X112,X110,X113] : (sP3(X112,'const.pred_set.SUBSET_2'(s(fun(X110,bool),X111),s(fun(X110,bool),X111)),X113)) ) | $spl490),
  inference(cnf_transformation,[],[f7466_D])).
fof(f7466_D,plain,(
  ( ! [X111,X112,X110,X113] : (sP3(X112,'const.pred_set.SUBSET_2'(s(fun(X110,bool),X111),s(fun(X110,bool),X111)),X113)) ) <=> ~$spl490),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl490])])).
fof(f7411,plain,(
  ( ! [X111,X112,X110,X113] : (~p(s(bool,'T_0')) | sP3(X112,'const.pred_set.SUBSET_2'(s(fun(X110,bool),X111),s(fun(X110,bool),X111)),X113)) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f378,f7327])).
fof(f7465,plain,(
  ( ! [X109,X107,X108,X106] : (sP3(X108,X109,'const.pred_set.SUBSET_2'(s(fun(X106,bool),X107),s(fun(X106,bool),X107)))) ) | $spl488),
  inference(cnf_transformation,[],[f7465_D])).
fof(f7465_D,plain,(
  ( ! [X109,X107,X108,X106] : (sP3(X108,X109,'const.pred_set.SUBSET_2'(s(fun(X106,bool),X107),s(fun(X106,bool),X107)))) ) <=> ~$spl488),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl488])])).
fof(f7410,plain,(
  ( ! [X109,X107,X108,X106] : (~p(s(bool,'T_0')) | sP3(X108,X109,'const.pred_set.SUBSET_2'(s(fun(X106,bool),X107),s(fun(X106,bool),X107)))) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f377,f7327])).
fof(f7464,plain,(
  ( ! [X94,X97,X95,X96] : (sP4(X96,X97,'const.pred_set.SUBSET_2'(s(fun(X94,bool),X95),s(fun(X94,bool),X95)))) ) | $spl486),
  inference(cnf_transformation,[],[f7464_D])).
fof(f7464_D,plain,(
  ( ! [X94,X97,X95,X96] : (sP4(X96,X97,'const.pred_set.SUBSET_2'(s(fun(X94,bool),X95),s(fun(X94,bool),X95)))) ) <=> ~$spl486),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl486])])).
fof(f7407,plain,(
  ( ! [X94,X97,X95,X96] : (~p(s(bool,'T_0')) | sP4(X96,X97,'const.pred_set.SUBSET_2'(s(fun(X94,bool),X95),s(fun(X94,bool),X95)))) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f373,f7327])).
fof(f7463,plain,(
  ( ! [X85,X83,X84,X82] : (sP5(X84,X85,'const.pred_set.SUBSET_2'(s(fun(X82,bool),X83),s(fun(X82,bool),X83)))) ) | $spl484),
  inference(cnf_transformation,[],[f7463_D])).
fof(f7463_D,plain,(
  ( ! [X85,X83,X84,X82] : (sP5(X84,X85,'const.pred_set.SUBSET_2'(s(fun(X82,bool),X83),s(fun(X82,bool),X83)))) ) <=> ~$spl484),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl484])])).
fof(f7404,plain,(
  ( ! [X85,X83,X84,X82] : (~p(s(bool,'T_0')) | sP5(X84,X85,'const.pred_set.SUBSET_2'(s(fun(X82,bool),X83),s(fun(X82,bool),X83)))) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f369,f7327])).
fof(f7462,plain,(
  ( ! [X80,X78,X81,X79] : (sP2('const.pred_set.SUBSET_2'(s(fun(X78,bool),X79),s(fun(X78,bool),X79)),X80,X81) | ~p(s(bool,X81))) ) | $spl482),
  inference(cnf_transformation,[],[f7462_D])).
fof(f7462_D,plain,(
  ( ! [X80,X78,X81,X79] : (sP2('const.pred_set.SUBSET_2'(s(fun(X78,bool),X79),s(fun(X78,bool),X79)),X80,X81) | ~p(s(bool,X81))) ) <=> ~$spl482),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl482])])).
fof(f7403,plain,(
  ( ! [X80,X78,X81,X79] : (~p(s(bool,'T_0')) | sP2('const.pred_set.SUBSET_2'(s(fun(X78,bool),X79),s(fun(X78,bool),X79)),X80,X81) | ~p(s(bool,X81))) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f352,f7327])).
fof(f7402,plain,(
  ( ! [X76,X74,X77,X75] : (~p(s(bool,'T_0')) | sP2(X76,'const.pred_set.SUBSET_2'(s(fun(X74,bool),X75),s(fun(X74,bool),X75)),X77)) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f351,f7327])).
fof(f7460,plain,(
  ( ! [X70,X72,X71,X73] : (sP2(X72,X73,'const.pred_set.SUBSET_2'(s(fun(X70,bool),X71),s(fun(X70,bool),X71))) | ~p(s(bool,X72))) ) | $spl478),
  inference(cnf_transformation,[],[f7460_D])).
fof(f7460_D,plain,(
  ( ! [X70,X72,X71,X73] : (sP2(X72,X73,'const.pred_set.SUBSET_2'(s(fun(X70,bool),X71),s(fun(X70,bool),X71))) | ~p(s(bool,X72))) ) <=> ~$spl478),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl478])])).
fof(f7401,plain,(
  ( ! [X70,X72,X71,X73] : (~p(s(bool,'T_0')) | ~p(s(bool,X72)) | sP2(X72,X73,'const.pred_set.SUBSET_2'(s(fun(X70,bool),X71),s(fun(X70,bool),X71)))) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f350,f7327])).
fof(f7459,plain,(
  ( ! [X68,X66,X69,X67] : (~p(s(bool,X68)) | sP2(X68,'const.pred_set.SUBSET_2'(s(fun(X66,bool),X67),s(fun(X66,bool),X67)),X69)) ) | $spl476),
  inference(cnf_transformation,[],[f7459_D])).
fof(f7459_D,plain,(
  ( ! [X68,X66,X69,X67] : (~p(s(bool,X68)) | sP2(X68,'const.pred_set.SUBSET_2'(s(fun(X66,bool),X67),s(fun(X66,bool),X67)),X69)) ) <=> ~$spl476),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl476])])).
fof(f7400,plain,(
  ( ! [X68,X66,X69,X67] : (~p(s(bool,'T_0')) | ~p(s(bool,X68)) | sP2(X68,'const.pred_set.SUBSET_2'(s(fun(X66,bool),X67),s(fun(X66,bool),X67)),X69)) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f349,f7327])).
fof(f7458,plain,(
  ( ! [X61,X59,X60] : (~sP1('const.pred_set.SUBSET_2'(s(fun(X59,bool),X60),s(fun(X59,bool),X60)),X61) | ~p(s(bool,X61))) ) | $spl474),
  inference(cnf_transformation,[],[f7458_D])).
fof(f7458_D,plain,(
  ( ! [X61,X59,X60] : (~sP1('const.pred_set.SUBSET_2'(s(fun(X59,bool),X60),s(fun(X59,bool),X60)),X61) | ~p(s(bool,X61))) ) <=> ~$spl474),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl474])])).
fof(f7398,plain,(
  ( ! [X61,X59,X60] : (~p(s(bool,'T_0')) | ~p(s(bool,X61)) | ~sP1('const.pred_set.SUBSET_2'(s(fun(X59,bool),X60),s(fun(X59,bool),X60)),X61)) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f335,f7327])).
fof(f7456,plain,(
  ( ! [X47,X48,X49] : (sP1('const.pred_set.SUBSET_2'(s(fun(X47,bool),X48),s(fun(X47,bool),X48)),X49) | p(s(bool,X49))) ) | $spl470),
  inference(cnf_transformation,[],[f7456_D])).
fof(f7456_D,plain,(
  ( ! [X47,X48,X49] : (sP1('const.pred_set.SUBSET_2'(s(fun(X47,bool),X48),s(fun(X47,bool),X48)),X49) | p(s(bool,X49))) ) <=> ~$spl470),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl470])])).
fof(f7397,plain,(
  ( ! [X57,X58,X56] : (~p(s(bool,'T_0')) | p(s(bool,X58)) | sP1('const.pred_set.SUBSET_2'(s(fun(X56,bool),X57),s(fun(X56,bool),X57)),X58)) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f334,f7327])).
fof(f7457,plain,(
  ( ! [X52,X50,X51] : (sP1(X52,'const.pred_set.SUBSET_2'(s(fun(X50,bool),X51),s(fun(X50,bool),X51))) | p(s(bool,X52))) ) | $spl472),
  inference(cnf_transformation,[],[f7457_D])).
fof(f7457_D,plain,(
  ( ! [X52,X50,X51] : (sP1(X52,'const.pred_set.SUBSET_2'(s(fun(X50,bool),X51),s(fun(X50,bool),X51))) | p(s(bool,X52))) ) <=> ~$spl472),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl472])])).
fof(f7396,plain,(
  ( ! [X54,X55,X53] : (~p(s(bool,'T_0')) | sP1(X55,'const.pred_set.SUBSET_2'(s(fun(X53,bool),X54),s(fun(X53,bool),X54))) | p(s(bool,X55))) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f333,f7327])).
fof(f7395,plain,(
  ( ! [X52,X50,X51] : (~p(s(bool,'T_0')) | p(s(bool,X52)) | sP1(X52,'const.pred_set.SUBSET_2'(s(fun(X50,bool),X51),s(fun(X50,bool),X51)))) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f332,f7327])).
fof(f7394,plain,(
  ( ! [X47,X48,X49] : (~p(s(bool,'T_0')) | p(s(bool,X49)) | sP1('const.pred_set.SUBSET_2'(s(fun(X47,bool),X48),s(fun(X47,bool),X48)),X49)) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f329,f7327])).
fof(f7455,plain,(
  ( ! [X45,X46,X44] : (~sP1(X46,'const.pred_set.SUBSET_2'(s(fun(X44,bool),X45),s(fun(X44,bool),X45))) | ~p(s(bool,X46))) ) | $spl468),
  inference(cnf_transformation,[],[f7455_D])).
fof(f7455_D,plain,(
  ( ! [X45,X46,X44] : (~sP1(X46,'const.pred_set.SUBSET_2'(s(fun(X44,bool),X45),s(fun(X44,bool),X45))) | ~p(s(bool,X46))) ) <=> ~$spl468),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl468])])).
fof(f7393,plain,(
  ( ! [X45,X46,X44] : (~p(s(bool,'T_0')) | ~p(s(bool,X46)) | ~sP1(X46,'const.pred_set.SUBSET_2'(s(fun(X44,bool),X45),s(fun(X44,bool),X45)))) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f328,f7327])).
fof(f7453,plain,(
  ( ! [X37,X35,X36] : (sP0('const.pred_set.SUBSET_2'(s(fun(X35,bool),X36),s(fun(X35,bool),X36)),X37)) ) | $spl466),
  inference(cnf_transformation,[],[f7453_D])).
fof(f7453_D,plain,(
  ( ! [X37,X35,X36] : (sP0('const.pred_set.SUBSET_2'(s(fun(X35,bool),X36),s(fun(X35,bool),X36)),X37)) ) <=> ~$spl466),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl466])])).
fof(f7392,plain,(
  ( ! [X43,X41,X42] : (~p(s(bool,'T_0')) | sP0('const.pred_set.SUBSET_2'(s(fun(X41,bool),X42),s(fun(X41,bool),X42)),X43)) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f319,f7327])).
fof(f7450,plain,(
  ( ! [X33,X34,X32] : (sP0(X34,'const.pred_set.SUBSET_2'(s(fun(X32,bool),X33),s(fun(X32,bool),X33)))) ) | $spl462),
  inference(cnf_transformation,[],[f7450_D])).
fof(f7450_D,plain,(
  ( ! [X33,X34,X32] : (sP0(X34,'const.pred_set.SUBSET_2'(s(fun(X32,bool),X33),s(fun(X32,bool),X33)))) ) <=> ~$spl462),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl462])])).
fof(f7454,plain,(
  ( ! [X39,X38,X40] : (~p(s(bool,'T_0')) | sP0(X40,'const.pred_set.SUBSET_2'(s(fun(X38,bool),X39),s(fun(X38,bool),X39)))) ) | ($spl4 | $spl80)),
  inference(global_subsumption,[],[f248,f247,f246,f447,f789,f788,f787,f786,f785,f784,f783,f782,f781,f780,f779,f778,f777,f776,f775,f774,f773,f772,f768,f767,f766,f763,f762,f761,f760,f759,f758,f757,f756,f755,f754,f753,f752,f751,f750,f749,f747,f746,f745,f744,f743,f742,f741,f737,f736,f735,f732,f731,f730,f729,f725,f724,f723,f720,f719,f718,f715,f714,f713,f712,f711,f710,f709,f708,f707,f706,f705,f704,f703,f702,f701,f700,f699,f698,f697,f696,f695,f694,f693,f692,f689,f688,f687,f686,f683,f682,f681,f680,f679,f678,f677,f674,f673,f672,f671,f668,f667,f666,f665,f662,f661,f660,f659,f658,f657,f656,f653,f652,f651,f650,f647,f646,f645,f644,f643,f642,f641,f638,f637,f636,f635,f632,f631,f630,f629,f626,f625,f624,f622,f621,f620,f616,f615,f614,f611,f610,f609,f608,f607,f606,f605,f602,f601,f600,f599,f596,f595,f594,f593,f592,f591,f590,f589,f588,f587,f586,f585,f584,f583,f582,f581,f578,f577,f576,f575,f574,f573,f572,f571,f570,f569,f568,f567,f566,f565,f564,f563,f562,f561,f560,f559,f558,f557,f556,f555,f554,f553,f552,f551,f550,f549,f548,f547,f546,f545,f544,f543,f542,f541,f540,f539,f538,f537,f536,f535,f534,f530,f529,f528,f505,f504,f503,f499,f498,f497,f488,f487,f486,f485,f484,f483,f482,f481,f480,f479,f478,f477,f476,f475,f474,f473,f472,f469,f468,f467,f466,f463,f462,f461,f460,f459,f451,f450,f449,f448,f790,f791,f800,f801,f802,f805,f808,f810,f817,f820,f821,f823,f824,f825,f827,f904,f903,f1493,f1502,f1495,f1521,f1520,f1519,f1518,f1517,f1639,f1658,f1681,f1715,f1784,f1799,f1830,f2122,f2121,f2120,f2119,f2118,f2116,f2115,f2114,f2113,f2112,f2166,f2167,f2168,f2169,f2702,f2710,f2718,f2726,f2734,f2742,f2750,f2758,f2802,f2817,f2832,f2847,f2867,f2882,f2897,f2912,f3054,f3073,f3141,f3160,f3242,f3261,f5649,f5669,f5822,f5858,f5875,f5892,f5909,f5926,f5943,f5960,f7263,f7273,f7284,f7283,f7297,f7296,f7327,f7334,f7371,f7370,f7367,f7366,f7389,f7390,f7391])).
fof(f7390,plain,(
  ( ! [X37,X35,X36] : (~p(s(bool,'T_0')) | sP0('const.pred_set.SUBSET_2'(s(fun(X35,bool),X36),s(fun(X35,bool),X36)),X37)) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f314,f7327])).
fof(f7389,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,'T_0')) | sP0(X34,'const.pred_set.SUBSET_2'(s(fun(X32,bool),X33),s(fun(X32,bool),X33)))) ) | ($spl4 | $spl80)),
  inference(superposition,[],[f313,f7327])).
fof(f6907,plain,(
  ~sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl461),
  inference(cnf_transformation,[],[f6907_D])).
fof(f6907_D,plain,(
  ~sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl461),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl461])])).
fof(f6903,plain,(
  ~sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl147),
  inference(unit_resulting_resolution,[],[f1378,f363])).
fof(f6895,plain,(
  ~sP6('F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl459),
  inference(cnf_transformation,[],[f6895_D])).
fof(f6895_D,plain,(
  ~sP6('F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl459),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl459])])).
fof(f6891,plain,(
  ~sP6('F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl141),
  inference(unit_resulting_resolution,[],[f1372,f362])).
fof(f6871,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl456),
  inference(cnf_transformation,[],[f6871_D])).
fof(f6871_D,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl456),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl456])])).
fof(f6864,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl124),
  inference(resolution,[],[f1356,f384])).
fof(f6858,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl124),
  inference(unit_resulting_resolution,[],[f1356,f384])).
fof(f5544,plain,(
  sP6('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl454),
  inference(cnf_transformation,[],[f5544_D])).
fof(f5544_D,plain,(
  sP6('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl454),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl454])])).
fof(f5453,plain,(
  sP6('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl168),
  inference(unit_resulting_resolution,[],[f437,f1400,f381])).
fof(f5454,plain,(
  sP6('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl168),
  inference(unit_resulting_resolution,[],[f250,f1400,f381])).
fof(f5455,plain,(
  sP6('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl168),
  inference(unit_resulting_resolution,[],[f249,f1400,f381])).
fof(f5542,plain,(
  sP6('T_0','T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl452),
  inference(cnf_transformation,[],[f5542_D])).
fof(f5542_D,plain,(
  sP6('T_0','T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl452),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl452])])).
fof(f5459,plain,(
  sP6('T_0','T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl168),
  inference(unit_resulting_resolution,[],[f437,f1400,f381])).
fof(f5460,plain,(
  sP6('T_0','T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl168),
  inference(unit_resulting_resolution,[],[f250,f1400,f381])).
fof(f5461,plain,(
  sP6('T_0','T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl168),
  inference(unit_resulting_resolution,[],[f249,f1400,f381])).
fof(f5398,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl448),
  inference(cnf_transformation,[],[f5398_D])).
fof(f5398_D,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl448),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl448])])).
fof(f5401,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl172),
  inference(global_subsumption,[],[f248,f247,f246,f1404,f5170,f5167,f5166,f5165,f5164,f5163,f5162,f5161,f5156,f5155,f5154,f5153,f5152,f5151,f5150,f5149,f5148,f5147,f5143,f5142,f5141,f5138,f5135,f5134,f5131,f5130,f5129,f5128,f5127,f5126,f5123,f5122,f5121,f5120,f5119,f5118,f5077,f5067,f5066,f5065,f5064,f5063,f5062,f5056,f5055,f5054,f5051,f5050,f5049,f5048,f5042,f5041,f5040,f5037,f5036,f5031,f5006,f5005,f5004,f5003,f5002,f5001,f5000,f4999,f4998,f4997,f4996,f4995,f4994,f4993,f4992,f4991,f4990,f4989,f4988,f4987,f4986,f4985,f4984,f4983,f4980,f4979,f4978,f4977,f4974,f4973,f4972,f4971,f4970,f4969,f4968,f4965,f4964,f4963,f4962,f4959,f4958,f4957,f4956,f4953,f4952,f4951,f4950,f4949,f4948,f4947,f4944,f4943,f4942,f4941,f4938,f4937,f4936,f4935,f4934,f4933,f4932,f4929,f4928,f4927,f4926,f4923,f4922,f4921,f4920,f4917,f4916,f4915,f4913,f4912,f4911,f4907,f4906,f4905,f4902,f4901,f4900,f4899,f4898,f4897,f4896,f4893,f4892,f4891,f4890,f4887,f4886,f4885,f4884,f4881,f4880,f4879,f4878,f4877,f4876,f4873,f4872,f4871,f4870,f4869,f4868,f4829,f4828,f4827,f4826,f4823,f4822,f4821,f4820,f4819,f4818,f4815,f4814,f4813,f4812,f4811,f4810,f4809,f4806,f4805,f4804,f4803,f4802,f4801,f4798,f4797,f4796,f4795,f4794,f4793,f4788,f4787,f4786,f4785,f4784,f4783,f4782,f4781,f4780,f4779,f4778,f4777,f4776,f4775,f4774,f4773,f4769,f4768,f4767,f4744,f4743,f4742,f4738,f4737,f4736,f4727,f4726,f4725,f4724,f4723,f4722,f4721,f4720,f4719,f4718,f4717,f4716,f4715,f4714,f4713,f4712,f4711,f4708,f4707,f4706,f4705,f4702,f4701,f4700,f4699,f4698,f4690,f4689,f4688,f4687,f4686,f5198,f5199,f5200,f5209,f5210,f5211,f5214,f5217,f5219,f5226,f5229,f5230,f5232,f5233,f5234,f5236])).
fof(f5236,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl172),
  inference(superposition,[],[f1404,f256])).
fof(f5234,plain,(
  ( ! [X50,X51] : (sP12('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X50,X51) | ~p(s(bool,X50))) ) | $spl172),
  inference(resolution,[],[f1404,f431])).
fof(f5233,plain,(
  ( ! [X48,X49] : (sP12('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X49,X48) | p(s(bool,X48))) ) | $spl172),
  inference(resolution,[],[f1404,f430])).
fof(f5232,plain,(
  ( ! [X47,X46] : (sP12(X47,X46,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | p(s(bool,X46)) | p(s(bool,X47))) ) | $spl172),
  inference(resolution,[],[f1404,f429])).
fof(f5230,plain,(
  ( ! [X43,X42] : (sP10('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X42,X43) | ~p(s(bool,X42))) ) | $spl172),
  inference(resolution,[],[f1404,f414])).
fof(f5229,plain,(
  ( ! [X41,X40] : (sP10('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X41,X40) | ~p(s(bool,X40))) ) | $spl172),
  inference(resolution,[],[f1404,f413])).
fof(f5226,plain,(
  ( ! [X35,X34] : (sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X34,X35) | ~p(s(bool,X35)) | ~p(s(bool,X34))) ) | $spl172),
  inference(resolution,[],[f1404,f397])).
fof(f5219,plain,(
  ( ! [X21,X20] : (sP2('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X20,X21) | ~p(s(bool,X21))) ) | $spl172),
  inference(resolution,[],[f1404,f352])).
fof(f5217,plain,(
  ( ! [X17,X16] : (sP2(X16,X17,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ~p(s(bool,X16))) ) | $spl172),
  inference(resolution,[],[f1404,f350])).
fof(f5214,plain,(
  ( ! [X11] : (~sP1('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X11) | ~p(s(bool,X11))) ) | $spl172),
  inference(resolution,[],[f1404,f335])).
fof(f5211,plain,(
  ( ! [X8] : (sP1(X8,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | p(s(bool,X8))) ) | $spl172),
  inference(resolution,[],[f1404,f332])).
fof(f5210,plain,(
  ( ! [X7] : (sP1('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X7) | p(s(bool,X7))) ) | $spl172),
  inference(resolution,[],[f1404,f329])).
fof(f5209,plain,(
  ( ! [X6] : (~sP1(X6,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ~p(s(bool,X6))) ) | $spl172),
  inference(resolution,[],[f1404,f328])).
fof(f4698,plain,(
  ( ! [X0] : (sP0(X0,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f255,f1404,f318])).
fof(f4699,plain,(
  ( ! [X0] : (sP0('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X0)) ) | $spl172),
  inference(unit_resulting_resolution,[],[f1404,f319])).
fof(f4719,plain,(
  ( ! [X0,X1] : (~sP1('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f304,f1404,f335])).
fof(f4720,plain,(
  ( ! [X0,X1] : (~sP1('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f305,f1404,f335])).
fof(f4725,plain,(
  ( ! [X0,X1] : (~sP1('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f304,f1404,f335])).
fof(f4726,plain,(
  ( ! [X0,X1] : (~sP1('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f305,f1404,f335])).
fof(f4736,plain,(
  ( ! [X0] : (sP2(X0,'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f249,f1404,f348])).
fof(f4737,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,'const.pred_set.SUBSET_2'(s(fun(X1,bool),X2),s(fun(X1,bool),X2)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f304,f1404,f348])).
fof(f4738,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,'const.pred_set.SUBSET_2'(s(fun(X1,bool),X2),s(fun(X1,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f305,f1404,f348])).
fof(f4742,plain,(
  ( ! [X0] : (sP2('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0',X0)) ) | $spl172),
  inference(unit_resulting_resolution,[],[f249,f1404,f349])).
fof(f4743,plain,(
  ( ! [X2,X0,X1] : (sP2('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),X2)) ) | $spl172),
  inference(unit_resulting_resolution,[],[f304,f1404,f349])).
fof(f4744,plain,(
  ( ! [X2,X0,X1] : (sP2('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),X2)) ) | $spl172),
  inference(unit_resulting_resolution,[],[f305,f1404,f349])).
fof(f4767,plain,(
  ( ! [X0] : (sP2('T_0',X0,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f249,f1404,f352])).
fof(f4768,plain,(
  ( ! [X2,X0,X1] : (sP2('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),X2,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f304,f1404,f352])).
fof(f4769,plain,(
  ( ! [X2,X0,X1] : (sP2('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),X2,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f305,f1404,f352])).
fof(f4773,plain,(
  ( ! [X0] : (sP2('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X0,'T_0')) ) | $spl172),
  inference(unit_resulting_resolution,[],[f249,f1404,f352])).
fof(f4774,plain,(
  ( ! [X2,X0,X1] : (sP2('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X0,'const.pred_set.SUBSET_2'(s(fun(X1,bool),X2),s(fun(X1,bool),X2)))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f304,f1404,f352])).
fof(f4775,plain,(
  ( ! [X2,X0,X1] : (sP2('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X0,'const.pred_set.SUBSET_2'(s(fun(X1,bool),X2),s(fun(X1,bool),'const.pred_set.UNIV_0')))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f305,f1404,f352])).
fof(f4776,plain,(
  ( ! [X0] : (sP2('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X0,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f1404,f1404,f352])).
fof(f4780,plain,(
  ( ! [X0,X1] : (~sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'F_0')) ) | $spl172),
  inference(unit_resulting_resolution,[],[f304,f255,f1404,f364])).
fof(f4781,plain,(
  ( ! [X0,X1] : (~sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'F_0')) ) | $spl172),
  inference(unit_resulting_resolution,[],[f305,f255,f1404,f364])).
fof(f4786,plain,(
  ( ! [X0,X1] : (~sP6('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0')) ) | $spl172),
  inference(unit_resulting_resolution,[],[f304,f255,f1404,f364])).
fof(f4787,plain,(
  ( ! [X0,X1] : (~sP6('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0')) ) | $spl172),
  inference(unit_resulting_resolution,[],[f305,f255,f1404,f364])).
fof(f4796,plain,(
  ( ! [X0,X1] : (~sP5('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'F_0')) ) | $spl172),
  inference(unit_resulting_resolution,[],[f304,f255,f1404,f368])).
fof(f4797,plain,(
  ( ! [X0,X1] : (~sP5('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'F_0')) ) | $spl172),
  inference(unit_resulting_resolution,[],[f305,f255,f1404,f368])).
fof(f4804,plain,(
  ( ! [X0,X1] : (~sP5('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0')) ) | $spl172),
  inference(unit_resulting_resolution,[],[f304,f255,f1404,f368])).
fof(f4805,plain,(
  ( ! [X0,X1] : (~sP5('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0')) ) | $spl172),
  inference(unit_resulting_resolution,[],[f305,f255,f1404,f368])).
fof(f4813,plain,(
  ( ! [X0,X1] : (~sP4('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'F_0')) ) | $spl172),
  inference(unit_resulting_resolution,[],[f304,f255,f1404,f372])).
fof(f4814,plain,(
  ( ! [X0,X1] : (~sP4('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'F_0')) ) | $spl172),
  inference(unit_resulting_resolution,[],[f305,f255,f1404,f372])).
fof(f4821,plain,(
  ( ! [X0,X1] : (~sP4('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0')) ) | $spl172),
  inference(unit_resulting_resolution,[],[f304,f255,f1404,f372])).
fof(f4822,plain,(
  ( ! [X0,X1] : (~sP4('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0')) ) | $spl172),
  inference(unit_resulting_resolution,[],[f305,f255,f1404,f372])).
fof(f4871,plain,(
  ( ! [X0,X1] : (~sP7('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'F_0')) ) | $spl172),
  inference(unit_resulting_resolution,[],[f304,f255,f1404,f391])).
fof(f4872,plain,(
  ( ! [X0,X1] : (~sP7('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'F_0')) ) | $spl172),
  inference(unit_resulting_resolution,[],[f305,f255,f1404,f391])).
fof(f4879,plain,(
  ( ! [X0,X1] : (~sP7('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0')) ) | $spl172),
  inference(unit_resulting_resolution,[],[f304,f255,f1404,f391])).
fof(f4880,plain,(
  ( ! [X0,X1] : (~sP7('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0')) ) | $spl172),
  inference(unit_resulting_resolution,[],[f305,f255,f1404,f391])).
fof(f4900,plain,(
  ( ! [X0,X1] : (sP8('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0')) ) | $spl172),
  inference(unit_resulting_resolution,[],[f304,f249,f1404,f397])).
fof(f4901,plain,(
  ( ! [X0,X1] : (sP8('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0')) ) | $spl172),
  inference(unit_resulting_resolution,[],[f305,f249,f1404,f397])).
fof(f4905,plain,(
  ( ! [X0,X1] : (sP8('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f249,f304,f1404,f397])).
fof(f4906,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X2,bool),X3),s(fun(X2,bool),X3)))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f304,f304,f1404,f397])).
fof(f4907,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X2,bool),X3),s(fun(X2,bool),X3)))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f305,f304,f1404,f397])).
fof(f4911,plain,(
  ( ! [X0,X1] : (sP8('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f249,f305,f1404,f397])).
fof(f4912,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X2,bool),X3),s(fun(X2,bool),'const.pred_set.UNIV_0')))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f304,f305,f1404,f397])).
fof(f4913,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X2,bool),X3),s(fun(X2,bool),'const.pred_set.UNIV_0')))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f305,f305,f1404,f397])).
fof(f4936,plain,(
  ( ! [X0,X1] : (sP8('T_0','const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f304,f249,f1404,f397])).
fof(f4937,plain,(
  ( ! [X0,X1] : (sP8('T_0','const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f305,f249,f1404,f397])).
fof(f4941,plain,(
  ( ! [X0,X1] : (sP8('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f249,f304,f1404,f397])).
fof(f4942,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.SUBSET_2'(s(fun(X2,bool),X3),s(fun(X2,bool),X3)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f304,f304,f1404,f397])).
fof(f4943,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.SUBSET_2'(s(fun(X2,bool),X3),s(fun(X2,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f305,f304,f1404,f397])).
fof(f4944,plain,(
  ( ! [X0,X1] : (sP8('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f1404,f304,f1404,f397])).
fof(f4947,plain,(
  ( ! [X0,X1] : (sP8('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f249,f305,f1404,f397])).
fof(f4948,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.SUBSET_2'(s(fun(X2,bool),X3),s(fun(X2,bool),X3)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f304,f305,f1404,f397])).
fof(f4949,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.SUBSET_2'(s(fun(X2,bool),X3),s(fun(X2,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f305,f305,f1404,f397])).
fof(f4950,plain,(
  ( ! [X0,X1] : (sP8('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f1404,f305,f1404,f397])).
fof(f4972,plain,(
  ( ! [X0,X1] : (sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'T_0')) ) | $spl172),
  inference(unit_resulting_resolution,[],[f304,f249,f1404,f397])).
fof(f4973,plain,(
  ( ! [X0,X1] : (sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'T_0')) ) | $spl172),
  inference(unit_resulting_resolution,[],[f305,f249,f1404,f397])).
fof(f4977,plain,(
  ( ! [X0,X1] : (sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f249,f304,f1404,f397])).
fof(f4978,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.SUBSET_2'(s(fun(X2,bool),X3),s(fun(X2,bool),X3)))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f304,f304,f1404,f397])).
fof(f4979,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.SUBSET_2'(s(fun(X2,bool),X3),s(fun(X2,bool),X3)))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f305,f304,f1404,f397])).
fof(f4980,plain,(
  ( ! [X0,X1] : (sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f1404,f304,f1404,f397])).
fof(f4983,plain,(
  ( ! [X0,X1] : (sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f249,f305,f1404,f397])).
fof(f4984,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.SUBSET_2'(s(fun(X2,bool),X3),s(fun(X2,bool),'const.pred_set.UNIV_0')))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f304,f305,f1404,f397])).
fof(f4985,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.SUBSET_2'(s(fun(X2,bool),X3),s(fun(X2,bool),'const.pred_set.UNIV_0')))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f305,f305,f1404,f397])).
fof(f4986,plain,(
  ( ! [X0,X1] : (sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f1404,f305,f1404,f397])).
fof(f4990,plain,(
  ( ! [X0,X1] : (sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f304,f1404,f1404,f397])).
fof(f4991,plain,(
  ( ! [X0,X1] : (sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f305,f1404,f1404,f397])).
fof(f4993,plain,(
  ( ! [X0] : (~sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X0,'F_0')) ) | $spl172),
  inference(unit_resulting_resolution,[],[f255,f1404,f398])).
fof(f4994,plain,(
  ( ! [X0] : (~sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0',X0)) ) | $spl172),
  inference(unit_resulting_resolution,[],[f255,f1404,f399])).
fof(f4998,plain,(
  ( ! [X0,X1] : (~sP8('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f255,f304,f1404,f400])).
fof(f4999,plain,(
  ( ! [X0,X1] : (~sP8('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f255,f305,f1404,f400])).
fof(f5004,plain,(
  ( ! [X0,X1] : (~sP8('F_0','const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f255,f304,f1404,f400])).
fof(f5005,plain,(
  ( ! [X0,X1] : (~sP8('F_0','const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f255,f305,f1404,f400])).
fof(f5036,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f1404,f409])).
fof(f5037,plain,(
  ( ! [X0,X1] : (sP9(X0,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X1)) ) | $spl172),
  inference(unit_resulting_resolution,[],[f1404,f410])).
fof(f5040,plain,(
  ( ! [X0] : (sP10('T_0',X0,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f249,f1404,f413])).
fof(f5041,plain,(
  ( ! [X2,X0,X1] : (sP10('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),X2,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f304,f1404,f413])).
fof(f5042,plain,(
  ( ! [X2,X0,X1] : (sP10('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),X2,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f305,f1404,f413])).
fof(f5048,plain,(
  ( ! [X0] : (sP10('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X0,'T_0')) ) | $spl172),
  inference(unit_resulting_resolution,[],[f249,f1404,f413])).
fof(f5049,plain,(
  ( ! [X2,X0,X1] : (sP10('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X0,'const.pred_set.SUBSET_2'(s(fun(X1,bool),X2),s(fun(X1,bool),X2)))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f304,f1404,f413])).
fof(f5050,plain,(
  ( ! [X2,X0,X1] : (sP10('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X0,'const.pred_set.SUBSET_2'(s(fun(X1,bool),X2),s(fun(X1,bool),'const.pred_set.UNIV_0')))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f305,f1404,f413])).
fof(f5051,plain,(
  ( ! [X0] : (sP10('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X0,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f1404,f1404,f413])).
fof(f5054,plain,(
  ( ! [X0] : (sP10('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X0)) ) | $spl172),
  inference(unit_resulting_resolution,[],[f249,f1404,f414])).
fof(f5055,plain,(
  ( ! [X2,X0,X1] : (sP10('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X2)) ) | $spl172),
  inference(unit_resulting_resolution,[],[f304,f1404,f414])).
fof(f5056,plain,(
  ( ! [X2,X0,X1] : (sP10('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X2)) ) | $spl172),
  inference(unit_resulting_resolution,[],[f305,f1404,f414])).
fof(f5062,plain,(
  ( ! [X0] : (sP10('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0',X0)) ) | $spl172),
  inference(unit_resulting_resolution,[],[f249,f1404,f414])).
fof(f5063,plain,(
  ( ! [X2,X0,X1] : (sP10('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),X2)) ) | $spl172),
  inference(unit_resulting_resolution,[],[f304,f1404,f414])).
fof(f5064,plain,(
  ( ! [X2,X0,X1] : (sP10('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),X2)) ) | $spl172),
  inference(unit_resulting_resolution,[],[f305,f1404,f414])).
fof(f5065,plain,(
  ( ! [X0] : (sP10('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X0)) ) | $spl172),
  inference(unit_resulting_resolution,[],[f1404,f1404,f414])).
fof(f5067,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f255,f1404,f416])).
fof(f5077,plain,(
  ( ! [X0] : (~sP10('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X0)) ) | $spl172),
  inference(unit_resulting_resolution,[],[f255,f1404,f417])).
fof(f5121,plain,(
  ( ! [X0,X1] : (~sP11('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f255,f304,f1404,f425])).
fof(f5122,plain,(
  ( ! [X0,X1] : (~sP11('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f255,f305,f1404,f425])).
fof(f5129,plain,(
  ( ! [X0,X1] : (~sP11('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f304,f255,f1404,f425])).
fof(f5130,plain,(
  ( ! [X0,X1] : (~sP11('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f305,f255,f1404,f425])).
fof(f5134,plain,(
  ( ! [X0,X1] : (sP11(X0,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X1)) ) | $spl172),
  inference(unit_resulting_resolution,[],[f1404,f427])).
fof(f5138,plain,(
  ( ! [X0] : (sP12('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X0,'F_0')) ) | $spl172),
  inference(unit_resulting_resolution,[],[f255,f1404,f430])).
fof(f5141,plain,(
  ( ! [X0] : (sP12('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X0)) ) | $spl172),
  inference(unit_resulting_resolution,[],[f249,f1404,f431])).
fof(f5142,plain,(
  ( ! [X2,X0,X1] : (sP12('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X2)) ) | $spl172),
  inference(unit_resulting_resolution,[],[f304,f1404,f431])).
fof(f5143,plain,(
  ( ! [X2,X0,X1] : (sP12('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X2)) ) | $spl172),
  inference(unit_resulting_resolution,[],[f305,f1404,f431])).
fof(f5147,plain,(
  ( ! [X0] : (sP12('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0',X0)) ) | $spl172),
  inference(unit_resulting_resolution,[],[f249,f1404,f431])).
fof(f5148,plain,(
  ( ! [X2,X0,X1] : (sP12('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),X2)) ) | $spl172),
  inference(unit_resulting_resolution,[],[f304,f1404,f431])).
fof(f5149,plain,(
  ( ! [X2,X0,X1] : (sP12('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),X2)) ) | $spl172),
  inference(unit_resulting_resolution,[],[f305,f1404,f431])).
fof(f5150,plain,(
  ( ! [X0] : (sP12('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X0)) ) | $spl172),
  inference(unit_resulting_resolution,[],[f1404,f1404,f431])).
fof(f5154,plain,(
  ( ! [X0,X1] : (~sP12('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f255,f304,f1404,f432])).
fof(f5155,plain,(
  ( ! [X0,X1] : (~sP12('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f255,f305,f1404,f432])).
fof(f5164,plain,(
  ( ! [X0,X1] : (~sP12('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f255,f304,f1404,f432])).
fof(f5165,plain,(
  ( ! [X0,X1] : (~sP12('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f255,f305,f1404,f432])).
fof(f5167,plain,(
  ( ! [X0] : (~sP12('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X0)) ) | $spl172),
  inference(unit_resulting_resolution,[],[f255,f1404,f434])).
fof(f5170,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X1)) ) | $spl172),
  inference(unit_resulting_resolution,[],[f1404,f438])).
fof(f5200,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl172),
  inference(resolution,[],[f1404,f265])).
fof(f5199,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl172),
  inference(resolution,[],[f1404,f263])).
fof(f5400,plain,(
  p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')))) | $spl450),
  inference(cnf_transformation,[],[f5400_D])).
fof(f5400_D,plain,(
  p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')))) <=> ~$spl450),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl450])])).
fof(f4686,plain,(
  p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),'const.pred_set.UNIV_0')))) | $spl172),
  inference(unit_resulting_resolution,[],[f305,f304,f1404,f344])).
fof(f4687,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl172),
  inference(unit_resulting_resolution,[],[f1404,f263])).
fof(f4688,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl172),
  inference(unit_resulting_resolution,[],[f1404,f265])).
fof(f5396,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl447),
  inference(cnf_transformation,[],[f5396_D])).
fof(f5396_D,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl447),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl447])])).
fof(f4689,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl172),
  inference(unit_resulting_resolution,[],[f1404,f266])).
fof(f4690,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl172),
  inference(unit_resulting_resolution,[],[f1404,f268])).
fof(f5388,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl439),
  inference(cnf_transformation,[],[f5388_D])).
fof(f5388_D,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl439),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl439])])).
fof(f4700,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f437,f1404,f328])).
fof(f4701,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f250,f1404,f328])).
fof(f4702,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f249,f1404,f328])).
fof(f5386,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl437),
  inference(cnf_transformation,[],[f5386_D])).
fof(f5386_D,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl437),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl437])])).
fof(f4705,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl172),
  inference(unit_resulting_resolution,[],[f1404,f1404,f328])).
fof(f5390,plain,(
  ~sP1('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl441),
  inference(cnf_transformation,[],[f5390_D])).
fof(f5390_D,plain,(
  ~sP1('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl441),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl441])])).
fof(f4706,plain,(
  ~sP1('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl172),
  inference(unit_resulting_resolution,[],[f437,f1404,f328])).
fof(f4707,plain,(
  ~sP1('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl172),
  inference(unit_resulting_resolution,[],[f250,f1404,f328])).
fof(f4708,plain,(
  ~sP1('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl172),
  inference(unit_resulting_resolution,[],[f249,f1404,f328])).
fof(f4711,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl172),
  inference(unit_resulting_resolution,[],[f1404,f1404,f328])).
fof(f5392,plain,(
  sP1('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl442),
  inference(cnf_transformation,[],[f5392_D])).
fof(f5392_D,plain,(
  sP1('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') <=> ~$spl442),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl442])])).
fof(f4712,plain,(
  sP1('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f255,f1404,f329])).
fof(f5394,plain,(
  sP1('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl444),
  inference(cnf_transformation,[],[f5394_D])).
fof(f5394_D,plain,(
  sP1('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl444),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl444])])).
fof(f4713,plain,(
  sP1('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl172),
  inference(unit_resulting_resolution,[],[f255,f1404,f332])).
fof(f4714,plain,(
  sP1('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl172),
  inference(unit_resulting_resolution,[],[f255,f1404,f333])).
fof(f4715,plain,(
  sP1('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f255,f1404,f334])).
fof(f4716,plain,(
  ~sP1('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl172),
  inference(unit_resulting_resolution,[],[f437,f1404,f335])).
fof(f4717,plain,(
  ~sP1('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl172),
  inference(unit_resulting_resolution,[],[f250,f1404,f335])).
fof(f4718,plain,(
  ~sP1('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl172),
  inference(unit_resulting_resolution,[],[f249,f1404,f335])).
fof(f4721,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl172),
  inference(unit_resulting_resolution,[],[f1404,f1404,f335])).
fof(f4722,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f437,f1404,f335])).
fof(f4723,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f250,f1404,f335])).
fof(f4724,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f249,f1404,f335])).
fof(f4727,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl172),
  inference(unit_resulting_resolution,[],[f1404,f1404,f335])).
fof(f5348,plain,(
  ~sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl407),
  inference(cnf_transformation,[],[f5348_D])).
fof(f5348_D,plain,(
  ~sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','F_0') <=> ~$spl407),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl407])])).
fof(f4777,plain,(
  ~sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f437,f255,f1404,f364])).
fof(f4778,plain,(
  ~sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f250,f255,f1404,f364])).
fof(f4779,plain,(
  ~sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f249,f255,f1404,f364])).
fof(f5382,plain,(
  ~sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl433),
  inference(cnf_transformation,[],[f5382_D])).
fof(f5382_D,plain,(
  ~sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') <=> ~$spl433),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl433])])).
fof(f4782,plain,(
  ~sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f1404,f255,f1404,f364])).
fof(f5384,plain,(
  ~sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl435),
  inference(cnf_transformation,[],[f5384_D])).
fof(f5384_D,plain,(
  ~sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') <=> ~$spl435),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl435])])).
fof(f4783,plain,(
  ~sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f437,f255,f1404,f364])).
fof(f4784,plain,(
  ~sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f250,f255,f1404,f364])).
fof(f4785,plain,(
  ~sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f249,f255,f1404,f364])).
fof(f4788,plain,(
  ~sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f1404,f255,f1404,f364])).
fof(f5380,plain,(
  sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl430),
  inference(cnf_transformation,[],[f5380_D])).
fof(f5380_D,plain,(
  sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl430),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl430])])).
fof(f4789,plain,(
  sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl4 | $spl172)),
  inference(unit_resulting_resolution,[],[f577,f575,f562,f1404,f365])).
fof(f5378,plain,(
  sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl428),
  inference(cnf_transformation,[],[f5378_D])).
fof(f5378_D,plain,(
  sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl428),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl428])])).
fof(f4790,plain,(
  sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl2 | $spl4 | $spl172)),
  inference(unit_resulting_resolution,[],[f577,f1066,f562,f1404,f365])).
fof(f5376,plain,(
  sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl426),
  inference(cnf_transformation,[],[f5376_D])).
fof(f5376_D,plain,(
  sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl426),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl426])])).
fof(f4791,plain,(
  sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl2 | $spl4 | $spl172)),
  inference(unit_resulting_resolution,[],[f1068,f575,f1053,f1404,f365])).
fof(f5374,plain,(
  sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl424),
  inference(cnf_transformation,[],[f5374_D])).
fof(f5374_D,plain,(
  sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl424),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl424])])).
fof(f4792,plain,(
  sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl2 | $spl172)),
  inference(unit_resulting_resolution,[],[f1068,f1066,f1053,f1404,f365])).
fof(f5372,plain,(
  ~sP5('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl423),
  inference(cnf_transformation,[],[f5372_D])).
fof(f5372_D,plain,(
  ~sP5('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','F_0') <=> ~$spl423),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl423])])).
fof(f4793,plain,(
  ~sP5('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f437,f255,f1404,f368])).
fof(f4794,plain,(
  ~sP5('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f250,f255,f1404,f368])).
fof(f4795,plain,(
  ~sP5('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f249,f255,f1404,f368])).
fof(f5366,plain,(
  ~sP5('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl419),
  inference(cnf_transformation,[],[f5366_D])).
fof(f5366_D,plain,(
  ~sP5('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') <=> ~$spl419),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl419])])).
fof(f4798,plain,(
  ~sP5('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f1404,f255,f1404,f368])).
fof(f5368,plain,(
  ~sP5('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl421),
  inference(cnf_transformation,[],[f5368_D])).
fof(f5368_D,plain,(
  ~sP5('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') <=> ~$spl421),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl421])])).
fof(f4801,plain,(
  ~sP5('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f437,f255,f1404,f368])).
fof(f4802,plain,(
  ~sP5('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f250,f255,f1404,f368])).
fof(f4803,plain,(
  ~sP5('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f249,f255,f1404,f368])).
fof(f4806,plain,(
  ~sP5('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f1404,f255,f1404,f368])).
fof(f5362,plain,(
  ~sP4('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl417),
  inference(cnf_transformation,[],[f5362_D])).
fof(f5362_D,plain,(
  ~sP4('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','F_0') <=> ~$spl417),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl417])])).
fof(f4810,plain,(
  ~sP4('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f437,f255,f1404,f372])).
fof(f4811,plain,(
  ~sP4('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f250,f255,f1404,f372])).
fof(f4812,plain,(
  ~sP4('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f249,f255,f1404,f372])).
fof(f5356,plain,(
  ~sP4('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl413),
  inference(cnf_transformation,[],[f5356_D])).
fof(f5356_D,plain,(
  ~sP4('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') <=> ~$spl413),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl413])])).
fof(f4815,plain,(
  ~sP4('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f1404,f255,f1404,f372])).
fof(f5358,plain,(
  ~sP4('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl415),
  inference(cnf_transformation,[],[f5358_D])).
fof(f5358_D,plain,(
  ~sP4('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') <=> ~$spl415),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl415])])).
fof(f4818,plain,(
  ~sP4('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f437,f255,f1404,f372])).
fof(f4819,plain,(
  ~sP4('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f250,f255,f1404,f372])).
fof(f4820,plain,(
  ~sP4('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f249,f255,f1404,f372])).
fof(f4823,plain,(
  ~sP4('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f1404,f255,f1404,f372])).
fof(f5352,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl410),
  inference(cnf_transformation,[],[f5352_D])).
fof(f5352_D,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) <=> ~$spl410),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl410])])).
fof(f4830,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | ($spl80 | $spl172)),
  inference(unit_resulting_resolution,[],[f903,f1404,f381])).
fof(f5350,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl408),
  inference(cnf_transformation,[],[f5350_D])).
fof(f5350_D,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl408),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl408])])).
fof(f4831,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | ($spl80 | $spl172)),
  inference(unit_resulting_resolution,[],[f903,f1404,f381])).
fof(f4832,plain,(
  ~sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','F_0') | ($spl83 | $spl172)),
  inference(unit_resulting_resolution,[],[f908,f1404,f382])).
fof(f5346,plain,(
  ~sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','T_0') | $spl405),
  inference(cnf_transformation,[],[f5346_D])).
fof(f5346_D,plain,(
  ~sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','T_0') <=> ~$spl405),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl405])])).
fof(f4833,plain,(
  ~sP6('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','T_0') | ($spl83 | $spl172)),
  inference(unit_resulting_resolution,[],[f908,f1404,f382])).
fof(f5242,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl348),
  inference(cnf_transformation,[],[f5242_D])).
fof(f5242_D,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl348),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl348])])).
fof(f4836,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | ($spl4 | $spl172)),
  inference(unit_resulting_resolution,[],[f437,f593,f1404,f387])).
fof(f4837,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | ($spl4 | $spl172)),
  inference(unit_resulting_resolution,[],[f250,f593,f1404,f387])).
fof(f4838,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | ($spl4 | $spl172)),
  inference(unit_resulting_resolution,[],[f249,f593,f1404,f387])).
fof(f5238,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl344),
  inference(cnf_transformation,[],[f5238_D])).
fof(f5238_D,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl344),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl344])])).
fof(f4841,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl4 | $spl172)),
  inference(unit_resulting_resolution,[],[f1404,f593,f1404,f387])).
fof(f5344,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl402),
  inference(cnf_transformation,[],[f5344_D])).
fof(f5344_D,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl402),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl402])])).
fof(f4843,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | ($spl2 | $spl172)),
  inference(unit_resulting_resolution,[],[f437,f1086,f1404,f387])).
fof(f4844,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | ($spl2 | $spl172)),
  inference(unit_resulting_resolution,[],[f250,f1086,f1404,f387])).
fof(f4845,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | ($spl2 | $spl172)),
  inference(unit_resulting_resolution,[],[f249,f1086,f1404,f387])).
fof(f5340,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl398),
  inference(cnf_transformation,[],[f5340_D])).
fof(f5340_D,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl398),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl398])])).
fof(f4848,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl2 | $spl172)),
  inference(unit_resulting_resolution,[],[f1404,f1086,f1404,f387])).
fof(f5240,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl346),
  inference(cnf_transformation,[],[f5240_D])).
fof(f5240_D,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl346),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl346])])).
fof(f4850,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl4 | $spl172)),
  inference(unit_resulting_resolution,[],[f437,f593,f1404,f387])).
fof(f4851,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl4 | $spl172)),
  inference(unit_resulting_resolution,[],[f250,f593,f1404,f387])).
fof(f4852,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl4 | $spl172)),
  inference(unit_resulting_resolution,[],[f249,f593,f1404,f387])).
fof(f4855,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl4 | $spl172)),
  inference(unit_resulting_resolution,[],[f1404,f593,f1404,f387])).
fof(f5342,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl400),
  inference(cnf_transformation,[],[f5342_D])).
fof(f5342_D,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl400),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl400])])).
fof(f4857,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl2 | $spl172)),
  inference(unit_resulting_resolution,[],[f437,f1086,f1404,f387])).
fof(f4858,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl2 | $spl172)),
  inference(unit_resulting_resolution,[],[f250,f1086,f1404,f387])).
fof(f4859,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl2 | $spl172)),
  inference(unit_resulting_resolution,[],[f249,f1086,f1404,f387])).
fof(f4862,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl2 | $spl172)),
  inference(unit_resulting_resolution,[],[f1404,f1086,f1404,f387])).
fof(f5334,plain,(
  ~sP7('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl397),
  inference(cnf_transformation,[],[f5334_D])).
fof(f5334_D,plain,(
  ~sP7('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','F_0') <=> ~$spl397),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl397])])).
fof(f4868,plain,(
  ~sP7('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f437,f255,f1404,f391])).
fof(f4869,plain,(
  ~sP7('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f250,f255,f1404,f391])).
fof(f4870,plain,(
  ~sP7('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f249,f255,f1404,f391])).
fof(f5328,plain,(
  ~sP7('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl393),
  inference(cnf_transformation,[],[f5328_D])).
fof(f5328_D,plain,(
  ~sP7('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') <=> ~$spl393),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl393])])).
fof(f4873,plain,(
  ~sP7('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f1404,f255,f1404,f391])).
fof(f5330,plain,(
  ~sP7('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl395),
  inference(cnf_transformation,[],[f5330_D])).
fof(f5330_D,plain,(
  ~sP7('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') <=> ~$spl395),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl395])])).
fof(f4876,plain,(
  ~sP7('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f437,f255,f1404,f391])).
fof(f4877,plain,(
  ~sP7('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f250,f255,f1404,f391])).
fof(f4878,plain,(
  ~sP7('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f249,f255,f1404,f391])).
fof(f4881,plain,(
  ~sP7('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f1404,f255,f1404,f391])).
fof(f5324,plain,(
  sP8('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl390),
  inference(cnf_transformation,[],[f5324_D])).
fof(f5324_D,plain,(
  sP8('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl390),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl390])])).
fof(f4885,plain,(
  sP8('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f437,f437,f1404,f397])).
fof(f4886,plain,(
  sP8('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f250,f437,f1404,f397])).
fof(f4887,plain,(
  sP8('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f249,f437,f1404,f397])).
fof(f5316,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl382),
  inference(cnf_transformation,[],[f5316_D])).
fof(f5316_D,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl382),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl382])])).
fof(f4890,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f1404,f437,f1404,f397])).
fof(f4891,plain,(
  sP8('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f437,f250,f1404,f397])).
fof(f4892,plain,(
  sP8('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f250,f250,f1404,f397])).
fof(f4893,plain,(
  sP8('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f249,f250,f1404,f397])).
fof(f4896,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f1404,f250,f1404,f397])).
fof(f4897,plain,(
  sP8('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f437,f249,f1404,f397])).
fof(f4898,plain,(
  sP8('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f250,f249,f1404,f397])).
fof(f4899,plain,(
  sP8('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f249,f249,f1404,f397])).
fof(f4902,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f1404,f249,f1404,f397])).
fof(f5320,plain,(
  sP8('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl386),
  inference(cnf_transformation,[],[f5320_D])).
fof(f5320_D,plain,(
  sP8('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl386),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl386])])).
fof(f4915,plain,(
  sP8('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl172),
  inference(unit_resulting_resolution,[],[f437,f1404,f1404,f397])).
fof(f4916,plain,(
  sP8('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl172),
  inference(unit_resulting_resolution,[],[f250,f1404,f1404,f397])).
fof(f4917,plain,(
  sP8('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl172),
  inference(unit_resulting_resolution,[],[f249,f1404,f1404,f397])).
fof(f5312,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl378),
  inference(cnf_transformation,[],[f5312_D])).
fof(f5312_D,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl378),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl378])])).
fof(f4920,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl172),
  inference(unit_resulting_resolution,[],[f1404,f1404,f1404,f397])).
fof(f5322,plain,(
  sP8('T_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl388),
  inference(cnf_transformation,[],[f5322_D])).
fof(f5322_D,plain,(
  sP8('T_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl388),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl388])])).
fof(f4921,plain,(
  sP8('T_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl172),
  inference(unit_resulting_resolution,[],[f437,f437,f1404,f397])).
fof(f4922,plain,(
  sP8('T_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl172),
  inference(unit_resulting_resolution,[],[f250,f437,f1404,f397])).
fof(f4923,plain,(
  sP8('T_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl172),
  inference(unit_resulting_resolution,[],[f249,f437,f1404,f397])).
fof(f4926,plain,(
  sP8('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl172),
  inference(unit_resulting_resolution,[],[f1404,f437,f1404,f397])).
fof(f4927,plain,(
  sP8('T_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl172),
  inference(unit_resulting_resolution,[],[f437,f250,f1404,f397])).
fof(f4928,plain,(
  sP8('T_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl172),
  inference(unit_resulting_resolution,[],[f250,f250,f1404,f397])).
fof(f4929,plain,(
  sP8('T_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl172),
  inference(unit_resulting_resolution,[],[f249,f250,f1404,f397])).
fof(f4932,plain,(
  sP8('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl172),
  inference(unit_resulting_resolution,[],[f1404,f250,f1404,f397])).
fof(f4933,plain,(
  sP8('T_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl172),
  inference(unit_resulting_resolution,[],[f437,f249,f1404,f397])).
fof(f4934,plain,(
  sP8('T_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl172),
  inference(unit_resulting_resolution,[],[f250,f249,f1404,f397])).
fof(f4935,plain,(
  sP8('T_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl172),
  inference(unit_resulting_resolution,[],[f249,f249,f1404,f397])).
fof(f4938,plain,(
  sP8('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl172),
  inference(unit_resulting_resolution,[],[f1404,f249,f1404,f397])).
fof(f5314,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl380),
  inference(cnf_transformation,[],[f5314_D])).
fof(f5314_D,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl380),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl380])])).
fof(f4951,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl172),
  inference(unit_resulting_resolution,[],[f437,f1404,f1404,f397])).
fof(f4952,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl172),
  inference(unit_resulting_resolution,[],[f250,f1404,f1404,f397])).
fof(f4953,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl172),
  inference(unit_resulting_resolution,[],[f249,f1404,f1404,f397])).
fof(f4956,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl172),
  inference(unit_resulting_resolution,[],[f1404,f1404,f1404,f397])).
fof(f5318,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','T_0') | $spl384),
  inference(cnf_transformation,[],[f5318_D])).
fof(f5318_D,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','T_0') <=> ~$spl384),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl384])])).
fof(f4957,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f437,f437,f1404,f397])).
fof(f4958,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f250,f437,f1404,f397])).
fof(f4959,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f249,f437,f1404,f397])).
fof(f4962,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f1404,f437,f1404,f397])).
fof(f4963,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f437,f250,f1404,f397])).
fof(f4964,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f250,f250,f1404,f397])).
fof(f4965,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f249,f250,f1404,f397])).
fof(f4968,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f1404,f250,f1404,f397])).
fof(f4969,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f437,f249,f1404,f397])).
fof(f4970,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f250,f249,f1404,f397])).
fof(f4971,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f249,f249,f1404,f397])).
fof(f4974,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f1404,f249,f1404,f397])).
fof(f4987,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl172),
  inference(unit_resulting_resolution,[],[f437,f1404,f1404,f397])).
fof(f4988,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl172),
  inference(unit_resulting_resolution,[],[f250,f1404,f1404,f397])).
fof(f4989,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl172),
  inference(unit_resulting_resolution,[],[f249,f1404,f1404,f397])).
fof(f4992,plain,(
  sP8('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl172),
  inference(unit_resulting_resolution,[],[f1404,f1404,f1404,f397])).
fof(f5310,plain,(
  ~sP8('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl377),
  inference(cnf_transformation,[],[f5310_D])).
fof(f5310_D,plain,(
  ~sP8('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl377),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl377])])).
fof(f4995,plain,(
  ~sP8('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f255,f437,f1404,f400])).
fof(f4996,plain,(
  ~sP8('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f255,f250,f1404,f400])).
fof(f4997,plain,(
  ~sP8('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f255,f249,f1404,f400])).
fof(f5306,plain,(
  ~sP8('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl373),
  inference(cnf_transformation,[],[f5306_D])).
fof(f5306_D,plain,(
  ~sP8('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl373),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl373])])).
fof(f5000,plain,(
  ~sP8('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl172),
  inference(unit_resulting_resolution,[],[f255,f1404,f1404,f400])).
fof(f5308,plain,(
  ~sP8('F_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl375),
  inference(cnf_transformation,[],[f5308_D])).
fof(f5308_D,plain,(
  ~sP8('F_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl375),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl375])])).
fof(f5001,plain,(
  ~sP8('F_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl172),
  inference(unit_resulting_resolution,[],[f255,f437,f1404,f400])).
fof(f5002,plain,(
  ~sP8('F_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl172),
  inference(unit_resulting_resolution,[],[f255,f250,f1404,f400])).
fof(f5003,plain,(
  ~sP8('F_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl172),
  inference(unit_resulting_resolution,[],[f255,f249,f1404,f400])).
fof(f5006,plain,(
  ~sP8('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl172),
  inference(unit_resulting_resolution,[],[f255,f1404,f1404,f400])).
fof(f5294,plain,(
  ~sP9('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','F_0') | $spl371),
  inference(cnf_transformation,[],[f5294_D])).
fof(f5294_D,plain,(
  ~sP9('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','F_0') <=> ~$spl371),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl371])])).
fof(f5031,plain,(
  ~sP9('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f255,f255,f1404,f408])).
fof(f5288,plain,(
  ~sP10('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','F_0') | $spl369),
  inference(cnf_transformation,[],[f5288_D])).
fof(f5288_D,plain,(
  ~sP10('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','F_0') <=> ~$spl369),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl369])])).
fof(f5066,plain,(
  ~sP10('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f255,f255,f1404,f415])).
fof(f5272,plain,(
  ~sP11('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','T_0') | $spl367),
  inference(cnf_transformation,[],[f5272_D])).
fof(f5272_D,plain,(
  ~sP11('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','T_0') <=> ~$spl367),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl367])])).
fof(f5118,plain,(
  ~sP11('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f255,f437,f1404,f425])).
fof(f5119,plain,(
  ~sP11('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f255,f250,f1404,f425])).
fof(f5120,plain,(
  ~sP11('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f255,f249,f1404,f425])).
fof(f5266,plain,(
  ~sP11('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl363),
  inference(cnf_transformation,[],[f5266_D])).
fof(f5266_D,plain,(
  ~sP11('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl363),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl363])])).
fof(f5123,plain,(
  ~sP11('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl172),
  inference(unit_resulting_resolution,[],[f255,f1404,f1404,f425])).
fof(f5268,plain,(
  ~sP11('T_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl365),
  inference(cnf_transformation,[],[f5268_D])).
fof(f5268_D,plain,(
  ~sP11('T_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl365),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl365])])).
fof(f5126,plain,(
  ~sP11('T_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl172),
  inference(unit_resulting_resolution,[],[f437,f255,f1404,f425])).
fof(f5127,plain,(
  ~sP11('T_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl172),
  inference(unit_resulting_resolution,[],[f250,f255,f1404,f425])).
fof(f5128,plain,(
  ~sP11('T_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl172),
  inference(unit_resulting_resolution,[],[f249,f255,f1404,f425])).
fof(f5131,plain,(
  ~sP11('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl172),
  inference(unit_resulting_resolution,[],[f1404,f255,f1404,f425])).
fof(f5262,plain,(
  sP12('F_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl360),
  inference(cnf_transformation,[],[f5262_D])).
fof(f5262_D,plain,(
  sP12('F_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl360),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl360])])).
fof(f5135,plain,(
  sP12('F_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl172),
  inference(unit_resulting_resolution,[],[f255,f255,f1404,f429])).
fof(f5258,plain,(
  ~sP12('T_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl359),
  inference(cnf_transformation,[],[f5258_D])).
fof(f5258_D,plain,(
  ~sP12('T_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl359),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl359])])).
fof(f5151,plain,(
  ~sP12('T_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl172),
  inference(unit_resulting_resolution,[],[f255,f437,f1404,f432])).
fof(f5152,plain,(
  ~sP12('T_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl172),
  inference(unit_resulting_resolution,[],[f255,f250,f1404,f432])).
fof(f5153,plain,(
  ~sP12('T_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl172),
  inference(unit_resulting_resolution,[],[f255,f249,f1404,f432])).
fof(f5250,plain,(
  ~sP12('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl355),
  inference(cnf_transformation,[],[f5250_D])).
fof(f5250_D,plain,(
  ~sP12('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl355),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl355])])).
fof(f5156,plain,(
  ~sP12('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl172),
  inference(unit_resulting_resolution,[],[f255,f1404,f1404,f432])).
fof(f5252,plain,(
  ~sP12('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','T_0') | $spl357),
  inference(cnf_transformation,[],[f5252_D])).
fof(f5252_D,plain,(
  ~sP12('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','T_0') <=> ~$spl357),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl357])])).
fof(f5161,plain,(
  ~sP12('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f255,f437,f1404,f432])).
fof(f5162,plain,(
  ~sP12('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f255,f250,f1404,f432])).
fof(f5163,plain,(
  ~sP12('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f255,f249,f1404,f432])).
fof(f5166,plain,(
  ~sP12('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl172),
  inference(unit_resulting_resolution,[],[f255,f1404,f1404,f432])).
fof(f5246,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl353),
  inference(cnf_transformation,[],[f5246_D])).
fof(f5246_D,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) <=> ~$spl353),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl353])])).
fof(f5177,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | ($spl4 | $spl172)),
  inference(unit_resulting_resolution,[],[f1404,f800])).
fof(f5244,plain,(
  ~sP1('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl351),
  inference(cnf_transformation,[],[f5244_D])).
fof(f5244_D,plain,(
  ~sP1('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl351),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl351])])).
fof(f5178,plain,(
  ~sP1('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl4 | $spl172)),
  inference(unit_resulting_resolution,[],[f1404,f805])).
fof(f5181,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | ($spl4 | $spl172)),
  inference(unit_resulting_resolution,[],[f437,f1404,f817])).
fof(f5182,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | ($spl4 | $spl172)),
  inference(unit_resulting_resolution,[],[f250,f1404,f817])).
fof(f5183,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | ($spl4 | $spl172)),
  inference(unit_resulting_resolution,[],[f249,f1404,f817])).
fof(f5186,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl4 | $spl172)),
  inference(unit_resulting_resolution,[],[f1404,f1404,f817])).
fof(f5188,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl4 | $spl172)),
  inference(unit_resulting_resolution,[],[f437,f1404,f817])).
fof(f5189,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl4 | $spl172)),
  inference(unit_resulting_resolution,[],[f250,f1404,f817])).
fof(f5190,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl4 | $spl172)),
  inference(unit_resulting_resolution,[],[f249,f1404,f817])).
fof(f5193,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl4 | $spl172)),
  inference(unit_resulting_resolution,[],[f1404,f1404,f817])).
fof(f4678,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl340),
  inference(cnf_transformation,[],[f4678_D])).
fof(f4678_D,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl340),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl340])])).
fof(f4682,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl170),
  inference(global_subsumption,[],[f248,f247,f246,f1402,f4449,f4446,f4445,f4444,f4443,f4442,f4441,f4440,f4435,f4434,f4433,f4432,f4431,f4430,f4429,f4428,f4427,f4426,f4422,f4421,f4420,f4417,f4414,f4413,f4410,f4409,f4408,f4407,f4406,f4405,f4402,f4401,f4400,f4399,f4398,f4397,f4356,f4346,f4345,f4344,f4343,f4342,f4341,f4335,f4334,f4333,f4330,f4329,f4328,f4327,f4321,f4320,f4319,f4316,f4315,f4310,f4285,f4284,f4283,f4282,f4281,f4280,f4279,f4278,f4277,f4276,f4275,f4274,f4273,f4272,f4271,f4270,f4269,f4268,f4267,f4266,f4265,f4264,f4263,f4262,f4259,f4258,f4257,f4256,f4253,f4252,f4251,f4250,f4249,f4248,f4247,f4244,f4243,f4242,f4241,f4238,f4237,f4236,f4235,f4232,f4231,f4230,f4229,f4228,f4227,f4226,f4223,f4222,f4221,f4220,f4217,f4216,f4215,f4214,f4213,f4212,f4211,f4208,f4207,f4206,f4205,f4202,f4201,f4200,f4199,f4196,f4195,f4194,f4192,f4191,f4190,f4186,f4185,f4184,f4181,f4180,f4179,f4178,f4177,f4176,f4175,f4172,f4171,f4170,f4169,f4166,f4165,f4164,f4163,f4160,f4159,f4158,f4157,f4156,f4155,f4152,f4151,f4150,f4149,f4148,f4147,f4108,f4107,f4106,f4105,f4102,f4101,f4100,f4099,f4098,f4097,f4094,f4093,f4092,f4091,f4090,f4089,f4088,f4085,f4084,f4083,f4082,f4081,f4080,f4077,f4076,f4075,f4074,f4073,f4072,f4067,f4066,f4065,f4064,f4063,f4062,f4061,f4060,f4059,f4058,f4057,f4056,f4055,f4054,f4053,f4052,f4048,f4047,f4046,f4023,f4022,f4021,f4017,f4016,f4015,f4006,f4005,f4004,f4003,f4002,f4001,f4000,f3999,f3998,f3997,f3996,f3995,f3994,f3993,f3992,f3991,f3990,f3987,f3986,f3985,f3984,f3981,f3980,f3979,f3978,f3977,f3969,f3968,f3967,f3966,f3965,f4478,f4479,f4480,f4489,f4490,f4491,f4494,f4497,f4499,f4506,f4509,f4510,f4512,f4513,f4514,f4516])).
fof(f4516,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl170),
  inference(superposition,[],[f1402,f256])).
fof(f4514,plain,(
  ( ! [X48,X49] : (sP12('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X48,X49) | ~p(s(bool,X48))) ) | $spl170),
  inference(resolution,[],[f1402,f431])).
fof(f4513,plain,(
  ( ! [X47,X46] : (sP12('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X47,X46) | p(s(bool,X46))) ) | $spl170),
  inference(resolution,[],[f1402,f430])).
fof(f4512,plain,(
  ( ! [X45,X44] : (sP12(X45,X44,'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | p(s(bool,X44)) | p(s(bool,X45))) ) | $spl170),
  inference(resolution,[],[f1402,f429])).
fof(f4510,plain,(
  ( ! [X41,X40] : (sP10('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X40,X41) | ~p(s(bool,X40))) ) | $spl170),
  inference(resolution,[],[f1402,f414])).
fof(f4509,plain,(
  ( ! [X39,X38] : (sP10('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X39,X38) | ~p(s(bool,X38))) ) | $spl170),
  inference(resolution,[],[f1402,f413])).
fof(f4506,plain,(
  ( ! [X33,X32] : (sP8('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X32,X33) | ~p(s(bool,X33)) | ~p(s(bool,X32))) ) | $spl170),
  inference(resolution,[],[f1402,f397])).
fof(f4499,plain,(
  ( ! [X19,X18] : (sP2('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X18,X19) | ~p(s(bool,X19))) ) | $spl170),
  inference(resolution,[],[f1402,f352])).
fof(f4497,plain,(
  ( ! [X14,X15] : (sP2(X14,X15,'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ~p(s(bool,X14))) ) | $spl170),
  inference(resolution,[],[f1402,f350])).
fof(f4494,plain,(
  ( ! [X9] : (~sP1('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X9) | ~p(s(bool,X9))) ) | $spl170),
  inference(resolution,[],[f1402,f335])).
fof(f4491,plain,(
  ( ! [X6] : (sP1(X6,'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | p(s(bool,X6))) ) | $spl170),
  inference(resolution,[],[f1402,f332])).
fof(f4490,plain,(
  ( ! [X5] : (sP1('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X5) | p(s(bool,X5))) ) | $spl170),
  inference(resolution,[],[f1402,f329])).
fof(f4489,plain,(
  ( ! [X4] : (~sP1(X4,'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ~p(s(bool,X4))) ) | $spl170),
  inference(resolution,[],[f1402,f328])).
fof(f3977,plain,(
  ( ! [X0] : (sP0(X0,'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl170),
  inference(unit_resulting_resolution,[],[f255,f1402,f318])).
fof(f3978,plain,(
  ( ! [X0] : (sP0('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X0)) ) | $spl170),
  inference(unit_resulting_resolution,[],[f1402,f319])).
fof(f3998,plain,(
  ( ! [X0,X1] : (~sP1('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl170),
  inference(unit_resulting_resolution,[],[f304,f1402,f335])).
fof(f3999,plain,(
  ( ! [X0,X1] : (~sP1('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl170),
  inference(unit_resulting_resolution,[],[f305,f1402,f335])).
fof(f4004,plain,(
  ( ! [X0,X1] : (~sP1('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)))) ) | $spl170),
  inference(unit_resulting_resolution,[],[f304,f1402,f335])).
fof(f4005,plain,(
  ( ! [X0,X1] : (~sP1('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) ) | $spl170),
  inference(unit_resulting_resolution,[],[f305,f1402,f335])).
fof(f4015,plain,(
  ( ! [X0] : (sP2(X0,'T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl170),
  inference(unit_resulting_resolution,[],[f249,f1402,f348])).
fof(f4016,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,'const.pred_set.SUBSET_2'(s(fun(X1,bool),X2),s(fun(X1,bool),X2)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl170),
  inference(unit_resulting_resolution,[],[f304,f1402,f348])).
fof(f4017,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,'const.pred_set.SUBSET_2'(s(fun(X1,bool),X2),s(fun(X1,bool),'const.pred_set.UNIV_0')),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl170),
  inference(unit_resulting_resolution,[],[f305,f1402,f348])).
fof(f4021,plain,(
  ( ! [X0] : (sP2('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0',X0)) ) | $spl170),
  inference(unit_resulting_resolution,[],[f249,f1402,f349])).
fof(f4022,plain,(
  ( ! [X2,X0,X1] : (sP2('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),X2)) ) | $spl170),
  inference(unit_resulting_resolution,[],[f304,f1402,f349])).
fof(f4023,plain,(
  ( ! [X2,X0,X1] : (sP2('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),X2)) ) | $spl170),
  inference(unit_resulting_resolution,[],[f305,f1402,f349])).
fof(f4046,plain,(
  ( ! [X0] : (sP2('T_0',X0,'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl170),
  inference(unit_resulting_resolution,[],[f249,f1402,f352])).
fof(f4047,plain,(
  ( ! [X2,X0,X1] : (sP2('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),X2,'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl170),
  inference(unit_resulting_resolution,[],[f304,f1402,f352])).
fof(f4048,plain,(
  ( ! [X2,X0,X1] : (sP2('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),X2,'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl170),
  inference(unit_resulting_resolution,[],[f305,f1402,f352])).
fof(f4052,plain,(
  ( ! [X0] : (sP2('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X0,'T_0')) ) | $spl170),
  inference(unit_resulting_resolution,[],[f249,f1402,f352])).
fof(f4053,plain,(
  ( ! [X2,X0,X1] : (sP2('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X0,'const.pred_set.SUBSET_2'(s(fun(X1,bool),X2),s(fun(X1,bool),X2)))) ) | $spl170),
  inference(unit_resulting_resolution,[],[f304,f1402,f352])).
fof(f4054,plain,(
  ( ! [X2,X0,X1] : (sP2('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X0,'const.pred_set.SUBSET_2'(s(fun(X1,bool),X2),s(fun(X1,bool),'const.pred_set.UNIV_0')))) ) | $spl170),
  inference(unit_resulting_resolution,[],[f305,f1402,f352])).
fof(f4055,plain,(
  ( ! [X0] : (sP2('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X0,'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl170),
  inference(unit_resulting_resolution,[],[f1402,f1402,f352])).
fof(f4059,plain,(
  ( ! [X0,X1] : (~sP6('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'F_0')) ) | $spl170),
  inference(unit_resulting_resolution,[],[f304,f255,f1402,f364])).
fof(f4060,plain,(
  ( ! [X0,X1] : (~sP6('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'F_0')) ) | $spl170),
  inference(unit_resulting_resolution,[],[f305,f255,f1402,f364])).
fof(f4065,plain,(
  ( ! [X0,X1] : (~sP6('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0')) ) | $spl170),
  inference(unit_resulting_resolution,[],[f304,f255,f1402,f364])).
fof(f4066,plain,(
  ( ! [X0,X1] : (~sP6('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0')) ) | $spl170),
  inference(unit_resulting_resolution,[],[f305,f255,f1402,f364])).
fof(f4075,plain,(
  ( ! [X0,X1] : (~sP5('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'F_0')) ) | $spl170),
  inference(unit_resulting_resolution,[],[f304,f255,f1402,f368])).
fof(f4076,plain,(
  ( ! [X0,X1] : (~sP5('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'F_0')) ) | $spl170),
  inference(unit_resulting_resolution,[],[f305,f255,f1402,f368])).
fof(f4083,plain,(
  ( ! [X0,X1] : (~sP5('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0')) ) | $spl170),
  inference(unit_resulting_resolution,[],[f304,f255,f1402,f368])).
fof(f4084,plain,(
  ( ! [X0,X1] : (~sP5('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0')) ) | $spl170),
  inference(unit_resulting_resolution,[],[f305,f255,f1402,f368])).
fof(f4092,plain,(
  ( ! [X0,X1] : (~sP4('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'F_0')) ) | $spl170),
  inference(unit_resulting_resolution,[],[f304,f255,f1402,f372])).
fof(f4093,plain,(
  ( ! [X0,X1] : (~sP4('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'F_0')) ) | $spl170),
  inference(unit_resulting_resolution,[],[f305,f255,f1402,f372])).
fof(f4100,plain,(
  ( ! [X0,X1] : (~sP4('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0')) ) | $spl170),
  inference(unit_resulting_resolution,[],[f304,f255,f1402,f372])).
fof(f4101,plain,(
  ( ! [X0,X1] : (~sP4('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0')) ) | $spl170),
  inference(unit_resulting_resolution,[],[f305,f255,f1402,f372])).
fof(f4150,plain,(
  ( ! [X0,X1] : (~sP7('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'F_0')) ) | $spl170),
  inference(unit_resulting_resolution,[],[f304,f255,f1402,f391])).
fof(f4151,plain,(
  ( ! [X0,X1] : (~sP7('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'F_0')) ) | $spl170),
  inference(unit_resulting_resolution,[],[f305,f255,f1402,f391])).
fof(f4158,plain,(
  ( ! [X0,X1] : (~sP7('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0')) ) | $spl170),
  inference(unit_resulting_resolution,[],[f304,f255,f1402,f391])).
fof(f4159,plain,(
  ( ! [X0,X1] : (~sP7('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0')) ) | $spl170),
  inference(unit_resulting_resolution,[],[f305,f255,f1402,f391])).
fof(f4179,plain,(
  ( ! [X0,X1] : (sP8('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0')) ) | $spl170),
  inference(unit_resulting_resolution,[],[f304,f249,f1402,f397])).
fof(f4180,plain,(
  ( ! [X0,X1] : (sP8('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0')) ) | $spl170),
  inference(unit_resulting_resolution,[],[f305,f249,f1402,f397])).
fof(f4184,plain,(
  ( ! [X0,X1] : (sP8('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)))) ) | $spl170),
  inference(unit_resulting_resolution,[],[f249,f304,f1402,f397])).
fof(f4185,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X2,bool),X3),s(fun(X2,bool),X3)))) ) | $spl170),
  inference(unit_resulting_resolution,[],[f304,f304,f1402,f397])).
fof(f4186,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X2,bool),X3),s(fun(X2,bool),X3)))) ) | $spl170),
  inference(unit_resulting_resolution,[],[f305,f304,f1402,f397])).
fof(f4190,plain,(
  ( ! [X0,X1] : (sP8('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) ) | $spl170),
  inference(unit_resulting_resolution,[],[f249,f305,f1402,f397])).
fof(f4191,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X2,bool),X3),s(fun(X2,bool),'const.pred_set.UNIV_0')))) ) | $spl170),
  inference(unit_resulting_resolution,[],[f304,f305,f1402,f397])).
fof(f4192,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X2,bool),X3),s(fun(X2,bool),'const.pred_set.UNIV_0')))) ) | $spl170),
  inference(unit_resulting_resolution,[],[f305,f305,f1402,f397])).
fof(f4215,plain,(
  ( ! [X0,X1] : (sP8('T_0','const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl170),
  inference(unit_resulting_resolution,[],[f304,f249,f1402,f397])).
fof(f4216,plain,(
  ( ! [X0,X1] : (sP8('T_0','const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl170),
  inference(unit_resulting_resolution,[],[f305,f249,f1402,f397])).
fof(f4220,plain,(
  ( ! [X0,X1] : (sP8('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl170),
  inference(unit_resulting_resolution,[],[f249,f304,f1402,f397])).
fof(f4221,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.SUBSET_2'(s(fun(X2,bool),X3),s(fun(X2,bool),X3)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl170),
  inference(unit_resulting_resolution,[],[f304,f304,f1402,f397])).
fof(f4222,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.SUBSET_2'(s(fun(X2,bool),X3),s(fun(X2,bool),'const.pred_set.UNIV_0')),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl170),
  inference(unit_resulting_resolution,[],[f305,f304,f1402,f397])).
fof(f4223,plain,(
  ( ! [X0,X1] : (sP8('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl170),
  inference(unit_resulting_resolution,[],[f1402,f304,f1402,f397])).
fof(f4226,plain,(
  ( ! [X0,X1] : (sP8('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl170),
  inference(unit_resulting_resolution,[],[f249,f305,f1402,f397])).
fof(f4227,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.SUBSET_2'(s(fun(X2,bool),X3),s(fun(X2,bool),X3)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl170),
  inference(unit_resulting_resolution,[],[f304,f305,f1402,f397])).
fof(f4228,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.SUBSET_2'(s(fun(X2,bool),X3),s(fun(X2,bool),'const.pred_set.UNIV_0')),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl170),
  inference(unit_resulting_resolution,[],[f305,f305,f1402,f397])).
fof(f4229,plain,(
  ( ! [X0,X1] : (sP8('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl170),
  inference(unit_resulting_resolution,[],[f1402,f305,f1402,f397])).
fof(f4251,plain,(
  ( ! [X0,X1] : (sP8('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'T_0')) ) | $spl170),
  inference(unit_resulting_resolution,[],[f304,f249,f1402,f397])).
fof(f4252,plain,(
  ( ! [X0,X1] : (sP8('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'T_0')) ) | $spl170),
  inference(unit_resulting_resolution,[],[f305,f249,f1402,f397])).
fof(f4256,plain,(
  ( ! [X0,X1] : (sP8('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)))) ) | $spl170),
  inference(unit_resulting_resolution,[],[f249,f304,f1402,f397])).
fof(f4257,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.SUBSET_2'(s(fun(X2,bool),X3),s(fun(X2,bool),X3)))) ) | $spl170),
  inference(unit_resulting_resolution,[],[f304,f304,f1402,f397])).
fof(f4258,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.SUBSET_2'(s(fun(X2,bool),X3),s(fun(X2,bool),X3)))) ) | $spl170),
  inference(unit_resulting_resolution,[],[f305,f304,f1402,f397])).
fof(f4259,plain,(
  ( ! [X0,X1] : (sP8('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)))) ) | $spl170),
  inference(unit_resulting_resolution,[],[f1402,f304,f1402,f397])).
fof(f4262,plain,(
  ( ! [X0,X1] : (sP8('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) ) | $spl170),
  inference(unit_resulting_resolution,[],[f249,f305,f1402,f397])).
fof(f4263,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.SUBSET_2'(s(fun(X2,bool),X3),s(fun(X2,bool),'const.pred_set.UNIV_0')))) ) | $spl170),
  inference(unit_resulting_resolution,[],[f304,f305,f1402,f397])).
fof(f4264,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.SUBSET_2'(s(fun(X2,bool),X3),s(fun(X2,bool),'const.pred_set.UNIV_0')))) ) | $spl170),
  inference(unit_resulting_resolution,[],[f305,f305,f1402,f397])).
fof(f4265,plain,(
  ( ! [X0,X1] : (sP8('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) ) | $spl170),
  inference(unit_resulting_resolution,[],[f1402,f305,f1402,f397])).
fof(f4269,plain,(
  ( ! [X0,X1] : (sP8('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl170),
  inference(unit_resulting_resolution,[],[f304,f1402,f1402,f397])).
fof(f4270,plain,(
  ( ! [X0,X1] : (sP8('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl170),
  inference(unit_resulting_resolution,[],[f305,f1402,f1402,f397])).
fof(f4272,plain,(
  ( ! [X0] : (~sP8('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X0,'F_0')) ) | $spl170),
  inference(unit_resulting_resolution,[],[f255,f1402,f398])).
fof(f4273,plain,(
  ( ! [X0] : (~sP8('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0',X0)) ) | $spl170),
  inference(unit_resulting_resolution,[],[f255,f1402,f399])).
fof(f4277,plain,(
  ( ! [X0,X1] : (~sP8('F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)))) ) | $spl170),
  inference(unit_resulting_resolution,[],[f255,f304,f1402,f400])).
fof(f4278,plain,(
  ( ! [X0,X1] : (~sP8('F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) ) | $spl170),
  inference(unit_resulting_resolution,[],[f255,f305,f1402,f400])).
fof(f4283,plain,(
  ( ! [X0,X1] : (~sP8('F_0','const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl170),
  inference(unit_resulting_resolution,[],[f255,f304,f1402,f400])).
fof(f4284,plain,(
  ( ! [X0,X1] : (~sP8('F_0','const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl170),
  inference(unit_resulting_resolution,[],[f255,f305,f1402,f400])).
fof(f4315,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl170),
  inference(unit_resulting_resolution,[],[f1402,f409])).
fof(f4316,plain,(
  ( ! [X0,X1] : (sP9(X0,'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X1)) ) | $spl170),
  inference(unit_resulting_resolution,[],[f1402,f410])).
fof(f4319,plain,(
  ( ! [X0] : (sP10('T_0',X0,'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl170),
  inference(unit_resulting_resolution,[],[f249,f1402,f413])).
fof(f4320,plain,(
  ( ! [X2,X0,X1] : (sP10('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),X2,'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl170),
  inference(unit_resulting_resolution,[],[f304,f1402,f413])).
fof(f4321,plain,(
  ( ! [X2,X0,X1] : (sP10('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),X2,'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl170),
  inference(unit_resulting_resolution,[],[f305,f1402,f413])).
fof(f4327,plain,(
  ( ! [X0] : (sP10('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X0,'T_0')) ) | $spl170),
  inference(unit_resulting_resolution,[],[f249,f1402,f413])).
fof(f4328,plain,(
  ( ! [X2,X0,X1] : (sP10('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X0,'const.pred_set.SUBSET_2'(s(fun(X1,bool),X2),s(fun(X1,bool),X2)))) ) | $spl170),
  inference(unit_resulting_resolution,[],[f304,f1402,f413])).
fof(f4329,plain,(
  ( ! [X2,X0,X1] : (sP10('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X0,'const.pred_set.SUBSET_2'(s(fun(X1,bool),X2),s(fun(X1,bool),'const.pred_set.UNIV_0')))) ) | $spl170),
  inference(unit_resulting_resolution,[],[f305,f1402,f413])).
fof(f4330,plain,(
  ( ! [X0] : (sP10('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X0,'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl170),
  inference(unit_resulting_resolution,[],[f1402,f1402,f413])).
fof(f4333,plain,(
  ( ! [X0] : (sP10('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X0)) ) | $spl170),
  inference(unit_resulting_resolution,[],[f249,f1402,f414])).
fof(f4334,plain,(
  ( ! [X2,X0,X1] : (sP10('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X2)) ) | $spl170),
  inference(unit_resulting_resolution,[],[f304,f1402,f414])).
fof(f4335,plain,(
  ( ! [X2,X0,X1] : (sP10('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X2)) ) | $spl170),
  inference(unit_resulting_resolution,[],[f305,f1402,f414])).
fof(f4341,plain,(
  ( ! [X0] : (sP10('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0',X0)) ) | $spl170),
  inference(unit_resulting_resolution,[],[f249,f1402,f414])).
fof(f4342,plain,(
  ( ! [X2,X0,X1] : (sP10('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),X2)) ) | $spl170),
  inference(unit_resulting_resolution,[],[f304,f1402,f414])).
fof(f4343,plain,(
  ( ! [X2,X0,X1] : (sP10('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),X2)) ) | $spl170),
  inference(unit_resulting_resolution,[],[f305,f1402,f414])).
fof(f4344,plain,(
  ( ! [X0] : (sP10('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X0)) ) | $spl170),
  inference(unit_resulting_resolution,[],[f1402,f1402,f414])).
fof(f4346,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl170),
  inference(unit_resulting_resolution,[],[f255,f1402,f416])).
fof(f4356,plain,(
  ( ! [X0] : (~sP10('F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X0)) ) | $spl170),
  inference(unit_resulting_resolution,[],[f255,f1402,f417])).
fof(f4400,plain,(
  ( ! [X0,X1] : (~sP11('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)))) ) | $spl170),
  inference(unit_resulting_resolution,[],[f255,f304,f1402,f425])).
fof(f4401,plain,(
  ( ! [X0,X1] : (~sP11('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) ) | $spl170),
  inference(unit_resulting_resolution,[],[f255,f305,f1402,f425])).
fof(f4408,plain,(
  ( ! [X0,X1] : (~sP11('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl170),
  inference(unit_resulting_resolution,[],[f304,f255,f1402,f425])).
fof(f4409,plain,(
  ( ! [X0,X1] : (~sP11('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl170),
  inference(unit_resulting_resolution,[],[f305,f255,f1402,f425])).
fof(f4413,plain,(
  ( ! [X0,X1] : (sP11(X0,'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X1)) ) | $spl170),
  inference(unit_resulting_resolution,[],[f1402,f427])).
fof(f4417,plain,(
  ( ! [X0] : (sP12('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X0,'F_0')) ) | $spl170),
  inference(unit_resulting_resolution,[],[f255,f1402,f430])).
fof(f4420,plain,(
  ( ! [X0] : (sP12('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X0)) ) | $spl170),
  inference(unit_resulting_resolution,[],[f249,f1402,f431])).
fof(f4421,plain,(
  ( ! [X2,X0,X1] : (sP12('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X2)) ) | $spl170),
  inference(unit_resulting_resolution,[],[f304,f1402,f431])).
fof(f4422,plain,(
  ( ! [X2,X0,X1] : (sP12('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X2)) ) | $spl170),
  inference(unit_resulting_resolution,[],[f305,f1402,f431])).
fof(f4426,plain,(
  ( ! [X0] : (sP12('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0',X0)) ) | $spl170),
  inference(unit_resulting_resolution,[],[f249,f1402,f431])).
fof(f4427,plain,(
  ( ! [X2,X0,X1] : (sP12('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),X2)) ) | $spl170),
  inference(unit_resulting_resolution,[],[f304,f1402,f431])).
fof(f4428,plain,(
  ( ! [X2,X0,X1] : (sP12('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),X2)) ) | $spl170),
  inference(unit_resulting_resolution,[],[f305,f1402,f431])).
fof(f4429,plain,(
  ( ! [X0] : (sP12('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X0)) ) | $spl170),
  inference(unit_resulting_resolution,[],[f1402,f1402,f431])).
fof(f4433,plain,(
  ( ! [X0,X1] : (~sP12('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl170),
  inference(unit_resulting_resolution,[],[f255,f304,f1402,f432])).
fof(f4434,plain,(
  ( ! [X0,X1] : (~sP12('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl170),
  inference(unit_resulting_resolution,[],[f255,f305,f1402,f432])).
fof(f4443,plain,(
  ( ! [X0,X1] : (~sP12('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)))) ) | $spl170),
  inference(unit_resulting_resolution,[],[f255,f304,f1402,f432])).
fof(f4444,plain,(
  ( ! [X0,X1] : (~sP12('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) ) | $spl170),
  inference(unit_resulting_resolution,[],[f255,f305,f1402,f432])).
fof(f4446,plain,(
  ( ! [X0] : (~sP12('F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X0)) ) | $spl170),
  inference(unit_resulting_resolution,[],[f255,f1402,f434])).
fof(f4449,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X1)) ) | $spl170),
  inference(unit_resulting_resolution,[],[f1402,f438])).
fof(f4480,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl170),
  inference(resolution,[],[f1402,f265])).
fof(f4479,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl170),
  inference(resolution,[],[f1402,f263])).
fof(f4680,plain,(
  s(fun(sK13,bool),sK14) = s(fun(sK13,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14))) | $spl342),
  inference(cnf_transformation,[],[f4680_D])).
fof(f4680_D,plain,(
  s(fun(sK13,bool),sK14) = s(fun(sK13,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14))) <=> ~$spl342),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl342])])).
fof(f4478,plain,(
  s(fun(sK13,bool),sK14) = s(fun(sK13,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14))) | $spl170),
  inference(resolution,[],[f1402,f339])).
fof(f339,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | s(fun(X1,bool),X4) = s(fun(X1,bool),'const.pred_set.IMAGE_2'(s(fun(X0,X1),X2),s(fun(X0,bool),X3)))) )),
  inference(cnf_transformation,[],[f201])).
fof(f201,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | s(fun(X1,bool),X4) = s(fun(X1,bool),'const.pred_set.IMAGE_2'(s(fun(X0,X1),X2),s(fun(X0,bool),X3)))) & (s(fun(X1,bool),X4) != s(fun(X1,bool),'const.pred_set.IMAGE_2'(s(fun(X0,X1),X2),s(fun(X0,bool),X3))) | p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))))),
  inference(nnf_transformation,[],[f99])).
fof(f99,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) <=> s(fun(X1,bool),X4) = s(fun(X1,bool),'const.pred_set.IMAGE_2'(s(fun(X0,X1),X2),s(fun(X0,bool),X3))))),
  inference(flattening,[],[f98])).
fof(f98,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) <=> s(fun(X1,bool),X4) = s(fun(X1,bool),'const.pred_set.IMAGE_2'(s(fun(X0,X1),X2),s(fun(X0,bool),X3))))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X6,X13] : ! [X9,X7,X10] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(X6,X13),X9),s(fun(X6,bool),X7),s(fun(X13,bool),X10)))) <=> s(fun(X13,bool),X10) = s(fun(X13,bool),'const.pred_set.IMAGE_2'(s(fun(X6,X13),X9),s(fun(X6,bool),X7))))),
  file('thm.pred_set.SUM_IMAGE_PERMUTES','thm.pred_set.IMAGE_SURJ')).
fof(f3965,plain,(
  s(fun(sK13,bool),sK14) = s(fun(sK13,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14))) | $spl170),
  inference(unit_resulting_resolution,[],[f1402,f339])).
fof(f3966,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl170),
  inference(unit_resulting_resolution,[],[f1402,f263])).
fof(f3967,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl170),
  inference(unit_resulting_resolution,[],[f1402,f265])).
fof(f4676,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl339),
  inference(cnf_transformation,[],[f4676_D])).
fof(f4676_D,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl339),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl339])])).
fof(f3968,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl170),
  inference(unit_resulting_resolution,[],[f1402,f266])).
fof(f3969,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl170),
  inference(unit_resulting_resolution,[],[f1402,f268])).
fof(f4668,plain,(
  ~sP1('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl331),
  inference(cnf_transformation,[],[f4668_D])).
fof(f4668_D,plain,(
  ~sP1('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl331),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl331])])).
fof(f3979,plain,(
  ~sP1('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl170),
  inference(unit_resulting_resolution,[],[f437,f1402,f328])).
fof(f3980,plain,(
  ~sP1('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl170),
  inference(unit_resulting_resolution,[],[f250,f1402,f328])).
fof(f3981,plain,(
  ~sP1('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl170),
  inference(unit_resulting_resolution,[],[f249,f1402,f328])).
fof(f4666,plain,(
  ~sP1('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl329),
  inference(cnf_transformation,[],[f4666_D])).
fof(f4666_D,plain,(
  ~sP1('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl329),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl329])])).
fof(f3984,plain,(
  ~sP1('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl170),
  inference(unit_resulting_resolution,[],[f1402,f1402,f328])).
fof(f4670,plain,(
  ~sP1('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl333),
  inference(cnf_transformation,[],[f4670_D])).
fof(f4670_D,plain,(
  ~sP1('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl333),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl333])])).
fof(f3985,plain,(
  ~sP1('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl170),
  inference(unit_resulting_resolution,[],[f437,f1402,f328])).
fof(f3986,plain,(
  ~sP1('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl170),
  inference(unit_resulting_resolution,[],[f250,f1402,f328])).
fof(f3987,plain,(
  ~sP1('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl170),
  inference(unit_resulting_resolution,[],[f249,f1402,f328])).
fof(f3990,plain,(
  ~sP1('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl170),
  inference(unit_resulting_resolution,[],[f1402,f1402,f328])).
fof(f4672,plain,(
  sP1('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl334),
  inference(cnf_transformation,[],[f4672_D])).
fof(f4672_D,plain,(
  sP1('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') <=> ~$spl334),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl334])])).
fof(f3991,plain,(
  sP1('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl170),
  inference(unit_resulting_resolution,[],[f255,f1402,f329])).
fof(f4674,plain,(
  sP1('F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl336),
  inference(cnf_transformation,[],[f4674_D])).
fof(f4674_D,plain,(
  sP1('F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl336),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl336])])).
fof(f3992,plain,(
  sP1('F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl170),
  inference(unit_resulting_resolution,[],[f255,f1402,f332])).
fof(f3993,plain,(
  sP1('F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl170),
  inference(unit_resulting_resolution,[],[f255,f1402,f333])).
fof(f3994,plain,(
  sP1('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl170),
  inference(unit_resulting_resolution,[],[f255,f1402,f334])).
fof(f3995,plain,(
  ~sP1('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl170),
  inference(unit_resulting_resolution,[],[f437,f1402,f335])).
fof(f3996,plain,(
  ~sP1('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl170),
  inference(unit_resulting_resolution,[],[f250,f1402,f335])).
fof(f3997,plain,(
  ~sP1('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl170),
  inference(unit_resulting_resolution,[],[f249,f1402,f335])).
fof(f4000,plain,(
  ~sP1('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl170),
  inference(unit_resulting_resolution,[],[f1402,f1402,f335])).
fof(f4001,plain,(
  ~sP1('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl170),
  inference(unit_resulting_resolution,[],[f437,f1402,f335])).
fof(f4002,plain,(
  ~sP1('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl170),
  inference(unit_resulting_resolution,[],[f250,f1402,f335])).
fof(f4003,plain,(
  ~sP1('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl170),
  inference(unit_resulting_resolution,[],[f249,f1402,f335])).
fof(f4006,plain,(
  ~sP1('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl170),
  inference(unit_resulting_resolution,[],[f1402,f1402,f335])).
fof(f4628,plain,(
  ~sP6('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl299),
  inference(cnf_transformation,[],[f4628_D])).
fof(f4628_D,plain,(
  ~sP6('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','F_0') <=> ~$spl299),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl299])])).
fof(f4056,plain,(
  ~sP6('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl170),
  inference(unit_resulting_resolution,[],[f437,f255,f1402,f364])).
fof(f4057,plain,(
  ~sP6('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl170),
  inference(unit_resulting_resolution,[],[f250,f255,f1402,f364])).
fof(f4058,plain,(
  ~sP6('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl170),
  inference(unit_resulting_resolution,[],[f249,f255,f1402,f364])).
fof(f4662,plain,(
  ~sP6('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl325),
  inference(cnf_transformation,[],[f4662_D])).
fof(f4662_D,plain,(
  ~sP6('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') <=> ~$spl325),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl325])])).
fof(f4061,plain,(
  ~sP6('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl170),
  inference(unit_resulting_resolution,[],[f1402,f255,f1402,f364])).
fof(f4664,plain,(
  ~sP6('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl327),
  inference(cnf_transformation,[],[f4664_D])).
fof(f4664_D,plain,(
  ~sP6('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') <=> ~$spl327),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl327])])).
fof(f4062,plain,(
  ~sP6('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl170),
  inference(unit_resulting_resolution,[],[f437,f255,f1402,f364])).
fof(f4063,plain,(
  ~sP6('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl170),
  inference(unit_resulting_resolution,[],[f250,f255,f1402,f364])).
fof(f4064,plain,(
  ~sP6('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl170),
  inference(unit_resulting_resolution,[],[f249,f255,f1402,f364])).
fof(f4067,plain,(
  ~sP6('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl170),
  inference(unit_resulting_resolution,[],[f1402,f255,f1402,f364])).
fof(f4660,plain,(
  sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl322),
  inference(cnf_transformation,[],[f4660_D])).
fof(f4660_D,plain,(
  sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl322),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl322])])).
fof(f4068,plain,(
  sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl4 | $spl170)),
  inference(unit_resulting_resolution,[],[f577,f575,f562,f1402,f365])).
fof(f4658,plain,(
  sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl320),
  inference(cnf_transformation,[],[f4658_D])).
fof(f4658_D,plain,(
  sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl320),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl320])])).
fof(f4069,plain,(
  sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl2 | $spl4 | $spl170)),
  inference(unit_resulting_resolution,[],[f577,f1066,f562,f1402,f365])).
fof(f4656,plain,(
  sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl318),
  inference(cnf_transformation,[],[f4656_D])).
fof(f4656_D,plain,(
  sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl318),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl318])])).
fof(f4070,plain,(
  sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl2 | $spl4 | $spl170)),
  inference(unit_resulting_resolution,[],[f1068,f575,f1053,f1402,f365])).
fof(f4654,plain,(
  sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl316),
  inference(cnf_transformation,[],[f4654_D])).
fof(f4654_D,plain,(
  sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl316),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl316])])).
fof(f4071,plain,(
  sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl2 | $spl170)),
  inference(unit_resulting_resolution,[],[f1068,f1066,f1053,f1402,f365])).
fof(f4652,plain,(
  ~sP5('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl315),
  inference(cnf_transformation,[],[f4652_D])).
fof(f4652_D,plain,(
  ~sP5('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','F_0') <=> ~$spl315),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl315])])).
fof(f4072,plain,(
  ~sP5('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl170),
  inference(unit_resulting_resolution,[],[f437,f255,f1402,f368])).
fof(f4073,plain,(
  ~sP5('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl170),
  inference(unit_resulting_resolution,[],[f250,f255,f1402,f368])).
fof(f4074,plain,(
  ~sP5('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl170),
  inference(unit_resulting_resolution,[],[f249,f255,f1402,f368])).
fof(f4646,plain,(
  ~sP5('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl311),
  inference(cnf_transformation,[],[f4646_D])).
fof(f4646_D,plain,(
  ~sP5('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') <=> ~$spl311),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl311])])).
fof(f4077,plain,(
  ~sP5('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl170),
  inference(unit_resulting_resolution,[],[f1402,f255,f1402,f368])).
fof(f4648,plain,(
  ~sP5('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl313),
  inference(cnf_transformation,[],[f4648_D])).
fof(f4648_D,plain,(
  ~sP5('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') <=> ~$spl313),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl313])])).
fof(f4080,plain,(
  ~sP5('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl170),
  inference(unit_resulting_resolution,[],[f437,f255,f1402,f368])).
fof(f4081,plain,(
  ~sP5('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl170),
  inference(unit_resulting_resolution,[],[f250,f255,f1402,f368])).
fof(f4082,plain,(
  ~sP5('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl170),
  inference(unit_resulting_resolution,[],[f249,f255,f1402,f368])).
fof(f4085,plain,(
  ~sP5('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl170),
  inference(unit_resulting_resolution,[],[f1402,f255,f1402,f368])).
fof(f4642,plain,(
  ~sP4('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl309),
  inference(cnf_transformation,[],[f4642_D])).
fof(f4642_D,plain,(
  ~sP4('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','F_0') <=> ~$spl309),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl309])])).
fof(f4089,plain,(
  ~sP4('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl170),
  inference(unit_resulting_resolution,[],[f437,f255,f1402,f372])).
fof(f4090,plain,(
  ~sP4('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl170),
  inference(unit_resulting_resolution,[],[f250,f255,f1402,f372])).
fof(f4091,plain,(
  ~sP4('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl170),
  inference(unit_resulting_resolution,[],[f249,f255,f1402,f372])).
fof(f4636,plain,(
  ~sP4('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl305),
  inference(cnf_transformation,[],[f4636_D])).
fof(f4636_D,plain,(
  ~sP4('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') <=> ~$spl305),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl305])])).
fof(f4094,plain,(
  ~sP4('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl170),
  inference(unit_resulting_resolution,[],[f1402,f255,f1402,f372])).
fof(f4638,plain,(
  ~sP4('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl307),
  inference(cnf_transformation,[],[f4638_D])).
fof(f4638_D,plain,(
  ~sP4('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') <=> ~$spl307),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl307])])).
fof(f4097,plain,(
  ~sP4('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl170),
  inference(unit_resulting_resolution,[],[f437,f255,f1402,f372])).
fof(f4098,plain,(
  ~sP4('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl170),
  inference(unit_resulting_resolution,[],[f250,f255,f1402,f372])).
fof(f4099,plain,(
  ~sP4('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl170),
  inference(unit_resulting_resolution,[],[f249,f255,f1402,f372])).
fof(f4102,plain,(
  ~sP4('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl170),
  inference(unit_resulting_resolution,[],[f1402,f255,f1402,f372])).
fof(f4632,plain,(
  sP6('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl302),
  inference(cnf_transformation,[],[f4632_D])).
fof(f4632_D,plain,(
  sP6('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) <=> ~$spl302),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl302])])).
fof(f4109,plain,(
  sP6('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | ($spl80 | $spl170)),
  inference(unit_resulting_resolution,[],[f903,f1402,f381])).
fof(f4630,plain,(
  sP6('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl300),
  inference(cnf_transformation,[],[f4630_D])).
fof(f4630_D,plain,(
  sP6('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl300),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl300])])).
fof(f4110,plain,(
  sP6('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | ($spl80 | $spl170)),
  inference(unit_resulting_resolution,[],[f903,f1402,f381])).
fof(f4111,plain,(
  ~sP6('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','F_0') | ($spl83 | $spl170)),
  inference(unit_resulting_resolution,[],[f908,f1402,f382])).
fof(f4626,plain,(
  ~sP6('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','T_0') | $spl297),
  inference(cnf_transformation,[],[f4626_D])).
fof(f4626_D,plain,(
  ~sP6('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','T_0') <=> ~$spl297),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl297])])).
fof(f4112,plain,(
  ~sP6('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','T_0') | ($spl83 | $spl170)),
  inference(unit_resulting_resolution,[],[f908,f1402,f382])).
fof(f4522,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl240),
  inference(cnf_transformation,[],[f4522_D])).
fof(f4522_D,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl240),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl240])])).
fof(f4115,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | ($spl4 | $spl170)),
  inference(unit_resulting_resolution,[],[f437,f593,f1402,f387])).
fof(f4116,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | ($spl4 | $spl170)),
  inference(unit_resulting_resolution,[],[f250,f593,f1402,f387])).
fof(f4117,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | ($spl4 | $spl170)),
  inference(unit_resulting_resolution,[],[f249,f593,f1402,f387])).
fof(f4518,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl236),
  inference(cnf_transformation,[],[f4518_D])).
fof(f4518_D,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl236),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl236])])).
fof(f4120,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl4 | $spl170)),
  inference(unit_resulting_resolution,[],[f1402,f593,f1402,f387])).
fof(f4624,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl294),
  inference(cnf_transformation,[],[f4624_D])).
fof(f4624_D,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl294),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl294])])).
fof(f4122,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | ($spl2 | $spl170)),
  inference(unit_resulting_resolution,[],[f437,f1086,f1402,f387])).
fof(f4123,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | ($spl2 | $spl170)),
  inference(unit_resulting_resolution,[],[f250,f1086,f1402,f387])).
fof(f4124,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | ($spl2 | $spl170)),
  inference(unit_resulting_resolution,[],[f249,f1086,f1402,f387])).
fof(f4620,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl290),
  inference(cnf_transformation,[],[f4620_D])).
fof(f4620_D,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl290),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl290])])).
fof(f4127,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl2 | $spl170)),
  inference(unit_resulting_resolution,[],[f1402,f1086,f1402,f387])).
fof(f4520,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl238),
  inference(cnf_transformation,[],[f4520_D])).
fof(f4520_D,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl238),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl238])])).
fof(f4129,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl4 | $spl170)),
  inference(unit_resulting_resolution,[],[f437,f593,f1402,f387])).
fof(f4130,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl4 | $spl170)),
  inference(unit_resulting_resolution,[],[f250,f593,f1402,f387])).
fof(f4131,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl4 | $spl170)),
  inference(unit_resulting_resolution,[],[f249,f593,f1402,f387])).
fof(f4134,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl4 | $spl170)),
  inference(unit_resulting_resolution,[],[f1402,f593,f1402,f387])).
fof(f4622,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl292),
  inference(cnf_transformation,[],[f4622_D])).
fof(f4622_D,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl292),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl292])])).
fof(f4136,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl2 | $spl170)),
  inference(unit_resulting_resolution,[],[f437,f1086,f1402,f387])).
fof(f4137,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl2 | $spl170)),
  inference(unit_resulting_resolution,[],[f250,f1086,f1402,f387])).
fof(f4138,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl2 | $spl170)),
  inference(unit_resulting_resolution,[],[f249,f1086,f1402,f387])).
fof(f4141,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl2 | $spl170)),
  inference(unit_resulting_resolution,[],[f1402,f1086,f1402,f387])).
fof(f4614,plain,(
  ~sP7('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl289),
  inference(cnf_transformation,[],[f4614_D])).
fof(f4614_D,plain,(
  ~sP7('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','F_0') <=> ~$spl289),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl289])])).
fof(f4147,plain,(
  ~sP7('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl170),
  inference(unit_resulting_resolution,[],[f437,f255,f1402,f391])).
fof(f4148,plain,(
  ~sP7('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl170),
  inference(unit_resulting_resolution,[],[f250,f255,f1402,f391])).
fof(f4149,plain,(
  ~sP7('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl170),
  inference(unit_resulting_resolution,[],[f249,f255,f1402,f391])).
fof(f4608,plain,(
  ~sP7('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl285),
  inference(cnf_transformation,[],[f4608_D])).
fof(f4608_D,plain,(
  ~sP7('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') <=> ~$spl285),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl285])])).
fof(f4152,plain,(
  ~sP7('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl170),
  inference(unit_resulting_resolution,[],[f1402,f255,f1402,f391])).
fof(f4610,plain,(
  ~sP7('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl287),
  inference(cnf_transformation,[],[f4610_D])).
fof(f4610_D,plain,(
  ~sP7('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') <=> ~$spl287),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl287])])).
fof(f4155,plain,(
  ~sP7('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl170),
  inference(unit_resulting_resolution,[],[f437,f255,f1402,f391])).
fof(f4156,plain,(
  ~sP7('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl170),
  inference(unit_resulting_resolution,[],[f250,f255,f1402,f391])).
fof(f4157,plain,(
  ~sP7('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl170),
  inference(unit_resulting_resolution,[],[f249,f255,f1402,f391])).
fof(f4160,plain,(
  ~sP7('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl170),
  inference(unit_resulting_resolution,[],[f1402,f255,f1402,f391])).
fof(f4604,plain,(
  sP8('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl282),
  inference(cnf_transformation,[],[f4604_D])).
fof(f4604_D,plain,(
  sP8('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl282),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl282])])).
fof(f4164,plain,(
  sP8('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl170),
  inference(unit_resulting_resolution,[],[f437,f437,f1402,f397])).
fof(f4165,plain,(
  sP8('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl170),
  inference(unit_resulting_resolution,[],[f250,f437,f1402,f397])).
fof(f4166,plain,(
  sP8('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl170),
  inference(unit_resulting_resolution,[],[f249,f437,f1402,f397])).
fof(f4596,plain,(
  sP8('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl274),
  inference(cnf_transformation,[],[f4596_D])).
fof(f4596_D,plain,(
  sP8('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl274),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl274])])).
fof(f4169,plain,(
  sP8('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl170),
  inference(unit_resulting_resolution,[],[f1402,f437,f1402,f397])).
fof(f4170,plain,(
  sP8('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl170),
  inference(unit_resulting_resolution,[],[f437,f250,f1402,f397])).
fof(f4171,plain,(
  sP8('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl170),
  inference(unit_resulting_resolution,[],[f250,f250,f1402,f397])).
fof(f4172,plain,(
  sP8('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl170),
  inference(unit_resulting_resolution,[],[f249,f250,f1402,f397])).
fof(f4175,plain,(
  sP8('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl170),
  inference(unit_resulting_resolution,[],[f1402,f250,f1402,f397])).
fof(f4176,plain,(
  sP8('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl170),
  inference(unit_resulting_resolution,[],[f437,f249,f1402,f397])).
fof(f4177,plain,(
  sP8('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl170),
  inference(unit_resulting_resolution,[],[f250,f249,f1402,f397])).
fof(f4178,plain,(
  sP8('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl170),
  inference(unit_resulting_resolution,[],[f249,f249,f1402,f397])).
fof(f4181,plain,(
  sP8('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl170),
  inference(unit_resulting_resolution,[],[f1402,f249,f1402,f397])).
fof(f4600,plain,(
  sP8('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl278),
  inference(cnf_transformation,[],[f4600_D])).
fof(f4600_D,plain,(
  sP8('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl278),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl278])])).
fof(f4194,plain,(
  sP8('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl170),
  inference(unit_resulting_resolution,[],[f437,f1402,f1402,f397])).
fof(f4195,plain,(
  sP8('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl170),
  inference(unit_resulting_resolution,[],[f250,f1402,f1402,f397])).
fof(f4196,plain,(
  sP8('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl170),
  inference(unit_resulting_resolution,[],[f249,f1402,f1402,f397])).
fof(f4592,plain,(
  sP8('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl270),
  inference(cnf_transformation,[],[f4592_D])).
fof(f4592_D,plain,(
  sP8('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl270),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl270])])).
fof(f4199,plain,(
  sP8('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl170),
  inference(unit_resulting_resolution,[],[f1402,f1402,f1402,f397])).
fof(f4602,plain,(
  sP8('T_0','T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl280),
  inference(cnf_transformation,[],[f4602_D])).
fof(f4602_D,plain,(
  sP8('T_0','T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl280),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl280])])).
fof(f4200,plain,(
  sP8('T_0','T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl170),
  inference(unit_resulting_resolution,[],[f437,f437,f1402,f397])).
fof(f4201,plain,(
  sP8('T_0','T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl170),
  inference(unit_resulting_resolution,[],[f250,f437,f1402,f397])).
fof(f4202,plain,(
  sP8('T_0','T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl170),
  inference(unit_resulting_resolution,[],[f249,f437,f1402,f397])).
fof(f4205,plain,(
  sP8('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl170),
  inference(unit_resulting_resolution,[],[f1402,f437,f1402,f397])).
fof(f4206,plain,(
  sP8('T_0','T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl170),
  inference(unit_resulting_resolution,[],[f437,f250,f1402,f397])).
fof(f4207,plain,(
  sP8('T_0','T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl170),
  inference(unit_resulting_resolution,[],[f250,f250,f1402,f397])).
fof(f4208,plain,(
  sP8('T_0','T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl170),
  inference(unit_resulting_resolution,[],[f249,f250,f1402,f397])).
fof(f4211,plain,(
  sP8('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl170),
  inference(unit_resulting_resolution,[],[f1402,f250,f1402,f397])).
fof(f4212,plain,(
  sP8('T_0','T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl170),
  inference(unit_resulting_resolution,[],[f437,f249,f1402,f397])).
fof(f4213,plain,(
  sP8('T_0','T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl170),
  inference(unit_resulting_resolution,[],[f250,f249,f1402,f397])).
fof(f4214,plain,(
  sP8('T_0','T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl170),
  inference(unit_resulting_resolution,[],[f249,f249,f1402,f397])).
fof(f4217,plain,(
  sP8('T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl170),
  inference(unit_resulting_resolution,[],[f1402,f249,f1402,f397])).
fof(f4594,plain,(
  sP8('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl272),
  inference(cnf_transformation,[],[f4594_D])).
fof(f4594_D,plain,(
  sP8('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl272),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl272])])).
fof(f4230,plain,(
  sP8('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl170),
  inference(unit_resulting_resolution,[],[f437,f1402,f1402,f397])).
fof(f4231,plain,(
  sP8('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl170),
  inference(unit_resulting_resolution,[],[f250,f1402,f1402,f397])).
fof(f4232,plain,(
  sP8('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl170),
  inference(unit_resulting_resolution,[],[f249,f1402,f1402,f397])).
fof(f4235,plain,(
  sP8('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl170),
  inference(unit_resulting_resolution,[],[f1402,f1402,f1402,f397])).
fof(f4598,plain,(
  sP8('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','T_0') | $spl276),
  inference(cnf_transformation,[],[f4598_D])).
fof(f4598_D,plain,(
  sP8('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','T_0') <=> ~$spl276),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl276])])).
fof(f4236,plain,(
  sP8('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','T_0') | $spl170),
  inference(unit_resulting_resolution,[],[f437,f437,f1402,f397])).
fof(f4237,plain,(
  sP8('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','T_0') | $spl170),
  inference(unit_resulting_resolution,[],[f250,f437,f1402,f397])).
fof(f4238,plain,(
  sP8('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','T_0') | $spl170),
  inference(unit_resulting_resolution,[],[f249,f437,f1402,f397])).
fof(f4241,plain,(
  sP8('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl170),
  inference(unit_resulting_resolution,[],[f1402,f437,f1402,f397])).
fof(f4242,plain,(
  sP8('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','T_0') | $spl170),
  inference(unit_resulting_resolution,[],[f437,f250,f1402,f397])).
fof(f4243,plain,(
  sP8('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','T_0') | $spl170),
  inference(unit_resulting_resolution,[],[f250,f250,f1402,f397])).
fof(f4244,plain,(
  sP8('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','T_0') | $spl170),
  inference(unit_resulting_resolution,[],[f249,f250,f1402,f397])).
fof(f4247,plain,(
  sP8('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl170),
  inference(unit_resulting_resolution,[],[f1402,f250,f1402,f397])).
fof(f4248,plain,(
  sP8('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','T_0') | $spl170),
  inference(unit_resulting_resolution,[],[f437,f249,f1402,f397])).
fof(f4249,plain,(
  sP8('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','T_0') | $spl170),
  inference(unit_resulting_resolution,[],[f250,f249,f1402,f397])).
fof(f4250,plain,(
  sP8('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','T_0') | $spl170),
  inference(unit_resulting_resolution,[],[f249,f249,f1402,f397])).
fof(f4253,plain,(
  sP8('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl170),
  inference(unit_resulting_resolution,[],[f1402,f249,f1402,f397])).
fof(f4266,plain,(
  sP8('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl170),
  inference(unit_resulting_resolution,[],[f437,f1402,f1402,f397])).
fof(f4267,plain,(
  sP8('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl170),
  inference(unit_resulting_resolution,[],[f250,f1402,f1402,f397])).
fof(f4268,plain,(
  sP8('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl170),
  inference(unit_resulting_resolution,[],[f249,f1402,f1402,f397])).
fof(f4271,plain,(
  sP8('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl170),
  inference(unit_resulting_resolution,[],[f1402,f1402,f1402,f397])).
fof(f4590,plain,(
  ~sP8('F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl269),
  inference(cnf_transformation,[],[f4590_D])).
fof(f4590_D,plain,(
  ~sP8('F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl269),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl269])])).
fof(f4274,plain,(
  ~sP8('F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl170),
  inference(unit_resulting_resolution,[],[f255,f437,f1402,f400])).
fof(f4275,plain,(
  ~sP8('F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl170),
  inference(unit_resulting_resolution,[],[f255,f250,f1402,f400])).
fof(f4276,plain,(
  ~sP8('F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl170),
  inference(unit_resulting_resolution,[],[f255,f249,f1402,f400])).
fof(f4586,plain,(
  ~sP8('F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl265),
  inference(cnf_transformation,[],[f4586_D])).
fof(f4586_D,plain,(
  ~sP8('F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl265),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl265])])).
fof(f4279,plain,(
  ~sP8('F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl170),
  inference(unit_resulting_resolution,[],[f255,f1402,f1402,f400])).
fof(f4588,plain,(
  ~sP8('F_0','T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl267),
  inference(cnf_transformation,[],[f4588_D])).
fof(f4588_D,plain,(
  ~sP8('F_0','T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl267),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl267])])).
fof(f4280,plain,(
  ~sP8('F_0','T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl170),
  inference(unit_resulting_resolution,[],[f255,f437,f1402,f400])).
fof(f4281,plain,(
  ~sP8('F_0','T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl170),
  inference(unit_resulting_resolution,[],[f255,f250,f1402,f400])).
fof(f4282,plain,(
  ~sP8('F_0','T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl170),
  inference(unit_resulting_resolution,[],[f255,f249,f1402,f400])).
fof(f4285,plain,(
  ~sP8('F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl170),
  inference(unit_resulting_resolution,[],[f255,f1402,f1402,f400])).
fof(f4574,plain,(
  ~sP9('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','F_0') | $spl263),
  inference(cnf_transformation,[],[f4574_D])).
fof(f4574_D,plain,(
  ~sP9('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','F_0') <=> ~$spl263),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl263])])).
fof(f4310,plain,(
  ~sP9('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','F_0') | $spl170),
  inference(unit_resulting_resolution,[],[f255,f255,f1402,f408])).
fof(f4568,plain,(
  ~sP10('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','F_0') | $spl261),
  inference(cnf_transformation,[],[f4568_D])).
fof(f4568_D,plain,(
  ~sP10('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','F_0') <=> ~$spl261),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl261])])).
fof(f4345,plain,(
  ~sP10('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','F_0') | $spl170),
  inference(unit_resulting_resolution,[],[f255,f255,f1402,f415])).
fof(f4552,plain,(
  ~sP11('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','T_0') | $spl259),
  inference(cnf_transformation,[],[f4552_D])).
fof(f4552_D,plain,(
  ~sP11('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','T_0') <=> ~$spl259),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl259])])).
fof(f4397,plain,(
  ~sP11('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','T_0') | $spl170),
  inference(unit_resulting_resolution,[],[f255,f437,f1402,f425])).
fof(f4398,plain,(
  ~sP11('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','T_0') | $spl170),
  inference(unit_resulting_resolution,[],[f255,f250,f1402,f425])).
fof(f4399,plain,(
  ~sP11('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','T_0') | $spl170),
  inference(unit_resulting_resolution,[],[f255,f249,f1402,f425])).
fof(f4546,plain,(
  ~sP11('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl255),
  inference(cnf_transformation,[],[f4546_D])).
fof(f4546_D,plain,(
  ~sP11('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl255),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl255])])).
fof(f4402,plain,(
  ~sP11('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl170),
  inference(unit_resulting_resolution,[],[f255,f1402,f1402,f425])).
fof(f4548,plain,(
  ~sP11('T_0','F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl257),
  inference(cnf_transformation,[],[f4548_D])).
fof(f4548_D,plain,(
  ~sP11('T_0','F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl257),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl257])])).
fof(f4405,plain,(
  ~sP11('T_0','F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl170),
  inference(unit_resulting_resolution,[],[f437,f255,f1402,f425])).
fof(f4406,plain,(
  ~sP11('T_0','F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl170),
  inference(unit_resulting_resolution,[],[f250,f255,f1402,f425])).
fof(f4407,plain,(
  ~sP11('T_0','F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl170),
  inference(unit_resulting_resolution,[],[f249,f255,f1402,f425])).
fof(f4410,plain,(
  ~sP11('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl170),
  inference(unit_resulting_resolution,[],[f1402,f255,f1402,f425])).
fof(f4542,plain,(
  sP12('F_0','F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl252),
  inference(cnf_transformation,[],[f4542_D])).
fof(f4542_D,plain,(
  sP12('F_0','F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl252),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl252])])).
fof(f4414,plain,(
  sP12('F_0','F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl170),
  inference(unit_resulting_resolution,[],[f255,f255,f1402,f429])).
fof(f4538,plain,(
  ~sP12('T_0','F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl251),
  inference(cnf_transformation,[],[f4538_D])).
fof(f4538_D,plain,(
  ~sP12('T_0','F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl251),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl251])])).
fof(f4430,plain,(
  ~sP12('T_0','F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl170),
  inference(unit_resulting_resolution,[],[f255,f437,f1402,f432])).
fof(f4431,plain,(
  ~sP12('T_0','F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl170),
  inference(unit_resulting_resolution,[],[f255,f250,f1402,f432])).
fof(f4432,plain,(
  ~sP12('T_0','F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl170),
  inference(unit_resulting_resolution,[],[f255,f249,f1402,f432])).
fof(f4530,plain,(
  ~sP12('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl247),
  inference(cnf_transformation,[],[f4530_D])).
fof(f4530_D,plain,(
  ~sP12('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl247),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl247])])).
fof(f4435,plain,(
  ~sP12('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl170),
  inference(unit_resulting_resolution,[],[f255,f1402,f1402,f432])).
fof(f4532,plain,(
  ~sP12('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','T_0') | $spl249),
  inference(cnf_transformation,[],[f4532_D])).
fof(f4532_D,plain,(
  ~sP12('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','T_0') <=> ~$spl249),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl249])])).
fof(f4440,plain,(
  ~sP12('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','T_0') | $spl170),
  inference(unit_resulting_resolution,[],[f255,f437,f1402,f432])).
fof(f4441,plain,(
  ~sP12('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','T_0') | $spl170),
  inference(unit_resulting_resolution,[],[f255,f250,f1402,f432])).
fof(f4442,plain,(
  ~sP12('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','T_0') | $spl170),
  inference(unit_resulting_resolution,[],[f255,f249,f1402,f432])).
fof(f4445,plain,(
  ~sP12('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl170),
  inference(unit_resulting_resolution,[],[f255,f1402,f1402,f432])).
fof(f4526,plain,(
  ~sP1('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl245),
  inference(cnf_transformation,[],[f4526_D])).
fof(f4526_D,plain,(
  ~sP1('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) <=> ~$spl245),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl245])])).
fof(f4456,plain,(
  ~sP1('const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | ($spl4 | $spl170)),
  inference(unit_resulting_resolution,[],[f1402,f800])).
fof(f4524,plain,(
  ~sP1('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl243),
  inference(cnf_transformation,[],[f4524_D])).
fof(f4524_D,plain,(
  ~sP1('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl243),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl243])])).
fof(f4457,plain,(
  ~sP1('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl4 | $spl170)),
  inference(unit_resulting_resolution,[],[f1402,f805])).
fof(f4460,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | ($spl4 | $spl170)),
  inference(unit_resulting_resolution,[],[f437,f1402,f817])).
fof(f4461,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | ($spl4 | $spl170)),
  inference(unit_resulting_resolution,[],[f250,f1402,f817])).
fof(f4462,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | ($spl4 | $spl170)),
  inference(unit_resulting_resolution,[],[f249,f1402,f817])).
fof(f4465,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl4 | $spl170)),
  inference(unit_resulting_resolution,[],[f1402,f1402,f817])).
fof(f4467,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl4 | $spl170)),
  inference(unit_resulting_resolution,[],[f437,f1402,f817])).
fof(f4468,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl4 | $spl170)),
  inference(unit_resulting_resolution,[],[f250,f1402,f817])).
fof(f4469,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl4 | $spl170)),
  inference(unit_resulting_resolution,[],[f249,f1402,f817])).
fof(f4472,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ($spl4 | $spl170)),
  inference(unit_resulting_resolution,[],[f1402,f1402,f817])).
fof(f3951,plain,(
  ~sP6('T_0','F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl235),
  inference(cnf_transformation,[],[f3951_D])).
fof(f3951_D,plain,(
  ~sP6('T_0','F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl235),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl235])])).
fof(f3947,plain,(
  ~sP6('T_0','F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl149),
  inference(unit_resulting_resolution,[],[f1380,f363])).
fof(f3941,plain,(
  ~sP6('F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl233),
  inference(cnf_transformation,[],[f3941_D])).
fof(f3941_D,plain,(
  ~sP6('F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl233),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl233])])).
fof(f3937,plain,(
  ~sP6('F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl145),
  inference(unit_resulting_resolution,[],[f1376,f362])).
fof(f3931,plain,(
  ~sP6('F_0','T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl231),
  inference(cnf_transformation,[],[f3931_D])).
fof(f3931_D,plain,(
  ~sP6('F_0','T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl231),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl231])])).
fof(f3927,plain,(
  ~sP6('F_0','T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl143),
  inference(unit_resulting_resolution,[],[f1374,f362])).
fof(f3900,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','T_0') | $spl228),
  inference(cnf_transformation,[],[f3900_D])).
fof(f3900_D,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','T_0') <=> ~$spl228),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl228])])).
fof(f3895,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','T_0') | $spl128),
  inference(resolution,[],[f1360,f384])).
fof(f3889,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','T_0') | $spl128),
  inference(unit_resulting_resolution,[],[f1360,f384])).
fof(f3883,plain,(
  sP7('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl226),
  inference(cnf_transformation,[],[f3883_D])).
fof(f3883_D,plain,(
  sP7('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl226),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl226])])).
fof(f3878,plain,(
  sP7('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl126),
  inference(resolution,[],[f1358,f384])).
fof(f3872,plain,(
  sP7('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl126),
  inference(unit_resulting_resolution,[],[f1358,f384])).
fof(f3829,plain,(
  sP4('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','T_0') | $spl224),
  inference(cnf_transformation,[],[f3829_D])).
fof(f3829_D,plain,(
  sP4('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','T_0') <=> ~$spl224),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl224])])).
fof(f3825,plain,(
  sP4('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','T_0') | $spl204),
  inference(resolution,[],[f2201,f362])).
fof(f3821,plain,(
  sP4('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','T_0') | $spl204),
  inference(unit_resulting_resolution,[],[f2201,f362])).
fof(f3797,plain,(
  sP11('F_0','F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl222),
  inference(cnf_transformation,[],[f3797_D])).
fof(f3797_D,plain,(
  sP11('F_0','F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl222),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl222])])).
fof(f3795,plain,(
  sP11('F_0','F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl98),
  inference(resolution,[],[f1330,f420])).
fof(f3789,plain,(
  sP11('F_0','F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl98),
  inference(unit_resulting_resolution,[],[f1330,f420])).
fof(f3727,plain,(
  sP5('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl218),
  inference(cnf_transformation,[],[f3727_D])).
fof(f3727_D,plain,(
  sP5('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl218),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl218])])).
fof(f3722,plain,(
  sP5('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl202),
  inference(resolution,[],[f2199,f363])).
fof(f3729,plain,(
  sP4('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl220),
  inference(cnf_transformation,[],[f3729_D])).
fof(f3729_D,plain,(
  sP4('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl220),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl220])])).
fof(f3721,plain,(
  sP4('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl202),
  inference(resolution,[],[f2199,f362])).
fof(f3716,plain,(
  sP4('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl202),
  inference(unit_resulting_resolution,[],[f2199,f362])).
fof(f3717,plain,(
  sP5('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl202),
  inference(unit_resulting_resolution,[],[f2199,f363])).
fof(f3415,plain,(
  sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl216),
  inference(cnf_transformation,[],[f3415_D])).
fof(f3415_D,plain,(
  sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl216),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl216])])).
fof(f3408,plain,(
  sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f437,f1053,f1066,f1067,f365])).
fof(f3409,plain,(
  sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f250,f1053,f1066,f1067,f365])).
fof(f3410,plain,(
  sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f249,f1053,f1066,f1067,f365])).
fof(f3395,plain,(
  sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl214),
  inference(cnf_transformation,[],[f3395_D])).
fof(f3395_D,plain,(
  sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl214),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl214])])).
fof(f3388,plain,(
  sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | ($spl2 | $spl4)),
  inference(unit_resulting_resolution,[],[f437,f577,f562,f1066,f365])).
fof(f3389,plain,(
  sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | ($spl2 | $spl4)),
  inference(unit_resulting_resolution,[],[f250,f577,f562,f1066,f365])).
fof(f3390,plain,(
  sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | ($spl2 | $spl4)),
  inference(unit_resulting_resolution,[],[f249,f577,f562,f1066,f365])).
fof(f3374,plain,(
  sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl212),
  inference(cnf_transformation,[],[f3374_D])).
fof(f3374_D,plain,(
  sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl212),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl212])])).
fof(f3364,plain,(
  sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | ($spl2 | $spl4)),
  inference(unit_resulting_resolution,[],[f437,f576,f575,f1053,f365])).
fof(f3365,plain,(
  sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | ($spl2 | $spl4)),
  inference(unit_resulting_resolution,[],[f250,f576,f575,f1053,f365])).
fof(f3366,plain,(
  sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | ($spl2 | $spl4)),
  inference(unit_resulting_resolution,[],[f249,f576,f575,f1053,f365])).
fof(f2460,plain,(
  ~sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl211),
  inference(cnf_transformation,[],[f2460_D])).
fof(f2460_D,plain,(
  ~sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) <=> ~$spl211),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl211])])).
fof(f2456,plain,(
  ~sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl57),
  inference(unit_resulting_resolution,[],[f879,f363])).
fof(f2448,plain,(
  ~sP6('F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl209),
  inference(cnf_transformation,[],[f2448_D])).
fof(f2448_D,plain,(
  ~sP6('F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) <=> ~$spl209),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl209])])).
fof(f2444,plain,(
  ~sP6('F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl51),
  inference(unit_resulting_resolution,[],[f873,f362])).
fof(f2424,plain,(
  sP7('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl206),
  inference(cnf_transformation,[],[f2424_D])).
fof(f2424_D,plain,(
  sP7('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f2417,plain,(
  sP7('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl38),
  inference(resolution,[],[f861,f384])).
fof(f2411,plain,(
  sP7('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl38),
  inference(unit_resulting_resolution,[],[f861,f384])).
fof(f2201,plain,(
  sP6('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl204),
  inference(cnf_transformation,[],[f2201_D])).
fof(f2201_D,plain,(
  sP6('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f2112,plain,(
  sP6('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl80),
  inference(unit_resulting_resolution,[],[f437,f903,f381])).
fof(f2113,plain,(
  sP6('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl80),
  inference(unit_resulting_resolution,[],[f250,f903,f381])).
fof(f2114,plain,(
  sP6('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl80),
  inference(unit_resulting_resolution,[],[f249,f903,f381])).
fof(f2199,plain,(
  sP6('T_0','T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl202),
  inference(cnf_transformation,[],[f2199_D])).
fof(f2199_D,plain,(
  sP6('T_0','T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f2118,plain,(
  sP6('T_0','T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl80),
  inference(unit_resulting_resolution,[],[f437,f903,f381])).
fof(f2119,plain,(
  sP6('T_0','T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl80),
  inference(unit_resulting_resolution,[],[f250,f903,f381])).
fof(f2120,plain,(
  sP6('T_0','T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl80),
  inference(unit_resulting_resolution,[],[f249,f903,f381])).
fof(f2060,plain,(
  ~sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0','T_0') | $spl201),
  inference(cnf_transformation,[],[f2060_D])).
fof(f2060_D,plain,(
  ~sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0','T_0') <=> ~$spl201),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl201])])).
fof(f2056,plain,(
  ~sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0','T_0') | $spl61),
  inference(unit_resulting_resolution,[],[f883,f363])).
fof(f2050,plain,(
  ~sP6('T_0','F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl199),
  inference(cnf_transformation,[],[f2050_D])).
fof(f2050_D,plain,(
  ~sP6('T_0','F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) <=> ~$spl199),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl199])])).
fof(f2046,plain,(
  ~sP6('T_0','F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl59),
  inference(unit_resulting_resolution,[],[f881,f363])).
fof(f2040,plain,(
  ~sP6('F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl197),
  inference(cnf_transformation,[],[f2040_D])).
fof(f2040_D,plain,(
  ~sP6('F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl197),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl197])])).
fof(f2036,plain,(
  ~sP6('F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl55),
  inference(unit_resulting_resolution,[],[f877,f362])).
fof(f2030,plain,(
  ~sP6('F_0','T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl195),
  inference(cnf_transformation,[],[f2030_D])).
fof(f2030_D,plain,(
  ~sP6('F_0','T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) <=> ~$spl195),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl195])])).
fof(f2026,plain,(
  ~sP6('F_0','T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl53),
  inference(unit_resulting_resolution,[],[f875,f362])).
fof(f2001,plain,(
  sP7('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','T_0') | $spl192),
  inference(cnf_transformation,[],[f2001_D])).
fof(f2001_D,plain,(
  sP7('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','T_0') <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f1996,plain,(
  sP7('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','T_0') | $spl42),
  inference(resolution,[],[f865,f384])).
fof(f1990,plain,(
  sP7('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','T_0') | $spl42),
  inference(unit_resulting_resolution,[],[f865,f384])).
fof(f1984,plain,(
  sP7('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl190),
  inference(cnf_transformation,[],[f1984_D])).
fof(f1984_D,plain,(
  sP7('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f1979,plain,(
  sP7('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl40),
  inference(resolution,[],[f863,f384])).
fof(f1973,plain,(
  sP7('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl40),
  inference(unit_resulting_resolution,[],[f863,f384])).
fof(f1905,plain,(
  sP11('F_0','F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl188),
  inference(cnf_transformation,[],[f1905_D])).
fof(f1905_D,plain,(
  sP11('F_0','F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f1903,plain,(
  sP11('F_0','F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl12),
  inference(resolution,[],[f835,f420])).
fof(f1897,plain,(
  sP11('F_0','F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl12),
  inference(unit_resulting_resolution,[],[f835,f420])).
fof(f1857,plain,(
  sP5('F_0','T_0','F_0') | $spl182),
  inference(cnf_transformation,[],[f1857_D])).
fof(f1857_D,plain,(
  sP5('F_0','T_0','F_0') <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f1853,plain,(
  sP5('F_0','T_0','F_0') | $spl90),
  inference(resolution,[],[f936,f363])).
fof(f1859,plain,(
  sP4('F_0','T_0','F_0') | $spl184),
  inference(cnf_transformation,[],[f1859_D])).
fof(f1859_D,plain,(
  sP4('F_0','T_0','F_0') <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f1852,plain,(
  sP4('F_0','T_0','F_0') | $spl90),
  inference(resolution,[],[f936,f362])).
fof(f1861,plain,(
  sP3('T_0','F_0','F_0') | $spl186),
  inference(cnf_transformation,[],[f1861_D])).
fof(f1861_D,plain,(
  sP3('T_0','F_0','F_0') <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f1851,plain,(
  sP3('T_0','F_0','F_0') | $spl90),
  inference(resolution,[],[f936,f361])).
fof(f361,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | sP3(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f213])).
fof(f1847,plain,(
  sP3('T_0','F_0','F_0') | $spl90),
  inference(unit_resulting_resolution,[],[f936,f361])).
fof(f1848,plain,(
  sP4('F_0','T_0','F_0') | $spl90),
  inference(unit_resulting_resolution,[],[f936,f362])).
fof(f1849,plain,(
  sP5('F_0','T_0','F_0') | $spl90),
  inference(unit_resulting_resolution,[],[f936,f363])).
fof(f1548,plain,(
  sP5('F_0','F_0','T_0') | $spl176),
  inference(cnf_transformation,[],[f1548_D])).
fof(f1548_D,plain,(
  sP5('F_0','F_0','T_0') <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f1545,plain,(
  sP5('F_0','F_0','T_0') | $spl86),
  inference(resolution,[],[f932,f363])).
fof(f1550,plain,(
  sP4('T_0','F_0','F_0') | $spl178),
  inference(cnf_transformation,[],[f1550_D])).
fof(f1550_D,plain,(
  sP4('T_0','F_0','F_0') <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f1544,plain,(
  sP4('T_0','F_0','F_0') | $spl86),
  inference(resolution,[],[f932,f362])).
fof(f1552,plain,(
  sP3('F_0','T_0','F_0') | $spl180),
  inference(cnf_transformation,[],[f1552_D])).
fof(f1552_D,plain,(
  sP3('F_0','T_0','F_0') <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f1543,plain,(
  sP3('F_0','T_0','F_0') | $spl86),
  inference(resolution,[],[f932,f361])).
fof(f1538,plain,(
  sP3('F_0','T_0','F_0') | $spl86),
  inference(unit_resulting_resolution,[],[f932,f361])).
fof(f1539,plain,(
  sP4('T_0','F_0','F_0') | $spl86),
  inference(unit_resulting_resolution,[],[f932,f362])).
fof(f1540,plain,(
  sP5('F_0','F_0','T_0') | $spl86),
  inference(unit_resulting_resolution,[],[f932,f363])).
fof(f1524,plain,(
  sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl174),
  inference(cnf_transformation,[],[f1524_D])).
fof(f1524_D,plain,(
  sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f1517,plain,(
  sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f437,f562,f575,f576,f365])).
fof(f1518,plain,(
  sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f250,f562,f575,f576,f365])).
fof(f1519,plain,(
  sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f249,f562,f575,f576,f365])).
fof(f1400,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl168),
  inference(cnf_transformation,[],[f1400_D])).
fof(f1400_D,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f1405,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl2),
  inference(global_subsumption,[],[f248,f247,f246,f445,f1282,f1281,f1280,f1279,f1278,f1277,f1276,f1275,f1274,f1273,f1272,f1271,f1270,f1269,f1268,f1267,f1266,f1265,f1261,f1260,f1259,f1256,f1255,f1254,f1253,f1252,f1251,f1250,f1249,f1248,f1247,f1246,f1245,f1244,f1243,f1242,f1240,f1239,f1238,f1237,f1236,f1235,f1234,f1230,f1229,f1228,f1225,f1224,f1223,f1222,f1218,f1217,f1216,f1213,f1212,f1211,f1208,f1207,f1206,f1205,f1204,f1203,f1202,f1201,f1200,f1199,f1198,f1197,f1196,f1195,f1194,f1193,f1192,f1191,f1190,f1189,f1188,f1187,f1186,f1185,f1182,f1181,f1180,f1179,f1176,f1175,f1174,f1173,f1172,f1171,f1170,f1167,f1166,f1165,f1164,f1161,f1160,f1159,f1158,f1155,f1154,f1153,f1152,f1151,f1150,f1149,f1146,f1145,f1144,f1143,f1140,f1139,f1138,f1137,f1136,f1135,f1134,f1131,f1130,f1129,f1128,f1125,f1124,f1123,f1122,f1119,f1118,f1117,f1115,f1114,f1113,f1109,f1108,f1107,f1104,f1103,f1102,f1101,f1100,f1099,f1098,f1095,f1094,f1093,f1092,f1089,f1088,f1087,f1086,f1085,f1084,f1083,f1082,f1081,f1080,f1079,f1078,f1077,f1076,f1075,f1074,f1069,f1068,f1067,f1066,f1065,f1064,f1063,f1062,f1061,f1060,f1059,f1058,f1057,f1056,f1055,f1054,f1053,f1052,f1051,f1050,f1049,f1048,f1047,f1046,f1045,f1044,f1043,f1042,f1041,f1040,f1039,f1038,f1037,f1036,f1035,f1034,f1033,f1032,f1031,f1030,f1029,f1028,f1027,f1026,f1025,f1021,f1020,f1019,f996,f995,f994,f990,f989,f988,f979,f978,f977,f976,f975,f974,f973,f972,f971,f970,f969,f968,f967,f966,f965,f964,f963,f960,f959,f958,f957,f954,f953,f952,f951,f950,f942,f941,f940,f939,f938,f937,f1283,f1284,f1285,f1286,f1295,f1296,f1297,f1300,f1303,f1305,f1312,f1315,f1316,f1318,f1319,f1320,f1322])).
fof(f1322,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl2),
  inference(superposition,[],[f445,f256])).
fof(f1320,plain,(
  ( ! [X48,X49] : (sP12('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X48,X49) | ~p(s(bool,X48))) ) | $spl2),
  inference(resolution,[],[f445,f431])).
fof(f1319,plain,(
  ( ! [X47,X46] : (sP12('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X47,X46) | p(s(bool,X46))) ) | $spl2),
  inference(resolution,[],[f445,f430])).
fof(f1318,plain,(
  ( ! [X45,X44] : (sP12(X45,X44,'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | p(s(bool,X44)) | p(s(bool,X45))) ) | $spl2),
  inference(resolution,[],[f445,f429])).
fof(f1316,plain,(
  ( ! [X41,X40] : (sP10('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X40,X41) | ~p(s(bool,X40))) ) | $spl2),
  inference(resolution,[],[f445,f414])).
fof(f1315,plain,(
  ( ! [X39,X38] : (sP10('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X39,X38) | ~p(s(bool,X38))) ) | $spl2),
  inference(resolution,[],[f445,f413])).
fof(f1305,plain,(
  ( ! [X19,X18] : (sP2('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X18,X19) | ~p(s(bool,X19))) ) | $spl2),
  inference(resolution,[],[f445,f352])).
fof(f1303,plain,(
  ( ! [X14,X15] : (sP2(X14,X15,'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | ~p(s(bool,X14))) ) | $spl2),
  inference(resolution,[],[f445,f350])).
fof(f1297,plain,(
  ( ! [X6] : (sP1(X6,'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | p(s(bool,X6))) ) | $spl2),
  inference(resolution,[],[f445,f332])).
fof(f1296,plain,(
  ( ! [X5] : (sP1('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X5) | p(s(bool,X5))) ) | $spl2),
  inference(resolution,[],[f445,f329])).
fof(f950,plain,(
  ( ! [X0] : (sP0(X0,'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f255,f445,f318])).
fof(f951,plain,(
  ( ! [X0] : (sP0('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f445,f319])).
fof(f971,plain,(
  ( ! [X0,X1] : (~sP1('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f304,f445,f335])).
fof(f972,plain,(
  ( ! [X0,X1] : (~sP1('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f305,f445,f335])).
fof(f977,plain,(
  ( ! [X0,X1] : (~sP1('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f304,f445,f335])).
fof(f978,plain,(
  ( ! [X0,X1] : (~sP1('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f305,f445,f335])).
fof(f988,plain,(
  ( ! [X0] : (sP2(X0,'T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f249,f445,f348])).
fof(f989,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,'const.pred_set.SUBSET_2'(s(fun(X1,bool),X2),s(fun(X1,bool),X2)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f304,f445,f348])).
fof(f990,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,'const.pred_set.SUBSET_2'(s(fun(X1,bool),X2),s(fun(X1,bool),'const.pred_set.UNIV_0')),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f305,f445,f348])).
fof(f994,plain,(
  ( ! [X0] : (sP2('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0',X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f249,f445,f349])).
fof(f995,plain,(
  ( ! [X2,X0,X1] : (sP2('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),X2)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f304,f445,f349])).
fof(f996,plain,(
  ( ! [X2,X0,X1] : (sP2('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),X2)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f305,f445,f349])).
fof(f1019,plain,(
  ( ! [X0] : (sP2('T_0',X0,'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f249,f445,f352])).
fof(f1020,plain,(
  ( ! [X2,X0,X1] : (sP2('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),X2,'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f304,f445,f352])).
fof(f1021,plain,(
  ( ! [X2,X0,X1] : (sP2('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),X2,'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f305,f445,f352])).
fof(f1025,plain,(
  ( ! [X0] : (sP2('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X0,'T_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f249,f445,f352])).
fof(f1026,plain,(
  ( ! [X2,X0,X1] : (sP2('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X0,'const.pred_set.SUBSET_2'(s(fun(X1,bool),X2),s(fun(X1,bool),X2)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f304,f445,f352])).
fof(f1027,plain,(
  ( ! [X2,X0,X1] : (sP2('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X0,'const.pred_set.SUBSET_2'(s(fun(X1,bool),X2),s(fun(X1,bool),'const.pred_set.UNIV_0')))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f305,f445,f352])).
fof(f1028,plain,(
  ( ! [X0] : (sP2('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X0,'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f445,f445,f352])).
fof(f1032,plain,(
  ( ! [X0,X1] : (~sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'F_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f304,f255,f445,f364])).
fof(f1033,plain,(
  ( ! [X0,X1] : (~sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'F_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f305,f255,f445,f364])).
fof(f1038,plain,(
  ( ! [X0,X1] : (~sP6('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f304,f255,f445,f364])).
fof(f1039,plain,(
  ( ! [X0,X1] : (~sP6('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f305,f255,f445,f364])).
fof(f1044,plain,(
  ( ! [X0,X1] : (~sP5('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'F_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f304,f255,f445,f368])).
fof(f1045,plain,(
  ( ! [X0,X1] : (~sP5('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'F_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f305,f255,f445,f368])).
fof(f1050,plain,(
  ( ! [X0,X1] : (~sP5('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f304,f255,f445,f368])).
fof(f1051,plain,(
  ( ! [X0,X1] : (~sP5('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f305,f255,f445,f368])).
fof(f1057,plain,(
  ( ! [X0,X1] : (~sP4('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'F_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f304,f255,f445,f372])).
fof(f1058,plain,(
  ( ! [X0,X1] : (~sP4('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'F_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f305,f255,f445,f372])).
fof(f1063,plain,(
  ( ! [X0,X1] : (~sP4('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f304,f255,f445,f372])).
fof(f1064,plain,(
  ( ! [X0,X1] : (~sP4('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f305,f255,f445,f372])).
fof(f1077,plain,(
  ( ! [X0,X1] : (~sP7('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'F_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f304,f255,f445,f391])).
fof(f1078,plain,(
  ( ! [X0,X1] : (~sP7('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'F_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f305,f255,f445,f391])).
fof(f1083,plain,(
  ( ! [X0,X1] : (~sP7('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f304,f255,f445,f391])).
fof(f1084,plain,(
  ( ! [X0,X1] : (~sP7('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f305,f255,f445,f391])).
fof(f1102,plain,(
  ( ! [X0,X1] : (sP8('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f304,f249,f445,f397])).
fof(f1103,plain,(
  ( ! [X0,X1] : (sP8('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f305,f249,f445,f397])).
fof(f1107,plain,(
  ( ! [X0,X1] : (sP8('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f249,f304,f445,f397])).
fof(f1108,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X2,bool),X3),s(fun(X2,bool),X3)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f304,f304,f445,f397])).
fof(f1109,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X2,bool),X3),s(fun(X2,bool),X3)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f305,f304,f445,f397])).
fof(f1113,plain,(
  ( ! [X0,X1] : (sP8('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f249,f305,f445,f397])).
fof(f1114,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X2,bool),X3),s(fun(X2,bool),'const.pred_set.UNIV_0')))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f304,f305,f445,f397])).
fof(f1115,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X2,bool),X3),s(fun(X2,bool),'const.pred_set.UNIV_0')))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f305,f305,f445,f397])).
fof(f1138,plain,(
  ( ! [X0,X1] : (sP8('T_0','const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f304,f249,f445,f397])).
fof(f1139,plain,(
  ( ! [X0,X1] : (sP8('T_0','const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f305,f249,f445,f397])).
fof(f1143,plain,(
  ( ! [X0,X1] : (sP8('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f249,f304,f445,f397])).
fof(f1144,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.SUBSET_2'(s(fun(X2,bool),X3),s(fun(X2,bool),X3)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f304,f304,f445,f397])).
fof(f1145,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.SUBSET_2'(s(fun(X2,bool),X3),s(fun(X2,bool),'const.pred_set.UNIV_0')),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f305,f304,f445,f397])).
fof(f1146,plain,(
  ( ! [X0,X1] : (sP8('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f445,f304,f445,f397])).
fof(f1149,plain,(
  ( ! [X0,X1] : (sP8('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f249,f305,f445,f397])).
fof(f1150,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.SUBSET_2'(s(fun(X2,bool),X3),s(fun(X2,bool),X3)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f304,f305,f445,f397])).
fof(f1151,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.SUBSET_2'(s(fun(X2,bool),X3),s(fun(X2,bool),'const.pred_set.UNIV_0')),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f305,f305,f445,f397])).
fof(f1152,plain,(
  ( ! [X0,X1] : (sP8('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f445,f305,f445,f397])).
fof(f1174,plain,(
  ( ! [X0,X1] : (sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'T_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f304,f249,f445,f397])).
fof(f1175,plain,(
  ( ! [X0,X1] : (sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'T_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f305,f249,f445,f397])).
fof(f1179,plain,(
  ( ! [X0,X1] : (sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f249,f304,f445,f397])).
fof(f1180,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.SUBSET_2'(s(fun(X2,bool),X3),s(fun(X2,bool),X3)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f304,f304,f445,f397])).
fof(f1181,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.SUBSET_2'(s(fun(X2,bool),X3),s(fun(X2,bool),X3)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f305,f304,f445,f397])).
fof(f1182,plain,(
  ( ! [X0,X1] : (sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f445,f304,f445,f397])).
fof(f1185,plain,(
  ( ! [X0,X1] : (sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f249,f305,f445,f397])).
fof(f1186,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.SUBSET_2'(s(fun(X2,bool),X3),s(fun(X2,bool),'const.pred_set.UNIV_0')))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f304,f305,f445,f397])).
fof(f1187,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.SUBSET_2'(s(fun(X2,bool),X3),s(fun(X2,bool),'const.pred_set.UNIV_0')))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f305,f305,f445,f397])).
fof(f1188,plain,(
  ( ! [X0,X1] : (sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f445,f305,f445,f397])).
fof(f1192,plain,(
  ( ! [X0,X1] : (sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f304,f445,f445,f397])).
fof(f1193,plain,(
  ( ! [X0,X1] : (sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f305,f445,f445,f397])).
fof(f1195,plain,(
  ( ! [X0] : (~sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X0,'F_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f255,f445,f398])).
fof(f1196,plain,(
  ( ! [X0] : (~sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0',X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f255,f445,f399])).
fof(f1200,plain,(
  ( ! [X0,X1] : (~sP8('F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f255,f304,f445,f400])).
fof(f1201,plain,(
  ( ! [X0,X1] : (~sP8('F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f255,f305,f445,f400])).
fof(f1206,plain,(
  ( ! [X0,X1] : (~sP8('F_0','const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f255,f304,f445,f400])).
fof(f1207,plain,(
  ( ! [X0,X1] : (~sP8('F_0','const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f255,f305,f445,f400])).
fof(f1212,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f445,f409])).
fof(f1213,plain,(
  ( ! [X0,X1] : (sP9(X0,'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f445,f410])).
fof(f1216,plain,(
  ( ! [X0] : (sP10('T_0',X0,'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f249,f445,f413])).
fof(f1217,plain,(
  ( ! [X2,X0,X1] : (sP10('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),X2,'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f304,f445,f413])).
fof(f1218,plain,(
  ( ! [X2,X0,X1] : (sP10('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),X2,'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f305,f445,f413])).
fof(f1222,plain,(
  ( ! [X0] : (sP10('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X0,'T_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f249,f445,f413])).
fof(f1223,plain,(
  ( ! [X2,X0,X1] : (sP10('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X0,'const.pred_set.SUBSET_2'(s(fun(X1,bool),X2),s(fun(X1,bool),X2)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f304,f445,f413])).
fof(f1224,plain,(
  ( ! [X2,X0,X1] : (sP10('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X0,'const.pred_set.SUBSET_2'(s(fun(X1,bool),X2),s(fun(X1,bool),'const.pred_set.UNIV_0')))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f305,f445,f413])).
fof(f1225,plain,(
  ( ! [X0] : (sP10('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X0,'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f445,f445,f413])).
fof(f1228,plain,(
  ( ! [X0] : (sP10('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f249,f445,f414])).
fof(f1229,plain,(
  ( ! [X2,X0,X1] : (sP10('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X2)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f304,f445,f414])).
fof(f1230,plain,(
  ( ! [X2,X0,X1] : (sP10('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X2)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f305,f445,f414])).
fof(f1234,plain,(
  ( ! [X0] : (sP10('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0',X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f249,f445,f414])).
fof(f1235,plain,(
  ( ! [X2,X0,X1] : (sP10('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),X2)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f304,f445,f414])).
fof(f1236,plain,(
  ( ! [X2,X0,X1] : (sP10('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),X2)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f305,f445,f414])).
fof(f1237,plain,(
  ( ! [X0] : (sP10('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f445,f445,f414])).
fof(f1239,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f255,f445,f416])).
fof(f1240,plain,(
  ( ! [X0] : (~sP10('F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f255,f445,f417])).
fof(f1245,plain,(
  ( ! [X0,X1] : (~sP11('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f255,f304,f445,f425])).
fof(f1246,plain,(
  ( ! [X0,X1] : (~sP11('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f255,f305,f445,f425])).
fof(f1251,plain,(
  ( ! [X0,X1] : (~sP11('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f304,f255,f445,f425])).
fof(f1252,plain,(
  ( ! [X0,X1] : (~sP11('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f305,f255,f445,f425])).
fof(f1254,plain,(
  ( ! [X0,X1] : (sP11(X0,'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f445,f427])).
fof(f1256,plain,(
  ( ! [X0] : (sP12('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X0,'F_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f255,f445,f430])).
fof(f1259,plain,(
  ( ! [X0] : (sP12('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f249,f445,f431])).
fof(f1260,plain,(
  ( ! [X2,X0,X1] : (sP12('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X2)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f304,f445,f431])).
fof(f1261,plain,(
  ( ! [X2,X0,X1] : (sP12('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X2)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f305,f445,f431])).
fof(f1265,plain,(
  ( ! [X0] : (sP12('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0',X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f249,f445,f431])).
fof(f1266,plain,(
  ( ! [X2,X0,X1] : (sP12('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),X2)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f304,f445,f431])).
fof(f1267,plain,(
  ( ! [X2,X0,X1] : (sP12('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),X2)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f305,f445,f431])).
fof(f1268,plain,(
  ( ! [X0] : (sP12('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f445,f445,f431])).
fof(f1272,plain,(
  ( ! [X0,X1] : (~sP12('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)),'F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f255,f304,f445,f432])).
fof(f1273,plain,(
  ( ! [X0,X1] : (~sP12('const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f255,f305,f445,f432])).
fof(f1278,plain,(
  ( ! [X0,X1] : (~sP12('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X1)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f255,f304,f445,f432])).
fof(f1279,plain,(
  ( ! [X0,X1] : (~sP12('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f255,f305,f445,f432])).
fof(f1281,plain,(
  ( ! [X0] : (~sP12('F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f255,f445,f434])).
fof(f1282,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f445,f438])).
fof(f1286,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl2),
  inference(resolution,[],[f445,f265])).
fof(f1285,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl2),
  inference(resolution,[],[f445,f263])).
fof(f1402,plain,(
  p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) | $spl170),
  inference(cnf_transformation,[],[f1402_D])).
fof(f1402_D,plain,(
  p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f1284,plain,(
  p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) | $spl2),
  inference(resolution,[],[f445,f342])).
fof(f342,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))) )),
  inference(cnf_transformation,[],[f203])).
fof(f203,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | (p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) & p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))))) & (~p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | ~p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))))),
  inference(flattening,[],[f202])).
fof(f202,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | (p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) & p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))))) & ((~p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | ~p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))) | p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))))),
  inference(nnf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) <=> (p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) & p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))))),
  inference(flattening,[],[f100])).
fof(f100,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) <=> (p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) & p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X6,X13] : ! [X9,X7,X10] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(X6,X13),X9),s(fun(X6,bool),X7),s(fun(X13,bool),X10)))) <=> (p(s(bool,'const.pred_set.INJ_3'(s(fun(X6,X13),X9),s(fun(X6,bool),X7),s(fun(X13,bool),X10)))) & p(s(bool,'const.pred_set.SURJ_3'(s(fun(X6,X13),X9),s(fun(X6,bool),X7),s(fun(X13,bool),X10))))))),
  file('thm.pred_set.SUM_IMAGE_PERMUTES','thm.pred_set.BIJ_DEF')).
fof(f1404,plain,(
  p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) | $spl172),
  inference(cnf_transformation,[],[f1404_D])).
fof(f1404_D,plain,(
  p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f1283,plain,(
  p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) | $spl2),
  inference(resolution,[],[f445,f341])).
fof(f341,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))) )),
  inference(cnf_transformation,[],[f203])).
fof(f937,plain,(
  p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) | $spl2),
  inference(unit_resulting_resolution,[],[f445,f341])).
fof(f938,plain,(
  p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) | $spl2),
  inference(unit_resulting_resolution,[],[f445,f342])).
fof(f939,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f445,f263])).
fof(f940,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f445,f265])).
fof(f1398,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl167),
  inference(cnf_transformation,[],[f1398_D])).
fof(f1398_D,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl167),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl167])])).
fof(f941,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f445,f266])).
fof(f942,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f445,f268])).
fof(f1390,plain,(
  ~sP1('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl159),
  inference(cnf_transformation,[],[f1390_D])).
fof(f1390_D,plain,(
  ~sP1('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl159),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl159])])).
fof(f952,plain,(
  ~sP1('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f437,f445,f328])).
fof(f953,plain,(
  ~sP1('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f250,f445,f328])).
fof(f954,plain,(
  ~sP1('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f249,f445,f328])).
fof(f1388,plain,(
  ~sP1('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl157),
  inference(cnf_transformation,[],[f1388_D])).
fof(f1388_D,plain,(
  ~sP1('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl157),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl157])])).
fof(f957,plain,(
  ~sP1('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f445,f445,f328])).
fof(f1392,plain,(
  ~sP1('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl161),
  inference(cnf_transformation,[],[f1392_D])).
fof(f1392_D,plain,(
  ~sP1('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl161),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl161])])).
fof(f958,plain,(
  ~sP1('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f437,f445,f328])).
fof(f959,plain,(
  ~sP1('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f250,f445,f328])).
fof(f960,plain,(
  ~sP1('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f249,f445,f328])).
fof(f963,plain,(
  ~sP1('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f445,f445,f328])).
fof(f1394,plain,(
  sP1('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl162),
  inference(cnf_transformation,[],[f1394_D])).
fof(f1394_D,plain,(
  sP1('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f964,plain,(
  sP1('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f255,f445,f329])).
fof(f1396,plain,(
  sP1('F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl164),
  inference(cnf_transformation,[],[f1396_D])).
fof(f1396_D,plain,(
  sP1('F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f965,plain,(
  sP1('F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f255,f445,f332])).
fof(f966,plain,(
  sP1('F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f255,f445,f333])).
fof(f967,plain,(
  sP1('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f255,f445,f334])).
fof(f968,plain,(
  ~sP1('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f437,f445,f335])).
fof(f969,plain,(
  ~sP1('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f250,f445,f335])).
fof(f970,plain,(
  ~sP1('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f249,f445,f335])).
fof(f973,plain,(
  ~sP1('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f445,f445,f335])).
fof(f974,plain,(
  ~sP1('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f437,f445,f335])).
fof(f975,plain,(
  ~sP1('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f250,f445,f335])).
fof(f976,plain,(
  ~sP1('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f249,f445,f335])).
fof(f979,plain,(
  ~sP1('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f445,f445,f335])).
fof(f1370,plain,(
  ~sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl139),
  inference(cnf_transformation,[],[f1370_D])).
fof(f1370_D,plain,(
  ~sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','F_0') <=> ~$spl139),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl139])])).
fof(f1029,plain,(
  ~sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f437,f255,f445,f364])).
fof(f1030,plain,(
  ~sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f250,f255,f445,f364])).
fof(f1031,plain,(
  ~sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f249,f255,f445,f364])).
fof(f1384,plain,(
  ~sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl153),
  inference(cnf_transformation,[],[f1384_D])).
fof(f1384_D,plain,(
  ~sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') <=> ~$spl153),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl153])])).
fof(f1034,plain,(
  ~sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f445,f255,f445,f364])).
fof(f1386,plain,(
  ~sP6('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl155),
  inference(cnf_transformation,[],[f1386_D])).
fof(f1386_D,plain,(
  ~sP6('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') <=> ~$spl155),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl155])])).
fof(f1035,plain,(
  ~sP6('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f437,f255,f445,f364])).
fof(f1036,plain,(
  ~sP6('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f250,f255,f445,f364])).
fof(f1037,plain,(
  ~sP6('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f249,f255,f445,f364])).
fof(f1040,plain,(
  ~sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f445,f255,f445,f364])).
fof(f1382,plain,(
  ~sP5('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl151),
  inference(cnf_transformation,[],[f1382_D])).
fof(f1382_D,plain,(
  ~sP5('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','F_0') <=> ~$spl151),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl151])])).
fof(f1041,plain,(
  ~sP5('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f437,f255,f445,f368])).
fof(f1042,plain,(
  ~sP5('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f250,f255,f445,f368])).
fof(f1043,plain,(
  ~sP5('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f249,f255,f445,f368])).
fof(f1378,plain,(
  ~sP5('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl147),
  inference(cnf_transformation,[],[f1378_D])).
fof(f1378_D,plain,(
  ~sP5('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') <=> ~$spl147),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl147])])).
fof(f1046,plain,(
  ~sP5('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f445,f255,f445,f368])).
fof(f1380,plain,(
  ~sP5('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl149),
  inference(cnf_transformation,[],[f1380_D])).
fof(f1380_D,plain,(
  ~sP5('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') <=> ~$spl149),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl149])])).
fof(f1047,plain,(
  ~sP5('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f437,f255,f445,f368])).
fof(f1048,plain,(
  ~sP5('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f250,f255,f445,f368])).
fof(f1049,plain,(
  ~sP5('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f249,f255,f445,f368])).
fof(f1052,plain,(
  ~sP5('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f445,f255,f445,f368])).
fof(f1376,plain,(
  ~sP4('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl145),
  inference(cnf_transformation,[],[f1376_D])).
fof(f1376_D,plain,(
  ~sP4('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','F_0') <=> ~$spl145),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl145])])).
fof(f1054,plain,(
  ~sP4('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f437,f255,f445,f372])).
fof(f1055,plain,(
  ~sP4('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f250,f255,f445,f372])).
fof(f1056,plain,(
  ~sP4('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f249,f255,f445,f372])).
fof(f1372,plain,(
  ~sP4('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl141),
  inference(cnf_transformation,[],[f1372_D])).
fof(f1372_D,plain,(
  ~sP4('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') <=> ~$spl141),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl141])])).
fof(f1059,plain,(
  ~sP4('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f445,f255,f445,f372])).
fof(f1374,plain,(
  ~sP4('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl143),
  inference(cnf_transformation,[],[f1374_D])).
fof(f1374_D,plain,(
  ~sP4('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') <=> ~$spl143),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl143])])).
fof(f1060,plain,(
  ~sP4('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f437,f255,f445,f372])).
fof(f1061,plain,(
  ~sP4('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f250,f255,f445,f372])).
fof(f1062,plain,(
  ~sP4('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f249,f255,f445,f372])).
fof(f1065,plain,(
  ~sP4('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f445,f255,f445,f372])).
fof(f1070,plain,(
  ~sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','F_0') | ($spl2 | $spl83)),
  inference(unit_resulting_resolution,[],[f908,f445,f382])).
fof(f1368,plain,(
  ~sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','T_0') | $spl137),
  inference(cnf_transformation,[],[f1368_D])).
fof(f1368_D,plain,(
  ~sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','T_0') <=> ~$spl137),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl137])])).
fof(f1071,plain,(
  ~sP6('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','T_0') | ($spl2 | $spl83)),
  inference(unit_resulting_resolution,[],[f908,f445,f382])).
fof(f1366,plain,(
  ~sP7('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl135),
  inference(cnf_transformation,[],[f1366_D])).
fof(f1366_D,plain,(
  ~sP7('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','F_0') <=> ~$spl135),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl135])])).
fof(f1074,plain,(
  ~sP7('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f437,f255,f445,f391])).
fof(f1075,plain,(
  ~sP7('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f250,f255,f445,f391])).
fof(f1076,plain,(
  ~sP7('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f249,f255,f445,f391])).
fof(f1362,plain,(
  ~sP7('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl131),
  inference(cnf_transformation,[],[f1362_D])).
fof(f1362_D,plain,(
  ~sP7('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') <=> ~$spl131),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl131])])).
fof(f1079,plain,(
  ~sP7('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f445,f255,f445,f391])).
fof(f1364,plain,(
  ~sP7('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl133),
  inference(cnf_transformation,[],[f1364_D])).
fof(f1364_D,plain,(
  ~sP7('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') <=> ~$spl133),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl133])])).
fof(f1080,plain,(
  ~sP7('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f437,f255,f445,f391])).
fof(f1081,plain,(
  ~sP7('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f250,f255,f445,f391])).
fof(f1082,plain,(
  ~sP7('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f249,f255,f445,f391])).
fof(f1085,plain,(
  ~sP7('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f445,f255,f445,f391])).
fof(f1360,plain,(
  sP8('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl128),
  inference(cnf_transformation,[],[f1360_D])).
fof(f1360_D,plain,(
  sP8('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f1087,plain,(
  sP8('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f437,f437,f445,f397])).
fof(f1088,plain,(
  sP8('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f250,f437,f445,f397])).
fof(f1089,plain,(
  sP8('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f249,f437,f445,f397])).
fof(f1352,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl120),
  inference(cnf_transformation,[],[f1352_D])).
fof(f1352_D,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f1092,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f445,f437,f445,f397])).
fof(f1093,plain,(
  sP8('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f437,f250,f445,f397])).
fof(f1094,plain,(
  sP8('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f250,f250,f445,f397])).
fof(f1095,plain,(
  sP8('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f249,f250,f445,f397])).
fof(f1098,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f445,f250,f445,f397])).
fof(f1099,plain,(
  sP8('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f437,f249,f445,f397])).
fof(f1100,plain,(
  sP8('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f250,f249,f445,f397])).
fof(f1101,plain,(
  sP8('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f249,f249,f445,f397])).
fof(f1104,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f445,f249,f445,f397])).
fof(f1356,plain,(
  sP8('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl124),
  inference(cnf_transformation,[],[f1356_D])).
fof(f1356_D,plain,(
  sP8('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f1117,plain,(
  sP8('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f437,f445,f445,f397])).
fof(f1118,plain,(
  sP8('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f250,f445,f445,f397])).
fof(f1119,plain,(
  sP8('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f249,f445,f445,f397])).
fof(f1348,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl116),
  inference(cnf_transformation,[],[f1348_D])).
fof(f1348_D,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f1122,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f445,f445,f445,f397])).
fof(f1358,plain,(
  sP8('T_0','T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl126),
  inference(cnf_transformation,[],[f1358_D])).
fof(f1358_D,plain,(
  sP8('T_0','T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f1123,plain,(
  sP8('T_0','T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f437,f437,f445,f397])).
fof(f1124,plain,(
  sP8('T_0','T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f250,f437,f445,f397])).
fof(f1125,plain,(
  sP8('T_0','T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f249,f437,f445,f397])).
fof(f1128,plain,(
  sP8('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f445,f437,f445,f397])).
fof(f1129,plain,(
  sP8('T_0','T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f437,f250,f445,f397])).
fof(f1130,plain,(
  sP8('T_0','T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f250,f250,f445,f397])).
fof(f1131,plain,(
  sP8('T_0','T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f249,f250,f445,f397])).
fof(f1134,plain,(
  sP8('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f445,f250,f445,f397])).
fof(f1135,plain,(
  sP8('T_0','T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f437,f249,f445,f397])).
fof(f1136,plain,(
  sP8('T_0','T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f250,f249,f445,f397])).
fof(f1137,plain,(
  sP8('T_0','T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f249,f249,f445,f397])).
fof(f1140,plain,(
  sP8('T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f445,f249,f445,f397])).
fof(f1350,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl118),
  inference(cnf_transformation,[],[f1350_D])).
fof(f1350_D,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f1153,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f437,f445,f445,f397])).
fof(f1154,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f250,f445,f445,f397])).
fof(f1155,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f249,f445,f445,f397])).
fof(f1158,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f445,f445,f445,f397])).
fof(f1354,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','T_0') | $spl122),
  inference(cnf_transformation,[],[f1354_D])).
fof(f1354_D,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','T_0') <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f1159,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f437,f437,f445,f397])).
fof(f1160,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f250,f437,f445,f397])).
fof(f1161,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f249,f437,f445,f397])).
fof(f1164,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f445,f437,f445,f397])).
fof(f1165,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f437,f250,f445,f397])).
fof(f1166,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f250,f250,f445,f397])).
fof(f1167,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f249,f250,f445,f397])).
fof(f1170,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f445,f250,f445,f397])).
fof(f1171,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f437,f249,f445,f397])).
fof(f1172,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f250,f249,f445,f397])).
fof(f1173,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f249,f249,f445,f397])).
fof(f1176,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f445,f249,f445,f397])).
fof(f1189,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f437,f445,f445,f397])).
fof(f1190,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f250,f445,f445,f397])).
fof(f1191,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f249,f445,f445,f397])).
fof(f1194,plain,(
  sP8('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f445,f445,f445,f397])).
fof(f1346,plain,(
  ~sP8('F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl115),
  inference(cnf_transformation,[],[f1346_D])).
fof(f1346_D,plain,(
  ~sP8('F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl115),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl115])])).
fof(f1197,plain,(
  ~sP8('F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f255,f437,f445,f400])).
fof(f1198,plain,(
  ~sP8('F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f255,f250,f445,f400])).
fof(f1199,plain,(
  ~sP8('F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f255,f249,f445,f400])).
fof(f1342,plain,(
  ~sP8('F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl111),
  inference(cnf_transformation,[],[f1342_D])).
fof(f1342_D,plain,(
  ~sP8('F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl111),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl111])])).
fof(f1202,plain,(
  ~sP8('F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f255,f445,f445,f400])).
fof(f1344,plain,(
  ~sP8('F_0','T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl113),
  inference(cnf_transformation,[],[f1344_D])).
fof(f1344_D,plain,(
  ~sP8('F_0','T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl113),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl113])])).
fof(f1203,plain,(
  ~sP8('F_0','T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f255,f437,f445,f400])).
fof(f1204,plain,(
  ~sP8('F_0','T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f255,f250,f445,f400])).
fof(f1205,plain,(
  ~sP8('F_0','T_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f255,f249,f445,f400])).
fof(f1208,plain,(
  ~sP8('F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f255,f445,f445,f400])).
fof(f1340,plain,(
  ~sP9('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','F_0') | $spl109),
  inference(cnf_transformation,[],[f1340_D])).
fof(f1340_D,plain,(
  ~sP9('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','F_0') <=> ~$spl109),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl109])])).
fof(f1211,plain,(
  ~sP9('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f255,f255,f445,f408])).
fof(f1338,plain,(
  ~sP10('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','F_0') | $spl107),
  inference(cnf_transformation,[],[f1338_D])).
fof(f1338_D,plain,(
  ~sP10('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','F_0') <=> ~$spl107),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl107])])).
fof(f1238,plain,(
  ~sP10('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f255,f255,f445,f415])).
fof(f1336,plain,(
  ~sP11('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','T_0') | $spl105),
  inference(cnf_transformation,[],[f1336_D])).
fof(f1336_D,plain,(
  ~sP11('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','T_0') <=> ~$spl105),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl105])])).
fof(f1242,plain,(
  ~sP11('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f255,f437,f445,f425])).
fof(f1243,plain,(
  ~sP11('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f255,f250,f445,f425])).
fof(f1244,plain,(
  ~sP11('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f255,f249,f445,f425])).
fof(f1332,plain,(
  ~sP11('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl101),
  inference(cnf_transformation,[],[f1332_D])).
fof(f1332_D,plain,(
  ~sP11('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl101),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl101])])).
fof(f1247,plain,(
  ~sP11('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f255,f445,f445,f425])).
fof(f1334,plain,(
  ~sP11('T_0','F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl103),
  inference(cnf_transformation,[],[f1334_D])).
fof(f1334_D,plain,(
  ~sP11('T_0','F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl103),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl103])])).
fof(f1248,plain,(
  ~sP11('T_0','F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f437,f255,f445,f425])).
fof(f1249,plain,(
  ~sP11('T_0','F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f250,f255,f445,f425])).
fof(f1250,plain,(
  ~sP11('T_0','F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f249,f255,f445,f425])).
fof(f1253,plain,(
  ~sP11('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f445,f255,f445,f425])).
fof(f1330,plain,(
  sP12('F_0','F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl98),
  inference(cnf_transformation,[],[f1330_D])).
fof(f1330_D,plain,(
  sP12('F_0','F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f1255,plain,(
  sP12('F_0','F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f255,f255,f445,f429])).
fof(f1328,plain,(
  ~sP12('T_0','F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl97),
  inference(cnf_transformation,[],[f1328_D])).
fof(f1328_D,plain,(
  ~sP12('T_0','F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl97),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl97])])).
fof(f1269,plain,(
  ~sP12('T_0','F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f255,f437,f445,f432])).
fof(f1270,plain,(
  ~sP12('T_0','F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f255,f250,f445,f432])).
fof(f1271,plain,(
  ~sP12('T_0','F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f255,f249,f445,f432])).
fof(f1324,plain,(
  ~sP12('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl93),
  inference(cnf_transformation,[],[f1324_D])).
fof(f1324_D,plain,(
  ~sP12('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) <=> ~$spl93),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl93])])).
fof(f1274,plain,(
  ~sP12('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f255,f445,f445,f432])).
fof(f1326,plain,(
  ~sP12('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','T_0') | $spl95),
  inference(cnf_transformation,[],[f1326_D])).
fof(f1326_D,plain,(
  ~sP12('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','T_0') <=> ~$spl95),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl95])])).
fof(f1275,plain,(
  ~sP12('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f255,f437,f445,f432])).
fof(f1276,plain,(
  ~sP12('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f255,f250,f445,f432])).
fof(f1277,plain,(
  ~sP12('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f255,f249,f445,f432])).
fof(f1280,plain,(
  ~sP12('const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)),'F_0','const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))) | $spl2),
  inference(unit_resulting_resolution,[],[f255,f445,f445,f432])).
fof(f936,plain,(
  sP6('F_0','F_0','T_0') | $spl90),
  inference(cnf_transformation,[],[f936_D])).
fof(f936_D,plain,(
  sP6('F_0','F_0','T_0') <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f912,plain,(
  sP6('F_0','F_0','T_0') | $spl83),
  inference(unit_resulting_resolution,[],[f255,f908,f380])).
fof(f934,plain,(
  ~sP6('T_0','F_0','T_0') | $spl89),
  inference(cnf_transformation,[],[f934_D])).
fof(f934_D,plain,(
  ~sP6('T_0','F_0','T_0') <=> ~$spl89),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl89])])).
fof(f913,plain,(
  ~sP6('T_0','F_0','T_0') | $spl83),
  inference(unit_resulting_resolution,[],[f437,f908,f382])).
fof(f914,plain,(
  ~sP6('T_0','F_0','T_0') | $spl83),
  inference(unit_resulting_resolution,[],[f250,f908,f382])).
fof(f915,plain,(
  ~sP6('T_0','F_0','T_0') | $spl83),
  inference(unit_resulting_resolution,[],[f249,f908,f382])).
fof(f932,plain,(
  sP6('F_0','T_0','F_0') | $spl86),
  inference(cnf_transformation,[],[f932_D])).
fof(f932_D,plain,(
  sP6('F_0','T_0','F_0') <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f921,plain,(
  sP6('F_0','T_0','F_0') | $spl83),
  inference(unit_resulting_resolution,[],[f255,f908,f380])).
fof(f930,plain,(
  ~sP6('T_0','T_0','F_0') | $spl85),
  inference(cnf_transformation,[],[f930_D])).
fof(f930_D,plain,(
  ~sP6('T_0','T_0','F_0') <=> ~$spl85),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl85])])).
fof(f922,plain,(
  ~sP6('T_0','T_0','F_0') | $spl83),
  inference(unit_resulting_resolution,[],[f437,f908,f382])).
fof(f923,plain,(
  ~sP6('T_0','T_0','F_0') | $spl83),
  inference(unit_resulting_resolution,[],[f250,f908,f382])).
fof(f924,plain,(
  ~sP6('T_0','T_0','F_0') | $spl83),
  inference(unit_resulting_resolution,[],[f249,f908,f382])).
fof(f908,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl83),
  inference(cnf_transformation,[],[f908_D])).
fof(f908_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl83),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl83])])).
fof(f906,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ($spl79 | $spl80)),
  inference(forward_demodulation,[],[f901,f903])).
fof(f903,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl80),
  inference(cnf_transformation,[],[f903_D])).
fof(f903_D,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f904,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl4),
  inference(global_subsumption,[],[f248,f247,f246,f447,f789,f788,f787,f786,f785,f784,f783,f782,f781,f780,f779,f778,f777,f776,f775,f774,f773,f772,f768,f767,f766,f763,f762,f761,f760,f759,f758,f757,f756,f755,f754,f753,f752,f751,f750,f749,f747,f746,f745,f744,f743,f742,f741,f737,f736,f735,f732,f731,f730,f729,f725,f724,f723,f720,f719,f718,f715,f714,f713,f712,f711,f710,f709,f708,f707,f706,f705,f704,f703,f702,f701,f700,f699,f698,f697,f696,f695,f694,f693,f692,f689,f688,f687,f686,f683,f682,f681,f680,f679,f678,f677,f674,f673,f672,f671,f668,f667,f666,f665,f662,f661,f660,f659,f658,f657,f656,f653,f652,f651,f650,f647,f646,f645,f644,f643,f642,f641,f638,f637,f636,f635,f632,f631,f630,f629,f626,f625,f624,f622,f621,f620,f616,f615,f614,f611,f610,f609,f608,f607,f606,f605,f602,f601,f600,f599,f596,f595,f594,f593,f592,f591,f590,f589,f588,f587,f586,f585,f584,f583,f582,f581,f578,f577,f576,f575,f574,f573,f572,f571,f570,f569,f568,f567,f566,f565,f564,f563,f562,f561,f560,f559,f558,f557,f556,f555,f554,f553,f552,f551,f550,f549,f548,f547,f546,f545,f544,f543,f542,f541,f540,f539,f538,f537,f536,f535,f534,f530,f529,f528,f505,f504,f503,f499,f498,f497,f488,f487,f486,f485,f484,f483,f482,f481,f480,f479,f478,f477,f476,f475,f474,f473,f472,f469,f468,f467,f466,f463,f462,f461,f460,f459,f451,f450,f449,f448,f790,f791,f800,f801,f802,f805,f808,f810,f817,f820,f821,f823,f824,f825,f827])).
fof(f791,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl4),
  inference(resolution,[],[f447,f265])).
fof(f790,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl4),
  inference(resolution,[],[f447,f263])).
fof(f448,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f447,f263])).
fof(f449,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f447,f265])).
fof(f901,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl79),
  inference(cnf_transformation,[],[f901_D])).
fof(f901_D,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) <=> ~$spl79),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl79])])).
fof(f450,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f447,f266])).
fof(f451,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f447,f268])).
fof(f893,plain,(
  ~sP1('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl71),
  inference(cnf_transformation,[],[f893_D])).
fof(f893_D,plain,(
  ~sP1('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl71),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl71])])).
fof(f461,plain,(
  ~sP1('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f437,f447,f328])).
fof(f462,plain,(
  ~sP1('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f250,f447,f328])).
fof(f463,plain,(
  ~sP1('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f249,f447,f328])).
fof(f891,plain,(
  ~sP1('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl69),
  inference(cnf_transformation,[],[f891_D])).
fof(f891_D,plain,(
  ~sP1('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) <=> ~$spl69),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl69])])).
fof(f466,plain,(
  ~sP1('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f447,f447,f328])).
fof(f895,plain,(
  ~sP1('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl73),
  inference(cnf_transformation,[],[f895_D])).
fof(f895_D,plain,(
  ~sP1('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) <=> ~$spl73),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl73])])).
fof(f467,plain,(
  ~sP1('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f437,f447,f328])).
fof(f468,plain,(
  ~sP1('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f250,f447,f328])).
fof(f469,plain,(
  ~sP1('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f249,f447,f328])).
fof(f472,plain,(
  ~sP1('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f447,f447,f328])).
fof(f897,plain,(
  sP1('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0') | $spl74),
  inference(cnf_transformation,[],[f897_D])).
fof(f897_D,plain,(
  sP1('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0') <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f473,plain,(
  sP1('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f255,f447,f329])).
fof(f899,plain,(
  sP1('F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl76),
  inference(cnf_transformation,[],[f899_D])).
fof(f899_D,plain,(
  sP1('F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f474,plain,(
  sP1('F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f255,f447,f332])).
fof(f475,plain,(
  sP1('F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f255,f447,f333])).
fof(f476,plain,(
  sP1('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f255,f447,f334])).
fof(f477,plain,(
  ~sP1('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f437,f447,f335])).
fof(f478,plain,(
  ~sP1('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f250,f447,f335])).
fof(f479,plain,(
  ~sP1('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f249,f447,f335])).
fof(f482,plain,(
  ~sP1('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f447,f447,f335])).
fof(f483,plain,(
  ~sP1('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f437,f447,f335])).
fof(f484,plain,(
  ~sP1('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f250,f447,f335])).
fof(f485,plain,(
  ~sP1('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f249,f447,f335])).
fof(f488,plain,(
  ~sP1('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f447,f447,f335])).
fof(f889,plain,(
  ~sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl67),
  inference(cnf_transformation,[],[f889_D])).
fof(f889_D,plain,(
  ~sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','F_0') <=> ~$spl67),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl67])])).
fof(f538,plain,(
  ~sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f437,f255,f447,f364])).
fof(f539,plain,(
  ~sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f250,f255,f447,f364])).
fof(f540,plain,(
  ~sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f249,f255,f447,f364])).
fof(f885,plain,(
  ~sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0') | $spl63),
  inference(cnf_transformation,[],[f885_D])).
fof(f885_D,plain,(
  ~sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0') <=> ~$spl63),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl63])])).
fof(f543,plain,(
  ~sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f447,f255,f447,f364])).
fof(f887,plain,(
  ~sP6('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0') | $spl65),
  inference(cnf_transformation,[],[f887_D])).
fof(f887_D,plain,(
  ~sP6('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0') <=> ~$spl65),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl65])])).
fof(f544,plain,(
  ~sP6('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f437,f255,f447,f364])).
fof(f545,plain,(
  ~sP6('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f250,f255,f447,f364])).
fof(f546,plain,(
  ~sP6('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f249,f255,f447,f364])).
fof(f549,plain,(
  ~sP6('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f447,f255,f447,f364])).
fof(f883,plain,(
  ~sP5('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl61),
  inference(cnf_transformation,[],[f883_D])).
fof(f883_D,plain,(
  ~sP5('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','F_0') <=> ~$spl61),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl61])])).
fof(f550,plain,(
  ~sP5('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f437,f255,f447,f368])).
fof(f551,plain,(
  ~sP5('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f250,f255,f447,f368])).
fof(f552,plain,(
  ~sP5('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f249,f255,f447,f368])).
fof(f879,plain,(
  ~sP5('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0') | $spl57),
  inference(cnf_transformation,[],[f879_D])).
fof(f879_D,plain,(
  ~sP5('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0') <=> ~$spl57),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl57])])).
fof(f555,plain,(
  ~sP5('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f447,f255,f447,f368])).
fof(f881,plain,(
  ~sP5('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0') | $spl59),
  inference(cnf_transformation,[],[f881_D])).
fof(f881_D,plain,(
  ~sP5('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0') <=> ~$spl59),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl59])])).
fof(f556,plain,(
  ~sP5('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f437,f255,f447,f368])).
fof(f557,plain,(
  ~sP5('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f250,f255,f447,f368])).
fof(f558,plain,(
  ~sP5('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f249,f255,f447,f368])).
fof(f561,plain,(
  ~sP5('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f447,f255,f447,f368])).
fof(f877,plain,(
  ~sP4('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl55),
  inference(cnf_transformation,[],[f877_D])).
fof(f877_D,plain,(
  ~sP4('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','F_0') <=> ~$spl55),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl55])])).
fof(f563,plain,(
  ~sP4('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f437,f255,f447,f372])).
fof(f564,plain,(
  ~sP4('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f250,f255,f447,f372])).
fof(f565,plain,(
  ~sP4('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f249,f255,f447,f372])).
fof(f873,plain,(
  ~sP4('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0') | $spl51),
  inference(cnf_transformation,[],[f873_D])).
fof(f873_D,plain,(
  ~sP4('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0') <=> ~$spl51),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl51])])).
fof(f568,plain,(
  ~sP4('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f447,f255,f447,f372])).
fof(f875,plain,(
  ~sP4('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0') | $spl53),
  inference(cnf_transformation,[],[f875_D])).
fof(f875_D,plain,(
  ~sP4('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0') <=> ~$spl53),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl53])])).
fof(f569,plain,(
  ~sP4('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f437,f255,f447,f372])).
fof(f570,plain,(
  ~sP4('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f250,f255,f447,f372])).
fof(f571,plain,(
  ~sP4('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f249,f255,f447,f372])).
fof(f574,plain,(
  ~sP4('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f447,f255,f447,f372])).
fof(f871,plain,(
  ~sP7('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl49),
  inference(cnf_transformation,[],[f871_D])).
fof(f871_D,plain,(
  ~sP7('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','F_0') <=> ~$spl49),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl49])])).
fof(f581,plain,(
  ~sP7('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f437,f255,f447,f391])).
fof(f582,plain,(
  ~sP7('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f250,f255,f447,f391])).
fof(f583,plain,(
  ~sP7('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f249,f255,f447,f391])).
fof(f867,plain,(
  ~sP7('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0') | $spl45),
  inference(cnf_transformation,[],[f867_D])).
fof(f867_D,plain,(
  ~sP7('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0') <=> ~$spl45),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl45])])).
fof(f586,plain,(
  ~sP7('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f447,f255,f447,f391])).
fof(f869,plain,(
  ~sP7('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0') | $spl47),
  inference(cnf_transformation,[],[f869_D])).
fof(f869_D,plain,(
  ~sP7('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0') <=> ~$spl47),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl47])])).
fof(f587,plain,(
  ~sP7('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f437,f255,f447,f391])).
fof(f588,plain,(
  ~sP7('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f250,f255,f447,f391])).
fof(f589,plain,(
  ~sP7('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f249,f255,f447,f391])).
fof(f592,plain,(
  ~sP7('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f447,f255,f447,f391])).
fof(f865,plain,(
  sP8('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl42),
  inference(cnf_transformation,[],[f865_D])).
fof(f865_D,plain,(
  sP8('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f594,plain,(
  sP8('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f437,f437,f447,f397])).
fof(f595,plain,(
  sP8('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f250,f437,f447,f397])).
fof(f596,plain,(
  sP8('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f249,f437,f447,f397])).
fof(f857,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl34),
  inference(cnf_transformation,[],[f857_D])).
fof(f857_D,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f599,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f447,f437,f447,f397])).
fof(f600,plain,(
  sP8('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f437,f250,f447,f397])).
fof(f601,plain,(
  sP8('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f250,f250,f447,f397])).
fof(f602,plain,(
  sP8('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f249,f250,f447,f397])).
fof(f605,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f447,f250,f447,f397])).
fof(f606,plain,(
  sP8('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f437,f249,f447,f397])).
fof(f607,plain,(
  sP8('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f250,f249,f447,f397])).
fof(f608,plain,(
  sP8('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f249,f249,f447,f397])).
fof(f611,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f447,f249,f447,f397])).
fof(f861,plain,(
  sP8('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl38),
  inference(cnf_transformation,[],[f861_D])).
fof(f861_D,plain,(
  sP8('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f624,plain,(
  sP8('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f437,f447,f447,f397])).
fof(f625,plain,(
  sP8('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f250,f447,f447,f397])).
fof(f626,plain,(
  sP8('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f249,f447,f447,f397])).
fof(f853,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl30),
  inference(cnf_transformation,[],[f853_D])).
fof(f853_D,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f629,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f447,f447,f447,f397])).
fof(f863,plain,(
  sP8('T_0','T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl40),
  inference(cnf_transformation,[],[f863_D])).
fof(f863_D,plain,(
  sP8('T_0','T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f630,plain,(
  sP8('T_0','T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f437,f437,f447,f397])).
fof(f631,plain,(
  sP8('T_0','T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f250,f437,f447,f397])).
fof(f632,plain,(
  sP8('T_0','T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f249,f437,f447,f397])).
fof(f635,plain,(
  sP8('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f447,f437,f447,f397])).
fof(f636,plain,(
  sP8('T_0','T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f437,f250,f447,f397])).
fof(f637,plain,(
  sP8('T_0','T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f250,f250,f447,f397])).
fof(f638,plain,(
  sP8('T_0','T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f249,f250,f447,f397])).
fof(f641,plain,(
  sP8('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f447,f250,f447,f397])).
fof(f642,plain,(
  sP8('T_0','T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f437,f249,f447,f397])).
fof(f643,plain,(
  sP8('T_0','T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f250,f249,f447,f397])).
fof(f644,plain,(
  sP8('T_0','T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f249,f249,f447,f397])).
fof(f647,plain,(
  sP8('T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f447,f249,f447,f397])).
fof(f855,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl32),
  inference(cnf_transformation,[],[f855_D])).
fof(f855_D,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f660,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f437,f447,f447,f397])).
fof(f661,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f250,f447,f447,f397])).
fof(f662,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f249,f447,f447,f397])).
fof(f665,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f447,f447,f447,f397])).
fof(f859,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','T_0') | $spl36),
  inference(cnf_transformation,[],[f859_D])).
fof(f859_D,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','T_0') <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f666,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f437,f437,f447,f397])).
fof(f667,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f250,f437,f447,f397])).
fof(f668,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f249,f437,f447,f397])).
fof(f671,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f447,f437,f447,f397])).
fof(f672,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f437,f250,f447,f397])).
fof(f673,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f250,f250,f447,f397])).
fof(f674,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f249,f250,f447,f397])).
fof(f677,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f447,f250,f447,f397])).
fof(f678,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f437,f249,f447,f397])).
fof(f679,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f250,f249,f447,f397])).
fof(f680,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f249,f249,f447,f397])).
fof(f683,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f447,f249,f447,f397])).
fof(f696,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f437,f447,f447,f397])).
fof(f697,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f250,f447,f447,f397])).
fof(f698,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f249,f447,f447,f397])).
fof(f701,plain,(
  sP8('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f447,f447,f447,f397])).
fof(f851,plain,(
  ~sP8('F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl29),
  inference(cnf_transformation,[],[f851_D])).
fof(f851_D,plain,(
  ~sP8('F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f704,plain,(
  ~sP8('F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f255,f437,f447,f400])).
fof(f705,plain,(
  ~sP8('F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f255,f250,f447,f400])).
fof(f706,plain,(
  ~sP8('F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f255,f249,f447,f400])).
fof(f847,plain,(
  ~sP8('F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl25),
  inference(cnf_transformation,[],[f847_D])).
fof(f847_D,plain,(
  ~sP8('F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f709,plain,(
  ~sP8('F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f255,f447,f447,f400])).
fof(f849,plain,(
  ~sP8('F_0','T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl27),
  inference(cnf_transformation,[],[f849_D])).
fof(f849_D,plain,(
  ~sP8('F_0','T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) <=> ~$spl27),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl27])])).
fof(f710,plain,(
  ~sP8('F_0','T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f255,f437,f447,f400])).
fof(f711,plain,(
  ~sP8('F_0','T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f255,f250,f447,f400])).
fof(f712,plain,(
  ~sP8('F_0','T_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f255,f249,f447,f400])).
fof(f715,plain,(
  ~sP8('F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f255,f447,f447,f400])).
fof(f845,plain,(
  ~sP9('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0','F_0') | $spl23),
  inference(cnf_transformation,[],[f845_D])).
fof(f845_D,plain,(
  ~sP9('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0','F_0') <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f718,plain,(
  ~sP9('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f255,f255,f447,f408])).
fof(f843,plain,(
  ~sP10('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0','F_0') | $spl21),
  inference(cnf_transformation,[],[f843_D])).
fof(f843_D,plain,(
  ~sP10('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0','F_0') <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f745,plain,(
  ~sP10('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f255,f255,f447,f415])).
fof(f841,plain,(
  ~sP11('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0','T_0') | $spl19),
  inference(cnf_transformation,[],[f841_D])).
fof(f841_D,plain,(
  ~sP11('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0','T_0') <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f749,plain,(
  ~sP11('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f255,f437,f447,f425])).
fof(f750,plain,(
  ~sP11('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f255,f250,f447,f425])).
fof(f751,plain,(
  ~sP11('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f255,f249,f447,f425])).
fof(f837,plain,(
  ~sP11('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl15),
  inference(cnf_transformation,[],[f837_D])).
fof(f837_D,plain,(
  ~sP11('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f754,plain,(
  ~sP11('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f255,f447,f447,f425])).
fof(f839,plain,(
  ~sP11('T_0','F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl17),
  inference(cnf_transformation,[],[f839_D])).
fof(f839_D,plain,(
  ~sP11('T_0','F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f755,plain,(
  ~sP11('T_0','F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f437,f255,f447,f425])).
fof(f756,plain,(
  ~sP11('T_0','F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f250,f255,f447,f425])).
fof(f757,plain,(
  ~sP11('T_0','F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f249,f255,f447,f425])).
fof(f760,plain,(
  ~sP11('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f447,f255,f447,f425])).
fof(f835,plain,(
  sP12('F_0','F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl12),
  inference(cnf_transformation,[],[f835_D])).
fof(f835_D,plain,(
  sP12('F_0','F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f762,plain,(
  sP12('F_0','F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f255,f255,f447,f429])).
fof(f833,plain,(
  ~sP12('T_0','F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl11),
  inference(cnf_transformation,[],[f833_D])).
fof(f833_D,plain,(
  ~sP12('T_0','F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f776,plain,(
  ~sP12('T_0','F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f255,f437,f447,f432])).
fof(f777,plain,(
  ~sP12('T_0','F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f255,f250,f447,f432])).
fof(f778,plain,(
  ~sP12('T_0','F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f255,f249,f447,f432])).
fof(f829,plain,(
  ~sP12('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl7),
  inference(cnf_transformation,[],[f829_D])).
fof(f829_D,plain,(
  ~sP12('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f781,plain,(
  ~sP12('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f255,f447,f447,f432])).
fof(f831,plain,(
  ~sP12('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0','T_0') | $spl9),
  inference(cnf_transformation,[],[f831_D])).
fof(f831_D,plain,(
  ~sP12('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0','T_0') <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f782,plain,(
  ~sP12('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f255,f437,f447,f432])).
fof(f783,plain,(
  ~sP12('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f255,f250,f447,f432])).
fof(f784,plain,(
  ~sP12('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f255,f249,f447,f432])).
fof(f787,plain,(
  ~sP12('const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)),'F_0','const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f255,f447,f447,f432])).
fof(f447,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) | $spl4),
  inference(cnf_transformation,[],[f447_D])).
fof(f447_D,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f246,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14))))),
  inference(cnf_transformation,[],[f159])).
fof(f159,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK13,bool),sK14)))) & (p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) & s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun(sK13,'type.num.num'),sK16),s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK13,'type.num.num'),sK16),s(fun(sK13,sK13),sK15))),s(fun(sK13,bool),sK14))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13,sK14,sK15,sK16])],[f114])).
fof(f114,plain,(
  ? [X0,X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) & ? [X2] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X0),X2),s(fun(X0,bool),X1),s(fun(X0,bool),X1)))) & ? [X3] : s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun(X0,'type.num.num'),X3),s(fun(X0,bool),X1))) != s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun(X0,'type.num.num'),'const.combin.o_2'(s(fun(X0,'type.num.num'),X3),s(fun(X0,X0),X2))),s(fun(X0,bool),X1)))))),
  inference(ennf_transformation,[],[f48])).
fof(f48,plain,(
  ~! [X0,X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) => ! [X2] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X0),X2),s(fun(X0,bool),X1),s(fun(X0,bool),X1)))) => ! [X3] : s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun(X0,'type.num.num'),X3),s(fun(X0,bool),X1))) = s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun(X0,'type.num.num'),'const.combin.o_2'(s(fun(X0,'type.num.num'),X3),s(fun(X0,X0),X2))),s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f47])).
fof(f47,plain,(
  ~! [X0] : ! [X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) => ! [X2] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X0),X2),s(fun(X0,bool),X1),s(fun(X0,bool),X1)))) => ! [X3] : s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun(X0,'type.num.num'),X3),s(fun(X0,bool),X1))) = s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun(X0,'type.num.num'),'const.combin.o_2'(s(fun(X0,'type.num.num'),X3),s(fun(X0,X0),X2))),s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X7)))) => ! [X8] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(X6,X6),X8),s(fun(X6,bool),X7),s(fun(X6,bool),X7)))) => ! [X9] : s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun(X6,'type.num.num'),X9),s(fun(X6,bool),X7))) = s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun(X6,'type.num.num'),'const.combin.o_2'(s(fun(X6,'type.num.num'),X9),s(fun(X6,X6),X8))),s(fun(X6,bool),X7)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X7)))) => ! [X8] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(X6,X6),X8),s(fun(X6,bool),X7),s(fun(X6,bool),X7)))) => ! [X9] : s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun(X6,'type.num.num'),X9),s(fun(X6,bool),X7))) = s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun(X6,'type.num.num'),'const.combin.o_2'(s(fun(X6,'type.num.num'),X9),s(fun(X6,X6),X8))),s(fun(X6,bool),X7)))))),
  file('thm.pred_set.SUM_IMAGE_PERMUTES','thm.pred_set.SUM_IMAGE_PERMUTES')).
fof(f445,plain,(
  p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) | $spl2),
  inference(cnf_transformation,[],[f445_D])).
fof(f445_D,plain,(
  p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f247,plain,(
  p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK13,sK13),sK15),s(fun(sK13,bool),sK14),s(fun(sK13,bool),sK14))))),
  inference(cnf_transformation,[],[f159])).
fof(f443,plain,(
  s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun(sK13,'type.num.num'),sK16),s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK13,'type.num.num'),sK16),s(fun(sK13,sK13),sK15))),s(fun(sK13,bool),sK14))) | $spl1),
  inference(cnf_transformation,[],[f443_D])).
fof(f443_D,plain,(
  s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun(sK13,'type.num.num'),sK16),s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK13,'type.num.num'),sK16),s(fun(sK13,sK13),sK15))),s(fun(sK13,bool),sK14))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f248,plain,(
  s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun(sK13,'type.num.num'),sK16),s(fun(sK13,bool),sK14))) != s('type.num.num','const.pred_set.SUM_IMAGE_2'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK13,'type.num.num'),sK16),s(fun(sK13,sK13),sK15))),s(fun(sK13,bool),sK14)))),
  inference(cnf_transformation,[],[f159])).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 1330
% Passive clauses: 1550
% Generated clauses: 19765
% Final active clauses: 1298
% Final passive clauses: 301
% Input formulas: 45
% Initial clauses: 189
% 
% Duplicate literals: 338
% Trivial inequalities: 129
% Fw demodulations: 4821
% Bw demodulations: 7
% Global subsumptions: 12
% 
% Simple tautologies: 1067
% Equational tautologies: 17
% Forward subsumptions: 10654
% Backward subsumptions: 83
% Fw demodulations to eq. taut.: 1
% 
% Binary resolution: 4993
% Unit resulting resolution: 7748
% Forward superposition: 469
% Backward superposition: 750
% Equality resolution: 17
% 
% Split clauses: 62
% Split components: 124
% SAT solver clauses: 4033
% SAT solver unit clauses: 6
% SAT solver binary clauses: 2878
% 
% Sat splits: 62
% Sat splitting refutations: 2
% 
% Memory used [KB]: 10746
% Time elapsed: 0.448 s
% ------------------------------
% ------------------------------
% Success in time 18.503 s
