Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.transc
% (77016)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 165
% Passive clauses: 1954
% Generated clauses: 21792
% Final active clauses: 152
% Final passive clauses: 1597
% Input formulas: 39
% Initial clauses: 55
% 
% Duplicate literals: 1453
% Trivial inequalities: 106
% Fw subsumption resolutions: 2832
% Fw demodulations: 3353
% Bw demodulations: 11
% 
% Simple tautologies: 352
% Equational tautologies: 20
% Forward subsumptions: 10155
% 
% Binary resolution: 345
% Forward superposition: 8107
% Backward superposition: 5156
% Self superposition: 150
% Equality factoring: 34
% Equality resolution: 27
% 
% Split clauses: 724
% Split components: 1454
% SAT solver clauses: 1415
% SAT solver unit clauses: 6
% SAT solver binary clauses: 56
% 
% Sat splits: 724
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 1414
% TWLsolver calls for satisfiability: 199
% 
% Memory used [KB]: 14072
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.transc
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.transc
% SZS output start Proof for thm.transc
fof(f3029,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f3016,f120,f180,f122,f182,f226,f227,f155,f200,f168,f213,f236,f239,f118,f176,f133,f192,f237,f241,f117,f178,f259,f262,f137,f194,f232,f234,f119,f174])).
fof(f174,plain,(
  ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1x1_3'(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.real.real_sub_0'),s(fun('type.realax.real','type.realax.real'),sK4),s(fun('type.realax.real','type.realax.real'),sK3))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),happ(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.real.real_sub_0'),s('type.realax.real',sK2))),s('type.realax.real',sK1))),s('type.realax.real',sK0)))) | $spl1),
  inference(cnf_transformation,[],[f174_D])).
fof(f174_D,plain,(
  ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1x1_3'(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.real.real_sub_0'),s(fun('type.realax.real','type.realax.real'),sK4),s(fun('type.realax.real','type.realax.real'),sK3))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),happ(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.real.real_sub_0'),s('type.realax.real',sK2))),s('type.realax.real',sK1))),s('type.realax.real',sK0)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f119,plain,(
  ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1x1_3'(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.real.real_sub_0'),s(fun('type.realax.real','type.realax.real'),sK4),s(fun('type.realax.real','type.realax.real'),sK3))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),happ(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.real.real_sub_0'),s('type.realax.real',sK2))),s('type.realax.real',sK1))),s('type.realax.real',sK0))))),
  inference(cnf_transformation,[],[f101])).
fof(f101,plain,(
  p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),sK4),s('type.realax.real',sK2),s('type.realax.real',sK0)))) & p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),sK3),s('type.realax.real',sK1),s('type.realax.real',sK0)))) & ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1x1_3'(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.real.real_sub_0'),s(fun('type.realax.real','type.realax.real'),sK4),s(fun('type.realax.real','type.realax.real'),sK3))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),happ(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.real.real_sub_0'),s('type.realax.real',sK2))),s('type.realax.real',sK1))),s('type.realax.real',sK0))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4])],[f86])).
fof(f86,plain,(
  ? [X0,X1,X2,X3,X4] : (p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X4),s('type.realax.real',X2),s('type.realax.real',X0)))) & p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X3),s('type.realax.real',X1),s('type.realax.real',X0)))) & ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1x1_3'(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.real.real_sub_0'),s(fun('type.realax.real','type.realax.real'),X4),s(fun('type.realax.real','type.realax.real'),X3))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),happ(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.real.real_sub_0'),s('type.realax.real',X2))),s('type.realax.real',X1))),s('type.realax.real',X0)))))),
  inference(flattening,[],[f85])).
fof(f85,plain,(
  ? [X0,X1,X2,X3,X4] : ((p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X4),s('type.realax.real',X2),s('type.realax.real',X0)))) & p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X3),s('type.realax.real',X1),s('type.realax.real',X0))))) & ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1x1_3'(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.real.real_sub_0'),s(fun('type.realax.real','type.realax.real'),X4),s(fun('type.realax.real','type.realax.real'),X3))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),happ(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.real.real_sub_0'),s('type.realax.real',X2))),s('type.realax.real',X1))),s('type.realax.real',X0)))))),
  inference(ennf_transformation,[],[f41])).
fof(f41,plain,(
  ~! [X0,X1,X2,X3,X4] : ((p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X4),s('type.realax.real',X2),s('type.realax.real',X0)))) & p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X3),s('type.realax.real',X1),s('type.realax.real',X0))))) => p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1x1_3'(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.real.real_sub_0'),s(fun('type.realax.real','type.realax.real'),X4),s(fun('type.realax.real','type.realax.real'),X3))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),happ(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.real.real_sub_0'),s('type.realax.real',X2))),s('type.realax.real',X1))),s('type.realax.real',X0)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8,X9,X10] : ((p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X10),s('type.realax.real',X8),s('type.realax.real',X6)))) & p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X9),s('type.realax.real',X7),s('type.realax.real',X6))))) => p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1x1_3'(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.real.real_sub_0'),s(fun('type.realax.real','type.realax.real'),X10),s(fun('type.realax.real','type.realax.real'),X9))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),happ(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.real.real_sub_0'),s('type.realax.real',X8))),s('type.realax.real',X7))),s('type.realax.real',X6)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8,X9,X10] : ((p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X10),s('type.realax.real',X8),s('type.realax.real',X6)))) & p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X9),s('type.realax.real',X7),s('type.realax.real',X6))))) => p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1x1_3'(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.real.real_sub_0'),s(fun('type.realax.real','type.realax.real'),X10),s(fun('type.realax.real','type.realax.real'),X9))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),happ(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.real.real_sub_0'),s('type.realax.real',X8))),s('type.realax.real',X7))),s('type.realax.real',X6)))))),
  file('thm.transc.DIFF_COMPOSITE|split|4','thm.transc.DIFF_COMPOSITE|split|4')).
fof(f234,plain,(
  ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1x1_3'(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.real.real_sub_0'),s(fun('type.realax.real','type.realax.real'),sK4),s(fun('type.realax.real','type.realax.real'),sK3))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK2),s('type.realax.real',sK1))),s('type.realax.real',sK0)))) | $spl85),
  inference(cnf_transformation,[],[f234_D])).
fof(f234_D,plain,(
  ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1x1_3'(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.real.real_sub_0'),s(fun('type.realax.real','type.realax.real'),sK4),s(fun('type.realax.real','type.realax.real'),sK3))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK2),s('type.realax.real',sK1))),s('type.realax.real',sK0)))) <=> ~$spl85),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl85])])).
fof(f232,plain,(
  ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1x1_3'(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.real.real_sub_0'),s(fun('type.realax.real','type.realax.real'),sK4),s(fun('type.realax.real','type.realax.real'),sK3))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK2),s('type.realax.real',sK1))),s('type.realax.real',sK0)))) | ($spl1 | $spl36)),
  inference(forward_demodulation,[],[f174,f200])).
fof(f194,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) | $spl32),
  inference(cnf_transformation,[],[f194_D])).
fof(f194_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f137,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f104])).
fof(f104,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f53])).
fof(f53,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X11] : (s(bool,'F_0') = s(bool,X11) <=> ~p(s(bool,X11)))),
  file('thm.transc.DIFF_COMPOSITE|split|4','thm.bool.EQ_CLAUSES|split|2')).
fof(f262,plain,(
  s(bool,'F_0') = s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1x1_3'(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.real.real_sub_0'),s(fun('type.realax.real','type.realax.real'),sK4),s(fun('type.realax.real','type.realax.real'),sK3))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK2),s('type.realax.real',sK1))),s('type.realax.real',sK0))) | $spl92),
  inference(cnf_transformation,[],[f262_D])).
fof(f262_D,plain,(
  s(bool,'F_0') = s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1x1_3'(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.real.real_sub_0'),s(fun('type.realax.real','type.realax.real'),sK4),s(fun('type.realax.real','type.realax.real'),sK3))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK2),s('type.realax.real',sK1))),s('type.realax.real',sK0))) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f259,plain,(
  s(bool,'F_0') = s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1x1_3'(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.real.real_sub_0'),s(fun('type.realax.real','type.realax.real'),sK4),s(fun('type.realax.real','type.realax.real'),sK3))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK2),s('type.realax.real',sK1))),s('type.realax.real',sK0))) | ($spl32 | $spl85)),
  inference(resolution,[],[f234,f194])).
fof(f178,plain,(
  p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),sK4),s('type.realax.real',sK2),s('type.realax.real',sK0)))) | $spl4),
  inference(cnf_transformation,[],[f178_D])).
fof(f178_D,plain,(
  p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),sK4),s('type.realax.real',sK2),s('type.realax.real',sK0)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f117,plain,(
  p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),sK4),s('type.realax.real',sK2),s('type.realax.real',sK0))))),
  inference(cnf_transformation,[],[f101])).
fof(f241,plain,(
  s(bool,'T_0') = s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),sK4),s('type.realax.real',sK2),s('type.realax.real',sK0))) | $spl88),
  inference(cnf_transformation,[],[f241_D])).
fof(f241_D,plain,(
  s(bool,'T_0') = s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),sK4),s('type.realax.real',sK2),s('type.realax.real',sK0))) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f237,plain,(
  s(bool,'T_0') = s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),sK4),s('type.realax.real',sK2),s('type.realax.real',sK0))) | ($spl4 | $spl28)),
  inference(resolution,[],[f192,f178])).
fof(f192,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl28),
  inference(cnf_transformation,[],[f192_D])).
fof(f192_D,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f133,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X11] : (s(bool,'T_0') = s(bool,X11) <=> p(s(bool,X11)))),
  file('thm.transc.DIFF_COMPOSITE|split|4','thm.bool.EQ_CLAUSES|split|0')).
fof(f176,plain,(
  p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),sK3),s('type.realax.real',sK1),s('type.realax.real',sK0)))) | $spl2),
  inference(cnf_transformation,[],[f176_D])).
fof(f176_D,plain,(
  p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),sK3),s('type.realax.real',sK1),s('type.realax.real',sK0)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f118,plain,(
  p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),sK3),s('type.realax.real',sK1),s('type.realax.real',sK0))))),
  inference(cnf_transformation,[],[f101])).
fof(f239,plain,(
  s(bool,'T_0') = s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),sK3),s('type.realax.real',sK1),s('type.realax.real',sK0))) | $spl86),
  inference(cnf_transformation,[],[f239_D])).
fof(f239_D,plain,(
  s(bool,'T_0') = s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),sK3),s('type.realax.real',sK1),s('type.realax.real',sK0))) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f236,plain,(
  s(bool,'T_0') = s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),sK3),s('type.realax.real',sK1),s('type.realax.real',sK0))) | ($spl2 | $spl28)),
  inference(resolution,[],[f192,f176])).
fof(f213,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1x1_3'(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.real.real_sub_0'),s(fun('type.realax.real','type.realax.real'),X0),s(fun('type.realax.real','type.realax.real'),X1))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),happ(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.real.real_sub_0'),s('type.realax.real',X2))),s('type.realax.real',X3))),s('type.realax.real',X4)))) | ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X3),s('type.realax.real',X4)))) | ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',X2),s('type.realax.real',X4))))) ) | $spl62),
  inference(cnf_transformation,[],[f213_D])).
fof(f213_D,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1x1_3'(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.real.real_sub_0'),s(fun('type.realax.real','type.realax.real'),X0),s(fun('type.realax.real','type.realax.real'),X1))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),happ(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.real.real_sub_0'),s('type.realax.real',X2))),s('type.realax.real',X3))),s('type.realax.real',X4)))) | ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X3),s('type.realax.real',X4)))) | ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',X2),s('type.realax.real',X4))))) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f168,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1x1_3'(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.real.real_sub_0'),s(fun('type.realax.real','type.realax.real'),X0),s(fun('type.realax.real','type.realax.real'),X1))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),happ(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.real.real_sub_0'),s('type.realax.real',X2))),s('type.realax.real',X3))),s('type.realax.real',X4)))) | ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X3),s('type.realax.real',X4)))) | ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',X2),s('type.realax.real',X4))))) )),
  inference(cnf_transformation,[],[f94])).
fof(f94,plain,(
  ! [X0,X1,X2,X3,X4] : (~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',X2),s('type.realax.real',X4)))) | ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X3),s('type.realax.real',X4)))) | p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1x1_3'(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.real.real_sub_0'),s(fun('type.realax.real','type.realax.real'),X0),s(fun('type.realax.real','type.realax.real'),X1))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),happ(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.real.real_sub_0'),s('type.realax.real',X2))),s('type.realax.real',X3))),s('type.realax.real',X4)))))),
  inference(flattening,[],[f93])).
fof(f93,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',X2),s('type.realax.real',X4)))) | ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X3),s('type.realax.real',X4))))) | p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1x1_3'(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.real.real_sub_0'),s(fun('type.realax.real','type.realax.real'),X0),s(fun('type.realax.real','type.realax.real'),X1))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),happ(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.real.real_sub_0'),s('type.realax.real',X2))),s('type.realax.real',X3))),s('type.realax.real',X4)))))),
  inference(ennf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0,X1,X2,X3,X4] : ((p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',X2),s('type.realax.real',X4)))) & p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X3),s('type.realax.real',X4))))) => p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1x1_3'(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.real.real_sub_0'),s(fun('type.realax.real','type.realax.real'),X0),s(fun('type.realax.real','type.realax.real'),X1))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),happ(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.real.real_sub_0'),s('type.realax.real',X2))),s('type.realax.real',X3))),s('type.realax.real',X4)))))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X10,X9,X8,X7,X6] : ((p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X10),s('type.realax.real',X8),s('type.realax.real',X6)))) & p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X9),s('type.realax.real',X7),s('type.realax.real',X6))))) => p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1x1_3'(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.real.real_sub_0'),s(fun('type.realax.real','type.realax.real'),X10),s(fun('type.realax.real','type.realax.real'),X9))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),happ(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.real.real_sub_0'),s('type.realax.real',X8))),s('type.realax.real',X7))),s('type.realax.real',X6)))))),
  file('thm.transc.DIFF_COMPOSITE|split|4','thm.lim.DIFF_SUB')).
fof(f200,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),happ(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.real.real_sub_0'),s('type.realax.real',X0))),s('type.realax.real',X1)))) ) | $spl36),
  inference(cnf_transformation,[],[f200_D])).
fof(f200_D,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),happ(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.real.real_sub_0'),s('type.realax.real',X0))),s('type.realax.real',X1)))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f155,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),happ(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.real.real_sub_0'),s('type.realax.real',X0))),s('type.realax.real',X1)))) )),
  inference(cnf_transformation,[],[f62])).
fof(f62,plain,(
  ! [X0,X1] : s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),happ(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.real.real_sub_0'),s('type.realax.real',X0))),s('type.realax.real',X1)))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X13,X14] : s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X13),s('type.realax.real',X14))) = s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),happ(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.real.real_sub_0'),s('type.realax.real',X13))),s('type.realax.real',X14)))),
  file('thm.transc.DIFF_COMPOSITE|split|4','ARITY_0_2_const.real.real_sub')).
fof(f227,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1x1_3'(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.real.real_sub_0'),s(fun('type.realax.real','type.realax.real'),X0),s(fun('type.realax.real','type.realax.real'),X1))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X2),s('type.realax.real',X3))),s('type.realax.real',X4)))) | ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X3),s('type.realax.real',X4)))) | ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',X2),s('type.realax.real',X4))))) ) | $spl78),
  inference(cnf_transformation,[],[f227_D])).
fof(f227_D,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1x1_3'(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.real.real_sub_0'),s(fun('type.realax.real','type.realax.real'),X0),s(fun('type.realax.real','type.realax.real'),X1))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X2),s('type.realax.real',X3))),s('type.realax.real',X4)))) | ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X3),s('type.realax.real',X4)))) | ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',X2),s('type.realax.real',X4))))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f226,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1x1_3'(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.real.real_sub_0'),s(fun('type.realax.real','type.realax.real'),X0),s(fun('type.realax.real','type.realax.real'),X1))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X2),s('type.realax.real',X3))),s('type.realax.real',X4)))) | ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X3),s('type.realax.real',X4)))) | ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',X2),s('type.realax.real',X4))))) ) | ($spl36 | $spl62)),
  inference(backward_demodulation,[],[f200,f213])).
fof(f182,plain,(
  ~p(s(bool,'F_0')) | $spl9),
  inference(cnf_transformation,[],[f182_D])).
fof(f182_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f122,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f42])).
fof(f42,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.transc.DIFF_COMPOSITE|split|4',aHL_FALSITY)).
fof(f180,plain,(
  p(s(bool,'T_0')) | $spl6),
  inference(cnf_transformation,[],[f180_D])).
fof(f180_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f120,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.transc.DIFF_COMPOSITE|split|4',aHL_TRUTH)).
fof(f3016,plain,(
  $false | ($spl6 | $spl9 | $spl78 | $spl86 | $spl88 | $spl92)),
  inference(subsumption_resolution,[],[f3015,f180])).
fof(f3015,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl9 | $spl78 | $spl86 | $spl88 | $spl92)),
  inference(forward_demodulation,[],[f3014,f241])).
fof(f3014,plain,(
  ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),sK4),s('type.realax.real',sK2),s('type.realax.real',sK0)))) | ($spl6 | $spl9 | $spl78 | $spl86 | $spl92)),
  inference(subsumption_resolution,[],[f3013,f180])).
fof(f3013,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),sK4),s('type.realax.real',sK2),s('type.realax.real',sK0)))) | ($spl9 | $spl78 | $spl86 | $spl92)),
  inference(forward_demodulation,[],[f3012,f239])).
fof(f3012,plain,(
  ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),sK3),s('type.realax.real',sK1),s('type.realax.real',sK0)))) | ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),sK4),s('type.realax.real',sK2),s('type.realax.real',sK0)))) | ($spl9 | $spl78 | $spl92)),
  inference(subsumption_resolution,[],[f2935,f182])).
fof(f2935,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),sK3),s('type.realax.real',sK1),s('type.realax.real',sK0)))) | ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),sK4),s('type.realax.real',sK2),s('type.realax.real',sK0)))) | ($spl78 | $spl92)),
  inference(superposition,[],[f227,f262])).
% SZS output end Proof for thm.transc
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 65
% Passive clauses: 305
% Generated clauses: 2854
% Final active clauses: 53
% Final passive clauses: 218
% Input formulas: 39
% Initial clauses: 55
% 
% Duplicate literals: 136
% Trivial inequalities: 33
% Fw subsumption resolutions: 146
% Fw demodulations: 248
% Bw demodulations: 33
% Condensations: 129
% Global subsumptions: 4
% 
% Simple tautologies: 85
% Equational tautologies: 18
% Forward subsumptions: 1261
% 
% Binary resolution: 42
% Forward superposition: 1043
% Backward superposition: 573
% Self superposition: 67
% Equality factoring: 30
% Equality resolution: 9
% 
% Split clauses: 27
% Split components: 59
% SAT solver clauses: 547
% SAT solver unit clauses: 79
% SAT solver binary clauses: 29
% SAT solver learnt clauses: 9
% 
% Sat splits: 27
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 530
% TWLsolver calls for satisfiability: 229
% 
% Memory used [KB]: 4989
% Time elapsed: 0.135 s
% ------------------------------
% ------------------------------
% Success in time 0.754 s
