Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.ring
% (82459)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 126
% Passive clauses: 2700
% Generated clauses: 19758
% Final active clauses: 126
% Final passive clauses: 2489
% Input formulas: 69
% Initial clauses: 108
% 
% Duplicate literals: 557
% Trivial inequalities: 300
% Fw subsumption resolutions: 2052
% Fw demodulations: 2752
% 
% Simple tautologies: 323
% Equational tautologies: 44
% Forward subsumptions: 10726
% 
% Binary resolution: 176
% Factoring: 1
% Forward superposition: 7561
% Backward superposition: 6927
% Self superposition: 244
% Equality factoring: 85
% Equality resolution: 34
% 
% Split clauses: 217
% Split components: 450
% SAT solver clauses: 389
% SAT solver unit clauses: 28
% SAT solver binary clauses: 130
% 
% Sat splits: 217
% 
% TWLsolver clauses: 389
% TWLsolver calls for satisfiability: 91
% 
% Memory used [KB]: 14072
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.ring
% (82472)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 126
% Passive clauses: 951
% Generated clauses: 17304
% Final active clauses: 120
% Final passive clauses: 814
% Input formulas: 69
% Initial clauses: 108
% 
% Duplicate literals: 554
% Trivial inequalities: 257
% Fw subsumption resolutions: 1040
% Fw demodulations: 2132
% Bw demodulations: 8
% Condensations: 1217
% Global subsumptions: 4
% 
% Simple tautologies: 210
% Equational tautologies: 42
% Forward subsumptions: 8998
% 
% Binary resolution: 40
% Factoring: 1
% Forward superposition: 5725
% Backward superposition: 4944
% Self superposition: 184
% Equality factoring: 83
% Equality resolution: 34
% 
% Split clauses: 57
% Split components: 118
% SAT solver clauses: 1965
% SAT solver unit clauses: 154
% SAT solver binary clauses: 77
% SAT solver learnt clauses: 10
% SAT solver learnt literals: 2
% 
% Sat splits: 57
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 1947
% TWLsolver calls for satisfiability: 606
% 
% Memory used [KB]: 15863
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.ring
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.ring
% SZS output start Proof for thm.ring
fof(f9920,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f231,f347,f345,f343,f341,f339,f337,f230,f346,f229,f344,f228,f342,f227,f340,f226,f338,f336,f353,f1886,f1887,f2113,f1888,f2114,f2112,f3174,f3176,f3175,f3177,f3179,f3181,f3185,f3187,f3189,f3191,f3195,f3197,f3199,f3201,f3221,f3223,f3229,f3240,f3246,f3432,f3434,f3437,f3440,f3441,f3442,f3443,f3444,f3445,f3446,f3447,f7335,f7345,f7347,f7329,f7349,f7352,f7344,f7214,f7217,f7354,f7223,f7357,f7224,f7358,f7225,f7359,f7226,f7360,f7227,f7361,f7328,f7364,f7231,f7355,f7365,f7886,f7911,f7910,f7343,f8259,f8261,f8264,f8268,f8662,f8665,f9037,f9045,f9509,f9518,f9534,f9699,f9712,f9737,f9878,f9892,f9917,f9918,f9919])).
fof(f9919,plain,(
  $false | ($spl0 | $spl7)),
  inference(subsumption_resolution,[],[f343,f9883])).
fof(f9883,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK3) = s(fun(sK0,fun(sK0,sK0)),sK8) | $spl0),
  inference(forward_demodulation,[],[f9882,f309])).
fof(f309,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(fun(X0,fun(X0,X0)),X3) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),'const.ring.recordtype|46|ring_5'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4),s(fun(X0,X0),X5)))))) )),
  inference(cnf_transformation,[],[f162])).
fof(f162,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(fun(X0,fun(X0,X0)),X3) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),'const.ring.recordtype|46|ring_5'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4),s(fun(X0,X0),X5)))))),
  inference(flattening,[],[f161])).
fof(f161,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5] : s(fun(X0,fun(X0,X0)),X3) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),'const.ring.recordtype|46|ring_5'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4),s(fun(X0,X0),X5)))))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X6] : ! [X25,X7,X19,X41,X42] : s(fun(X6,fun(X6,X6)),X19) = s(fun(X6,fun(X6,X6)),'const.ring.ring_RP_1'(s('type.ring.ring'(X6),'const.ring.recordtype|46|ring_5'(s(X6,X25),s(X6,X7),s(fun(X6,fun(X6,X6)),X19),s(fun(X6,fun(X6,X6)),X41),s(fun(X6,X6),X42)))))),
  file('thm.ring.ring_|49||49|','thm.ring.ring_RP')).
fof(f9882,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK8) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RP_1'(s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))))) | $spl0),
  inference(forward_demodulation,[],[f9825,f309])).
fof(f9825,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RP_1'(s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RP_1'(s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK6),s(sK0,sK7),s(fun(sK0,fun(sK0,sK0)),sK8),s(fun(sK0,fun(sK0,sK0)),sK9),s(fun(sK0,sK0),sK10))))) | $spl0),
  inference(unit_resulting_resolution,[],[f336,f291])).
fof(f291,plain,(
  ( ! [X2,X0,X1] : (s('type.ring.ring'(X0),X1) != s('type.ring.ring'(X0),X2) | s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X2)))) )),
  inference(cnf_transformation,[],[f219])).
fof(f219,plain,(
  ! [X0,X1,X2] : ((s('type.ring.ring'(X0),X1) != s('type.ring.ring'(X0),X2) | (s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X2))) & s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X2))) & s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X2))) & s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X2))) & s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X2))))) & (s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))) != s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X2))) | s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))) != s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X2))) | s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))) != s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X2))) | s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))) != s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X2))) | s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X1))) != s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X2))) | s('type.ring.ring'(X0),X1) = s('type.ring.ring'(X0),X2)))),
  inference(flattening,[],[f218])).
fof(f218,plain,(
  ! [X0,X1,X2] : ((s('type.ring.ring'(X0),X1) != s('type.ring.ring'(X0),X2) | (s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X2))) & s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X2))) & s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X2))) & s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X2))) & s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X2))))) & ((s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))) != s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X2))) | s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))) != s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X2))) | s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))) != s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X2))) | s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))) != s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X2))) | s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X1))) != s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X2)))) | s('type.ring.ring'(X0),X1) = s('type.ring.ring'(X0),X2)))),
  inference(nnf_transformation,[],[f142])).
fof(f142,plain,(
  ! [X0,X1,X2] : (s('type.ring.ring'(X0),X1) = s('type.ring.ring'(X0),X2) <=> (s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X2))) & s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X2))) & s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X2))) & s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X2))) & s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X2)))))),
  inference(flattening,[],[f141])).
fof(f141,plain,(
  ! [X0] : ! [X1,X2] : (s('type.ring.ring'(X0),X1) = s('type.ring.ring'(X0),X2) <=> (s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X2))) & s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X2))) & s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X2))) & s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X2))) & s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X2)))))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X6] : ! [X43,X44] : (s('type.ring.ring'(X6),X43) = s('type.ring.ring'(X6),X44) <=> (s(X6,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X6),X43))) = s(X6,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X6),X44))) & s(X6,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X6),X43))) = s(X6,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X6),X44))) & s(fun(X6,fun(X6,X6)),'const.ring.ring_RP_1'(s('type.ring.ring'(X6),X43))) = s(fun(X6,fun(X6,X6)),'const.ring.ring_RP_1'(s('type.ring.ring'(X6),X44))) & s(fun(X6,fun(X6,X6)),'const.ring.ring_RM_1'(s('type.ring.ring'(X6),X43))) = s(fun(X6,fun(X6,X6)),'const.ring.ring_RM_1'(s('type.ring.ring'(X6),X44))) & s(fun(X6,X6),'const.ring.ring_RN_1'(s('type.ring.ring'(X6),X43))) = s(fun(X6,X6),'const.ring.ring_RN_1'(s('type.ring.ring'(X6),X44)))))),
  file('thm.ring.ring_|49||49|','thm.ring.ring_63_prove_ltb')).
fof(f9918,plain,(
  $false | $spl0),
  inference(global_subsumption,[],[f231,f230,f229,f228,f227,f226,f232,f235,f237,f236,f261,f262,f263,f264,f265,f266,f267,f268,f269,f270,f271,f272,f273,f274,f275,f276,f277,f278,f279,f280,f281,f282,f283,f284,f285,f286,f288,f294,f293,f292,f291,f290,f289,f301,f302,f303,f304,f305,f306,f307,f308,f309,f310,f311,f312,f313,f314,f315,f316,f322,f321,f320,f319,f318,f317,f326,f323,f327,f328,f329,f330,f331,f332,f333,f358,f389,f392,f426,f563,f765,f766,f767,f770,f786,f796,f797,f798,f801,f902,f1154,f1155,f1156,f1165,f1167,f1243,f1263,f1264,f1265,f1266,f1267,f1268,f1251,f1270,f1492,f1652,f1783,f1886,f1887,f1888,f1922,f1981,f2149,f2156,f2289,f2359,f2360,f2361,f2367,f2382,f3174,f5008,f5531,f5999,f6033,f6083,f6364,f6417,f6490,f7335,f7345,f7347,f7329,f7352,f7214,f7217,f7356,f7223,f7224,f7225,f7226,f7227,f7328,f7231,f7312,f7886,f7911,f6253,f6204,f6202,f5875,f5828,f5826,f2365,f2364,f1817,f1815,f800,f769,f7084,f7031,f7030,f4351,f4316,f4314,f2363,f2362,f1577,f1575,f799,f768,f1333,f1372,f2901,f7569,f2318,f2319,f2566,f2567,f7079,f6248,f6120,f5978,f5871,f5620,f2962,f2954,f2953,f2952,f2951,f2950,f2949,f2948,f2947,f2946,f1765,f1202,f1201,f1195,f1193,f1190,f1189,f6258,f6257,f5879,f5878,f1854,f1853,f1174,f631,f7092,f7091,f4355,f4354,f1611,f1610,f1173,f740,f5530,f5007,f3090,f3091,f6987,f6959,f5772,f5759,f2131,f1242,f1205,f1176,f1175,f1172,f1171,f1170,f561,f560,f559,f558,f287,f5623,f5483,f5481,f5626,f4964,f4962,f1269,f1277,f1314,f3086,f3084,f3042,f3040,f7924,f7884,f7882,f7881,f7880,f7879,f7878,f7877,f7876,f7875,f7874,f7873,f7872,f7871,f7870,f7869,f7847,f7846,f7845,f7844,f7843,f7842,f7841,f7840,f7839,f7838,f7837,f7836,f7835,f7834,f7833,f7832,f7353,f7249,f7247,f7246,f7244,f7243,f7242,f7240,f7239,f7238,f7237,f7236,f7235,f7234,f7212,f7210,f7209,f7208,f7207,f7206,f7205,f7204,f7203,f7202,f7201,f7200,f7199,f7198,f7197,f7082,f7080,f7074,f7073,f7072,f7069,f7029,f6939,f6938,f6937,f6936,f6932,f6363,f6300,f6251,f6249,f6243,f6242,f6241,f6238,f5919,f5874,f5872,f5866,f5865,f5864,f5861,f5775,f5774,f5773,f5771,f5770,f5722,f5719,f5716,f5567,f5042,f4855,f4691,f4561,f4379,f4350,f4345,f3173,f3172,f3168,f3166,f3165,f3164,f3163,f3162,f3161,f3160,f3159,f3158,f3157,f3116,f3085,f3079,f3078,f3077,f3074,f2976,f2965,f2964,f2963,f2961,f2960,f2619,f2618,f2613,f2612,f2611,f2515,f2488,f2341,f2190,f2188,f2187,f2186,f2185,f2155,f2154,f2153,f1964,f1940,f1907,f1848,f1847,f1846,f1843,f1819,f1785,f1784,f1782,f1781,f1740,f1739,f1684,f1683,f1682,f1681,f1680,f1679,f1678,f1677,f1606,f1605,f1602,f1579,f1554,f1553,f1479,f1478,f1477,f1476,f1475,f1474,f1473,f1459,f1458,f1411,f1410,f1409,f1408,f1407,f1406,f1405,f1244,f1203,f1200,f1192,f1191,f1169,f735,f718,f699,f686,f685,f649,f648,f647,f646,f645,f612,f534,f490,f450,f391,f241,f240,f234,f8641,f336,f9874,f9875,f9880,f9881,f9883,f9884,f9886,f9887,f9900,f9903,f9906,f9910,f9914,f9891,f9877])).
fof(f9877,plain,(
  s(sK0,sK2) = s(sK0,sK7) | $spl0),
  inference(forward_demodulation,[],[f9876,f307])).
fof(f9876,plain,(
  s(sK0,sK7) = s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))))) | $spl0),
  inference(forward_demodulation,[],[f9827,f307])).
fof(f9827,plain,(
  s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))))) = s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK6),s(sK0,sK7),s(fun(sK0,fun(sK0,sK0)),sK8),s(fun(sK0,fun(sK0,sK0)),sK9),s(fun(sK0,sK0),sK10))))) | $spl0),
  inference(unit_resulting_resolution,[],[f336,f290])).
fof(f9891,plain,(
  s(sK0,sK2) = s(sK0,sK7) | $spl0),
  inference(forward_demodulation,[],[f9890,f307])).
fof(f9890,plain,(
  s(sK0,sK7) = s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))))) | $spl0),
  inference(forward_demodulation,[],[f9822,f307])).
fof(f9822,plain,(
  s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))))) = s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK6),s(sK0,sK7),s(fun(sK0,fun(sK0,sK0)),sK8),s(fun(sK0,fun(sK0,sK0)),sK9),s(fun(sK0,sK0),sK10))))) | $spl0),
  inference(unit_resulting_resolution,[],[f336,f290])).
fof(f9914,plain,(
  ( ! [X24] : (s(fun(sK0,sK0),sK5) = s(fun(sK0,sK0),'const.ring.ring_RN_1'(s('type.ring.ring'(sK0),X24))) | s('type.ring.ring'(sK0),X24) != s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5)))) ) | $spl0),
  inference(forward_demodulation,[],[f9913,f9880])).
fof(f9913,plain,(
  ( ! [X24] : (s(fun(sK0,sK0),sK10) = s(fun(sK0,sK0),'const.ring.ring_RN_1'(s('type.ring.ring'(sK0),X24))) | s('type.ring.ring'(sK0),X24) != s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5)))) ) | $spl0),
  inference(forward_demodulation,[],[f9853,f308])).
fof(f9853,plain,(
  ( ! [X24] : (s('type.ring.ring'(sK0),X24) != s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))) | s(fun(sK0,sK0),'const.ring.ring_RN_1'(s('type.ring.ring'(sK0),X24))) = s(fun(sK0,sK0),'const.ring.ring_RN_1'(s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK6),s(sK0,sK7),s(fun(sK0,fun(sK0,sK0)),sK8),s(fun(sK0,fun(sK0,sK0)),sK9),s(fun(sK0,sK0),sK10)))))) ) | $spl0),
  inference(superposition,[],[f293,f336])).
fof(f9910,plain,(
  ( ! [X22] : (s(fun(sK0,fun(sK0,sK0)),sK4) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),X22))) | s('type.ring.ring'(sK0),X22) != s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5)))) ) | $spl0),
  inference(forward_demodulation,[],[f9909,f9886])).
fof(f9909,plain,(
  ( ! [X22] : (s(fun(sK0,fun(sK0,sK0)),sK9) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),X22))) | s('type.ring.ring'(sK0),X22) != s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5)))) ) | $spl0),
  inference(forward_demodulation,[],[f9851,f310])).
fof(f9851,plain,(
  ( ! [X22] : (s('type.ring.ring'(sK0),X22) != s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))) | s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),X22))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK6),s(sK0,sK7),s(fun(sK0,fun(sK0,sK0)),sK8),s(fun(sK0,fun(sK0,sK0)),sK9),s(fun(sK0,sK0),sK10)))))) ) | $spl0),
  inference(superposition,[],[f292,f336])).
fof(f9906,plain,(
  ( ! [X20] : (s(fun(sK0,fun(sK0,sK0)),sK3) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RP_1'(s('type.ring.ring'(sK0),X20))) | s('type.ring.ring'(sK0),X20) != s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5)))) ) | $spl0),
  inference(forward_demodulation,[],[f9905,f9883])).
fof(f9905,plain,(
  ( ! [X20] : (s(fun(sK0,fun(sK0,sK0)),sK8) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RP_1'(s('type.ring.ring'(sK0),X20))) | s('type.ring.ring'(sK0),X20) != s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5)))) ) | $spl0),
  inference(forward_demodulation,[],[f9849,f309])).
fof(f9849,plain,(
  ( ! [X20] : (s('type.ring.ring'(sK0),X20) != s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))) | s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RP_1'(s('type.ring.ring'(sK0),X20))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RP_1'(s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK6),s(sK0,sK7),s(fun(sK0,fun(sK0,sK0)),sK8),s(fun(sK0,fun(sK0,sK0)),sK9),s(fun(sK0,sK0),sK10)))))) ) | $spl0),
  inference(superposition,[],[f291,f336])).
fof(f9903,plain,(
  ( ! [X18] : (s(sK0,sK7) = s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),X18))) | s('type.ring.ring'(sK0),X18) != s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5)))) ) | $spl0),
  inference(forward_demodulation,[],[f9847,f307])).
fof(f9847,plain,(
  ( ! [X18] : (s('type.ring.ring'(sK0),X18) != s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))) | s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),X18))) = s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK6),s(sK0,sK7),s(fun(sK0,fun(sK0,sK0)),sK8),s(fun(sK0,fun(sK0,sK0)),sK9),s(fun(sK0,sK0),sK10)))))) ) | $spl0),
  inference(superposition,[],[f290,f336])).
fof(f9900,plain,(
  ( ! [X16] : (s(sK0,sK1) = s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),X16))) | s('type.ring.ring'(sK0),X16) != s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5)))) ) | $spl0),
  inference(forward_demodulation,[],[f9899,f9874])).
fof(f9899,plain,(
  ( ! [X16] : (s(sK0,sK6) = s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),X16))) | s('type.ring.ring'(sK0),X16) != s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5)))) ) | $spl0),
  inference(forward_demodulation,[],[f9845,f306])).
fof(f9845,plain,(
  ( ! [X16] : (s('type.ring.ring'(sK0),X16) != s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))) | s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),X16))) = s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK6),s(sK0,sK7),s(fun(sK0,fun(sK0,sK0)),sK8),s(fun(sK0,fun(sK0,sK0)),sK9),s(fun(sK0,sK0),sK10)))))) ) | $spl0),
  inference(superposition,[],[f289,f336])).
fof(f9887,plain,(
  s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))) = s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK7),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))) | $spl0),
  inference(backward_demodulation,[],[f9886,f9884])).
fof(f9886,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK4) = s(fun(sK0,fun(sK0,sK0)),sK9) | $spl0),
  inference(forward_demodulation,[],[f9885,f310])).
fof(f9885,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK9) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))))) | $spl0),
  inference(forward_demodulation,[],[f9824,f310])).
fof(f9824,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK6),s(sK0,sK7),s(fun(sK0,fun(sK0,sK0)),sK8),s(fun(sK0,fun(sK0,sK0)),sK9),s(fun(sK0,sK0),sK10))))) | $spl0),
  inference(unit_resulting_resolution,[],[f336,f292])).
fof(f9884,plain,(
  s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))) = s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK7),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK9),s(fun(sK0,sK0),sK5))) | $spl0),
  inference(backward_demodulation,[],[f9883,f9881])).
fof(f9881,plain,(
  s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))) = s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK7),s(fun(sK0,fun(sK0,sK0)),sK8),s(fun(sK0,fun(sK0,sK0)),sK9),s(fun(sK0,sK0),sK5))) | $spl0),
  inference(backward_demodulation,[],[f9880,f9875])).
fof(f9880,plain,(
  s(fun(sK0,sK0),sK5) = s(fun(sK0,sK0),sK10) | $spl0),
  inference(forward_demodulation,[],[f9879,f308])).
fof(f9879,plain,(
  s(fun(sK0,sK0),sK10) = s(fun(sK0,sK0),'const.ring.ring_RN_1'(s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))))) | $spl0),
  inference(forward_demodulation,[],[f9826,f308])).
fof(f9826,plain,(
  s(fun(sK0,sK0),'const.ring.ring_RN_1'(s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))))) = s(fun(sK0,sK0),'const.ring.ring_RN_1'(s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK6),s(sK0,sK7),s(fun(sK0,fun(sK0,sK0)),sK8),s(fun(sK0,fun(sK0,sK0)),sK9),s(fun(sK0,sK0),sK10))))) | $spl0),
  inference(unit_resulting_resolution,[],[f336,f293])).
fof(f9875,plain,(
  s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))) = s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK7),s(fun(sK0,fun(sK0,sK0)),sK8),s(fun(sK0,fun(sK0,sK0)),sK9),s(fun(sK0,sK0),sK10))) | $spl0),
  inference(backward_demodulation,[],[f9874,f336])).
fof(f9874,plain,(
  s(sK0,sK1) = s(sK0,sK6) | $spl0),
  inference(forward_demodulation,[],[f9873,f306])).
fof(f9873,plain,(
  s(sK0,sK6) = s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))))) | $spl0),
  inference(forward_demodulation,[],[f9828,f306])).
fof(f9828,plain,(
  s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))))) = s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK6),s(sK0,sK7),s(fun(sK0,fun(sK0,sK0)),sK8),s(fun(sK0,fun(sK0,sK0)),sK9),s(fun(sK0,sK0),sK10))))) | $spl0),
  inference(unit_resulting_resolution,[],[f336,f289])).
fof(f8641,plain,(
  ( ! [X68,X66,X67,X65] : (s('type.ring.ring'(X65),X66) != s('type.ring.ring'(X65),X68) | s(fun(X65,fun(X65,X65)),'const.ring.ring_RM_1'(s('type.ring.ring'(X65),X68))) = s(fun(X65,fun(X65,X65)),'const.ring.ring_RM_1'(s('type.ring.ring'(X65),sK11(X66,X67,'type.ring.ring'(X65))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.ring.ring'(X65),bool),X67)))) )),
  inference(superposition,[],[f292,f563])).
fof(f234,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f73])).
fof(f73,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.ring.ring_|49||49|',aHL_FALSITY)).
fof(f240,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f204])).
fof(f204,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f76])).
fof(f76,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X31] : (s(bool,'F_0') = s(bool,X31) <=> ~p(s(bool,X31)))),
  file('thm.ring.ring_|49||49|','thm.bool.EQ_CLAUSES|split|2')).
fof(f241,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f204])).
fof(f391,plain,(
  ( ! [X4,X2,X3] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X2),s(X3,X4))) | s(bool,'F_0') = s(bool,X2)) )),
  inference(superposition,[],[f235,f328])).
fof(f450,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X3),s(fun('type.num.num',bool),X4),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X3)) )),
  inference(superposition,[],[f261,f235])).
fof(f490,plain,(
  ( ! [X6,X4,X5] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X4),'const.pair.,_2'(s(bool,X5),s(X4,X6))))) | s(bool,'F_0') = s(bool,X5)) )),
  inference(superposition,[],[f329,f235])).
fof(f534,plain,(
  ( ! [X6,X4,X5] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X4,bool),'const.pair.,_2'(s(X4,X5),s(bool,X6))))) | s(bool,'F_0') = s(bool,X6)) )),
  inference(superposition,[],[f330,f235])).
fof(f612,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X4))) | s(bool,'T_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|48|_fupd_2'(s(fun(bool,bool),X3),s('type.ring.ring'(bool),X4)))))) )),
  inference(superposition,[],[f262,f235])).
fof(f645,plain,(
  ( ! [X14,X12,X15,X13,X11] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X12),s(X14,X15))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X11),'const.pair.,_2'(s(bool,X12),s(X11,X13)))))) )),
  inference(superposition,[],[f391,f329])).
fof(f646,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X16),s(X18,X19))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X16),s(fun('type.num.num',bool),X17),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f391,f261])).
fof(f647,plain,(
  ( ! [X24,X23,X21,X22,X20] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X22),s(X23,X24))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X20,bool),'const.pair.,_2'(s(X20,X21),s(bool,X22)))))) )),
  inference(superposition,[],[f391,f330])).
fof(f648,plain,(
  ( ! [X28,X26,X27,X25] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X26))),s(X27,X28))) | s(bool,'F_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|48|_fupd_2'(s(fun(bool,bool),X25),s('type.ring.ring'(bool),X26)))))) )),
  inference(superposition,[],[f391,f262])).
fof(f649,plain,(
  ( ! [X30,X33,X31,X29,X32] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X29),s(X32,X33))) | s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,X29),s(X30,X31)))) )),
  inference(superposition,[],[f391,f328])).
fof(f685,plain,(
  ( ! [X26,X23] : (~p(s(bool,X26)) | s(bool,X23) = s(bool,X26) | s(bool,'F_0') = s(bool,X23)) )),
  inference(forward_demodulation,[],[f684,f328])).
fof(f684,plain,(
  ( ! [X26,X24,X23,X25] : (~p(s(bool,X26)) | s(bool,X26) = s(bool,'const.combin.K_2'(s(bool,X23),s(X24,X25))) | s(bool,'F_0') = s(bool,X23)) )),
  inference(subsumption_resolution,[],[f671,f232])).
fof(f671,plain,(
  ( ! [X26,X24,X23,X25] : (~p(s(bool,'T_0')) | ~p(s(bool,X26)) | s(bool,X26) = s(bool,'const.combin.K_2'(s(bool,X23),s(X24,X25))) | s(bool,'F_0') = s(bool,X23)) )),
  inference(superposition,[],[f326,f391])).
fof(f686,plain,(
  ( ! [X30,X28,X29,X27] : (s(bool,X30) != s(bool,'const.combin.K_2'(s(bool,X27),s(X28,X29))) | s(bool,'T_0') = s(bool,X30) | s(bool,'F_0') = s(bool,X27)) )),
  inference(subsumption_resolution,[],[f672,f232])).
fof(f672,plain,(
  ( ! [X30,X28,X29,X27] : (~p(s(bool,'T_0')) | s(bool,X30) != s(bool,'const.combin.K_2'(s(bool,X27),s(X28,X29))) | s(bool,'T_0') = s(bool,X30) | s(bool,'F_0') = s(bool,X27)) )),
  inference(superposition,[],[f392,f391])).
fof(f699,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X7) = s(bool,'const.combin.K_2'(s(bool,X4),s(X5,X6))) | s(bool,'F_0') = s(bool,X7) | s(bool,'F_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f692,f232])).
fof(f692,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,'T_0')) | s(bool,X7) = s(bool,'const.combin.K_2'(s(bool,X4),s(X5,X6))) | s(bool,'F_0') = s(bool,X7) | s(bool,'F_0') = s(bool,X4)) )),
  inference(superposition,[],[f685,f391])).
fof(f718,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X4))) | s(bool,'T_0') = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|49|_fupd_2'(s(fun(bool,bool),X3),s('type.ring.ring'(bool),X4)))))) )),
  inference(superposition,[],[f263,f235])).
fof(f735,plain,(
  ( ! [X59,X57,X60,X58] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X58))),s(X59,X60))) | s(bool,'F_0') = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|49|_fupd_2'(s(fun(bool,bool),X57),s('type.ring.ring'(bool),X58)))))) )),
  inference(superposition,[],[f391,f263])).
fof(f1169,plain,(
  ( ! [X15,X16] : (s(bool,'F_0') = s(bool,X15) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X16))) | s(bool,'T_0') = s(bool,sK11(X15,X16,bool))) )),
  inference(superposition,[],[f563,f235])).
fof(f1191,plain,(
  ( ! [X45,X44] : (~p(s(bool,sK11(X44,X45,bool))) | s(bool,'F_0') != s(bool,X44) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X45)))) )),
  inference(superposition,[],[f240,f563])).
fof(f1192,plain,(
  ( ! [X47,X46] : (p(s(bool,X46)) | s(bool,'F_0') = s(bool,sK11(X46,X47,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X47)))) )),
  inference(superposition,[],[f241,f563])).
fof(f1200,plain,(
  ( ! [X70,X72,X71,X69] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X69),s(X71,X72))) | s(bool,'F_0') = s(bool,sK11(X69,X70,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X70)))) )),
  inference(superposition,[],[f391,f563])).
fof(f1203,plain,(
  ( ! [X80,X81,X79] : (~p(s(bool,X79)) | s(bool,X81) = s(bool,sK11(X79,X80,bool)) | s(bool,'F_0') = s(bool,X81) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X80)))) )),
  inference(superposition,[],[f685,f563])).
fof(f1244,plain,(
  ( ! [X59,X57,X58,X56] : (s(bool,X58) = s(bool,'|dollar|exists_1'(s(fun(X56,bool),X57))) | s(bool,'F_0') = s(bool,X58) | s(X56,X59) = s(X56,sK11(X59,X57,X56))) )),
  inference(subsumption_resolution,[],[f1224,f232])).
fof(f1224,plain,(
  ( ! [X59,X57,X58,X56] : (~p(s(bool,'T_0')) | s(bool,X58) = s(bool,'|dollar|exists_1'(s(fun(X56,bool),X57))) | s(bool,'F_0') = s(bool,X58) | s(X56,X59) = s(X56,sK11(X59,X57,X56))) )),
  inference(superposition,[],[f685,f563])).
fof(f1405,plain,(
  ( ! [X35,X33,X36,X34] : (s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X34),s(fun('type.num.num',bool),X36),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X33),'const.pair.,_2'(s(bool,X34),s(X33,X35)))))) )),
  inference(superposition,[],[f450,f329])).
fof(f1406,plain,(
  ( ! [X39,X37,X38] : (s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X37),s(fun('type.num.num',bool),X39),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X37),s(fun('type.num.num',bool),X38),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f450,f261])).
fof(f1407,plain,(
  ( ! [X43,X41,X42,X40] : (s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X42),s(fun('type.num.num',bool),X43),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X40,bool),'const.pair.,_2'(s(X40,X41),s(bool,X42)))))) )),
  inference(superposition,[],[f450,f330])).
fof(f1408,plain,(
  ( ! [X45,X46,X44] : (s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X45))),s(fun('type.num.num',bool),X46),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|49|_fupd_2'(s(fun(bool,bool),X44),s('type.ring.ring'(bool),X45)))))) )),
  inference(superposition,[],[f450,f263])).
fof(f1409,plain,(
  ( ! [X47,X48,X49] : (s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X48))),s(fun('type.num.num',bool),X49),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|48|_fupd_2'(s(fun(bool,bool),X47),s('type.ring.ring'(bool),X48)))))) )),
  inference(superposition,[],[f450,f262])).
fof(f1410,plain,(
  ( ! [X52,X50,X53,X51] : (s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X50),s(fun('type.num.num',bool),X53),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,X50),s(X51,X52)))) )),
  inference(superposition,[],[f450,f328])).
fof(f1411,plain,(
  ( ! [X54,X56,X55] : (s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X54),s(fun('type.num.num',bool),X56),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,sK11(X54,X55,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X55)))) )),
  inference(superposition,[],[f450,f563])).
fof(f1458,plain,(
  ( ! [X35,X33,X34] : (s(bool,X35) != s(bool,'const.ind_type.FCONS_3'(s(bool,X33),s(fun('type.num.num',bool),X34),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X35) | s(bool,'F_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f1434,f232])).
fof(f1434,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,'T_0')) | s(bool,X35) != s(bool,'const.ind_type.FCONS_3'(s(bool,X33),s(fun('type.num.num',bool),X34),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X35) | s(bool,'F_0') = s(bool,X33)) )),
  inference(superposition,[],[f392,f450])).
fof(f1459,plain,(
  ( ! [X39,X41,X40] : (s(bool,X41) = s(bool,'const.ind_type.FCONS_3'(s(bool,X39),s(fun('type.num.num',bool),X40),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X41) | s(bool,'F_0') = s(bool,X39)) )),
  inference(subsumption_resolution,[],[f1436,f232])).
fof(f1436,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,'T_0')) | s(bool,X41) = s(bool,'const.ind_type.FCONS_3'(s(bool,X39),s(fun('type.num.num',bool),X40),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X41) | s(bool,'F_0') = s(bool,X39)) )),
  inference(superposition,[],[f685,f450])).
fof(f1473,plain,(
  ( ! [X52,X50,X53,X51,X49] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X52),'const.pair.,_2'(s(bool,X50),s(X52,X53))))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X49),'const.pair.,_2'(s(bool,X50),s(X49,X51)))))) )),
  inference(superposition,[],[f490,f329])).
fof(f1474,plain,(
  ( ! [X57,X54,X56,X55] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X56),'const.pair.,_2'(s(bool,X54),s(X56,X57))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X54),s(fun('type.num.num',bool),X55),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f490,f261])).
fof(f1475,plain,(
  ( ! [X61,X59,X62,X60,X58] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X61),'const.pair.,_2'(s(bool,X60),s(X61,X62))))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X58,bool),'const.pair.,_2'(s(X58,X59),s(bool,X60)))))) )),
  inference(superposition,[],[f490,f330])).
fof(f1476,plain,(
  ( ! [X66,X64,X65,X63] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X65),'const.pair.,_2'(s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X64))),s(X65,X66))))) | s(bool,'F_0') = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|49|_fupd_2'(s(fun(bool,bool),X63),s('type.ring.ring'(bool),X64)))))) )),
  inference(superposition,[],[f490,f263])).
fof(f1477,plain,(
  ( ! [X70,X68,X69,X67] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X69),'const.pair.,_2'(s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X68))),s(X69,X70))))) | s(bool,'F_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|48|_fupd_2'(s(fun(bool,bool),X67),s('type.ring.ring'(bool),X68)))))) )),
  inference(superposition,[],[f490,f262])).
fof(f1478,plain,(
  ( ! [X74,X72,X71,X75,X73] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X74),'const.pair.,_2'(s(bool,X71),s(X74,X75))))) | s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,X71),s(X72,X73)))) )),
  inference(superposition,[],[f490,f328])).
fof(f1479,plain,(
  ( ! [X78,X76,X79,X77] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X78),'const.pair.,_2'(s(bool,X76),s(X78,X79))))) | s(bool,'F_0') = s(bool,sK11(X76,X77,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X77)))) )),
  inference(superposition,[],[f490,f563])).
fof(f1553,plain,(
  ( ! [X47,X48,X46,X49] : (s(bool,X49) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X46),'const.pair.,_2'(s(bool,X47),s(X46,X48))))) | s(bool,'T_0') = s(bool,X49) | s(bool,'F_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f1518,f232])).
fof(f1518,plain,(
  ( ! [X47,X48,X46,X49] : (~p(s(bool,'T_0')) | s(bool,X49) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X46),'const.pair.,_2'(s(bool,X47),s(X46,X48))))) | s(bool,'T_0') = s(bool,X49) | s(bool,'F_0') = s(bool,X47)) )),
  inference(superposition,[],[f392,f490])).
fof(f1554,plain,(
  ( ! [X61,X59,X60,X58] : (s(bool,X61) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X58),'const.pair.,_2'(s(bool,X59),s(X58,X60))))) | s(bool,'F_0') = s(bool,X61) | s(bool,'F_0') = s(bool,X59)) )),
  inference(subsumption_resolution,[],[f1521,f232])).
fof(f1521,plain,(
  ( ! [X61,X59,X60,X58] : (~p(s(bool,'T_0')) | s(bool,X61) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X58),'const.pair.,_2'(s(bool,X59),s(X58,X60))))) | s(bool,'F_0') = s(bool,X61) | s(bool,'F_0') = s(bool,X59)) )),
  inference(superposition,[],[f685,f490])).
fof(f1579,plain,(
  ( ! [X15,X16] : (s(bool,'F_0') = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X16))) | s(bool,'T_0') = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_RN_fupd_2'(s(fun(fun(bool,bool),fun(bool,bool)),X15),s('type.ring.ring'(bool),X16)))))) )),
  inference(superposition,[],[f264,f235])).
fof(f1602,plain,(
  ( ! [X80,X78,X79,X77] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X78))),s(X79,X80))) | s(bool,'F_0') = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_RN_fupd_2'(s(fun(fun(bool,bool),fun(bool,bool)),X77),s('type.ring.ring'(bool),X78)))))) )),
  inference(superposition,[],[f391,f264])).
fof(f1605,plain,(
  ( ! [X88,X87,X89] : (s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X88))),s(fun('type.num.num',bool),X89),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_RN_fupd_2'(s(fun(fun(bool,bool),fun(bool,bool)),X87),s('type.ring.ring'(bool),X88)))))) )),
  inference(superposition,[],[f450,f264])).
fof(f1606,plain,(
  ( ! [X92,X90,X93,X91] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X92),'const.pair.,_2'(s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X91))),s(X92,X93))))) | s(bool,'F_0') = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_RN_fupd_2'(s(fun(fun(bool,bool),fun(bool,bool)),X90),s('type.ring.ring'(bool),X91)))))) )),
  inference(superposition,[],[f490,f264])).
fof(f1677,plain,(
  ( ! [X57,X54,X58,X56,X55] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X57,bool),'const.pair.,_2'(s(X57,X58),s(bool,X55))))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X54),'const.pair.,_2'(s(bool,X55),s(X54,X56)))))) )),
  inference(superposition,[],[f534,f329])).
fof(f1678,plain,(
  ( ! [X61,X59,X62,X60] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X61,bool),'const.pair.,_2'(s(X61,X62),s(bool,X59))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X59),s(fun('type.num.num',bool),X60),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f534,f261])).
fof(f1679,plain,(
  ( ! [X66,X64,X67,X65,X63] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X66,bool),'const.pair.,_2'(s(X66,X67),s(bool,X65))))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X63,bool),'const.pair.,_2'(s(X63,X64),s(bool,X65)))))) )),
  inference(superposition,[],[f534,f330])).
fof(f1680,plain,(
  ( ! [X70,X68,X71,X69] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X70,bool),'const.pair.,_2'(s(X70,X71),s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X69))))))) | s(bool,'F_0') = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|49|_fupd_2'(s(fun(bool,bool),X68),s('type.ring.ring'(bool),X69)))))) )),
  inference(superposition,[],[f534,f263])).
fof(f1681,plain,(
  ( ! [X74,X72,X75,X73] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X74,bool),'const.pair.,_2'(s(X74,X75),s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X73))))))) | s(bool,'F_0') = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_RN_fupd_2'(s(fun(fun(bool,bool),fun(bool,bool)),X72),s('type.ring.ring'(bool),X73)))))) )),
  inference(superposition,[],[f534,f264])).
fof(f1682,plain,(
  ( ! [X78,X76,X79,X77] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X78,bool),'const.pair.,_2'(s(X78,X79),s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X77))))))) | s(bool,'F_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|48|_fupd_2'(s(fun(bool,bool),X76),s('type.ring.ring'(bool),X77)))))) )),
  inference(superposition,[],[f534,f262])).
fof(f1683,plain,(
  ( ! [X80,X83,X81,X84,X82] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X83,bool),'const.pair.,_2'(s(X83,X84),s(bool,X80))))) | s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,X80),s(X81,X82)))) )),
  inference(superposition,[],[f534,f328])).
fof(f1684,plain,(
  ( ! [X88,X87,X85,X86] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X87,bool),'const.pair.,_2'(s(X87,X88),s(bool,X85))))) | s(bool,'F_0') = s(bool,sK11(X85,X86,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X86)))) )),
  inference(superposition,[],[f534,f563])).
fof(f1739,plain,(
  ( ! [X47,X48,X46,X49] : (s(bool,X49) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X46,bool),'const.pair.,_2'(s(X46,X47),s(bool,X48))))) | s(bool,'T_0') = s(bool,X49) | s(bool,'F_0') = s(bool,X48)) )),
  inference(subsumption_resolution,[],[f1700,f232])).
fof(f1700,plain,(
  ( ! [X47,X48,X46,X49] : (~p(s(bool,'T_0')) | s(bool,X49) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X46,bool),'const.pair.,_2'(s(X46,X47),s(bool,X48))))) | s(bool,'T_0') = s(bool,X49) | s(bool,'F_0') = s(bool,X48)) )),
  inference(superposition,[],[f392,f534])).
fof(f1740,plain,(
  ( ! [X66,X64,X65,X63] : (s(bool,X66) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X63,bool),'const.pair.,_2'(s(X63,X64),s(bool,X65))))) | s(bool,'F_0') = s(bool,X66) | s(bool,'F_0') = s(bool,X65)) )),
  inference(subsumption_resolution,[],[f1704,f232])).
fof(f1704,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,'T_0')) | s(bool,X66) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X63,bool),'const.pair.,_2'(s(X63,X64),s(bool,X65))))) | s(bool,'F_0') = s(bool,X66) | s(bool,'F_0') = s(bool,X65)) )),
  inference(superposition,[],[f685,f534])).
fof(f1781,plain,(
  ( ! [X30,X33,X31,X34,X32] : (s(bool,X33) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X30),'const.pair.,_2'(s(bool,X31),s(X30,X32))))) | s(bool,X33) = s(bool,X34) | s(bool,'T_0') = s(bool,X34) | s(bool,'F_0') = s(bool,X31)) )),
  inference(subsumption_resolution,[],[f1753,f232])).
fof(f1753,plain,(
  ( ! [X30,X33,X31,X34,X32] : (~p(s(bool,'T_0')) | s(bool,X33) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X30),'const.pair.,_2'(s(bool,X31),s(X30,X32))))) | s(bool,X33) = s(bool,X34) | s(bool,'T_0') = s(bool,X34) | s(bool,'F_0') = s(bool,X31)) )),
  inference(superposition,[],[f1154,f490])).
fof(f1782,plain,(
  ( ! [X37,X35,X38,X36] : (s(bool,X37) = s(bool,'const.ind_type.FCONS_3'(s(bool,X35),s(fun('type.num.num',bool),X36),s('type.num.num','const.num.|48|_0'))) | s(bool,X37) = s(bool,X38) | s(bool,'T_0') = s(bool,X38) | s(bool,'F_0') = s(bool,X35)) )),
  inference(subsumption_resolution,[],[f1754,f232])).
fof(f1754,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'T_0')) | s(bool,X37) = s(bool,'const.ind_type.FCONS_3'(s(bool,X35),s(fun('type.num.num',bool),X36),s('type.num.num','const.num.|48|_0'))) | s(bool,X37) = s(bool,X38) | s(bool,'T_0') = s(bool,X38) | s(bool,'F_0') = s(bool,X35)) )),
  inference(superposition,[],[f1154,f450])).
fof(f1784,plain,(
  ( ! [X47,X45,X48,X46,X44] : (s(bool,X47) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X44,bool),'const.pair.,_2'(s(X44,X45),s(bool,X46))))) | s(bool,X47) = s(bool,X48) | s(bool,'T_0') = s(bool,X48) | s(bool,'F_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f1756,f232])).
fof(f1756,plain,(
  ( ! [X47,X45,X48,X46,X44] : (~p(s(bool,'T_0')) | s(bool,X47) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X44,bool),'const.pair.,_2'(s(X44,X45),s(bool,X46))))) | s(bool,X47) = s(bool,X48) | s(bool,'T_0') = s(bool,X48) | s(bool,'F_0') = s(bool,X46)) )),
  inference(superposition,[],[f1154,f534])).
fof(f1785,plain,(
  ( ! [X52,X50,X53,X51,X49] : (s(bool,X52) = s(bool,'const.combin.K_2'(s(bool,X49),s(X50,X51))) | s(bool,X52) = s(bool,X53) | s(bool,'T_0') = s(bool,X53) | s(bool,'F_0') = s(bool,X49)) )),
  inference(subsumption_resolution,[],[f1757,f232])).
fof(f1757,plain,(
  ( ! [X52,X50,X53,X51,X49] : (~p(s(bool,'T_0')) | s(bool,X52) = s(bool,'const.combin.K_2'(s(bool,X49),s(X50,X51))) | s(bool,X52) = s(bool,X53) | s(bool,'T_0') = s(bool,X53) | s(bool,'F_0') = s(bool,X49)) )),
  inference(superposition,[],[f1154,f391])).
fof(f1819,plain,(
  ( ! [X15,X16] : (s(bool,'F_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X16))) | s(bool,'T_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_RN_fupd_2'(s(fun(fun(bool,bool),fun(bool,bool)),X15),s('type.ring.ring'(bool),X16)))))) )),
  inference(superposition,[],[f265,f235])).
fof(f1843,plain,(
  ( ! [X83,X81,X84,X82] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X82))),s(X83,X84))) | s(bool,'F_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_RN_fupd_2'(s(fun(fun(bool,bool),fun(bool,bool)),X81),s('type.ring.ring'(bool),X82)))))) )),
  inference(superposition,[],[f391,f265])).
fof(f1846,plain,(
  ( ! [X92,X93,X91] : (s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X92))),s(fun('type.num.num',bool),X93),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_RN_fupd_2'(s(fun(fun(bool,bool),fun(bool,bool)),X91),s('type.ring.ring'(bool),X92)))))) )),
  inference(superposition,[],[f450,f265])).
fof(f1847,plain,(
  ( ! [X94,X97,X95,X96] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X96),'const.pair.,_2'(s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X95))),s(X96,X97))))) | s(bool,'F_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_RN_fupd_2'(s(fun(fun(bool,bool),fun(bool,bool)),X94),s('type.ring.ring'(bool),X95)))))) )),
  inference(superposition,[],[f490,f265])).
fof(f1848,plain,(
  ( ! [X101,X99,X100,X98] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X100,bool),'const.pair.,_2'(s(X100,X101),s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X99))))))) | s(bool,'F_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_RN_fupd_2'(s(fun(fun(bool,bool),fun(bool,bool)),X98),s('type.ring.ring'(bool),X99)))))) )),
  inference(superposition,[],[f534,f265])).
fof(f1907,plain,(
  ( ! [X154,X152,X151,X149,X153,X150] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X149),s(X151,X152))) | s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,sK11(X149,X150,bool)),s(X153,X154))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X150)))) )),
  inference(superposition,[],[f649,f563])).
fof(f1940,plain,(
  ( ! [X39,X36,X40] : (s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,X36),s(X39,X40))) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f649,f328])).
fof(f1964,plain,(
  ( ! [X154,X152,X151,X149,X153,X150] : (s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,X149),s(X151,X152))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,sK11(X149,X150,bool)),s(X153,X154))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X150)))) )),
  inference(superposition,[],[f649,f563])).
fof(f2153,plain,(
  ( ! [X76,X74,X72,X71,X75,X73] : (s(bool,X74) != s(bool,'const.combin.K_2'(s(bool,X71),s(X72,X73))) | s(bool,'T_0') = s(bool,X74) | s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,X71),s(X75,X76)))) )),
  inference(subsumption_resolution,[],[f2011,f232])).
fof(f2011,plain,(
  ( ! [X76,X74,X72,X71,X75,X73] : (~p(s(bool,'T_0')) | s(bool,X74) != s(bool,'const.combin.K_2'(s(bool,X71),s(X72,X73))) | s(bool,'T_0') = s(bool,X74) | s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,X71),s(X75,X76)))) )),
  inference(superposition,[],[f392,f649])).
fof(f2154,plain,(
  ( ! [X103,X107,X105,X108,X106,X104] : (s(bool,X106) = s(bool,'const.combin.K_2'(s(bool,X103),s(X104,X105))) | s(bool,'F_0') = s(bool,X106) | s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,X103),s(X107,X108)))) )),
  inference(subsumption_resolution,[],[f2016,f232])).
fof(f2016,plain,(
  ( ! [X103,X107,X105,X108,X106,X104] : (~p(s(bool,'T_0')) | s(bool,X106) = s(bool,'const.combin.K_2'(s(bool,X103),s(X104,X105))) | s(bool,'F_0') = s(bool,X106) | s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,X103),s(X107,X108)))) )),
  inference(superposition,[],[f685,f649])).
fof(f2155,plain,(
  ( ! [X111,X109,X114,X112,X110,X115,X113] : (s(bool,X112) = s(bool,'const.combin.K_2'(s(bool,X109),s(X110,X111))) | s(bool,X112) = s(bool,X113) | s(bool,'T_0') = s(bool,X113) | s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,X109),s(X114,X115)))) )),
  inference(subsumption_resolution,[],[f2017,f232])).
fof(f2017,plain,(
  ( ! [X111,X109,X114,X112,X110,X115,X113] : (~p(s(bool,'T_0')) | s(bool,X112) = s(bool,'const.combin.K_2'(s(bool,X109),s(X110,X111))) | s(bool,X112) = s(bool,X113) | s(bool,'T_0') = s(bool,X113) | s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,X109),s(X114,X115)))) )),
  inference(superposition,[],[f1154,f649])).
fof(f2185,plain,(
  ( ! [X88,X87,X85,X86] : (s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X86))),s(X87,X88))) | s(bool,'T_0') = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|49|_fupd_2'(s(fun(bool,bool),X85),s('type.ring.ring'(bool),X86)))))) )),
  inference(superposition,[],[f1940,f263])).
fof(f2186,plain,(
  ( ! [X92,X90,X91,X89] : (s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X90))),s(X91,X92))) | s(bool,'T_0') = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_RN_fupd_2'(s(fun(fun(bool,bool),fun(bool,bool)),X89),s('type.ring.ring'(bool),X90)))))) )),
  inference(superposition,[],[f1940,f264])).
fof(f2187,plain,(
  ( ! [X94,X95,X93,X96] : (s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X94))),s(X95,X96))) | s(bool,'T_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|48|_fupd_2'(s(fun(bool,bool),X93),s('type.ring.ring'(bool),X94)))))) )),
  inference(superposition,[],[f1940,f262])).
fof(f2188,plain,(
  ( ! [X99,X97,X100,X98] : (s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X98))),s(X99,X100))) | s(bool,'T_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_RN_fupd_2'(s(fun(fun(bool,bool),fun(bool,bool)),X97),s('type.ring.ring'(bool),X98)))))) )),
  inference(superposition,[],[f1940,f265])).
fof(f2190,plain,(
  ( ! [X109,X107,X108,X106] : (s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,X106),s(X108,X109))) | s(bool,'T_0') = s(bool,sK11(X106,X107,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X107)))) )),
  inference(superposition,[],[f1940,f563])).
fof(f2341,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X0,X1)))) )),
  inference(unit_resulting_resolution,[],[f358,f2289])).
fof(f2488,plain,(
  ( ! [X39,X40] : (s(bool,X39) != s(bool,'const.combin.K_2'(s(bool,'F_0'),s(bool,X39))) | s(bool,X39) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f2341,f389])).
fof(f2515,plain,(
  ( ! [X14,X12,X13] : (s(bool,X14) = s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X12,X13))) | s(bool,'T_0') = s(bool,X14)) )),
  inference(trivial_inequality_removal,[],[f2508])).
fof(f2508,plain,(
  ( ! [X14,X12,X13] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X14) = s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X12,X13))) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f2341,f389])).
fof(f2611,plain,(
  ( ! [X85,X83,X81,X84,X82] : (s(bool,X82) = s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X84,X85))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X81),'const.pair.,_2'(s(bool,X82),s(X81,X83)))))) )),
  inference(superposition,[],[f2515,f329])).
fof(f2612,plain,(
  ( ! [X88,X87,X89,X86] : (s(bool,X86) = s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X88,X89))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X86),s(fun('type.num.num',bool),X87),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f2515,f261])).
fof(f2613,plain,(
  ( ! [X94,X92,X90,X93,X91] : (s(bool,X92) = s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X93,X94))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X90,bool),'const.pair.,_2'(s(X90,X91),s(bool,X92)))))) )),
  inference(superposition,[],[f2515,f330])).
fof(f2618,plain,(
  ( ! [X111,X114,X112,X115,X113] : (s(bool,X111) = s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X114,X115))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X111),s(X112,X113)))) )),
  inference(superposition,[],[f2515,f328])).
fof(f2619,plain,(
  ( ! [X118,X116,X119,X117] : (s(bool,X116) = s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X118,X119))) | s(bool,'T_0') = s(bool,sK11(X116,X117,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X117)))) )),
  inference(superposition,[],[f2515,f563])).
fof(f2960,plain,(
  ( ! [X28,X26,X29,X27] : (s(bool,'T_0') = s(bool,sK11('const.pair.FST_1'(s('type.pair.prod'(bool,X26),'const.pair.,_2'(s(bool,X27),s(X26,X28)))),X29,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X29))) | s(bool,'F_0') = s(bool,X27)) )),
  inference(subsumption_resolution,[],[f2937,f232])).
fof(f2937,plain,(
  ( ! [X28,X26,X29,X27] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,sK11('const.pair.FST_1'(s('type.pair.prod'(bool,X26),'const.pair.,_2'(s(bool,X27),s(X26,X28)))),X29,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X29))) | s(bool,'F_0') = s(bool,X27)) )),
  inference(superposition,[],[f1190,f490])).
fof(f2961,plain,(
  ( ! [X30,X31,X32] : (s(bool,'T_0') = s(bool,sK11('const.ind_type.FCONS_3'(s(bool,X30),s(fun('type.num.num',bool),X31),s('type.num.num','const.num.|48|_0')),X32,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X32))) | s(bool,'F_0') = s(bool,X30)) )),
  inference(subsumption_resolution,[],[f2938,f232])).
fof(f2938,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,sK11('const.ind_type.FCONS_3'(s(bool,X30),s(fun('type.num.num',bool),X31),s('type.num.num','const.num.|48|_0')),X32,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X32))) | s(bool,'F_0') = s(bool,X30)) )),
  inference(superposition,[],[f1190,f450])).
fof(f2963,plain,(
  ( ! [X39,X37,X38,X40] : (s(bool,'T_0') = s(bool,sK11('const.pair.SND_1'(s('type.pair.prod'(X37,bool),'const.pair.,_2'(s(X37,X38),s(bool,X39)))),X40,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X40))) | s(bool,'F_0') = s(bool,X39)) )),
  inference(subsumption_resolution,[],[f2940,f232])).
fof(f2940,plain,(
  ( ! [X39,X37,X38,X40] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,sK11('const.pair.SND_1'(s('type.pair.prod'(X37,bool),'const.pair.,_2'(s(X37,X38),s(bool,X39)))),X40,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X40))) | s(bool,'F_0') = s(bool,X39)) )),
  inference(superposition,[],[f1190,f534])).
fof(f2964,plain,(
  ( ! [X54,X52,X56,X55,X53,X51] : (s(bool,'T_0') = s(bool,sK11('const.combin.K_2'(s(bool,X51),s(X52,X53)),X54,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X54))) | s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,X51),s(X55,X56)))) )),
  inference(subsumption_resolution,[],[f2943,f232])).
fof(f2943,plain,(
  ( ! [X54,X52,X56,X55,X53,X51] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,sK11('const.combin.K_2'(s(bool,X51),s(X52,X53)),X54,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X54))) | s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,X51),s(X55,X56)))) )),
  inference(superposition,[],[f1190,f649])).
fof(f2965,plain,(
  ( ! [X59,X57,X60,X58] : (s(bool,'T_0') = s(bool,sK11('const.combin.K_2'(s(bool,X57),s(X58,X59)),X60,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X60))) | s(bool,'F_0') = s(bool,X57)) )),
  inference(subsumption_resolution,[],[f2944,f232])).
fof(f2944,plain,(
  ( ! [X59,X57,X60,X58] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,sK11('const.combin.K_2'(s(bool,X57),s(X58,X59)),X60,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X60))) | s(bool,'F_0') = s(bool,X57)) )),
  inference(superposition,[],[f1190,f391])).
fof(f2976,plain,(
  ( ! [X6,X4,X5] : (s(bool,'F_0') != s(bool,X4) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X5))) | s(bool,X6) = s(bool,sK11(X4,X5,bool)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f2967,f232])).
fof(f2967,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X4) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X5))) | s(bool,X6) = s(bool,sK11(X4,X5,bool)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f1191,f389])).
fof(f3074,plain,(
  ( ! [X140,X138,X141,X139,X137] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,happ(s(fun('type.num.num',bool),X138),s('type.num.num',X139))),s(X140,X141))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X137),s(fun('type.num.num',bool),X138),s('type.num.num','const.num.SUC_1'(s('type.num.num',X139)))))) )),
  inference(superposition,[],[f391,f301])).
fof(f3077,plain,(
  ( ! [X152,X151,X153,X150] : (s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,happ(s(fun('type.num.num',bool),X151),s('type.num.num',X152))),s(fun('type.num.num',bool),X153),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X150),s(fun('type.num.num',bool),X151),s('type.num.num','const.num.SUC_1'(s('type.num.num',X152)))))) )),
  inference(superposition,[],[f450,f301])).
fof(f3078,plain,(
  ( ! [X158,X156,X154,X157,X155] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X157),'const.pair.,_2'(s(bool,happ(s(fun('type.num.num',bool),X155),s('type.num.num',X156))),s(X157,X158))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X154),s(fun('type.num.num',bool),X155),s('type.num.num','const.num.SUC_1'(s('type.num.num',X156)))))) )),
  inference(superposition,[],[f490,f301])).
fof(f3079,plain,(
  ( ! [X163,X161,X159,X162,X160] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X162,bool),'const.pair.,_2'(s(X162,X163),s(bool,happ(s(fun('type.num.num',bool),X160),s('type.num.num',X161))))))) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X159),s(fun('type.num.num',bool),X160),s('type.num.num','const.num.SUC_1'(s('type.num.num',X161)))))) )),
  inference(superposition,[],[f534,f301])).
fof(f3085,plain,(
  ( ! [X191,X194,X192,X195,X193] : (s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,happ(s(fun('type.num.num',bool),X192),s('type.num.num',X193))),s(X194,X195))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X191),s(fun('type.num.num',bool),X192),s('type.num.num','const.num.SUC_1'(s('type.num.num',X193)))))) )),
  inference(superposition,[],[f1940,f301])).
fof(f3116,plain,(
  ( ! [X6,X4,X5] : (s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),X5),s('type.num.num',X6))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X4),s(fun('type.num.num',bool),X5),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))) )),
  inference(forward_demodulation,[],[f3039,f328])).
fof(f3039,plain,(
  ( ! [X6,X4,X8,X7,X5] : (s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X7,X8))) = s(bool,happ(s(fun('type.num.num',bool),X5),s('type.num.num',X6))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X4),s(fun('type.num.num',bool),X5),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))) )),
  inference(superposition,[],[f301,f2515])).
fof(f3157,plain,(
  ( ! [X68,X66,X67,X65] : (p(s(bool,X66)) | s(bool,'F_0') = s(bool,sK11('const.pair.FST_1'(s('type.pair.prod'(bool,X65),'const.pair.,_2'(s(bool,X66),s(X65,X67)))),X68,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X68)))) )),
  inference(superposition,[],[f1192,f329])).
fof(f3158,plain,(
  ( ! [X70,X71,X69] : (p(s(bool,X69)) | s(bool,'F_0') = s(bool,sK11('const.ind_type.FCONS_3'(s(bool,X69),s(fun('type.num.num',bool),X70),s('type.num.num','const.num.|48|_0')),X71,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X71)))) )),
  inference(superposition,[],[f1192,f261])).
fof(f3159,plain,(
  ( ! [X74,X72,X75,X73] : (p(s(bool,happ(s(fun('type.num.num',bool),X73),s('type.num.num',X74)))) | s(bool,'F_0') = s(bool,sK11('const.ind_type.FCONS_3'(s(bool,X72),s(fun('type.num.num',bool),X73),s('type.num.num','const.num.SUC_1'(s('type.num.num',X74)))),X75,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X75)))) )),
  inference(superposition,[],[f1192,f301])).
fof(f3160,plain,(
  ( ! [X78,X76,X79,X77] : (p(s(bool,X78)) | s(bool,'F_0') = s(bool,sK11('const.pair.SND_1'(s('type.pair.prod'(X76,bool),'const.pair.,_2'(s(X76,X77),s(bool,X78)))),X79,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X79)))) )),
  inference(superposition,[],[f1192,f330])).
fof(f3161,plain,(
  ( ! [X80,X81,X82] : (p(s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X81)))) | s(bool,'F_0') = s(bool,sK11('const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|49|_fupd_2'(s(fun(bool,bool),X80),s('type.ring.ring'(bool),X81)))),X82,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X82)))) )),
  inference(superposition,[],[f1192,f263])).
fof(f3162,plain,(
  ( ! [X85,X83,X84] : (p(s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X84)))) | s(bool,'F_0') = s(bool,sK11('const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_RN_fupd_2'(s(fun(fun(bool,bool),fun(bool,bool)),X83),s('type.ring.ring'(bool),X84)))),X85,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X85)))) )),
  inference(superposition,[],[f1192,f264])).
fof(f3163,plain,(
  ( ! [X88,X87,X86] : (p(s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X87)))) | s(bool,'F_0') = s(bool,sK11('const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|48|_fupd_2'(s(fun(bool,bool),X86),s('type.ring.ring'(bool),X87)))),X88,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X88)))) )),
  inference(superposition,[],[f1192,f262])).
fof(f3164,plain,(
  ( ! [X90,X91,X89] : (p(s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X90)))) | s(bool,'F_0') = s(bool,sK11('const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_RN_fupd_2'(s(fun(fun(bool,bool),fun(bool,bool)),X89),s('type.ring.ring'(bool),X90)))),X91,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X91)))) )),
  inference(superposition,[],[f1192,f265])).
fof(f3165,plain,(
  ( ! [X94,X92,X95,X93] : (p(s(bool,X92)) | s(bool,'F_0') = s(bool,sK11('const.combin.K_2'(s(bool,X92),s(X93,X94)),X95,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X95)))) )),
  inference(superposition,[],[f1192,f328])).
fof(f3166,plain,(
  ( ! [X97,X98,X96] : (p(s(bool,X96)) | s(bool,'F_0') = s(bool,sK11(sK11(X96,X97,bool),X98,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X98))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X97)))) )),
  inference(superposition,[],[f1192,f563])).
fof(f3168,plain,(
  ( ! [X0,X3] : (s(bool,'F_0') = s(bool,sK11(X0,X3,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X3))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f3167,f234])).
fof(f3167,plain,(
  ( ! [X0,X3] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,sK11(X0,X3,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X3))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(forward_demodulation,[],[f3139,f328])).
fof(f3139,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X1,X2)))) | s(bool,'F_0') = s(bool,sK11(X0,X3,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X3))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1192,f2515])).
fof(f3172,plain,(
  ( ! [X43,X41,X44,X42] : (s(bool,'F_0') = s(bool,sK11('const.combin.K_2'(s(bool,X41),s(X42,X43)),X44,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X44))) | s(bool,'T_0') = s(bool,X41)) )),
  inference(subsumption_resolution,[],[f3152,f234])).
fof(f3152,plain,(
  ( ! [X43,X41,X44,X42] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,sK11('const.combin.K_2'(s(bool,X41),s(X42,X43)),X44,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X44))) | s(bool,'T_0') = s(bool,X41)) )),
  inference(superposition,[],[f1192,f1940])).
fof(f3173,plain,(
  ( ! [X47,X45,X50,X48,X46,X49] : (s(bool,'F_0') = s(bool,sK11('const.combin.K_2'(s(bool,X45),s(X46,X47)),X48,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X48))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X45),s(X49,X50)))) )),
  inference(subsumption_resolution,[],[f3153,f234])).
fof(f3153,plain,(
  ( ! [X47,X45,X50,X48,X46,X49] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,sK11('const.combin.K_2'(s(bool,X45),s(X46,X47)),X48,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X48))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X45),s(X49,X50)))) )),
  inference(superposition,[],[f1192,f649])).
fof(f4345,plain,(
  ( ! [X94,X95,X93,X96] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X94))),s(X95,X96))) | s(bool,'F_0') = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_RP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X93),s('type.ring.ring'(bool),X94)))))) )),
  inference(superposition,[],[f391,f269])).
fof(f4350,plain,(
  ( ! [X111,X112,X110,X113] : (s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X111))),s(X112,X113))) | s(bool,'T_0') = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_RP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X110),s('type.ring.ring'(bool),X111)))))) )),
  inference(superposition,[],[f1940,f269])).
fof(f4379,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X4))) | s(bool,'T_0') = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_RP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X3),s('type.ring.ring'(bool),X4)))))) )),
  inference(forward_demodulation,[],[f4313,f328])).
fof(f4313,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X4))) = s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X5,X6))) | s(bool,'T_0') = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_RP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X3),s('type.ring.ring'(bool),X4)))))) )),
  inference(superposition,[],[f269,f2515])).
fof(f4561,plain,(
  ( ! [X68,X66,X67] : (s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X67))),s(fun('type.num.num',bool),X68),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_RP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X66),s('type.ring.ring'(bool),X67)))))) )),
  inference(superposition,[],[f450,f269])).
fof(f4691,plain,(
  ( ! [X94,X92,X93,X91] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X93),'const.pair.,_2'(s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X92))),s(X93,X94))))) | s(bool,'F_0') = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_RP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X91),s('type.ring.ring'(bool),X92)))))) )),
  inference(superposition,[],[f490,f269])).
fof(f4855,plain,(
  ( ! [X99,X97,X98,X96] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X98,bool),'const.pair.,_2'(s(X98,X99),s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X97))))))) | s(bool,'F_0') = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_RP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X96),s('type.ring.ring'(bool),X97)))))) )),
  inference(superposition,[],[f534,f269])).
fof(f5042,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X3),s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X4))))) | s(bool,'T_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|49|_fupd_2'(s(fun(bool,bool),X3),s('type.ring.ring'(bool),X4)))))) )),
  inference(forward_demodulation,[],[f4961,f328])).
fof(f4961,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X5,X6))) = s(bool,happ(s(fun(bool,bool),X3),s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X4))))) | s(bool,'T_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|49|_fupd_2'(s(fun(bool,bool),X3),s('type.ring.ring'(bool),X4)))))) )),
  inference(superposition,[],[f272,f2515])).
fof(f5567,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X3),s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X4))))) | s(bool,'T_0') = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|48|_fupd_2'(s(fun(bool,bool),X3),s('type.ring.ring'(bool),X4)))))) )),
  inference(forward_demodulation,[],[f5480,f328])).
fof(f5480,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X5,X6))) = s(bool,happ(s(fun(bool,bool),X3),s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X4))))) | s(bool,'T_0') = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|48|_fupd_2'(s(fun(bool,bool),X3),s('type.ring.ring'(bool),X4)))))) )),
  inference(superposition,[],[f273,f2515])).
fof(f5716,plain,(
  ( ! [X80,X81,X82] : (p(s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X81)))) | s(bool,'F_0') = s(bool,sK11('const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_RP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X80),s('type.ring.ring'(bool),X81)))),X82,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X82)))) )),
  inference(superposition,[],[f1192,f269])).
fof(f5719,plain,(
  ( ! [X90,X91,X89] : (p(s(bool,happ(s(fun(bool,bool),X89),s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X90)))))) | s(bool,'F_0') = s(bool,sK11('const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|48|_fupd_2'(s(fun(bool,bool),X89),s('type.ring.ring'(bool),X90)))),X91,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X91)))) )),
  inference(superposition,[],[f1192,f273])).
fof(f5722,plain,(
  ( ! [X99,X100,X98] : (p(s(bool,happ(s(fun(bool,bool),X98),s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X99)))))) | s(bool,'F_0') = s(bool,sK11('const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|49|_fupd_2'(s(fun(bool,bool),X98),s('type.ring.ring'(bool),X99)))),X100,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X100)))) )),
  inference(superposition,[],[f1192,f272])).
fof(f5770,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (s(bool,'|dollar|exists_1'(s(fun(X47,bool),X48))) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X44),'const.pair.,_2'(s(bool,X45),s(X44,X46))))) | s(X47,X49) = s(X47,sK11(X49,X48,X47)) | s(bool,'F_0') = s(bool,X45)) )),
  inference(subsumption_resolution,[],[f5738,f232])).
fof(f5738,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (~p(s(bool,'T_0')) | s(bool,'|dollar|exists_1'(s(fun(X47,bool),X48))) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X44),'const.pair.,_2'(s(bool,X45),s(X44,X46))))) | s(X47,X49) = s(X47,sK11(X49,X48,X47)) | s(bool,'F_0') = s(bool,X45)) )),
  inference(superposition,[],[f1242,f490])).
fof(f5771,plain,(
  ( ! [X54,X52,X50,X53,X51] : (s(bool,'|dollar|exists_1'(s(fun(X52,bool),X53))) = s(bool,'const.ind_type.FCONS_3'(s(bool,X50),s(fun('type.num.num',bool),X51),s('type.num.num','const.num.|48|_0'))) | s(X52,X54) = s(X52,sK11(X54,X53,X52)) | s(bool,'F_0') = s(bool,X50)) )),
  inference(subsumption_resolution,[],[f5739,f232])).
fof(f5739,plain,(
  ( ! [X54,X52,X50,X53,X51] : (~p(s(bool,'T_0')) | s(bool,'|dollar|exists_1'(s(fun(X52,bool),X53))) = s(bool,'const.ind_type.FCONS_3'(s(bool,X50),s(fun('type.num.num',bool),X51),s('type.num.num','const.num.|48|_0'))) | s(X52,X54) = s(X52,sK11(X54,X53,X52)) | s(bool,'F_0') = s(bool,X50)) )),
  inference(superposition,[],[f1242,f450])).
fof(f5773,plain,(
  ( ! [X61,X66,X64,X62,X65,X63] : (s(bool,'|dollar|exists_1'(s(fun(X64,bool),X65))) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X61,bool),'const.pair.,_2'(s(X61,X62),s(bool,X63))))) | s(X64,X66) = s(X64,sK11(X66,X65,X64)) | s(bool,'F_0') = s(bool,X63)) )),
  inference(subsumption_resolution,[],[f5741,f232])).
fof(f5741,plain,(
  ( ! [X61,X66,X64,X62,X65,X63] : (~p(s(bool,'T_0')) | s(bool,'|dollar|exists_1'(s(fun(X64,bool),X65))) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X61,bool),'const.pair.,_2'(s(X61,X62),s(bool,X63))))) | s(X64,X66) = s(X64,sK11(X66,X65,X64)) | s(bool,'F_0') = s(bool,X63)) )),
  inference(superposition,[],[f1242,f534])).
fof(f5774,plain,(
  ( ! [X80,X78,X76,X81,X79,X77,X75,X82] : (s(bool,'|dollar|exists_1'(s(fun(X78,bool),X79))) = s(bool,'const.combin.K_2'(s(bool,X75),s(X76,X77))) | s(X78,X80) = s(X78,sK11(X80,X79,X78)) | s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,X75),s(X81,X82)))) )),
  inference(subsumption_resolution,[],[f5743,f232])).
fof(f5743,plain,(
  ( ! [X80,X78,X76,X81,X79,X77,X75,X82] : (~p(s(bool,'T_0')) | s(bool,'|dollar|exists_1'(s(fun(X78,bool),X79))) = s(bool,'const.combin.K_2'(s(bool,X75),s(X76,X77))) | s(X78,X80) = s(X78,sK11(X80,X79,X78)) | s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,X75),s(X81,X82)))) )),
  inference(superposition,[],[f1242,f649])).
fof(f5775,plain,(
  ( ! [X94,X92,X90,X93,X91,X89] : (s(bool,'|dollar|exists_1'(s(fun(X92,bool),X93))) = s(bool,'const.combin.K_2'(s(bool,X89),s(X90,X91))) | s(X92,X94) = s(X92,sK11(X94,X93,X92)) | s(bool,'F_0') = s(bool,X89)) )),
  inference(subsumption_resolution,[],[f5745,f232])).
fof(f5745,plain,(
  ( ! [X94,X92,X90,X93,X91,X89] : (~p(s(bool,'T_0')) | s(bool,'|dollar|exists_1'(s(fun(X92,bool),X93))) = s(bool,'const.combin.K_2'(s(bool,X89),s(X90,X91))) | s(X92,X94) = s(X92,sK11(X94,X93,X92)) | s(bool,'F_0') = s(bool,X89)) )),
  inference(superposition,[],[f1242,f391])).
fof(f5861,plain,(
  ( ! [X116,X114,X115,X113] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X114))),s(X115,X116))) | s(bool,'F_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_RP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X113),s('type.ring.ring'(bool),X114)))))) )),
  inference(superposition,[],[f391,f275])).
fof(f5864,plain,(
  ( ! [X125,X123,X124] : (s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X124))),s(fun('type.num.num',bool),X125),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_RP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X123),s('type.ring.ring'(bool),X124)))))) )),
  inference(superposition,[],[f450,f275])).
fof(f5865,plain,(
  ( ! [X127,X128,X126,X129] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X128),'const.pair.,_2'(s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X127))),s(X128,X129))))) | s(bool,'F_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_RP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X126),s('type.ring.ring'(bool),X127)))))) )),
  inference(superposition,[],[f490,f275])).
fof(f5866,plain,(
  ( ! [X132,X130,X133,X131] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X132,bool),'const.pair.,_2'(s(X132,X133),s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X131))))))) | s(bool,'F_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_RP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X130),s('type.ring.ring'(bool),X131)))))) )),
  inference(superposition,[],[f534,f275])).
fof(f5872,plain,(
  ( ! [X158,X156,X157] : (p(s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X157)))) | s(bool,'F_0') = s(bool,sK11('const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_RP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X156),s('type.ring.ring'(bool),X157)))),X158,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X158)))) )),
  inference(superposition,[],[f1192,f275])).
fof(f5874,plain,(
  ( ! [X167,X165,X166,X164] : (s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X165))),s(X166,X167))) | s(bool,'T_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_RP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X164),s('type.ring.ring'(bool),X165)))))) )),
  inference(superposition,[],[f1940,f275])).
fof(f5919,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X4))) | s(bool,'T_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_RP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X3),s('type.ring.ring'(bool),X4)))))) )),
  inference(forward_demodulation,[],[f5825,f328])).
fof(f5825,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X4))) = s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X5,X6))) | s(bool,'T_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_RP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X3),s('type.ring.ring'(bool),X4)))))) )),
  inference(superposition,[],[f275,f2515])).
fof(f6238,plain,(
  ( ! [X121,X122,X120,X119] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X120))),s(X121,X122))) | s(bool,'F_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_RM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X119),s('type.ring.ring'(bool),X120)))))) )),
  inference(superposition,[],[f391,f276])).
fof(f6241,plain,(
  ( ! [X130,X131,X129] : (s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X130))),s(fun('type.num.num',bool),X131),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_RM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X129),s('type.ring.ring'(bool),X130)))))) )),
  inference(superposition,[],[f450,f276])).
fof(f6242,plain,(
  ( ! [X134,X132,X135,X133] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X134),'const.pair.,_2'(s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X133))),s(X134,X135))))) | s(bool,'F_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_RM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X132),s('type.ring.ring'(bool),X133)))))) )),
  inference(superposition,[],[f490,f276])).
fof(f6243,plain,(
  ( ! [X138,X136,X139,X137] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X138,bool),'const.pair.,_2'(s(X138,X139),s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X137))))))) | s(bool,'F_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_RM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X136),s('type.ring.ring'(bool),X137)))))) )),
  inference(superposition,[],[f534,f276])).
fof(f6249,plain,(
  ( ! [X163,X164,X162] : (p(s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X163)))) | s(bool,'F_0') = s(bool,sK11('const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_RM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X162),s('type.ring.ring'(bool),X163)))),X164,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X164)))) )),
  inference(superposition,[],[f1192,f276])).
fof(f6251,plain,(
  ( ! [X173,X171,X172,X170] : (s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X171))),s(X172,X173))) | s(bool,'T_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_RM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X170),s('type.ring.ring'(bool),X171)))))) )),
  inference(superposition,[],[f1940,f276])).
fof(f6300,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X4))) | s(bool,'T_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_RM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X3),s('type.ring.ring'(bool),X4)))))) )),
  inference(forward_demodulation,[],[f6201,f328])).
fof(f6201,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X4))) = s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X5,X6))) | s(bool,'T_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_RM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X3),s('type.ring.ring'(bool),X4)))))) )),
  inference(superposition,[],[f276,f2515])).
fof(f6363,plain,(
  ( ! [X231,X229,X233,X230,X228,X232] : (s(bool,X228) = s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X230,X231))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,sK11(X228,X229,bool)),s(X232,X233))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X229)))) )),
  inference(superposition,[],[f2618,f563])).
fof(f6932,plain,(
  ( ! [X111,X109,X107,X112,X110,X108] : (s(bool,X110) != s(bool,'const.combin.K_2'(s(bool,X107),s(X108,X109))) | s(bool,'T_0') = s(bool,X110) | s(bool,X107) = s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X111,X112)))) )),
  inference(subsumption_resolution,[],[f6680,f232])).
fof(f6680,plain,(
  ( ! [X111,X109,X107,X112,X110,X108] : (~p(s(bool,'T_0')) | s(bool,X110) != s(bool,'const.combin.K_2'(s(bool,X107),s(X108,X109))) | s(bool,'T_0') = s(bool,X110) | s(bool,X107) = s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X111,X112)))) )),
  inference(superposition,[],[f392,f2618])).
fof(f6936,plain,(
  ( ! [X158,X161,X159,X157,X162,X160] : (s(bool,X160) = s(bool,'const.combin.K_2'(s(bool,X157),s(X158,X159))) | s(bool,'F_0') = s(bool,X160) | s(bool,X157) = s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X161,X162)))) )),
  inference(subsumption_resolution,[],[f6687,f232])).
fof(f6687,plain,(
  ( ! [X158,X161,X159,X157,X162,X160] : (~p(s(bool,'T_0')) | s(bool,X160) = s(bool,'const.combin.K_2'(s(bool,X157),s(X158,X159))) | s(bool,'F_0') = s(bool,X160) | s(bool,X157) = s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X161,X162)))) )),
  inference(superposition,[],[f685,f2618])).
fof(f6937,plain,(
  ( ! [X167,X165,X163,X169,X166,X164,X168] : (s(bool,X166) = s(bool,'const.combin.K_2'(s(bool,X163),s(X164,X165))) | s(bool,X166) = s(bool,X167) | s(bool,'T_0') = s(bool,X167) | s(bool,X163) = s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X168,X169)))) )),
  inference(subsumption_resolution,[],[f6688,f232])).
fof(f6688,plain,(
  ( ! [X167,X165,X163,X169,X166,X164,X168] : (~p(s(bool,'T_0')) | s(bool,X166) = s(bool,'const.combin.K_2'(s(bool,X163),s(X164,X165))) | s(bool,X166) = s(bool,X167) | s(bool,'T_0') = s(bool,X167) | s(bool,X163) = s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X168,X169)))) )),
  inference(superposition,[],[f1154,f2618])).
fof(f6938,plain,(
  ( ! [X175,X173,X171,X174,X172,X170] : (s(bool,'T_0') = s(bool,sK11('const.combin.K_2'(s(bool,X170),s(X171,X172)),X173,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X173))) | s(bool,X170) = s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X174,X175)))) )),
  inference(subsumption_resolution,[],[f6689,f232])).
fof(f6689,plain,(
  ( ! [X175,X173,X171,X174,X172,X170] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,sK11('const.combin.K_2'(s(bool,X170),s(X171,X172)),X173,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X173))) | s(bool,X170) = s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X174,X175)))) )),
  inference(superposition,[],[f1190,f2618])).
fof(f6939,plain,(
  ( ! [X189,X187,X185,X182,X188,X186,X184,X183] : (s(bool,'|dollar|exists_1'(s(fun(X185,bool),X186))) = s(bool,'const.combin.K_2'(s(bool,X182),s(X183,X184))) | s(X185,X187) = s(X185,sK11(X187,X186,X185)) | s(bool,X182) = s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X188,X189)))) )),
  inference(subsumption_resolution,[],[f6691,f232])).
fof(f6691,plain,(
  ( ! [X189,X187,X185,X182,X188,X186,X184,X183] : (~p(s(bool,'T_0')) | s(bool,'|dollar|exists_1'(s(fun(X185,bool),X186))) = s(bool,'const.combin.K_2'(s(bool,X182),s(X183,X184))) | s(X185,X187) = s(X185,sK11(X187,X186,X185)) | s(bool,X182) = s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X188,X189)))) )),
  inference(superposition,[],[f1242,f2618])).
fof(f7029,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X4))) | s(bool,'T_0') = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_RM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X3),s('type.ring.ring'(bool),X4)))))) )),
  inference(superposition,[],[f277,f235])).
fof(f7069,plain,(
  ( ! [X134,X132,X133,X131] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X132))),s(X133,X134))) | s(bool,'F_0') = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_RM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X131),s('type.ring.ring'(bool),X132)))))) )),
  inference(superposition,[],[f391,f277])).
fof(f7072,plain,(
  ( ! [X142,X143,X141] : (s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X142))),s(fun('type.num.num',bool),X143),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_RM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X141),s('type.ring.ring'(bool),X142)))))) )),
  inference(superposition,[],[f450,f277])).
fof(f7073,plain,(
  ( ! [X146,X144,X147,X145] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X146),'const.pair.,_2'(s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X145))),s(X146,X147))))) | s(bool,'F_0') = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_RM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X144),s('type.ring.ring'(bool),X145)))))) )),
  inference(superposition,[],[f490,f277])).
fof(f7074,plain,(
  ( ! [X151,X149,X150,X148] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X150,bool),'const.pair.,_2'(s(X150,X151),s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X149))))))) | s(bool,'F_0') = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_RM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X148),s('type.ring.ring'(bool),X149)))))) )),
  inference(superposition,[],[f534,f277])).
fof(f7080,plain,(
  ( ! [X175,X176,X174] : (p(s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X175)))) | s(bool,'F_0') = s(bool,sK11('const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_RM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X174),s('type.ring.ring'(bool),X175)))),X176,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X176)))) )),
  inference(superposition,[],[f1192,f277])).
fof(f7082,plain,(
  ( ! [X185,X182,X184,X183] : (s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X183))),s(X184,X185))) | s(bool,'T_0') = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_RM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X182),s('type.ring.ring'(bool),X183)))))) )),
  inference(superposition,[],[f1940,f277])).
fof(f7197,plain,(
  ( ! [X127,X125,X123,X128,X126,X124] : (s(bool,X124) != s(bool,'const.combin.K_2'(s(bool,X126),s(X127,X128))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X123),'const.pair.,_2'(s(bool,X124),s(X123,X125))))) | s(bool,'F_0') = s(bool,X126)) )),
  inference(superposition,[],[f686,f329])).
fof(f7198,plain,(
  ( ! [X132,X130,X133,X131,X129] : (s(bool,X129) != s(bool,'const.combin.K_2'(s(bool,X131),s(X132,X133))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X129),s(fun('type.num.num',bool),X130),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X131)) )),
  inference(superposition,[],[f686,f261])).
fof(f7199,plain,(
  ( ! [X134,X138,X136,X135,X139,X137] : (s(bool,'const.combin.K_2'(s(bool,X137),s(X138,X139))) != s(bool,happ(s(fun('type.num.num',bool),X135),s('type.num.num',X136))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X134),s(fun('type.num.num',bool),X135),s('type.num.num','const.num.SUC_1'(s('type.num.num',X136))))) | s(bool,'F_0') = s(bool,X137)) )),
  inference(superposition,[],[f686,f301])).
fof(f7200,plain,(
  ( ! [X144,X142,X140,X145,X143,X141] : (s(bool,X142) != s(bool,'const.combin.K_2'(s(bool,X143),s(X144,X145))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X140,bool),'const.pair.,_2'(s(X140,X141),s(bool,X142))))) | s(bool,'F_0') = s(bool,X143)) )),
  inference(superposition,[],[f686,f330])).
fof(f7201,plain,(
  ( ! [X146,X149,X147,X150,X148] : (s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X147))) != s(bool,'const.combin.K_2'(s(bool,X148),s(X149,X150))) | s(bool,'T_0') = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_RP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X146),s('type.ring.ring'(bool),X147))))) | s(bool,'F_0') = s(bool,X148)) )),
  inference(superposition,[],[f686,f269])).
fof(f7202,plain,(
  ( ! [X154,X152,X151,X155,X153] : (s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X152))) != s(bool,'const.combin.K_2'(s(bool,X153),s(X154,X155))) | s(bool,'T_0') = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_RN_fupd_2'(s(fun(fun(bool,bool),fun(bool,bool)),X151),s('type.ring.ring'(bool),X152))))) | s(bool,'F_0') = s(bool,X153)) )),
  inference(superposition,[],[f686,f264])).
fof(f7203,plain,(
  ( ! [X158,X156,X159,X157,X160] : (s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X157))) != s(bool,'const.combin.K_2'(s(bool,X158),s(X159,X160))) | s(bool,'T_0') = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|49|_fupd_2'(s(fun(bool,bool),X156),s('type.ring.ring'(bool),X157))))) | s(bool,'F_0') = s(bool,X158)) )),
  inference(superposition,[],[f686,f263])).
fof(f7204,plain,(
  ( ! [X165,X163,X161,X164,X162] : (s(bool,'const.combin.K_2'(s(bool,X163),s(X164,X165))) != s(bool,happ(s(fun(bool,bool),X161),s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X162))))) | s(bool,'T_0') = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|48|_fupd_2'(s(fun(bool,bool),X161),s('type.ring.ring'(bool),X162))))) | s(bool,'F_0') = s(bool,X163)) )),
  inference(superposition,[],[f686,f273])).
fof(f7205,plain,(
  ( ! [X167,X169,X166,X170,X168] : (s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X167))) != s(bool,'const.combin.K_2'(s(bool,X168),s(X169,X170))) | s(bool,'T_0') = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_RM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X166),s('type.ring.ring'(bool),X167))))) | s(bool,'F_0') = s(bool,X168)) )),
  inference(superposition,[],[f686,f277])).
fof(f7206,plain,(
  ( ! [X175,X173,X171,X174,X172] : (s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X172))) != s(bool,'const.combin.K_2'(s(bool,X173),s(X174,X175))) | s(bool,'T_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_RP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X171),s('type.ring.ring'(bool),X172))))) | s(bool,'F_0') = s(bool,X173)) )),
  inference(superposition,[],[f686,f275])).
fof(f7207,plain,(
  ( ! [X177,X180,X178,X176,X179] : (s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X177))) != s(bool,'const.combin.K_2'(s(bool,X178),s(X179,X180))) | s(bool,'T_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_RN_fupd_2'(s(fun(fun(bool,bool),fun(bool,bool)),X176),s('type.ring.ring'(bool),X177))))) | s(bool,'F_0') = s(bool,X178)) )),
  inference(superposition,[],[f686,f265])).
fof(f7208,plain,(
  ( ! [X185,X182,X184,X183,X181] : (s(bool,'const.combin.K_2'(s(bool,X183),s(X184,X185))) != s(bool,happ(s(fun(bool,bool),X181),s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X182))))) | s(bool,'T_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|49|_fupd_2'(s(fun(bool,bool),X181),s('type.ring.ring'(bool),X182))))) | s(bool,'F_0') = s(bool,X183)) )),
  inference(superposition,[],[f686,f272])).
fof(f7209,plain,(
  ( ! [X189,X187,X190,X188,X186] : (s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X187))) != s(bool,'const.combin.K_2'(s(bool,X188),s(X189,X190))) | s(bool,'T_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|48|_fupd_2'(s(fun(bool,bool),X186),s('type.ring.ring'(bool),X187))))) | s(bool,'F_0') = s(bool,X188)) )),
  inference(superposition,[],[f686,f262])).
fof(f7210,plain,(
  ( ! [X191,X194,X192,X195,X193] : (s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X192))) != s(bool,'const.combin.K_2'(s(bool,X193),s(X194,X195))) | s(bool,'T_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_RM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X191),s('type.ring.ring'(bool),X192))))) | s(bool,'F_0') = s(bool,X193)) )),
  inference(superposition,[],[f686,f276])).
fof(f7212,plain,(
  ( ! [X206,X204,X202,X205,X203] : (s(bool,X202) != s(bool,'const.combin.K_2'(s(bool,X204),s(X205,X206))) | s(bool,'T_0') = s(bool,sK11(X202,X203,bool)) | s(bool,'F_0') = s(bool,X204) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X203)))) )),
  inference(superposition,[],[f686,f563])).
fof(f7234,plain,(
  ( ! [X127,X125,X130,X128,X126,X129] : (s(bool,X128) != s(bool,'const.combin.K_2'(s(bool,X126),s(X129,X130))) | s(bool,'T_0') = s(bool,X128) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X125),'const.pair.,_2'(s(bool,X126),s(X125,X127)))))) )),
  inference(superposition,[],[f686,f329])).
fof(f7235,plain,(
  ( ! [X134,X132,X135,X133,X131] : (s(bool,X133) != s(bool,'const.combin.K_2'(s(bool,X131),s(X134,X135))) | s(bool,'T_0') = s(bool,X133) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X131),s(fun('type.num.num',bool),X132),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f686,f261])).
fof(f7236,plain,(
  ( ! [X140,X138,X136,X141,X139,X137] : (s(bool,X139) != s(bool,'const.combin.K_2'(s(bool,happ(s(fun('type.num.num',bool),X137),s('type.num.num',X138))),s(X140,X141))) | s(bool,'T_0') = s(bool,X139) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X136),s(fun('type.num.num',bool),X137),s('type.num.num','const.num.SUC_1'(s('type.num.num',X138)))))) )),
  inference(superposition,[],[f686,f301])).
fof(f7237,plain,(
  ( ! [X146,X144,X142,X147,X145,X143] : (s(bool,X145) != s(bool,'const.combin.K_2'(s(bool,X144),s(X146,X147))) | s(bool,'T_0') = s(bool,X145) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X142,bool),'const.pair.,_2'(s(X142,X143),s(bool,X144)))))) )),
  inference(superposition,[],[f686,f330])).
fof(f7238,plain,(
  ( ! [X152,X151,X149,X150,X148] : (s(bool,X150) != s(bool,'const.combin.K_2'(s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X149))),s(X151,X152))) | s(bool,'T_0') = s(bool,X150) | s(bool,'F_0') = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_RP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X148),s('type.ring.ring'(bool),X149)))))) )),
  inference(superposition,[],[f686,f269])).
fof(f7239,plain,(
  ( ! [X156,X154,X157,X155,X153] : (s(bool,X155) != s(bool,'const.combin.K_2'(s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X154))),s(X156,X157))) | s(bool,'T_0') = s(bool,X155) | s(bool,'F_0') = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_RN_fupd_2'(s(fun(fun(bool,bool),fun(bool,bool)),X153),s('type.ring.ring'(bool),X154)))))) )),
  inference(superposition,[],[f686,f264])).
fof(f7240,plain,(
  ( ! [X158,X161,X159,X162,X160] : (s(bool,X160) != s(bool,'const.combin.K_2'(s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X159))),s(X161,X162))) | s(bool,'T_0') = s(bool,X160) | s(bool,'F_0') = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|49|_fupd_2'(s(fun(bool,bool),X158),s('type.ring.ring'(bool),X159)))))) )),
  inference(superposition,[],[f686,f263])).
fof(f7242,plain,(
  ( ! [X171,X169,X172,X170,X168] : (s(bool,X170) != s(bool,'const.combin.K_2'(s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X169))),s(X171,X172))) | s(bool,'T_0') = s(bool,X170) | s(bool,'F_0') = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_RM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X168),s('type.ring.ring'(bool),X169)))))) )),
  inference(superposition,[],[f686,f277])).
fof(f7243,plain,(
  ( ! [X177,X175,X173,X176,X174] : (s(bool,X175) != s(bool,'const.combin.K_2'(s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X174))),s(X176,X177))) | s(bool,'T_0') = s(bool,X175) | s(bool,'F_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_RP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X173),s('type.ring.ring'(bool),X174)))))) )),
  inference(superposition,[],[f686,f275])).
fof(f7244,plain,(
  ( ! [X182,X180,X178,X181,X179] : (s(bool,X180) != s(bool,'const.combin.K_2'(s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X179))),s(X181,X182))) | s(bool,'T_0') = s(bool,X180) | s(bool,'F_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_RN_fupd_2'(s(fun(fun(bool,bool),fun(bool,bool)),X178),s('type.ring.ring'(bool),X179)))))) )),
  inference(superposition,[],[f686,f265])).
fof(f7246,plain,(
  ( ! [X191,X189,X192,X190,X188] : (s(bool,X190) != s(bool,'const.combin.K_2'(s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X189))),s(X191,X192))) | s(bool,'T_0') = s(bool,X190) | s(bool,'F_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|48|_fupd_2'(s(fun(bool,bool),X188),s('type.ring.ring'(bool),X189)))))) )),
  inference(superposition,[],[f686,f262])).
fof(f7247,plain,(
  ( ! [X196,X194,X197,X195,X193] : (s(bool,X195) != s(bool,'const.combin.K_2'(s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X194))),s(X196,X197))) | s(bool,'T_0') = s(bool,X195) | s(bool,'F_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_RM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X193),s('type.ring.ring'(bool),X194)))))) )),
  inference(superposition,[],[f686,f276])).
fof(f7249,plain,(
  ( ! [X208,X206,X204,X207,X205] : (s(bool,X206) != s(bool,'const.combin.K_2'(s(bool,X204),s(X207,X208))) | s(bool,'T_0') = s(bool,X206) | s(bool,'F_0') = s(bool,sK11(X204,X205,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X205)))) )),
  inference(superposition,[],[f686,f563])).
fof(f7353,plain,(
  ( ! [X111,X116,X114,X112,X115,X113] : (s(bool,X111) != s(bool,'const.combin.K_2'(s(bool,X112),s(X113,X114))) | s(bool,'F_0') = s(bool,X112) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X111),s(X115,X116)))) )),
  inference(subsumption_resolution,[],[f7195,f2341])).
fof(f7195,plain,(
  ( ! [X111,X109,X116,X114,X112,X110,X115,X113] : (s(bool,X111) != s(bool,'const.combin.K_2'(s(bool,X112),s(X113,X114))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X109,X110))) | s(bool,'F_0') = s(bool,X112) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X111),s(X115,X116)))) )),
  inference(superposition,[],[f686,f2618])).
fof(f7832,plain,(
  ( ! [X85,X83,X86,X84] : (s(bool,X84) != s(bool,X86) | s(bool,'T_0') = s(bool,X86) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X83),'const.pair.,_2'(s(bool,X84),s(X83,X85)))))) )),
  inference(superposition,[],[f7312,f329])).
fof(f7833,plain,(
  ( ! [X88,X87,X89] : (s(bool,X87) != s(bool,X89) | s(bool,'T_0') = s(bool,X89) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X87),s(fun('type.num.num',bool),X88),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f7312,f261])).
fof(f7834,plain,(
  ( ! [X92,X90,X93,X91] : (s(bool,X93) != s(bool,happ(s(fun('type.num.num',bool),X91),s('type.num.num',X92))) | s(bool,'T_0') = s(bool,X93) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X90),s(fun('type.num.num',bool),X91),s('type.num.num','const.num.SUC_1'(s('type.num.num',X92)))))) )),
  inference(superposition,[],[f7312,f301])).
fof(f7835,plain,(
  ( ! [X94,X97,X95,X96] : (s(bool,X96) != s(bool,X97) | s(bool,'T_0') = s(bool,X97) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X94,bool),'const.pair.,_2'(s(X94,X95),s(bool,X96)))))) )),
  inference(superposition,[],[f7312,f330])).
fof(f7836,plain,(
  ( ! [X99,X100,X98] : (s(bool,X100) != s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X99))) | s(bool,'T_0') = s(bool,X100) | s(bool,'F_0') = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_RP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X98),s('type.ring.ring'(bool),X99)))))) )),
  inference(superposition,[],[f7312,f269])).
fof(f7837,plain,(
  ( ! [X103,X101,X102] : (s(bool,X103) != s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X102))) | s(bool,'T_0') = s(bool,X103) | s(bool,'F_0') = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_RN_fupd_2'(s(fun(fun(bool,bool),fun(bool,bool)),X101),s('type.ring.ring'(bool),X102)))))) )),
  inference(superposition,[],[f7312,f264])).
fof(f7838,plain,(
  ( ! [X105,X106,X104] : (s(bool,X106) != s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X105))) | s(bool,'T_0') = s(bool,X106) | s(bool,'F_0') = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|49|_fupd_2'(s(fun(bool,bool),X104),s('type.ring.ring'(bool),X105)))))) )),
  inference(superposition,[],[f7312,f263])).
fof(f7839,plain,(
  ( ! [X109,X107,X108] : (s(bool,X109) != s(bool,happ(s(fun(bool,bool),X107),s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X108))))) | s(bool,'T_0') = s(bool,X109) | s(bool,'F_0') = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|48|_fupd_2'(s(fun(bool,bool),X107),s('type.ring.ring'(bool),X108)))))) )),
  inference(superposition,[],[f7312,f273])).
fof(f7840,plain,(
  ( ! [X111,X112,X110] : (s(bool,X112) != s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X111))) | s(bool,'T_0') = s(bool,X112) | s(bool,'F_0') = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_RM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X110),s('type.ring.ring'(bool),X111)))))) )),
  inference(superposition,[],[f7312,f277])).
fof(f7841,plain,(
  ( ! [X114,X115,X113] : (s(bool,X115) != s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X114))) | s(bool,'T_0') = s(bool,X115) | s(bool,'F_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_RP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X113),s('type.ring.ring'(bool),X114)))))) )),
  inference(superposition,[],[f7312,f275])).
fof(f7842,plain,(
  ( ! [X118,X116,X117] : (s(bool,X118) != s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X117))) | s(bool,'T_0') = s(bool,X118) | s(bool,'F_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_RN_fupd_2'(s(fun(fun(bool,bool),fun(bool,bool)),X116),s('type.ring.ring'(bool),X117)))))) )),
  inference(superposition,[],[f7312,f265])).
fof(f7843,plain,(
  ( ! [X121,X120,X119] : (s(bool,X121) != s(bool,happ(s(fun(bool,bool),X119),s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X120))))) | s(bool,'T_0') = s(bool,X121) | s(bool,'F_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|49|_fupd_2'(s(fun(bool,bool),X119),s('type.ring.ring'(bool),X120)))))) )),
  inference(superposition,[],[f7312,f272])).
fof(f7844,plain,(
  ( ! [X123,X124,X122] : (s(bool,X124) != s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X123))) | s(bool,'T_0') = s(bool,X124) | s(bool,'F_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|48|_fupd_2'(s(fun(bool,bool),X122),s('type.ring.ring'(bool),X123)))))) )),
  inference(superposition,[],[f7312,f262])).
fof(f7845,plain,(
  ( ! [X127,X125,X126] : (s(bool,X127) != s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X126))) | s(bool,'T_0') = s(bool,X127) | s(bool,'F_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_RM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X125),s('type.ring.ring'(bool),X126)))))) )),
  inference(superposition,[],[f7312,f276])).
fof(f7846,plain,(
  ( ! [X130,X128,X131,X129] : (s(bool,X128) != s(bool,X131) | s(bool,'T_0') = s(bool,X131) | s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,X128),s(X129,X130)))) )),
  inference(superposition,[],[f7312,f328])).
fof(f7847,plain,(
  ( ! [X134,X132,X133] : (s(bool,X132) != s(bool,X134) | s(bool,'T_0') = s(bool,X134) | s(bool,'F_0') = s(bool,sK11(X132,X133,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X133)))) )),
  inference(superposition,[],[f7312,f563])).
fof(f7869,plain,(
  ( ! [X83,X81,X84,X82] : (s(bool,X82) != s(bool,X84) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X81),'const.pair.,_2'(s(bool,X82),s(X81,X83))))) | s(bool,'F_0') = s(bool,X84)) )),
  inference(superposition,[],[f7312,f329])).
fof(f7870,plain,(
  ( ! [X87,X85,X86] : (s(bool,X85) != s(bool,X87) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X85),s(fun('type.num.num',bool),X86),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X87)) )),
  inference(superposition,[],[f7312,f261])).
fof(f7871,plain,(
  ( ! [X90,X88,X91,X89] : (s(bool,X91) != s(bool,happ(s(fun('type.num.num',bool),X89),s('type.num.num',X90))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X88),s(fun('type.num.num',bool),X89),s('type.num.num','const.num.SUC_1'(s('type.num.num',X90))))) | s(bool,'F_0') = s(bool,X91)) )),
  inference(superposition,[],[f7312,f301])).
fof(f7872,plain,(
  ( ! [X94,X92,X95,X93] : (s(bool,X94) != s(bool,X95) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X92,bool),'const.pair.,_2'(s(X92,X93),s(bool,X94))))) | s(bool,'F_0') = s(bool,X95)) )),
  inference(superposition,[],[f7312,f330])).
fof(f7873,plain,(
  ( ! [X97,X98,X96] : (s(bool,X98) != s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X97))) | s(bool,'T_0') = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_RP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X96),s('type.ring.ring'(bool),X97))))) | s(bool,'F_0') = s(bool,X98)) )),
  inference(superposition,[],[f7312,f269])).
fof(f7874,plain,(
  ( ! [X101,X99,X100] : (s(bool,X101) != s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X100))) | s(bool,'T_0') = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_RN_fupd_2'(s(fun(fun(bool,bool),fun(bool,bool)),X99),s('type.ring.ring'(bool),X100))))) | s(bool,'F_0') = s(bool,X101)) )),
  inference(superposition,[],[f7312,f264])).
fof(f7875,plain,(
  ( ! [X103,X102,X104] : (s(bool,X104) != s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X103))) | s(bool,'T_0') = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|49|_fupd_2'(s(fun(bool,bool),X102),s('type.ring.ring'(bool),X103))))) | s(bool,'F_0') = s(bool,X104)) )),
  inference(superposition,[],[f7312,f263])).
fof(f7876,plain,(
  ( ! [X107,X105,X106] : (s(bool,X107) != s(bool,happ(s(fun(bool,bool),X105),s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X106))))) | s(bool,'T_0') = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|48|_fupd_2'(s(fun(bool,bool),X105),s('type.ring.ring'(bool),X106))))) | s(bool,'F_0') = s(bool,X107)) )),
  inference(superposition,[],[f7312,f273])).
fof(f7877,plain,(
  ( ! [X109,X110,X108] : (s(bool,X110) != s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X109))) | s(bool,'T_0') = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_RM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X108),s('type.ring.ring'(bool),X109))))) | s(bool,'F_0') = s(bool,X110)) )),
  inference(superposition,[],[f7312,f277])).
fof(f7878,plain,(
  ( ! [X111,X112,X113] : (s(bool,X113) != s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X112))) | s(bool,'T_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_RP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X111),s('type.ring.ring'(bool),X112))))) | s(bool,'F_0') = s(bool,X113)) )),
  inference(superposition,[],[f7312,f275])).
fof(f7879,plain,(
  ( ! [X116,X114,X115] : (s(bool,X116) != s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X115))) | s(bool,'T_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_RN_fupd_2'(s(fun(fun(bool,bool),fun(bool,bool)),X114),s('type.ring.ring'(bool),X115))))) | s(bool,'F_0') = s(bool,X116)) )),
  inference(superposition,[],[f7312,f265])).
fof(f7880,plain,(
  ( ! [X118,X119,X117] : (s(bool,X119) != s(bool,happ(s(fun(bool,bool),X117),s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X118))))) | s(bool,'T_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|49|_fupd_2'(s(fun(bool,bool),X117),s('type.ring.ring'(bool),X118))))) | s(bool,'F_0') = s(bool,X119)) )),
  inference(superposition,[],[f7312,f272])).
fof(f7881,plain,(
  ( ! [X121,X122,X120] : (s(bool,X122) != s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X121))) | s(bool,'T_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|48|_fupd_2'(s(fun(bool,bool),X120),s('type.ring.ring'(bool),X121))))) | s(bool,'F_0') = s(bool,X122)) )),
  inference(superposition,[],[f7312,f262])).
fof(f7882,plain,(
  ( ! [X125,X123,X124] : (s(bool,X125) != s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X124))) | s(bool,'T_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_RM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X123),s('type.ring.ring'(bool),X124))))) | s(bool,'F_0') = s(bool,X125)) )),
  inference(superposition,[],[f7312,f276])).
fof(f7884,plain,(
  ( ! [X132,X130,X131] : (s(bool,X130) != s(bool,X132) | s(bool,'T_0') = s(bool,sK11(X130,X131,bool)) | s(bool,'F_0') = s(bool,X132) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X131)))) )),
  inference(superposition,[],[f7312,f563])).
fof(f7924,plain,(
  ( ! [X76,X74,X75,X73] : (s(bool,X73) != s(bool,X74) | s(bool,'F_0') = s(bool,X74) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X73),s(X75,X76)))) )),
  inference(subsumption_resolution,[],[f7867,f2341])).
fof(f7867,plain,(
  ( ! [X76,X74,X72,X71,X75,X73] : (s(bool,X73) != s(bool,X74) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X71,X72))) | s(bool,'F_0') = s(bool,X74) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X73),s(X75,X76)))) )),
  inference(superposition,[],[f7312,f2618])).
fof(f3040,plain,(
  ( ! [X12,X10,X11,X9] : (s(bool,'T_0') = s(bool,happ(s(fun('type.num.num',bool),X10),s('type.num.num',X11))) | s(bool,X12) = s(bool,'const.ind_type.FCONS_3'(s(bool,X9),s(fun('type.num.num',bool),X10),s('type.num.num','const.num.SUC_1'(s('type.num.num',X11))))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f301,f389])).
fof(f3042,plain,(
  ( ! [X19,X17,X20,X18] : (s(bool,X20) = s(bool,happ(s(fun('type.num.num',bool),X18),s('type.num.num',X19))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X17),s(fun('type.num.num',bool),X18),s('type.num.num','const.num.SUC_1'(s('type.num.num',X19))))) | s(bool,'T_0') = s(bool,X20)) )),
  inference(superposition,[],[f301,f389])).
fof(f3084,plain,(
  ( ! [X189,X187,X190,X188] : (~p(s(bool,happ(s(fun('type.num.num',bool),X188),s('type.num.num',X189)))) | s(bool,'T_0') = s(bool,sK11('const.ind_type.FCONS_3'(s(bool,X187),s(fun('type.num.num',bool),X188),s('type.num.num','const.num.SUC_1'(s('type.num.num',X189)))),X190,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X190)))) )),
  inference(superposition,[],[f1190,f301])).
fof(f3086,plain,(
  ( ! [X198,X196,X200,X199,X197] : (s(bool,'T_0') != s(bool,'const.combin.K_2'(s(bool,happ(s(fun('type.num.num',bool),X197),s('type.num.num',X198))),s(X199,X200))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X196),s(fun('type.num.num',bool),X197),s('type.num.num','const.num.SUC_1'(s('type.num.num',X198)))))) )),
  inference(superposition,[],[f2289,f301])).
fof(f1314,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(bool,X3) = s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f1299,f232])).
fof(f1299,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(bool,X3) = s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1269,f389])).
fof(f1277,plain,(
  ( ! [X30,X28,X29,X27] : (~p(s(bool,happ(s(fun(X27,bool),X28),s(X27,X30)))) | p(s(bool,'|dollar|exists_1'(s(fun(X27,bool),sK11(X28,X29,fun(X27,bool)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X27,bool),bool),X29)))) )),
  inference(superposition,[],[f1269,f563])).
fof(f1269,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X1)))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X2))))) )),
  inference(subsumption_resolution,[],[f1252,f236])).
fof(f1252,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X1)))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X2)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X2)))) )),
  inference(superposition,[],[f288,f563])).
fof(f4962,plain,(
  ( ! [X8,X7,X9] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X7),s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X8))))) | s(bool,X9) = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|49|_fupd_2'(s(fun(bool,bool),X7),s('type.ring.ring'(bool),X8))))) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f272,f389])).
fof(f4964,plain,(
  ( ! [X14,X15,X13] : (s(bool,X15) = s(bool,happ(s(fun(bool,bool),X13),s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X14))))) | s(bool,'T_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|49|_fupd_2'(s(fun(bool,bool),X13),s('type.ring.ring'(bool),X14))))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f272,f389])).
fof(f5626,plain,(
  ( ! [X99,X100,X98] : (~p(s(bool,happ(s(fun(bool,bool),X98),s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X99)))))) | s(bool,'T_0') = s(bool,sK11('const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|49|_fupd_2'(s(fun(bool,bool),X98),s('type.ring.ring'(bool),X99)))),X100,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X100)))) )),
  inference(superposition,[],[f1190,f272])).
fof(f5481,plain,(
  ( ! [X8,X7,X9] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X7),s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X8))))) | s(bool,X9) = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|48|_fupd_2'(s(fun(bool,bool),X7),s('type.ring.ring'(bool),X8))))) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f273,f389])).
fof(f5483,plain,(
  ( ! [X14,X15,X13] : (s(bool,X15) = s(bool,happ(s(fun(bool,bool),X13),s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X14))))) | s(bool,'T_0') = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|48|_fupd_2'(s(fun(bool,bool),X13),s('type.ring.ring'(bool),X14))))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f273,f389])).
fof(f5623,plain,(
  ( ! [X90,X91,X89] : (~p(s(bool,happ(s(fun(bool,bool),X89),s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X90)))))) | s(bool,'T_0') = s(bool,sK11('const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|48|_fupd_2'(s(fun(bool,bool),X89),s('type.ring.ring'(bool),X90)))),X91,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X91)))) )),
  inference(superposition,[],[f1190,f273])).
fof(f287,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(X0,X2) = s(X0,sK11(X2,X1,X0))) )),
  inference(cnf_transformation,[],[f217])).
fof(f217,plain,(
  ! [X0,X1,X2] : ((s(X0,X2) = s(X0,sK11(X2,X1,X0)) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK11(X2,X1,X0)))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK11])],[f195])).
fof(f195,plain,(
  ! [X0,X1,X2] : (? [X3] : (s(X0,X2) = s(X0,X3) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(ennf_transformation,[],[f140])).
fof(f140,plain,(
  ! [X0,X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f139])).
fof(f139,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X6] : ! [X30,X31] : (! [X23] : (s(X6,X23) = s(X6,X31) => p(s(bool,happ(s(fun(X6,bool),X30),s(X6,X23))))) => p(s(bool,'|dollar|exists_1'(s(fun(X6,bool),X30)))))),
  file('thm.ring.ring_|49||49|','thm.bool.JRH_INDUCT_UTIL')).
fof(f558,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X2)))) | s(X0,X4) = s(X0,sK11(X4,'const.pair.FST_1'(s('type.pair.prod'(fun(X0,bool),X1),'const.pair.,_2'(s(fun(X0,bool),X2),s(X1,X3)))),X0))) )),
  inference(superposition,[],[f287,f329])).
fof(f559,plain,(
  ( ! [X6,X8,X7,X5] : (p(s(bool,'|dollar|exists_1'(s(fun(X5,bool),X6)))) | s(X5,X8) = s(X5,sK11(X8,'const.ind_type.FCONS_3'(s(fun(X5,bool),X6),s(fun('type.num.num',fun(X5,bool)),X7),s('type.num.num','const.num.|48|_0')),X5))) )),
  inference(superposition,[],[f287,f261])).
fof(f560,plain,(
  ( ! [X12,X10,X13,X11,X9] : (p(s(bool,'|dollar|exists_1'(s(fun(X9,bool),X12)))) | s(X9,X13) = s(X9,sK11(X13,'const.pair.SND_1'(s('type.pair.prod'(X10,fun(X9,bool)),'const.pair.,_2'(s(X10,X11),s(fun(X9,bool),X12)))),X9))) )),
  inference(superposition,[],[f287,f330])).
fof(f561,plain,(
  ( ! [X14,X17,X15,X18,X16] : (p(s(bool,'|dollar|exists_1'(s(fun(X14,bool),X15)))) | s(X14,X18) = s(X14,sK11(X18,'const.combin.K_2'(s(fun(X14,bool),X15),s(X16,X17)),X14))) )),
  inference(superposition,[],[f287,f328])).
fof(f1170,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X2))) | s(X0,X4) = s(X0,sK11(X4,'const.pair.FST_1'(s('type.pair.prod'(fun(X0,bool),X1),'const.pair.,_2'(s(fun(X0,bool),X2),s(X1,X3)))),X0))) )),
  inference(superposition,[],[f563,f329])).
fof(f1171,plain,(
  ( ! [X6,X8,X7,X5] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X5,bool),X6))) | s(X5,X8) = s(X5,sK11(X8,'const.ind_type.FCONS_3'(s(fun(X5,bool),X6),s(fun('type.num.num',fun(X5,bool)),X7),s('type.num.num','const.num.|48|_0')),X5))) )),
  inference(superposition,[],[f563,f261])).
fof(f1172,plain,(
  ( ! [X12,X10,X13,X11,X9] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X9,bool),X12))) | s(X9,X13) = s(X9,sK11(X13,'const.pair.SND_1'(s('type.pair.prod'(X10,fun(X9,bool)),'const.pair.,_2'(s(X10,X11),s(fun(X9,bool),X12)))),X9))) )),
  inference(superposition,[],[f563,f330])).
fof(f1175,plain,(
  ( ! [X26,X24,X23,X25,X22] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X22,bool),X23))) | s(X22,X26) = s(X22,sK11(X26,'const.combin.K_2'(s(fun(X22,bool),X23),s(X24,X25)),X22))) )),
  inference(superposition,[],[f563,f328])).
fof(f1176,plain,(
  ( ! [X30,X28,X29,X27] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X27,bool),X28))) | s(X27,X30) = s(X27,sK11(X30,sK11(X28,X29,fun(X27,bool)),X27)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X27,bool),bool),X29)))) )),
  inference(superposition,[],[f563,f563])).
fof(f1205,plain,(
  ( ! [X88,X87,X89,X86] : (p(s(bool,'|dollar|exists_1'(s(fun(X86,bool),X87)))) | s(X86,X89) = s(X86,sK11(X89,sK11(X87,X88,fun(X86,bool)),X86)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X86,bool),bool),X88)))) )),
  inference(superposition,[],[f287,f563])).
fof(f1242,plain,(
  ( ! [X30,X28,X29,X27] : (~p(s(bool,X29)) | s(bool,X29) = s(bool,'|dollar|exists_1'(s(fun(X27,bool),X28))) | s(X27,X30) = s(X27,sK11(X30,X28,X27))) )),
  inference(subsumption_resolution,[],[f1218,f232])).
fof(f1218,plain,(
  ( ! [X30,X28,X29,X27] : (~p(s(bool,'T_0')) | ~p(s(bool,X29)) | s(bool,X29) = s(bool,'|dollar|exists_1'(s(fun(X27,bool),X28))) | s(X27,X30) = s(X27,sK11(X30,X28,X27))) )),
  inference(superposition,[],[f326,f563])).
fof(f2131,plain,(
  ( ! [X59,X64,X60,X65,X63] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,'|dollar|exists_1'(s(fun(X59,bool),X60))),s(X63,X64))) | s(X59,X65) = s(X59,sK11(X65,X60,X59))) )),
  inference(subsumption_resolution,[],[f2130,f358])).
fof(f2130,plain,(
  ( ! [X59,X64,X60,X65,X63] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,'|dollar|exists_1'(s(fun(X59,bool),X60))),s(X63,X64))) | s(X59,X65) = s(X59,sK11(X65,X60,X59))) )),
  inference(forward_demodulation,[],[f1951,f328])).
fof(f1951,plain,(
  ( ! [X61,X59,X64,X62,X60,X65,X63] : (s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,'T_0'),s(X61,X62))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,'|dollar|exists_1'(s(fun(X59,bool),X60))),s(X63,X64))) | s(X59,X65) = s(X59,sK11(X65,X60,X59))) )),
  inference(superposition,[],[f649,f563])).
fof(f5759,plain,(
  ( ! [X167,X165,X169,X166,X168] : (~p(s(bool,X165)) | s(bool,sK11(X165,X166,bool)) = s(bool,'|dollar|exists_1'(s(fun(X167,bool),X168))) | s(X167,X169) = s(X167,sK11(X169,X168,X167)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X166)))) )),
  inference(superposition,[],[f1242,f563])).
fof(f5772,plain,(
  ( ! [X59,X57,X60,X58,X56,X55] : (s(bool,'|dollar|exists_1'(s(fun(X55,bool),X56))) = s(bool,'|dollar|exists_1'(s(fun(X57,bool),X58))) | s(X57,X59) = s(X57,sK11(X59,X58,X57)) | s(X55,X60) = s(X55,sK11(X60,X56,X55))) )),
  inference(subsumption_resolution,[],[f5740,f232])).
fof(f5740,plain,(
  ( ! [X59,X57,X60,X58,X56,X55] : (~p(s(bool,'T_0')) | s(bool,'|dollar|exists_1'(s(fun(X55,bool),X56))) = s(bool,'|dollar|exists_1'(s(fun(X57,bool),X58))) | s(X57,X59) = s(X57,sK11(X59,X58,X57)) | s(X55,X60) = s(X55,sK11(X60,X56,X55))) )),
  inference(superposition,[],[f1242,f563])).
fof(f6959,plain,(
  ( ! [X103,X101,X99,X102,X100,X98] : (p(s(bool,'|dollar|exists_1'(s(fun(X98,bool),X99)))) | s(X98,X101) = s(X98,sK11(X101,'const.combin.K_2'(s(fun(X98,bool),sK11(X99,X100,fun(X98,bool))),s(X102,X103)),X98)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X98,bool),bool),X100)))) )),
  inference(superposition,[],[f561,f563])).
fof(f6987,plain,(
  ( ! [X47,X43,X50,X48,X46,X42,X49] : (s(X42,X46) = s(X42,sK11(X46,'const.combin.K_2'(s(fun(X42,bool),X43),s(X47,X48)),X42)) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,'|dollar|exists_1'(s(fun(X42,bool),X43))),s(X49,X50)))) )),
  inference(subsumption_resolution,[],[f6986,f234])).
fof(f6986,plain,(
  ( ! [X47,X43,X50,X48,X46,X42,X49] : (p(s(bool,'F_0')) | s(X42,X46) = s(X42,sK11(X46,'const.combin.K_2'(s(fun(X42,bool),X43),s(X47,X48)),X42)) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,'|dollar|exists_1'(s(fun(X42,bool),X43))),s(X49,X50)))) )),
  inference(forward_demodulation,[],[f6967,f328])).
fof(f6967,plain,(
  ( ! [X47,X45,X43,X50,X48,X46,X44,X42,X49] : (p(s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X44,X45)))) | s(X42,X46) = s(X42,sK11(X46,'const.combin.K_2'(s(fun(X42,bool),X43),s(X47,X48)),X42)) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,'|dollar|exists_1'(s(fun(X42,bool),X43))),s(X49,X50)))) )),
  inference(superposition,[],[f561,f2618])).
fof(f3091,plain,(
  ( ! [X222,X225,X223,X221,X224] : (p(s(bool,'|dollar|exists_1'(s(fun(X221,bool),happ(s(fun('type.num.num',fun(X221,bool)),X223),s('type.num.num',X224)))))) | s(X221,X225) = s(X221,sK11(X225,'const.ind_type.FCONS_3'(s(fun(X221,bool),X222),s(fun('type.num.num',fun(X221,bool)),X223),s('type.num.num','const.num.SUC_1'(s('type.num.num',X224)))),X221))) )),
  inference(superposition,[],[f287,f301])).
fof(f3090,plain,(
  ( ! [X218,X216,X220,X219,X217] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X216,bool),happ(s(fun('type.num.num',fun(X216,bool)),X218),s('type.num.num',X219))))) | s(X216,X220) = s(X216,sK11(X220,'const.ind_type.FCONS_3'(s(fun(X216,bool),X217),s(fun('type.num.num',fun(X216,bool)),X218),s('type.num.num','const.num.SUC_1'(s('type.num.num',X219)))),X216))) )),
  inference(superposition,[],[f563,f301])).
fof(f5007,plain,(
  ( ! [X146,X147,X145,X148] : (p(s(bool,'|dollar|exists_1'(s(fun(X145,bool),happ(s(fun(fun(X145,bool),fun(X145,bool)),X146),s(fun(X145,bool),'const.ring.ring_R|49|_1'(s('type.ring.ring'(fun(X145,bool)),X147)))))))) | s(X145,X148) = s(X145,sK11(X148,'const.ring.ring_R|49|_1'(s('type.ring.ring'(fun(X145,bool)),'const.ring.ring_R|49|_fupd_2'(s(fun(fun(X145,bool),fun(X145,bool)),X146),s('type.ring.ring'(fun(X145,bool)),X147)))),X145))) )),
  inference(superposition,[],[f287,f272])).
fof(f5530,plain,(
  ( ! [X171,X169,X172,X170] : (p(s(bool,'|dollar|exists_1'(s(fun(X169,bool),happ(s(fun(fun(X169,bool),fun(X169,bool)),X170),s(fun(X169,bool),'const.ring.ring_R|48|_1'(s('type.ring.ring'(fun(X169,bool)),X171)))))))) | s(X169,X172) = s(X169,sK11(X172,'const.ring.ring_R|48|_1'(s('type.ring.ring'(fun(X169,bool)),'const.ring.ring_R|48|_fupd_2'(s(fun(fun(X169,bool),fun(X169,bool)),X170),s('type.ring.ring'(fun(X169,bool)),X171)))),X169))) )),
  inference(superposition,[],[f287,f273])).
fof(f740,plain,(
  ( ! [X76,X74,X77,X75] : (p(s(bool,'|dollar|exists_1'(s(fun(X74,bool),'const.ring.ring_R|48|_1'(s('type.ring.ring'(fun(X74,bool)),X76)))))) | s(X74,X77) = s(X74,sK11(X77,'const.ring.ring_R|48|_1'(s('type.ring.ring'(fun(X74,bool)),'const.ring.ring_R|49|_fupd_2'(s(fun(fun(X74,bool),fun(X74,bool)),X75),s('type.ring.ring'(fun(X74,bool)),X76)))),X74))) )),
  inference(superposition,[],[f287,f263])).
fof(f1173,plain,(
  ( ! [X14,X17,X15,X16] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X14,bool),'const.ring.ring_R|48|_1'(s('type.ring.ring'(fun(X14,bool)),X16))))) | s(X14,X17) = s(X14,sK11(X17,'const.ring.ring_R|48|_1'(s('type.ring.ring'(fun(X14,bool)),'const.ring.ring_R|49|_fupd_2'(s(fun(fun(X14,bool),fun(X14,bool)),X15),s('type.ring.ring'(fun(X14,bool)),X16)))),X14))) )),
  inference(superposition,[],[f563,f263])).
fof(f1610,plain,(
  ( ! [X107,X105,X108,X106] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X105,bool),'const.ring.ring_R|48|_1'(s('type.ring.ring'(fun(X105,bool)),X107))))) | s(X105,X108) = s(X105,sK11(X108,'const.ring.ring_R|48|_1'(s('type.ring.ring'(fun(X105,bool)),'const.ring.ring_RN_fupd_2'(s(fun(fun(fun(X105,bool),fun(X105,bool)),fun(fun(X105,bool),fun(X105,bool))),X106),s('type.ring.ring'(fun(X105,bool)),X107)))),X105))) )),
  inference(superposition,[],[f563,f264])).
fof(f1611,plain,(
  ( ! [X111,X109,X112,X110] : (p(s(bool,'|dollar|exists_1'(s(fun(X109,bool),'const.ring.ring_R|48|_1'(s('type.ring.ring'(fun(X109,bool)),X111)))))) | s(X109,X112) = s(X109,sK11(X112,'const.ring.ring_R|48|_1'(s('type.ring.ring'(fun(X109,bool)),'const.ring.ring_RN_fupd_2'(s(fun(fun(fun(X109,bool),fun(X109,bool)),fun(fun(X109,bool),fun(X109,bool))),X110),s('type.ring.ring'(fun(X109,bool)),X111)))),X109))) )),
  inference(superposition,[],[f287,f264])).
fof(f4354,plain,(
  ( ! [X127,X128,X126,X129] : (p(s(bool,'|dollar|exists_1'(s(fun(X126,bool),'const.ring.ring_R|48|_1'(s('type.ring.ring'(fun(X126,bool)),X128)))))) | s(X126,X129) = s(X126,sK11(X129,'const.ring.ring_R|48|_1'(s('type.ring.ring'(fun(X126,bool)),'const.ring.ring_RP_fupd_2'(s(fun(fun(fun(X126,bool),fun(fun(X126,bool),fun(X126,bool))),fun(fun(X126,bool),fun(fun(X126,bool),fun(X126,bool)))),X127),s('type.ring.ring'(fun(X126,bool)),X128)))),X126))) )),
  inference(superposition,[],[f287,f269])).
fof(f4355,plain,(
  ( ! [X132,X130,X133,X131] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X130,bool),'const.ring.ring_R|48|_1'(s('type.ring.ring'(fun(X130,bool)),X132))))) | s(X130,X133) = s(X130,sK11(X133,'const.ring.ring_R|48|_1'(s('type.ring.ring'(fun(X130,bool)),'const.ring.ring_RP_fupd_2'(s(fun(fun(fun(X130,bool),fun(fun(X130,bool),fun(X130,bool))),fun(fun(X130,bool),fun(fun(X130,bool),fun(X130,bool)))),X131),s('type.ring.ring'(fun(X130,bool)),X132)))),X130))) )),
  inference(superposition,[],[f563,f269])).
fof(f7091,plain,(
  ( ! [X229,X227,X228,X226] : (p(s(bool,'|dollar|exists_1'(s(fun(X226,bool),'const.ring.ring_R|48|_1'(s('type.ring.ring'(fun(X226,bool)),X228)))))) | s(X226,X229) = s(X226,sK11(X229,'const.ring.ring_R|48|_1'(s('type.ring.ring'(fun(X226,bool)),'const.ring.ring_RM_fupd_2'(s(fun(fun(fun(X226,bool),fun(fun(X226,bool),fun(X226,bool))),fun(fun(X226,bool),fun(fun(X226,bool),fun(X226,bool)))),X227),s('type.ring.ring'(fun(X226,bool)),X228)))),X226))) )),
  inference(superposition,[],[f287,f277])).
fof(f7092,plain,(
  ( ! [X231,X233,X230,X232] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X230,bool),'const.ring.ring_R|48|_1'(s('type.ring.ring'(fun(X230,bool)),X232))))) | s(X230,X233) = s(X230,sK11(X233,'const.ring.ring_R|48|_1'(s('type.ring.ring'(fun(X230,bool)),'const.ring.ring_RM_fupd_2'(s(fun(fun(fun(X230,bool),fun(fun(X230,bool),fun(X230,bool))),fun(fun(X230,bool),fun(fun(X230,bool),fun(X230,bool)))),X231),s('type.ring.ring'(fun(X230,bool)),X232)))),X230))) )),
  inference(superposition,[],[f563,f277])).
fof(f631,plain,(
  ( ! [X66,X64,X65,X63] : (p(s(bool,'|dollar|exists_1'(s(fun(X63,bool),'const.ring.ring_R|49|_1'(s('type.ring.ring'(fun(X63,bool)),X65)))))) | s(X63,X66) = s(X63,sK11(X66,'const.ring.ring_R|49|_1'(s('type.ring.ring'(fun(X63,bool)),'const.ring.ring_R|48|_fupd_2'(s(fun(fun(X63,bool),fun(X63,bool)),X64),s('type.ring.ring'(fun(X63,bool)),X65)))),X63))) )),
  inference(superposition,[],[f287,f262])).
fof(f1174,plain,(
  ( ! [X21,X19,X20,X18] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X18,bool),'const.ring.ring_R|49|_1'(s('type.ring.ring'(fun(X18,bool)),X20))))) | s(X18,X21) = s(X18,sK11(X21,'const.ring.ring_R|49|_1'(s('type.ring.ring'(fun(X18,bool)),'const.ring.ring_R|48|_fupd_2'(s(fun(fun(X18,bool),fun(X18,bool)),X19),s('type.ring.ring'(fun(X18,bool)),X20)))),X18))) )),
  inference(superposition,[],[f563,f262])).
fof(f1853,plain,(
  ( ! [X118,X120,X119,X117] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X117,bool),'const.ring.ring_R|49|_1'(s('type.ring.ring'(fun(X117,bool)),X119))))) | s(X117,X120) = s(X117,sK11(X120,'const.ring.ring_R|49|_1'(s('type.ring.ring'(fun(X117,bool)),'const.ring.ring_RN_fupd_2'(s(fun(fun(fun(X117,bool),fun(X117,bool)),fun(fun(X117,bool),fun(X117,bool))),X118),s('type.ring.ring'(fun(X117,bool)),X119)))),X117))) )),
  inference(superposition,[],[f563,f265])).
fof(f1854,plain,(
  ( ! [X123,X121,X124,X122] : (p(s(bool,'|dollar|exists_1'(s(fun(X121,bool),'const.ring.ring_R|49|_1'(s('type.ring.ring'(fun(X121,bool)),X123)))))) | s(X121,X124) = s(X121,sK11(X124,'const.ring.ring_R|49|_1'(s('type.ring.ring'(fun(X121,bool)),'const.ring.ring_RN_fupd_2'(s(fun(fun(fun(X121,bool),fun(X121,bool)),fun(fun(X121,bool),fun(X121,bool))),X122),s('type.ring.ring'(fun(X121,bool)),X123)))),X121))) )),
  inference(superposition,[],[f287,f265])).
fof(f5878,plain,(
  ( ! [X182,X180,X183,X181] : (p(s(bool,'|dollar|exists_1'(s(fun(X180,bool),'const.ring.ring_R|49|_1'(s('type.ring.ring'(fun(X180,bool)),X182)))))) | s(X180,X183) = s(X180,sK11(X183,'const.ring.ring_R|49|_1'(s('type.ring.ring'(fun(X180,bool)),'const.ring.ring_RP_fupd_2'(s(fun(fun(fun(X180,bool),fun(fun(X180,bool),fun(X180,bool))),fun(fun(X180,bool),fun(fun(X180,bool),fun(X180,bool)))),X181),s('type.ring.ring'(fun(X180,bool)),X182)))),X180))) )),
  inference(superposition,[],[f287,f275])).
fof(f5879,plain,(
  ( ! [X187,X185,X186,X184] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X184,bool),'const.ring.ring_R|49|_1'(s('type.ring.ring'(fun(X184,bool)),X186))))) | s(X184,X187) = s(X184,sK11(X187,'const.ring.ring_R|49|_1'(s('type.ring.ring'(fun(X184,bool)),'const.ring.ring_RP_fupd_2'(s(fun(fun(fun(X184,bool),fun(fun(X184,bool),fun(X184,bool))),fun(fun(X184,bool),fun(fun(X184,bool),fun(X184,bool)))),X185),s('type.ring.ring'(fun(X184,bool)),X186)))),X184))) )),
  inference(superposition,[],[f563,f275])).
fof(f6257,plain,(
  ( ! [X198,X196,X199,X197] : (p(s(bool,'|dollar|exists_1'(s(fun(X196,bool),'const.ring.ring_R|49|_1'(s('type.ring.ring'(fun(X196,bool)),X198)))))) | s(X196,X199) = s(X196,sK11(X199,'const.ring.ring_R|49|_1'(s('type.ring.ring'(fun(X196,bool)),'const.ring.ring_RM_fupd_2'(s(fun(fun(fun(X196,bool),fun(fun(X196,bool),fun(X196,bool))),fun(fun(X196,bool),fun(fun(X196,bool),fun(X196,bool)))),X197),s('type.ring.ring'(fun(X196,bool)),X198)))),X196))) )),
  inference(superposition,[],[f287,f276])).
fof(f6258,plain,(
  ( ! [X202,X200,X203,X201] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X200,bool),'const.ring.ring_R|49|_1'(s('type.ring.ring'(fun(X200,bool)),X202))))) | s(X200,X203) = s(X200,sK11(X203,'const.ring.ring_R|49|_1'(s('type.ring.ring'(fun(X200,bool)),'const.ring.ring_RM_fupd_2'(s(fun(fun(fun(X200,bool),fun(fun(X200,bool),fun(X200,bool))),fun(fun(X200,bool),fun(fun(X200,bool),fun(X200,bool)))),X201),s('type.ring.ring'(fun(X200,bool)),X202)))),X200))) )),
  inference(superposition,[],[f563,f276])).
fof(f1189,plain,(
  ( ! [X41,X40] : (s(bool,'T_0') != s(bool,sK11(X40,X41,bool)) | p(s(bool,X40)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X41)))) )),
  inference(superposition,[],[f236,f563])).
fof(f1190,plain,(
  ( ! [X43,X42] : (~p(s(bool,X42)) | s(bool,'T_0') = s(bool,sK11(X42,X43,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X43)))) )),
  inference(superposition,[],[f237,f563])).
fof(f1193,plain,(
  ( ! [X50,X48,X49] : (p(s(bool,X48)) | p(s(bool,X50)) | s(bool,X50) = s(bool,sK11(X48,X49,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X49)))) )),
  inference(superposition,[],[f323,f563])).
fof(f1195,plain,(
  ( ! [X54,X56,X55] : (~p(s(bool,X54)) | ~p(s(bool,X56)) | s(bool,X56) = s(bool,sK11(X54,X55,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X55)))) )),
  inference(superposition,[],[f326,f563])).
fof(f1201,plain,(
  ( ! [X74,X75,X73] : (~p(s(bool,X73)) | s(bool,X75) != s(bool,sK11(X73,X74,bool)) | s(bool,'T_0') = s(bool,X75) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X74)))) )),
  inference(superposition,[],[f392,f563])).
fof(f1202,plain,(
  ( ! [X78,X76,X77] : (p(s(bool,X76)) | s(bool,X78) = s(bool,sK11(X76,X77,bool)) | s(bool,'T_0') = s(bool,X78) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X77)))) )),
  inference(superposition,[],[f426,f563])).
fof(f1765,plain,(
  ( ! [X88,X87,X85,X86] : (~p(s(bool,X85)) | s(bool,X87) = s(bool,sK11(X85,X86,bool)) | s(bool,X87) = s(bool,X88) | s(bool,'T_0') = s(bool,X88) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X86)))) )),
  inference(superposition,[],[f1154,f563])).
fof(f2946,plain,(
  ( ! [X68,X66,X67,X65] : (~p(s(bool,X66)) | s(bool,'T_0') = s(bool,sK11('const.pair.FST_1'(s('type.pair.prod'(bool,X65),'const.pair.,_2'(s(bool,X66),s(X65,X67)))),X68,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X68)))) )),
  inference(superposition,[],[f1190,f329])).
fof(f2947,plain,(
  ( ! [X70,X71,X69] : (~p(s(bool,X69)) | s(bool,'T_0') = s(bool,sK11('const.ind_type.FCONS_3'(s(bool,X69),s(fun('type.num.num',bool),X70),s('type.num.num','const.num.|48|_0')),X71,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X71)))) )),
  inference(superposition,[],[f1190,f261])).
fof(f2948,plain,(
  ( ! [X74,X72,X75,X73] : (~p(s(bool,X74)) | s(bool,'T_0') = s(bool,sK11('const.pair.SND_1'(s('type.pair.prod'(X72,bool),'const.pair.,_2'(s(X72,X73),s(bool,X74)))),X75,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X75)))) )),
  inference(superposition,[],[f1190,f330])).
fof(f2949,plain,(
  ( ! [X78,X76,X77] : (~p(s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X77)))) | s(bool,'T_0') = s(bool,sK11('const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|49|_fupd_2'(s(fun(bool,bool),X76),s('type.ring.ring'(bool),X77)))),X78,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X78)))) )),
  inference(superposition,[],[f1190,f263])).
fof(f2950,plain,(
  ( ! [X80,X81,X79] : (~p(s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X80)))) | s(bool,'T_0') = s(bool,sK11('const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_RN_fupd_2'(s(fun(fun(bool,bool),fun(bool,bool)),X79),s('type.ring.ring'(bool),X80)))),X81,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X81)))) )),
  inference(superposition,[],[f1190,f264])).
fof(f2951,plain,(
  ( ! [X83,X84,X82] : (~p(s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X83)))) | s(bool,'T_0') = s(bool,sK11('const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|48|_fupd_2'(s(fun(bool,bool),X82),s('type.ring.ring'(bool),X83)))),X84,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X84)))) )),
  inference(superposition,[],[f1190,f262])).
fof(f2952,plain,(
  ( ! [X87,X85,X86] : (~p(s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X86)))) | s(bool,'T_0') = s(bool,sK11('const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_RN_fupd_2'(s(fun(fun(bool,bool),fun(bool,bool)),X85),s('type.ring.ring'(bool),X86)))),X87,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X87)))) )),
  inference(superposition,[],[f1190,f265])).
fof(f2953,plain,(
  ( ! [X90,X88,X91,X89] : (~p(s(bool,X88)) | s(bool,'T_0') = s(bool,sK11('const.combin.K_2'(s(bool,X88),s(X89,X90)),X91,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X91)))) )),
  inference(superposition,[],[f1190,f328])).
fof(f2954,plain,(
  ( ! [X94,X92,X93] : (~p(s(bool,X92)) | s(bool,'T_0') = s(bool,sK11(sK11(X92,X93,bool),X94,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X94))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X93)))) )),
  inference(superposition,[],[f1190,f563])).
fof(f2962,plain,(
  ( ! [X35,X33,X36,X34] : (s(bool,'T_0') = s(bool,sK11('|dollar|exists_1'(s(fun(X33,bool),X34)),X35,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X35))) | s(X33,X36) = s(X33,sK11(X36,X34,X33))) )),
  inference(subsumption_resolution,[],[f2939,f232])).
fof(f2939,plain,(
  ( ! [X35,X33,X36,X34] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,sK11('|dollar|exists_1'(s(fun(X33,bool),X34)),X35,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X35))) | s(X33,X36) = s(X33,sK11(X36,X34,X33))) )),
  inference(superposition,[],[f1190,f563])).
fof(f5620,plain,(
  ( ! [X80,X81,X82] : (~p(s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X81)))) | s(bool,'T_0') = s(bool,sK11('const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_RP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X80),s('type.ring.ring'(bool),X81)))),X82,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X82)))) )),
  inference(superposition,[],[f1190,f269])).
fof(f5871,plain,(
  ( ! [X154,X155,X153] : (~p(s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X154)))) | s(bool,'T_0') = s(bool,sK11('const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_RP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X153),s('type.ring.ring'(bool),X154)))),X155,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X155)))) )),
  inference(superposition,[],[f1190,f275])).
fof(f5978,plain,(
  ( ! [X206,X204,X202,X205,X203,X201] : (s(bool,'T_0') != s(bool,'const.combin.K_2'(s(bool,X201),s(X203,X204))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,sK11(X201,X202,bool)),s(X205,X206))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X202)))) )),
  inference(superposition,[],[f2156,f563])).
fof(f6120,plain,(
  ( ! [X142,X140,X141,X139] : (s(bool,'T_0') != s(bool,X139) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,sK11(X139,X140,bool)),s(X141,X142))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X140)))) )),
  inference(superposition,[],[f6033,f563])).
fof(f6248,plain,(
  ( ! [X161,X159,X160] : (~p(s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X160)))) | s(bool,'T_0') = s(bool,sK11('const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_RM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X159),s('type.ring.ring'(bool),X160)))),X161,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X161)))) )),
  inference(superposition,[],[f1190,f276])).
fof(f7079,plain,(
  ( ! [X173,X171,X172] : (~p(s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X172)))) | s(bool,'T_0') = s(bool,sK11('const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_RM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X171),s('type.ring.ring'(bool),X172)))),X173,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X173)))) )),
  inference(superposition,[],[f1190,f277])).
fof(f2567,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'const.ring.ring_RN_1'(s('type.ring.ring'(bool),X10)))))) | s(bool,X11) = s(bool,sK11(X11,'const.ring.ring_RN_1'(s('type.ring.ring'(bool),'const.ring.ring_R|48|_fupd_2'(s(fun(bool,bool),X9),s('type.ring.ring'(bool),X10)))),bool))) )),
  inference(superposition,[],[f287,f267])).
fof(f2566,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'const.ring.ring_RN_1'(s('type.ring.ring'(bool),X7))))) | s(bool,X8) = s(bool,sK11(X8,'const.ring.ring_RN_1'(s('type.ring.ring'(bool),'const.ring.ring_R|48|_fupd_2'(s(fun(bool,bool),X6),s('type.ring.ring'(bool),X7)))),bool))) )),
  inference(superposition,[],[f563,f267])).
fof(f2319,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'const.ring.ring_RN_1'(s('type.ring.ring'(bool),X10)))))) | s(bool,X11) = s(bool,sK11(X11,'const.ring.ring_RN_1'(s('type.ring.ring'(bool),'const.ring.ring_R|49|_fupd_2'(s(fun(bool,bool),X9),s('type.ring.ring'(bool),X10)))),bool))) )),
  inference(superposition,[],[f287,f266])).
fof(f2318,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),'const.ring.ring_RN_1'(s('type.ring.ring'(bool),X7))))) | s(bool,X8) = s(bool,sK11(X8,'const.ring.ring_RN_1'(s('type.ring.ring'(bool),'const.ring.ring_R|49|_fupd_2'(s(fun(bool,bool),X6),s('type.ring.ring'(bool),X7)))),bool))) )),
  inference(superposition,[],[f563,f266])).
fof(f7569,plain,(
  ( ! [X68,X66,X67,X65] : (s('type.ring.ring'(X65),X66) != s('type.ring.ring'(X65),X68) | s(fun(X65,fun(X65,X65)),'const.ring.ring_RP_1'(s('type.ring.ring'(X65),X68))) = s(fun(X65,fun(X65,X65)),'const.ring.ring_RP_1'(s('type.ring.ring'(X65),sK11(X66,X67,'type.ring.ring'(X65))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.ring.ring'(X65),bool),X67)))) )),
  inference(superposition,[],[f291,f563])).
fof(f2901,plain,(
  ( ! [X37,X35,X38,X36] : (s('type.ring.ring'(X35),X36) != s('type.ring.ring'(X35),X38) | s(fun(X35,X35),'const.ring.ring_RN_1'(s('type.ring.ring'(X35),X38))) = s(fun(X35,X35),'const.ring.ring_RN_1'(s('type.ring.ring'(X35),sK11(X36,X37,'type.ring.ring'(X35))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.ring.ring'(X35),bool),X37)))) )),
  inference(superposition,[],[f293,f563])).
fof(f1372,plain,(
  ( ! [X30,X28,X29,X27] : (s('type.ring.ring'(X27),X28) != s('type.ring.ring'(X27),X30) | s(X27,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X27),X30))) = s(X27,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X27),sK11(X28,X29,'type.ring.ring'(X27))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.ring.ring'(X27),bool),X29)))) )),
  inference(superposition,[],[f290,f563])).
fof(f1333,plain,(
  ( ! [X30,X28,X29,X27] : (s('type.ring.ring'(X27),X28) != s('type.ring.ring'(X27),X30) | s(X27,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X27),X30))) = s(X27,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X27),sK11(X28,X29,'type.ring.ring'(X27))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.ring.ring'(X27),bool),X29)))) )),
  inference(superposition,[],[f289,f563])).
fof(f768,plain,(
  ( ! [X37,X35,X36] : (s(bool,X37) = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X36))) | s(bool,'T_0') = s(bool,X37) | s(bool,'T_0') = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|49|_fupd_2'(s(fun(bool,bool),X35),s('type.ring.ring'(bool),X36)))))) )),
  inference(superposition,[],[f389,f263])).
fof(f799,plain,(
  ( ! [X30,X28,X29] : (s(bool,'T_0') = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X29))) | s(bool,X30) = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|49|_fupd_2'(s(fun(bool,bool),X28),s('type.ring.ring'(bool),X29))))) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f389,f263])).
fof(f1575,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X4))) | s(bool,X5) = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_RN_fupd_2'(s(fun(fun(bool,bool),fun(bool,bool)),X3),s('type.ring.ring'(bool),X4))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f264,f389])).
fof(f1577,plain,(
  ( ! [X10,X11,X9] : (s(bool,X11) = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X10))) | s(bool,'T_0') = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_RN_fupd_2'(s(fun(fun(bool,bool),fun(bool,bool)),X9),s('type.ring.ring'(bool),X10))))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f264,f389])).
fof(f2362,plain,(
  ( ! [X88,X87,X85,X86] : (s(bool,'T_0') != s(bool,'const.combin.K_2'(s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X86))),s(X87,X88))) | s(bool,'T_0') = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|49|_fupd_2'(s(fun(bool,bool),X85),s('type.ring.ring'(bool),X86)))))) )),
  inference(superposition,[],[f2289,f263])).
fof(f2363,plain,(
  ( ! [X92,X90,X91,X89] : (s(bool,'T_0') != s(bool,'const.combin.K_2'(s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X90))),s(X91,X92))) | s(bool,'T_0') = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_RN_fupd_2'(s(fun(fun(bool,bool),fun(bool,bool)),X89),s('type.ring.ring'(bool),X90)))))) )),
  inference(superposition,[],[f2289,f264])).
fof(f4314,plain,(
  ( ! [X8,X7,X9] : (s(bool,'T_0') = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X8))) | s(bool,X9) = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_RP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X7),s('type.ring.ring'(bool),X8))))) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f269,f389])).
fof(f4316,plain,(
  ( ! [X14,X15,X13] : (s(bool,X15) = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X14))) | s(bool,'T_0') = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_RP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X13),s('type.ring.ring'(bool),X14))))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f269,f389])).
fof(f4351,plain,(
  ( ! [X116,X114,X117,X115] : (s(bool,'T_0') != s(bool,'const.combin.K_2'(s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X115))),s(X116,X117))) | s(bool,'T_0') = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_RP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X114),s('type.ring.ring'(bool),X115)))))) )),
  inference(superposition,[],[f2289,f269])).
fof(f7030,plain,(
  ( ! [X6,X7,X5] : (s(bool,'T_0') = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X6))) | s(bool,X7) = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_RM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X5),s('type.ring.ring'(bool),X6))))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f277,f389])).
fof(f7031,plain,(
  ( ! [X10,X8,X9] : (s(bool,X10) = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X9))) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_RM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X8),s('type.ring.ring'(bool),X9)))))) )),
  inference(superposition,[],[f277,f389])).
fof(f7084,plain,(
  ( ! [X194,X192,X195,X193] : (s(bool,'T_0') != s(bool,'const.combin.K_2'(s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),X193))),s(X194,X195))) | s(bool,'T_0') = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_RM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X192),s('type.ring.ring'(bool),X193)))))) )),
  inference(superposition,[],[f2289,f277])).
fof(f769,plain,(
  ( ! [X39,X38,X40] : (s(bool,X40) = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X39))) | s(bool,'T_0') = s(bool,X40) | s(bool,'T_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|48|_fupd_2'(s(fun(bool,bool),X38),s('type.ring.ring'(bool),X39)))))) )),
  inference(superposition,[],[f389,f262])).
fof(f800,plain,(
  ( ! [X33,X31,X32] : (s(bool,'T_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X32))) | s(bool,X33) = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|48|_fupd_2'(s(fun(bool,bool),X31),s('type.ring.ring'(bool),X32))))) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f389,f262])).
fof(f1815,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X4))) | s(bool,X5) = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_RN_fupd_2'(s(fun(fun(bool,bool),fun(bool,bool)),X3),s('type.ring.ring'(bool),X4))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f265,f389])).
fof(f1817,plain,(
  ( ! [X10,X11,X9] : (s(bool,X11) = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X10))) | s(bool,'T_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_RN_fupd_2'(s(fun(fun(bool,bool),fun(bool,bool)),X9),s('type.ring.ring'(bool),X10))))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f265,f389])).
fof(f2364,plain,(
  ( ! [X94,X95,X93,X96] : (s(bool,'T_0') != s(bool,'const.combin.K_2'(s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X94))),s(X95,X96))) | s(bool,'T_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|48|_fupd_2'(s(fun(bool,bool),X93),s('type.ring.ring'(bool),X94)))))) )),
  inference(superposition,[],[f2289,f262])).
fof(f2365,plain,(
  ( ! [X99,X97,X100,X98] : (s(bool,'T_0') != s(bool,'const.combin.K_2'(s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X98))),s(X99,X100))) | s(bool,'T_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_RN_fupd_2'(s(fun(fun(bool,bool),fun(bool,bool)),X97),s('type.ring.ring'(bool),X98)))))) )),
  inference(superposition,[],[f2289,f265])).
fof(f5826,plain,(
  ( ! [X8,X7,X9] : (s(bool,'T_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X8))) | s(bool,X9) = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_RP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X7),s('type.ring.ring'(bool),X8))))) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f275,f389])).
fof(f5828,plain,(
  ( ! [X14,X15,X13] : (s(bool,X15) = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X14))) | s(bool,'T_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_RP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X13),s('type.ring.ring'(bool),X14))))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f275,f389])).
fof(f5875,plain,(
  ( ! [X171,X169,X170,X168] : (s(bool,'T_0') != s(bool,'const.combin.K_2'(s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X169))),s(X170,X171))) | s(bool,'T_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_RP_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X168),s('type.ring.ring'(bool),X169)))))) )),
  inference(superposition,[],[f2289,f275])).
fof(f6202,plain,(
  ( ! [X8,X7,X9] : (s(bool,'T_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X8))) | s(bool,X9) = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_RM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X7),s('type.ring.ring'(bool),X8))))) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f276,f389])).
fof(f6204,plain,(
  ( ! [X14,X15,X13] : (s(bool,X15) = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X14))) | s(bool,'T_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_RM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X13),s('type.ring.ring'(bool),X14))))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f276,f389])).
fof(f6253,plain,(
  ( ! [X182,X180,X183,X181] : (s(bool,'T_0') != s(bool,'const.combin.K_2'(s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),X181))),s(X182,X183))) | s(bool,'T_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_RM_fupd_2'(s(fun(fun(bool,fun(bool,bool)),fun(bool,fun(bool,bool))),X180),s('type.ring.ring'(bool),X181)))))) )),
  inference(superposition,[],[f2289,f276])).
fof(f7312,plain,(
  ( ! [X72,X75] : (s(bool,X72) != s(bool,X75) | s(bool,'T_0') = s(bool,X75) | s(bool,'F_0') = s(bool,X72)) )),
  inference(superposition,[],[f686,f328])).
fof(f7356,plain,(
  ( ! [X39,X41,X38,X42,X40] : (s(bool,X39) != s(bool,'const.combin.K_2'(s(bool,X38),s(X40,X41))) | s(bool,'T_0') = s(bool,X39) | s(bool,X38) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(subsumption_resolution,[],[f7220,f358])).
fof(f7220,plain,(
  ( ! [X39,X41,X38,X42,X40] : (s(bool,X39) != s(bool,'const.combin.K_2'(s(bool,X38),s(X40,X41))) | s(bool,'T_0') = s(bool,X39) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X38) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f686,f389])).
fof(f6490,plain,(
  ( ! [X134,X132,X136,X135,X133] : (s(bool,X132) = s(bool,'const.combin.K_2'(s(bool,X133),s(bool,X132))) | s(bool,X133) = s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X134,X135))) | s(bool,X132) = s(bool,X136) | s(bool,'T_0') = s(bool,X136)) )),
  inference(superposition,[],[f2618,f389])).
fof(f6417,plain,(
  ( ! [X6,X4,X2,X7,X5,X3] : (s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X3),s(X4,X5))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X2),s(X6,X7)))) )),
  inference(superposition,[],[f2618,f2618])).
fof(f6364,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,X1) = s(bool,'const.combin.K_2'(s(bool,X0),s(X2,X3))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X1),s(X4,X5))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2618,f235])).
fof(f6083,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X1),s(X2,X3))) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f6033,f389])).
fof(f6033,plain,(
  ( ! [X68,X69,X65] : (s(bool,'T_0') != s(bool,X65) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X65),s(X68,X69)))) )),
  inference(superposition,[],[f2156,f328])).
fof(f5999,plain,(
  ( ! [X121,X118,X122,X120,X119] : (s(bool,X118) != s(bool,'const.combin.K_2'(s(bool,X119),s(bool,X118))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X119),s(X120,X121))) | s(bool,X118) = s(bool,X122) | s(bool,'T_0') = s(bool,X122)) )),
  inference(superposition,[],[f2156,f389])).
fof(f5531,plain,(
  ( ! [X175,X173,X176,X174] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X173,bool),happ(s(fun(fun(X173,bool),fun(X173,bool)),X174),s(fun(X173,bool),'const.ring.ring_R|48|_1'(s('type.ring.ring'(fun(X173,bool)),X175))))))) | s(X173,X176) = s(X173,sK11(X176,'const.ring.ring_R|48|_1'(s('type.ring.ring'(fun(X173,bool)),'const.ring.ring_R|48|_fupd_2'(s(fun(fun(X173,bool),fun(X173,bool)),X174),s('type.ring.ring'(fun(X173,bool)),X175)))),X173))) )),
  inference(superposition,[],[f563,f273])).
fof(f5008,plain,(
  ( ! [X152,X151,X149,X150] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X149,bool),happ(s(fun(fun(X149,bool),fun(X149,bool)),X150),s(fun(X149,bool),'const.ring.ring_R|49|_1'(s('type.ring.ring'(fun(X149,bool)),X151))))))) | s(X149,X152) = s(X149,sK11(X152,'const.ring.ring_R|49|_1'(s('type.ring.ring'(fun(X149,bool)),'const.ring.ring_R|49|_fupd_2'(s(fun(fun(X149,bool),fun(X149,bool)),X150),s('type.ring.ring'(fun(X149,bool)),X151)))),X149))) )),
  inference(superposition,[],[f563,f272])).
fof(f2382,plain,(
  ( ! [X54,X55,X53] : (s(bool,X53) != s(bool,'const.combin.K_2'(s(bool,X54),s(bool,X53))) | s(bool,'T_0') = s(bool,X54) | s(bool,X53) = s(bool,X55) | s(bool,'T_0') = s(bool,X55)) )),
  inference(superposition,[],[f2289,f389])).
fof(f2367,plain,(
  ( ! [X109,X107,X108,X106] : (s(bool,'T_0') != s(bool,'const.combin.K_2'(s(bool,X106),s(X108,X109))) | s(bool,'T_0') = s(bool,sK11(X106,X107,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X107)))) )),
  inference(superposition,[],[f2289,f563])).
fof(f2361,plain,(
  ( ! [X80,X83,X81,X84,X82] : (s(bool,'T_0') != s(bool,'const.combin.K_2'(s(bool,X82),s(X83,X84))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X80,bool),'const.pair.,_2'(s(X80,X81),s(bool,X82)))))) )),
  inference(superposition,[],[f2289,f330])).
fof(f2360,plain,(
  ( ! [X78,X76,X79,X77] : (s(bool,'T_0') != s(bool,'const.combin.K_2'(s(bool,X76),s(X78,X79))) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X76),s(fun('type.num.num',bool),X77),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f2289,f261])).
fof(f2359,plain,(
  ( ! [X74,X72,X71,X75,X73] : (s(bool,'T_0') != s(bool,'const.combin.K_2'(s(bool,X72),s(X74,X75))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X71),'const.pair.,_2'(s(bool,X72),s(X71,X73)))))) )),
  inference(superposition,[],[f2289,f329])).
fof(f2289,plain,(
  ( ! [X14,X15,X16] : (s(bool,'T_0') != s(bool,'const.combin.K_2'(s(bool,X14),s(X15,X16))) | s(bool,'T_0') = s(bool,X14)) )),
  inference(subsumption_resolution,[],[f2229,f234])).
fof(f2229,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,'const.combin.K_2'(s(bool,X14),s(X15,X16))) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f236,f1940])).
fof(f2156,plain,(
  ( ! [X12,X10,X8,X11,X9] : (s(bool,'T_0') != s(bool,'const.combin.K_2'(s(bool,X8),s(X9,X10))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X8),s(X11,X12)))) )),
  inference(subsumption_resolution,[],[f2031,f234])).
fof(f2031,plain,(
  ( ! [X12,X10,X8,X11,X9] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,'const.combin.K_2'(s(bool,X8),s(X9,X10))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X8),s(X11,X12)))) )),
  inference(superposition,[],[f236,f649])).
fof(f2149,plain,(
  ( ! [X19,X17,X15,X20,X18,X16] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X15),s(X18,X19))) | s(bool,X20) = s(bool,'const.combin.K_2'(s(bool,X15),s(X16,X17))) | s(bool,'T_0') = s(bool,X20)) )),
  inference(subsumption_resolution,[],[f1995,f358])).
fof(f1995,plain,(
  ( ! [X19,X17,X15,X20,X18,X16] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X15),s(X18,X19))) | s(bool,X20) = s(bool,'const.combin.K_2'(s(bool,X15),s(X16,X17))) | s(bool,'T_0') = s(bool,X20)) )),
  inference(superposition,[],[f649,f389])).
fof(f1981,plain,(
  ( ! [X94,X92,X93,X91] : (s(bool,X91) = s(bool,'const.combin.K_2'(s(bool,X92),s(bool,X91))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X92),s(X93,X94))) | s(bool,'T_0') = s(bool,X91)) )),
  inference(superposition,[],[f649,f235])).
fof(f1922,plain,(
  ( ! [X85,X83,X81,X84,X82] : (s(bool,X81) = s(bool,'const.combin.K_2'(s(bool,X82),s(bool,X81))) | s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,X82),s(X83,X84))) | s(bool,X81) = s(bool,X85) | s(bool,'T_0') = s(bool,X85)) )),
  inference(superposition,[],[f649,f389])).
fof(f1783,plain,(
  ( ! [X39,X43,X41,X42,X40] : (s(bool,X41) = s(bool,'|dollar|exists_1'(s(fun(X39,bool),X40))) | s(bool,X41) = s(bool,X42) | s(bool,'T_0') = s(bool,X42) | s(X39,X43) = s(X39,sK11(X43,X40,X39))) )),
  inference(subsumption_resolution,[],[f1755,f232])).
fof(f1755,plain,(
  ( ! [X39,X43,X41,X42,X40] : (~p(s(bool,'T_0')) | s(bool,X41) = s(bool,'|dollar|exists_1'(s(fun(X39,bool),X40))) | s(bool,X41) = s(bool,X42) | s(bool,'T_0') = s(bool,X42) | s(X39,X43) = s(X39,sK11(X43,X40,X39))) )),
  inference(superposition,[],[f1154,f563])).
fof(f1652,plain,(
  ( ! [X50,X51,X49] : (s(bool,X49) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X49),s(bool,X50))))) | s(bool,'F_0') = s(bool,X50) | s(bool,X49) = s(bool,X51) | s(bool,'T_0') = s(bool,X51)) )),
  inference(superposition,[],[f534,f389])).
fof(f1492,plain,(
  ( ! [X47,X45,X46] : (s(bool,X45) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X46),s(bool,X45))))) | s(bool,'F_0') = s(bool,X46) | s(bool,X45) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f490,f389])).
fof(f1270,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(bool,X3) = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK11(X2,X1,X0)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f1258,f232])).
fof(f1258,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(bool,X3) = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK11(X2,X1,X0)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f288,f389])).
fof(f1251,plain,(
  ( ! [X30,X28,X29,X27] : (~p(s(bool,happ(s(fun(X27,bool),X28),s(X27,sK11(X30,sK11(X28,X29,fun(X27,bool)),X27))))) | p(s(bool,'|dollar|exists_1'(s(fun(X27,bool),sK11(X28,X29,fun(X27,bool)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X27,bool),bool),X29)))) )),
  inference(superposition,[],[f288,f563])).
fof(f1268,plain,(
  ( ! [X26,X24,X23,X25,X22] : (p(s(bool,'|dollar|exists_1'(s(fun(X22,bool),X23)))) | ~p(s(bool,happ(s(fun(X22,bool),X23),s(X22,sK11(X26,'const.combin.K_2'(s(fun(X22,bool),X23),s(X24,X25)),X22)))))) )),
  inference(forward_demodulation,[],[f1250,f328])).
fof(f1250,plain,(
  ( ! [X26,X24,X23,X25,X22] : (~p(s(bool,happ(s(fun(X22,bool),X23),s(X22,sK11(X26,'const.combin.K_2'(s(fun(X22,bool),X23),s(X24,X25)),X22))))) | p(s(bool,'|dollar|exists_1'(s(fun(X22,bool),'const.combin.K_2'(s(fun(X22,bool),X23),s(X24,X25))))))) )),
  inference(superposition,[],[f288,f328])).
fof(f1267,plain,(
  ( ! [X21,X19,X20,X18] : (p(s(bool,'|dollar|exists_1'(s(fun(X18,bool),'const.ring.ring_R|49|_1'(s('type.ring.ring'(fun(X18,bool)),X20)))))) | ~p(s(bool,happ(s(fun(X18,bool),'const.ring.ring_R|49|_1'(s('type.ring.ring'(fun(X18,bool)),X20))),s(X18,sK11(X21,'const.ring.ring_R|49|_1'(s('type.ring.ring'(fun(X18,bool)),'const.ring.ring_R|48|_fupd_2'(s(fun(fun(X18,bool),fun(X18,bool)),X19),s('type.ring.ring'(fun(X18,bool)),X20)))),X18)))))) )),
  inference(forward_demodulation,[],[f1249,f262])).
fof(f1249,plain,(
  ( ! [X21,X19,X20,X18] : (~p(s(bool,happ(s(fun(X18,bool),'const.ring.ring_R|49|_1'(s('type.ring.ring'(fun(X18,bool)),X20))),s(X18,sK11(X21,'const.ring.ring_R|49|_1'(s('type.ring.ring'(fun(X18,bool)),'const.ring.ring_R|48|_fupd_2'(s(fun(fun(X18,bool),fun(X18,bool)),X19),s('type.ring.ring'(fun(X18,bool)),X20)))),X18))))) | p(s(bool,'|dollar|exists_1'(s(fun(X18,bool),'const.ring.ring_R|49|_1'(s('type.ring.ring'(fun(X18,bool)),'const.ring.ring_R|48|_fupd_2'(s(fun(fun(X18,bool),fun(X18,bool)),X19),s('type.ring.ring'(fun(X18,bool)),X20))))))))) )),
  inference(superposition,[],[f288,f262])).
fof(f1266,plain,(
  ( ! [X14,X17,X15,X16] : (p(s(bool,'|dollar|exists_1'(s(fun(X14,bool),'const.ring.ring_R|48|_1'(s('type.ring.ring'(fun(X14,bool)),X16)))))) | ~p(s(bool,happ(s(fun(X14,bool),'const.ring.ring_R|48|_1'(s('type.ring.ring'(fun(X14,bool)),X16))),s(X14,sK11(X17,'const.ring.ring_R|48|_1'(s('type.ring.ring'(fun(X14,bool)),'const.ring.ring_R|49|_fupd_2'(s(fun(fun(X14,bool),fun(X14,bool)),X15),s('type.ring.ring'(fun(X14,bool)),X16)))),X14)))))) )),
  inference(forward_demodulation,[],[f1248,f263])).
fof(f1248,plain,(
  ( ! [X14,X17,X15,X16] : (~p(s(bool,happ(s(fun(X14,bool),'const.ring.ring_R|48|_1'(s('type.ring.ring'(fun(X14,bool)),X16))),s(X14,sK11(X17,'const.ring.ring_R|48|_1'(s('type.ring.ring'(fun(X14,bool)),'const.ring.ring_R|49|_fupd_2'(s(fun(fun(X14,bool),fun(X14,bool)),X15),s('type.ring.ring'(fun(X14,bool)),X16)))),X14))))) | p(s(bool,'|dollar|exists_1'(s(fun(X14,bool),'const.ring.ring_R|48|_1'(s('type.ring.ring'(fun(X14,bool)),'const.ring.ring_R|49|_fupd_2'(s(fun(fun(X14,bool),fun(X14,bool)),X15),s('type.ring.ring'(fun(X14,bool)),X16))))))))) )),
  inference(superposition,[],[f288,f263])).
fof(f1265,plain,(
  ( ! [X12,X10,X13,X11,X9] : (p(s(bool,'|dollar|exists_1'(s(fun(X9,bool),X12)))) | ~p(s(bool,happ(s(fun(X9,bool),X12),s(X9,sK11(X13,'const.pair.SND_1'(s('type.pair.prod'(X10,fun(X9,bool)),'const.pair.,_2'(s(X10,X11),s(fun(X9,bool),X12)))),X9)))))) )),
  inference(forward_demodulation,[],[f1247,f330])).
fof(f1247,plain,(
  ( ! [X12,X10,X13,X11,X9] : (~p(s(bool,happ(s(fun(X9,bool),X12),s(X9,sK11(X13,'const.pair.SND_1'(s('type.pair.prod'(X10,fun(X9,bool)),'const.pair.,_2'(s(X10,X11),s(fun(X9,bool),X12)))),X9))))) | p(s(bool,'|dollar|exists_1'(s(fun(X9,bool),'const.pair.SND_1'(s('type.pair.prod'(X10,fun(X9,bool)),'const.pair.,_2'(s(X10,X11),s(fun(X9,bool),X12))))))))) )),
  inference(superposition,[],[f288,f330])).
fof(f1264,plain,(
  ( ! [X6,X8,X7,X5] : (p(s(bool,'|dollar|exists_1'(s(fun(X5,bool),X6)))) | ~p(s(bool,happ(s(fun(X5,bool),X6),s(X5,sK11(X8,'const.ind_type.FCONS_3'(s(fun(X5,bool),X6),s(fun('type.num.num',fun(X5,bool)),X7),s('type.num.num','const.num.|48|_0')),X5)))))) )),
  inference(forward_demodulation,[],[f1246,f261])).
fof(f1246,plain,(
  ( ! [X6,X8,X7,X5] : (~p(s(bool,happ(s(fun(X5,bool),X6),s(X5,sK11(X8,'const.ind_type.FCONS_3'(s(fun(X5,bool),X6),s(fun('type.num.num',fun(X5,bool)),X7),s('type.num.num','const.num.|48|_0')),X5))))) | p(s(bool,'|dollar|exists_1'(s(fun(X5,bool),'const.ind_type.FCONS_3'(s(fun(X5,bool),X6),s(fun('type.num.num',fun(X5,bool)),X7),s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f288,f261])).
fof(f1263,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK11(X4,'const.pair.FST_1'(s('type.pair.prod'(fun(X0,bool),X1),'const.pair.,_2'(s(fun(X0,bool),X2),s(X1,X3)))),X0)))))) )),
  inference(forward_demodulation,[],[f1245,f329])).
fof(f1245,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK11(X4,'const.pair.FST_1'(s('type.pair.prod'(fun(X0,bool),X1),'const.pair.,_2'(s(fun(X0,bool),X2),s(X1,X3)))),X0))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X0,bool),X1),'const.pair.,_2'(s(fun(X0,bool),X2),s(X1,X3))))))))) )),
  inference(superposition,[],[f288,f329])).
fof(f1243,plain,(
  ( ! [X50,X48,X51,X49] : (s(bool,X50) != s(bool,'|dollar|exists_1'(s(fun(X48,bool),X49))) | s(bool,'T_0') = s(bool,X50) | s(X48,X51) = s(X48,sK11(X51,X49,X48))) )),
  inference(subsumption_resolution,[],[f1222,f232])).
fof(f1222,plain,(
  ( ! [X50,X48,X51,X49] : (~p(s(bool,'T_0')) | s(bool,X50) != s(bool,'|dollar|exists_1'(s(fun(X48,bool),X49))) | s(bool,'T_0') = s(bool,X50) | s(X48,X51) = s(X48,sK11(X51,X49,X48))) )),
  inference(superposition,[],[f392,f563])).
fof(f1167,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) = s(bool,X11) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X10))) | s(bool,'T_0') = s(bool,sK11(X9,X10,bool)) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f563,f389])).
fof(f1165,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X4))) | s(bool,X5) = s(bool,sK11(X3,X4,bool)) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f563,f389])).
fof(f1156,plain,(
  ( ! [X30,X31,X29] : (s(bool,X29) = s(bool,X30) | s(bool,'F_0') = s(bool,X30) | s(bool,X29) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) )),
  inference(subsumption_resolution,[],[f914,f232])).
fof(f914,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,'T_0')) | s(bool,X29) = s(bool,X30) | s(bool,'F_0') = s(bool,X30) | s(bool,X29) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) )),
  inference(superposition,[],[f685,f389])).
fof(f1155,plain,(
  ( ! [X24,X23,X25] : (s(bool,X23) != s(bool,X24) | s(bool,'T_0') = s(bool,X24) | s(bool,X23) = s(bool,X25) | s(bool,'T_0') = s(bool,X25)) )),
  inference(subsumption_resolution,[],[f912,f232])).
fof(f912,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,'T_0')) | s(bool,X23) != s(bool,X24) | s(bool,'T_0') = s(bool,X24) | s(bool,X23) = s(bool,X25) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f392,f389])).
fof(f1154,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f910,f232])).
fof(f910,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f326,f389])).
fof(f902,plain,(
  ( ! [X61,X59,X60] : (s(bool,X59) = s(bool,'const.combin.K_2'(s(bool,X60),s(bool,X59))) | s(bool,'F_0') = s(bool,X60) | s(bool,X59) = s(bool,X61) | s(bool,'T_0') = s(bool,X61)) )),
  inference(superposition,[],[f391,f389])).
fof(f801,plain,(
  ( ! [X37,X35,X36,X34] : (s(bool,'T_0') = s(bool,X34) | s(bool,X37) = s(bool,'const.combin.K_2'(s(bool,X34),s(X35,X36))) | s(bool,'T_0') = s(bool,X37)) )),
  inference(superposition,[],[f389,f328])).
fof(f798,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,'T_0') = s(bool,X26) | s(bool,X27) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X24,bool),'const.pair.,_2'(s(X24,X25),s(bool,X26))))) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f389,f330])).
fof(f797,plain,(
  ( ! [X23,X21,X22] : (s(bool,'T_0') = s(bool,X21) | s(bool,X23) = s(bool,'const.ind_type.FCONS_3'(s(bool,X21),s(fun('type.num.num',bool),X22),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X23)) )),
  inference(superposition,[],[f389,f261])).
fof(f796,plain,(
  ( ! [X19,X17,X20,X18] : (s(bool,'T_0') = s(bool,X18) | s(bool,X20) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X17),'const.pair.,_2'(s(bool,X18),s(X17,X19))))) | s(bool,'T_0') = s(bool,X20)) )),
  inference(superposition,[],[f389,f329])).
fof(f786,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f389,f389])).
fof(f770,plain,(
  ( ! [X43,X41,X44,X42] : (s(bool,X41) = s(bool,X44) | s(bool,'T_0') = s(bool,X44) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X41),s(X42,X43)))) )),
  inference(superposition,[],[f389,f328])).
fof(f767,plain,(
  ( ! [X33,X31,X34,X32] : (s(bool,X33) = s(bool,X34) | s(bool,'T_0') = s(bool,X34) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X31,bool),'const.pair.,_2'(s(X31,X32),s(bool,X33)))))) )),
  inference(superposition,[],[f389,f330])).
fof(f766,plain,(
  ( ! [X30,X28,X29] : (s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X28),s(fun('type.num.num',bool),X29),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f389,f261])).
fof(f765,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,X25) = s(bool,X27) | s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X24),'const.pair.,_2'(s(bool,X25),s(X24,X26)))))) )),
  inference(superposition,[],[f389,f329])).
fof(f563,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,sK11(X2,X1,X0)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) )),
  inference(subsumption_resolution,[],[f562,f234])).
fof(f562,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | s(X0,X2) = s(X0,sK11(X2,X1,X0)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) )),
  inference(superposition,[],[f287,f235])).
fof(f426,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f421,f234])).
fof(f421,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f323,f235])).
fof(f392,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f240,f235])).
fof(f389,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f235,f235])).
fof(f358,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f234,f236])).
fof(f333,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))) )),
  inference(cnf_transformation,[],[f193])).
fof(f193,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(flattening,[],[f192])).
fof(f192,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X34,X35,X36] : ! [X37,X38,X39] : s(X34,happ(s(fun(X35,X34),X37),s(X35,happ(s(fun(X36,X35),X38),s(X36,X39))))) = s(X34,happ(s(fun(X36,X34),'_dst_x0x1_2'(s(fun(X35,X34),X37),s(fun(X36,X35),X38))),s(X36,X39)))),
  file('thm.ring.ring_|49||49|','_dst_x0x1_ax')).
fof(f332,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,'const.ring.ring_CASE_2'(s('type.ring.ring'(X1),'const.ring.recordtype|46|ring_5'(s(X1,X2),s(X1,X3),s(fun(X1,fun(X1,X1)),X4),s(fun(X1,fun(X1,X1)),X5),s(fun(X1,X1),X6))),s(fun(X1,fun(X1,fun(fun(X1,fun(X1,X1)),fun(fun(X1,fun(X1,X1)),fun(fun(X1,X1),X0))))),X7))) = s(X0,happ(s(fun(fun(X1,X1),X0),happ(s(fun(fun(X1,fun(X1,X1)),fun(fun(X1,X1),X0)),happ(s(fun(fun(X1,fun(X1,X1)),fun(fun(X1,fun(X1,X1)),fun(fun(X1,X1),X0))),happ(s(fun(X1,fun(fun(X1,fun(X1,X1)),fun(fun(X1,fun(X1,X1)),fun(fun(X1,X1),X0)))),happ(s(fun(X1,fun(X1,fun(fun(X1,fun(X1,X1)),fun(fun(X1,fun(X1,X1)),fun(fun(X1,X1),X0))))),X7),s(X1,X2))),s(X1,X3))),s(fun(X1,fun(X1,X1)),X4))),s(fun(X1,fun(X1,X1)),X5))),s(fun(X1,X1),X6)))) )),
  inference(cnf_transformation,[],[f191])).
fof(f191,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X0,'const.ring.ring_CASE_2'(s('type.ring.ring'(X1),'const.ring.recordtype|46|ring_5'(s(X1,X2),s(X1,X3),s(fun(X1,fun(X1,X1)),X4),s(fun(X1,fun(X1,X1)),X5),s(fun(X1,X1),X6))),s(fun(X1,fun(X1,fun(fun(X1,fun(X1,X1)),fun(fun(X1,fun(X1,X1)),fun(fun(X1,X1),X0))))),X7))) = s(X0,happ(s(fun(fun(X1,X1),X0),happ(s(fun(fun(X1,fun(X1,X1)),fun(fun(X1,X1),X0)),happ(s(fun(fun(X1,fun(X1,X1)),fun(fun(X1,fun(X1,X1)),fun(fun(X1,X1),X0))),happ(s(fun(X1,fun(fun(X1,fun(X1,X1)),fun(fun(X1,fun(X1,X1)),fun(fun(X1,X1),X0)))),happ(s(fun(X1,fun(X1,fun(fun(X1,fun(X1,X1)),fun(fun(X1,fun(X1,X1)),fun(fun(X1,X1),X0))))),X7),s(X1,X2))),s(X1,X3))),s(fun(X1,fun(X1,X1)),X4))),s(fun(X1,fun(X1,X1)),X5))),s(fun(X1,X1),X6)))),
  inference(flattening,[],[f190])).
fof(f190,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6,X7] : s(X0,'const.ring.ring_CASE_2'(s('type.ring.ring'(X1),'const.ring.recordtype|46|ring_5'(s(X1,X2),s(X1,X3),s(fun(X1,fun(X1,X1)),X4),s(fun(X1,fun(X1,X1)),X5),s(fun(X1,X1),X6))),s(fun(X1,fun(X1,fun(fun(X1,fun(X1,X1)),fun(fun(X1,fun(X1,X1)),fun(fun(X1,X1),X0))))),X7))) = s(X0,happ(s(fun(fun(X1,X1),X0),happ(s(fun(fun(X1,fun(X1,X1)),fun(fun(X1,X1),X0)),happ(s(fun(fun(X1,fun(X1,X1)),fun(fun(X1,fun(X1,X1)),fun(fun(X1,X1),X0))),happ(s(fun(X1,fun(fun(X1,fun(X1,X1)),fun(fun(X1,fun(X1,X1)),fun(fun(X1,X1),X0)))),happ(s(fun(X1,fun(X1,fun(fun(X1,fun(X1,X1)),fun(fun(X1,fun(X1,X1)),fun(fun(X1,X1),X0))))),X7),s(X1,X2))),s(X1,X3))),s(fun(X1,fun(X1,X1)),X4))),s(fun(X1,fun(X1,X1)),X5))),s(fun(X1,X1),X6)))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X17,X6] : ! [X7,X8,X9,X10,X11,X19] : s(X17,'const.ring.ring_CASE_2'(s('type.ring.ring'(X6),'const.ring.recordtype|46|ring_5'(s(X6,X7),s(X6,X8),s(fun(X6,fun(X6,X6)),X9),s(fun(X6,fun(X6,X6)),X10),s(fun(X6,X6),X11))),s(fun(X6,fun(X6,fun(fun(X6,fun(X6,X6)),fun(fun(X6,fun(X6,X6)),fun(fun(X6,X6),X17))))),X19))) = s(X17,happ(s(fun(fun(X6,X6),X17),happ(s(fun(fun(X6,fun(X6,X6)),fun(fun(X6,X6),X17)),happ(s(fun(fun(X6,fun(X6,X6)),fun(fun(X6,fun(X6,X6)),fun(fun(X6,X6),X17))),happ(s(fun(X6,fun(fun(X6,fun(X6,X6)),fun(fun(X6,fun(X6,X6)),fun(fun(X6,X6),X17)))),happ(s(fun(X6,fun(X6,fun(fun(X6,fun(X6,X6)),fun(fun(X6,fun(X6,X6)),fun(fun(X6,X6),X17))))),X19),s(X6,X7))),s(X6,X8))),s(fun(X6,fun(X6,X6)),X9))),s(fun(X6,fun(X6,X6)),X10))),s(fun(X6,X6),X11)))),
  file('thm.ring.ring_|49||49|','thm.ring.ring_case_def')).
fof(f331,plain,(
  ( ! [X2,X0,X3,X1] : (s(fun(X0,X1),X2) = s(fun(X0,X1),X3) | s(X1,f(s(X0,sK13(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK13(X3,X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f225])).
fof(f225,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK13(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK13(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13])],[f198])).
fof(f198,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f189])).
fof(f189,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f188])).
fof(f188,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.ring.ring_|49||49|',aHL_EXT)).
fof(f330,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))) )),
  inference(cnf_transformation,[],[f187])).
fof(f187,plain,(
  ! [X0,X1,X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(flattening,[],[f186])).
fof(f186,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X6,X17] : ! [X23,X24] : s(X17,X24) = s(X17,'const.pair.SND_1'(s('type.pair.prod'(X6,X17),'const.pair.,_2'(s(X6,X23),s(X17,X24)))))),
  file('thm.ring.ring_|49||49|','thm.pair.SND')).
fof(f329,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))) )),
  inference(cnf_transformation,[],[f185])).
fof(f185,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(flattening,[],[f184])).
fof(f184,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X17,X6] : ! [X23,X24] : s(X6,X23) = s(X6,'const.pair.FST_1'(s('type.pair.prod'(X6,X17),'const.pair.,_2'(s(X6,X23),s(X17,X24)))))),
  file('thm.ring.ring_|49||49|','thm.pair.FST')).
fof(f328,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.combin.K_2'(s(X1,X2),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f183])).
fof(f183,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,'const.combin.K_2'(s(X1,X2),s(X0,X3)))),
  inference(flattening,[],[f182])).
fof(f182,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,'const.combin.K_2'(s(X1,X2),s(X0,X3)))),
  inference(rectify,[],[f70])).
fof(f70,axiom,(
  ! [X17,X6] : ! [X23,X24] : s(X6,X23) = s(X6,'const.combin.K_2'(s(X6,X23),s(X17,X24)))),
  file('thm.ring.ring_|49||49|','thm.combin.K_THM')).
fof(f327,plain,(
  ( ! [X6,X4,X2,X0,X5,X1] : (s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),sK12(X2,X1,X0)),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),sK12(X2,X1,X0)),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))) )),
  inference(cnf_transformation,[],[f224])).
fof(f224,plain,(
  ! [X0,X1,X2] : ! [X4,X5,X6] : s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),sK12(X2,X1,X0)),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),sK12(X2,X1,X0)),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12])],[f181])).
fof(f181,plain,(
  ! [X0,X1,X2] : ? [X3] : ! [X4,X5,X6] : s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),X3),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),X3),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))),
  inference(flattening,[],[f180])).
fof(f180,plain,(
  ! [X0,X1] : ! [X2] : ? [X3] : ! [X4,X5,X6] : s(X0,happ(s(fun('type.ind_type.recspace'(X1),X0),X3),s('type.ind_type.recspace'(X1),'const.ind_type.CONSTR_3'(s('type.num.num',X4),s(X1,X5),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))))) = s(X0,happ(s(fun(fun('type.num.num',X0),X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)),happ(s(fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0))),happ(s(fun('type.num.num',fun(X1,fun(fun('type.num.num','type.ind_type.recspace'(X1)),fun(fun('type.num.num',X0),X0)))),X2),s('type.num.num',X4))),s(X1,X5))),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6))),s(fun('type.num.num',X0),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X1),X0),X3),s(fun('type.num.num','type.ind_type.recspace'(X1)),X6)))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X17,X6] : ! [X18] : ? [X19] : ! [X20,X21,X22] : s(X17,happ(s(fun('type.ind_type.recspace'(X6),X17),X19),s('type.ind_type.recspace'(X6),'const.ind_type.CONSTR_3'(s('type.num.num',X20),s(X6,X21),s(fun('type.num.num','type.ind_type.recspace'(X6)),X22))))) = s(X17,happ(s(fun(fun('type.num.num',X17),X17),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X6)),fun(fun('type.num.num',X17),X17)),happ(s(fun(X6,fun(fun('type.num.num','type.ind_type.recspace'(X6)),fun(fun('type.num.num',X17),X17))),happ(s(fun('type.num.num',fun(X6,fun(fun('type.num.num','type.ind_type.recspace'(X6)),fun(fun('type.num.num',X17),X17)))),X18),s('type.num.num',X20))),s(X6,X21))),s(fun('type.num.num','type.ind_type.recspace'(X6)),X22))),s(fun('type.num.num',X17),'_dst_x0x1_2'(s(fun('type.ind_type.recspace'(X6),X17),X19),s(fun('type.num.num','type.ind_type.recspace'(X6)),X22)))))),
  file('thm.ring.ring_|49||49|','thm.ind_type.CONSTR_REC')).
fof(f323,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f197])).
fof(f197,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f196])).
fof(f196,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f179])).
fof(f179,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X32,X33] : ((p(s(bool,X32)) => p(s(bool,X33))) => ((p(s(bool,X33)) => p(s(bool,X32))) => s(bool,X32) = s(bool,X33)))),
  file('thm.ring.ring_|49||49|','thm.bool.IMP_ANTISYM_AX')).
fof(f326,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f197])).
fof(f317,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (s(X0,X1) = s(X0,X6) | s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X1))),s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X2))),s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X3))),s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X4))),s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),'const.combin.K_1'(s(fun(X0,X0),X5))),s('type.ring.ring'(X0),'const.bool.ARB_0'))))))))))) != s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X6))),s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X7))),s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X8))),s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X9))),s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),'const.combin.K_1'(s(fun(X0,X0),X10))),s('type.ring.ring'(X0),'const.bool.ARB_0')))))))))))) )),
  inference(cnf_transformation,[],[f223])).
fof(f223,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : ((s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X1))),s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X2))),s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X3))),s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X4))),s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),'const.combin.K_1'(s(fun(X0,X0),X5))),s('type.ring.ring'(X0),'const.bool.ARB_0'))))))))))) != s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X6))),s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X7))),s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X8))),s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X9))),s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),'const.combin.K_1'(s(fun(X0,X0),X10))),s('type.ring.ring'(X0),'const.bool.ARB_0'))))))))))) | (s(X0,X1) = s(X0,X6) & s(X0,X2) = s(X0,X7) & s(fun(X0,fun(X0,X0)),X3) = s(fun(X0,fun(X0,X0)),X8) & s(fun(X0,fun(X0,X0)),X4) = s(fun(X0,fun(X0,X0)),X9) & s(fun(X0,X0),X5) = s(fun(X0,X0),X10))) & (s(X0,X1) != s(X0,X6) | s(X0,X2) != s(X0,X7) | s(fun(X0,fun(X0,X0)),X3) != s(fun(X0,fun(X0,X0)),X8) | s(fun(X0,fun(X0,X0)),X4) != s(fun(X0,fun(X0,X0)),X9) | s(fun(X0,X0),X5) != s(fun(X0,X0),X10) | s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X1))),s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X2))),s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X3))),s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X4))),s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),'const.combin.K_1'(s(fun(X0,X0),X5))),s('type.ring.ring'(X0),'const.bool.ARB_0'))))))))))) = s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X6))),s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X7))),s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X8))),s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X9))),s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),'const.combin.K_1'(s(fun(X0,X0),X10))),s('type.ring.ring'(X0),'const.bool.ARB_0')))))))))))))),
  inference(flattening,[],[f222])).
fof(f222,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : ((s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X1))),s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X2))),s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X3))),s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X4))),s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),'const.combin.K_1'(s(fun(X0,X0),X5))),s('type.ring.ring'(X0),'const.bool.ARB_0'))))))))))) != s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X6))),s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X7))),s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X8))),s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X9))),s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),'const.combin.K_1'(s(fun(X0,X0),X10))),s('type.ring.ring'(X0),'const.bool.ARB_0'))))))))))) | (s(X0,X1) = s(X0,X6) & s(X0,X2) = s(X0,X7) & s(fun(X0,fun(X0,X0)),X3) = s(fun(X0,fun(X0,X0)),X8) & s(fun(X0,fun(X0,X0)),X4) = s(fun(X0,fun(X0,X0)),X9) & s(fun(X0,X0),X5) = s(fun(X0,X0),X10))) & ((s(X0,X1) != s(X0,X6) | s(X0,X2) != s(X0,X7) | s(fun(X0,fun(X0,X0)),X3) != s(fun(X0,fun(X0,X0)),X8) | s(fun(X0,fun(X0,X0)),X4) != s(fun(X0,fun(X0,X0)),X9) | s(fun(X0,X0),X5) != s(fun(X0,X0),X10)) | s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X1))),s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X2))),s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X3))),s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X4))),s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),'const.combin.K_1'(s(fun(X0,X0),X5))),s('type.ring.ring'(X0),'const.bool.ARB_0'))))))))))) = s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X6))),s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X7))),s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X8))),s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X9))),s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),'const.combin.K_1'(s(fun(X0,X0),X10))),s('type.ring.ring'(X0),'const.bool.ARB_0')))))))))))))),
  inference(nnf_transformation,[],[f178])).
fof(f178,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : (s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X1))),s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X2))),s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X3))),s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X4))),s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),'const.combin.K_1'(s(fun(X0,X0),X5))),s('type.ring.ring'(X0),'const.bool.ARB_0'))))))))))) = s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X6))),s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X7))),s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X8))),s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X9))),s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),'const.combin.K_1'(s(fun(X0,X0),X10))),s('type.ring.ring'(X0),'const.bool.ARB_0'))))))))))) <=> (s(X0,X1) = s(X0,X6) & s(X0,X2) = s(X0,X7) & s(fun(X0,fun(X0,X0)),X3) = s(fun(X0,fun(X0,X0)),X8) & s(fun(X0,fun(X0,X0)),X4) = s(fun(X0,fun(X0,X0)),X9) & s(fun(X0,X0),X5) = s(fun(X0,X0),X10)))),
  inference(flattening,[],[f177])).
fof(f177,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : (s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X1))),s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X2))),s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X3))),s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X4))),s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),'const.combin.K_1'(s(fun(X0,X0),X5))),s('type.ring.ring'(X0),'const.bool.ARB_0'))))))))))) = s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X6))),s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X7))),s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X8))),s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X9))),s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),'const.combin.K_1'(s(fun(X0,X0),X10))),s('type.ring.ring'(X0),'const.bool.ARB_0'))))))))))) <=> (s(X0,X1) = s(X0,X6) & s(X0,X2) = s(X0,X7) & s(fun(X0,fun(X0,X0)),X3) = s(fun(X0,fun(X0,X0)),X8) & s(fun(X0,fun(X0,X0)),X4) = s(fun(X0,fun(X0,X0)),X9) & s(fun(X0,X0),X5) = s(fun(X0,X0),X10)))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X6] : ! [X45,X8,X46,X47,X42,X48,X9,X49,X50,X40] : (s('type.ring.ring'(X6),'const.ring.ring_R|48|_fupd_2'(s(fun(X6,X6),'const.combin.K_1'(s(X6,X45))),s('type.ring.ring'(X6),'const.ring.ring_R|49|_fupd_2'(s(fun(X6,X6),'const.combin.K_1'(s(X6,X8))),s('type.ring.ring'(X6),'const.ring.ring_RP_fupd_2'(s(fun(fun(X6,fun(X6,X6)),fun(X6,fun(X6,X6))),'const.combin.K_1'(s(fun(X6,fun(X6,X6)),X46))),s('type.ring.ring'(X6),'const.ring.ring_RM_fupd_2'(s(fun(fun(X6,fun(X6,X6)),fun(X6,fun(X6,X6))),'const.combin.K_1'(s(fun(X6,fun(X6,X6)),X47))),s('type.ring.ring'(X6),'const.ring.ring_RN_fupd_2'(s(fun(fun(X6,X6),fun(X6,X6)),'const.combin.K_1'(s(fun(X6,X6),X42))),s('type.ring.ring'(X6),'const.bool.ARB_0'))))))))))) = s('type.ring.ring'(X6),'const.ring.ring_R|48|_fupd_2'(s(fun(X6,X6),'const.combin.K_1'(s(X6,X48))),s('type.ring.ring'(X6),'const.ring.ring_R|49|_fupd_2'(s(fun(X6,X6),'const.combin.K_1'(s(X6,X9))),s('type.ring.ring'(X6),'const.ring.ring_RP_fupd_2'(s(fun(fun(X6,fun(X6,X6)),fun(X6,fun(X6,X6))),'const.combin.K_1'(s(fun(X6,fun(X6,X6)),X49))),s('type.ring.ring'(X6),'const.ring.ring_RM_fupd_2'(s(fun(fun(X6,fun(X6,X6)),fun(X6,fun(X6,X6))),'const.combin.K_1'(s(fun(X6,fun(X6,X6)),X50))),s('type.ring.ring'(X6),'const.ring.ring_RN_fupd_2'(s(fun(fun(X6,X6),fun(X6,X6)),'const.combin.K_1'(s(fun(X6,X6),X40))),s('type.ring.ring'(X6),'const.bool.ARB_0'))))))))))) <=> (s(X6,X45) = s(X6,X48) & s(X6,X8) = s(X6,X9) & s(fun(X6,fun(X6,X6)),X46) = s(fun(X6,fun(X6,X6)),X49) & s(fun(X6,fun(X6,X6)),X47) = s(fun(X6,fun(X6,X6)),X50) & s(fun(X6,X6),X40) = s(fun(X6,X6),X42)))),
  file('thm.ring.ring_|49||49|','thm.ring.ring_literal_|49||49|')).
fof(f318,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (s(X0,X2) = s(X0,X7) | s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X1))),s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X2))),s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X3))),s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X4))),s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),'const.combin.K_1'(s(fun(X0,X0),X5))),s('type.ring.ring'(X0),'const.bool.ARB_0'))))))))))) != s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X6))),s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X7))),s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X8))),s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X9))),s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),'const.combin.K_1'(s(fun(X0,X0),X10))),s('type.ring.ring'(X0),'const.bool.ARB_0')))))))))))) )),
  inference(cnf_transformation,[],[f223])).
fof(f319,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (s(fun(X0,fun(X0,X0)),X3) = s(fun(X0,fun(X0,X0)),X8) | s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X1))),s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X2))),s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X3))),s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X4))),s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),'const.combin.K_1'(s(fun(X0,X0),X5))),s('type.ring.ring'(X0),'const.bool.ARB_0'))))))))))) != s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X6))),s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X7))),s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X8))),s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X9))),s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),'const.combin.K_1'(s(fun(X0,X0),X10))),s('type.ring.ring'(X0),'const.bool.ARB_0')))))))))))) )),
  inference(cnf_transformation,[],[f223])).
fof(f320,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (s(fun(X0,fun(X0,X0)),X4) = s(fun(X0,fun(X0,X0)),X9) | s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X1))),s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X2))),s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X3))),s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X4))),s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),'const.combin.K_1'(s(fun(X0,X0),X5))),s('type.ring.ring'(X0),'const.bool.ARB_0'))))))))))) != s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X6))),s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X7))),s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X8))),s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X9))),s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),'const.combin.K_1'(s(fun(X0,X0),X10))),s('type.ring.ring'(X0),'const.bool.ARB_0')))))))))))) )),
  inference(cnf_transformation,[],[f223])).
fof(f321,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (s(fun(X0,X0),X5) = s(fun(X0,X0),X10) | s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X1))),s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X2))),s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X3))),s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X4))),s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),'const.combin.K_1'(s(fun(X0,X0),X5))),s('type.ring.ring'(X0),'const.bool.ARB_0'))))))))))) != s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X6))),s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X7))),s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X8))),s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X9))),s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),'const.combin.K_1'(s(fun(X0,X0),X10))),s('type.ring.ring'(X0),'const.bool.ARB_0')))))))))))) )),
  inference(cnf_transformation,[],[f223])).
fof(f322,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X1))),s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X2))),s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X3))),s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X4))),s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),'const.combin.K_1'(s(fun(X0,X0),X5))),s('type.ring.ring'(X0),'const.bool.ARB_0'))))))))))) = s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X6))),s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X7))),s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X8))),s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X9))),s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),'const.combin.K_1'(s(fun(X0,X0),X10))),s('type.ring.ring'(X0),'const.bool.ARB_0'))))))))))) | s(fun(X0,X0),X5) != s(fun(X0,X0),X10) | s(fun(X0,fun(X0,X0)),X4) != s(fun(X0,fun(X0,X0)),X9) | s(fun(X0,fun(X0,X0)),X3) != s(fun(X0,fun(X0,X0)),X8) | s(X0,X2) != s(X0,X7) | s(X0,X1) != s(X0,X6)) )),
  inference(cnf_transformation,[],[f223])).
fof(f316,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X1),s('type.ring.ring'(X0),'const.ring.recordtype|46|ring_5'(s(X0,X2),s(X0,X3),s(fun(X0,fun(X0,X0)),X4),s(fun(X0,fun(X0,X0)),X5),s(fun(X0,X0),X6))))) = s('type.ring.ring'(X0),'const.ring.recordtype|46|ring_5'(s(X0,X2),s(X0,X3),s(fun(X0,fun(X0,X0)),happ(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X1),s(fun(X0,fun(X0,X0)),X4))),s(fun(X0,fun(X0,X0)),X5),s(fun(X0,X0),X6)))) )),
  inference(cnf_transformation,[],[f176])).
fof(f176,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X1),s('type.ring.ring'(X0),'const.ring.recordtype|46|ring_5'(s(X0,X2),s(X0,X3),s(fun(X0,fun(X0,X0)),X4),s(fun(X0,fun(X0,X0)),X5),s(fun(X0,X0),X6))))) = s('type.ring.ring'(X0),'const.ring.recordtype|46|ring_5'(s(X0,X2),s(X0,X3),s(fun(X0,fun(X0,X0)),happ(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X1),s(fun(X0,fun(X0,X0)),X4))),s(fun(X0,fun(X0,X0)),X5),s(fun(X0,X0),X6)))),
  inference(flattening,[],[f175])).
fof(f175,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6] : s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X1),s('type.ring.ring'(X0),'const.ring.recordtype|46|ring_5'(s(X0,X2),s(X0,X3),s(fun(X0,fun(X0,X0)),X4),s(fun(X0,fun(X0,X0)),X5),s(fun(X0,X0),X6))))) = s('type.ring.ring'(X0),'const.ring.recordtype|46|ring_5'(s(X0,X2),s(X0,X3),s(fun(X0,fun(X0,X0)),happ(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X1),s(fun(X0,fun(X0,X0)),X4))),s(fun(X0,fun(X0,X0)),X5),s(fun(X0,X0),X6)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X6] : ! [X40,X25,X7,X19,X41,X42] : s('type.ring.ring'(X6),'const.ring.ring_RP_fupd_2'(s(fun(fun(X6,fun(X6,X6)),fun(X6,fun(X6,X6))),X40),s('type.ring.ring'(X6),'const.ring.recordtype|46|ring_5'(s(X6,X25),s(X6,X7),s(fun(X6,fun(X6,X6)),X19),s(fun(X6,fun(X6,X6)),X41),s(fun(X6,X6),X42))))) = s('type.ring.ring'(X6),'const.ring.recordtype|46|ring_5'(s(X6,X25),s(X6,X7),s(fun(X6,fun(X6,X6)),happ(s(fun(fun(X6,fun(X6,X6)),fun(X6,fun(X6,X6))),X40),s(fun(X6,fun(X6,X6)),X19))),s(fun(X6,fun(X6,X6)),X41),s(fun(X6,X6),X42)))),
  file('thm.ring.ring_|49||49|','thm.ring.ring_RP_fupd')).
fof(f315,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X1),s('type.ring.ring'(X0),'const.ring.recordtype|46|ring_5'(s(X0,X2),s(X0,X3),s(fun(X0,fun(X0,X0)),X4),s(fun(X0,fun(X0,X0)),X5),s(fun(X0,X0),X6))))) = s('type.ring.ring'(X0),'const.ring.recordtype|46|ring_5'(s(X0,X2),s(X0,X3),s(fun(X0,fun(X0,X0)),X4),s(fun(X0,fun(X0,X0)),happ(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X1),s(fun(X0,fun(X0,X0)),X5))),s(fun(X0,X0),X6)))) )),
  inference(cnf_transformation,[],[f174])).
fof(f174,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X1),s('type.ring.ring'(X0),'const.ring.recordtype|46|ring_5'(s(X0,X2),s(X0,X3),s(fun(X0,fun(X0,X0)),X4),s(fun(X0,fun(X0,X0)),X5),s(fun(X0,X0),X6))))) = s('type.ring.ring'(X0),'const.ring.recordtype|46|ring_5'(s(X0,X2),s(X0,X3),s(fun(X0,fun(X0,X0)),X4),s(fun(X0,fun(X0,X0)),happ(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X1),s(fun(X0,fun(X0,X0)),X5))),s(fun(X0,X0),X6)))),
  inference(flattening,[],[f173])).
fof(f173,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6] : s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X1),s('type.ring.ring'(X0),'const.ring.recordtype|46|ring_5'(s(X0,X2),s(X0,X3),s(fun(X0,fun(X0,X0)),X4),s(fun(X0,fun(X0,X0)),X5),s(fun(X0,X0),X6))))) = s('type.ring.ring'(X0),'const.ring.recordtype|46|ring_5'(s(X0,X2),s(X0,X3),s(fun(X0,fun(X0,X0)),X4),s(fun(X0,fun(X0,X0)),happ(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X1),s(fun(X0,fun(X0,X0)),X5))),s(fun(X0,X0),X6)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X6] : ! [X40,X25,X7,X19,X41,X42] : s('type.ring.ring'(X6),'const.ring.ring_RM_fupd_2'(s(fun(fun(X6,fun(X6,X6)),fun(X6,fun(X6,X6))),X40),s('type.ring.ring'(X6),'const.ring.recordtype|46|ring_5'(s(X6,X25),s(X6,X7),s(fun(X6,fun(X6,X6)),X19),s(fun(X6,fun(X6,X6)),X41),s(fun(X6,X6),X42))))) = s('type.ring.ring'(X6),'const.ring.recordtype|46|ring_5'(s(X6,X25),s(X6,X7),s(fun(X6,fun(X6,X6)),X19),s(fun(X6,fun(X6,X6)),happ(s(fun(fun(X6,fun(X6,X6)),fun(X6,fun(X6,X6))),X40),s(fun(X6,fun(X6,X6)),X41))),s(fun(X6,X6),X42)))),
  file('thm.ring.ring_|49||49|','thm.ring.ring_RM_fupd')).
fof(f314,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),X1),s('type.ring.ring'(X0),'const.ring.recordtype|46|ring_5'(s(X0,X2),s(X0,X3),s(fun(X0,fun(X0,X0)),X4),s(fun(X0,fun(X0,X0)),X5),s(fun(X0,X0),X6))))) = s('type.ring.ring'(X0),'const.ring.recordtype|46|ring_5'(s(X0,X2),s(X0,X3),s(fun(X0,fun(X0,X0)),X4),s(fun(X0,fun(X0,X0)),X5),s(fun(X0,X0),happ(s(fun(fun(X0,X0),fun(X0,X0)),X1),s(fun(X0,X0),X6)))))) )),
  inference(cnf_transformation,[],[f172])).
fof(f172,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),X1),s('type.ring.ring'(X0),'const.ring.recordtype|46|ring_5'(s(X0,X2),s(X0,X3),s(fun(X0,fun(X0,X0)),X4),s(fun(X0,fun(X0,X0)),X5),s(fun(X0,X0),X6))))) = s('type.ring.ring'(X0),'const.ring.recordtype|46|ring_5'(s(X0,X2),s(X0,X3),s(fun(X0,fun(X0,X0)),X4),s(fun(X0,fun(X0,X0)),X5),s(fun(X0,X0),happ(s(fun(fun(X0,X0),fun(X0,X0)),X1),s(fun(X0,X0),X6)))))),
  inference(flattening,[],[f171])).
fof(f171,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6] : s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),X1),s('type.ring.ring'(X0),'const.ring.recordtype|46|ring_5'(s(X0,X2),s(X0,X3),s(fun(X0,fun(X0,X0)),X4),s(fun(X0,fun(X0,X0)),X5),s(fun(X0,X0),X6))))) = s('type.ring.ring'(X0),'const.ring.recordtype|46|ring_5'(s(X0,X2),s(X0,X3),s(fun(X0,fun(X0,X0)),X4),s(fun(X0,fun(X0,X0)),X5),s(fun(X0,X0),happ(s(fun(fun(X0,X0),fun(X0,X0)),X1),s(fun(X0,X0),X6)))))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X6] : ! [X40,X25,X7,X19,X41,X42] : s('type.ring.ring'(X6),'const.ring.ring_RN_fupd_2'(s(fun(fun(X6,X6),fun(X6,X6)),X40),s('type.ring.ring'(X6),'const.ring.recordtype|46|ring_5'(s(X6,X25),s(X6,X7),s(fun(X6,fun(X6,X6)),X19),s(fun(X6,fun(X6,X6)),X41),s(fun(X6,X6),X42))))) = s('type.ring.ring'(X6),'const.ring.recordtype|46|ring_5'(s(X6,X25),s(X6,X7),s(fun(X6,fun(X6,X6)),X19),s(fun(X6,fun(X6,X6)),X41),s(fun(X6,X6),happ(s(fun(fun(X6,X6),fun(X6,X6)),X40),s(fun(X6,X6),X42)))))),
  file('thm.ring.ring_|49||49|','thm.ring.ring_RN_fupd')).
fof(f313,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.ring.ring'(X0),'const.ring.recordtype|46|ring_5'(s(X0,X2),s(X0,happ(s(fun(X0,X0),X1),s(X0,X3))),s(fun(X0,fun(X0,X0)),X4),s(fun(X0,fun(X0,X0)),X5),s(fun(X0,X0),X6))) = s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),X1),s('type.ring.ring'(X0),'const.ring.recordtype|46|ring_5'(s(X0,X2),s(X0,X3),s(fun(X0,fun(X0,X0)),X4),s(fun(X0,fun(X0,X0)),X5),s(fun(X0,X0),X6)))))) )),
  inference(cnf_transformation,[],[f170])).
fof(f170,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s('type.ring.ring'(X0),'const.ring.recordtype|46|ring_5'(s(X0,X2),s(X0,happ(s(fun(X0,X0),X1),s(X0,X3))),s(fun(X0,fun(X0,X0)),X4),s(fun(X0,fun(X0,X0)),X5),s(fun(X0,X0),X6))) = s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),X1),s('type.ring.ring'(X0),'const.ring.recordtype|46|ring_5'(s(X0,X2),s(X0,X3),s(fun(X0,fun(X0,X0)),X4),s(fun(X0,fun(X0,X0)),X5),s(fun(X0,X0),X6)))))),
  inference(flattening,[],[f169])).
fof(f169,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6] : s('type.ring.ring'(X0),'const.ring.recordtype|46|ring_5'(s(X0,X2),s(X0,happ(s(fun(X0,X0),X1),s(X0,X3))),s(fun(X0,fun(X0,X0)),X4),s(fun(X0,fun(X0,X0)),X5),s(fun(X0,X0),X6))) = s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),X1),s('type.ring.ring'(X0),'const.ring.recordtype|46|ring_5'(s(X0,X2),s(X0,X3),s(fun(X0,fun(X0,X0)),X4),s(fun(X0,fun(X0,X0)),X5),s(fun(X0,X0),X6)))))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X6] : ! [X40,X25,X7,X19,X41,X42] : s('type.ring.ring'(X6),'const.ring.recordtype|46|ring_5'(s(X6,X25),s(X6,happ(s(fun(X6,X6),X40),s(X6,X7))),s(fun(X6,fun(X6,X6)),X19),s(fun(X6,fun(X6,X6)),X41),s(fun(X6,X6),X42))) = s('type.ring.ring'(X6),'const.ring.ring_R|49|_fupd_2'(s(fun(X6,X6),X40),s('type.ring.ring'(X6),'const.ring.recordtype|46|ring_5'(s(X6,X25),s(X6,X7),s(fun(X6,fun(X6,X6)),X19),s(fun(X6,fun(X6,X6)),X41),s(fun(X6,X6),X42)))))),
  file('thm.ring.ring_|49||49|','thm.ring.ring_R|49|_fupd')).
fof(f312,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.ring.ring'(X0),'const.ring.recordtype|46|ring_5'(s(X0,happ(s(fun(X0,X0),X1),s(X0,X2))),s(X0,X3),s(fun(X0,fun(X0,X0)),X4),s(fun(X0,fun(X0,X0)),X5),s(fun(X0,X0),X6))) = s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),X1),s('type.ring.ring'(X0),'const.ring.recordtype|46|ring_5'(s(X0,X2),s(X0,X3),s(fun(X0,fun(X0,X0)),X4),s(fun(X0,fun(X0,X0)),X5),s(fun(X0,X0),X6)))))) )),
  inference(cnf_transformation,[],[f168])).
fof(f168,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s('type.ring.ring'(X0),'const.ring.recordtype|46|ring_5'(s(X0,happ(s(fun(X0,X0),X1),s(X0,X2))),s(X0,X3),s(fun(X0,fun(X0,X0)),X4),s(fun(X0,fun(X0,X0)),X5),s(fun(X0,X0),X6))) = s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),X1),s('type.ring.ring'(X0),'const.ring.recordtype|46|ring_5'(s(X0,X2),s(X0,X3),s(fun(X0,fun(X0,X0)),X4),s(fun(X0,fun(X0,X0)),X5),s(fun(X0,X0),X6)))))),
  inference(flattening,[],[f167])).
fof(f167,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6] : s('type.ring.ring'(X0),'const.ring.recordtype|46|ring_5'(s(X0,happ(s(fun(X0,X0),X1),s(X0,X2))),s(X0,X3),s(fun(X0,fun(X0,X0)),X4),s(fun(X0,fun(X0,X0)),X5),s(fun(X0,X0),X6))) = s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),X1),s('type.ring.ring'(X0),'const.ring.recordtype|46|ring_5'(s(X0,X2),s(X0,X3),s(fun(X0,fun(X0,X0)),X4),s(fun(X0,fun(X0,X0)),X5),s(fun(X0,X0),X6)))))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X6] : ! [X40,X25,X7,X19,X41,X42] : s('type.ring.ring'(X6),'const.ring.recordtype|46|ring_5'(s(X6,happ(s(fun(X6,X6),X40),s(X6,X25))),s(X6,X7),s(fun(X6,fun(X6,X6)),X19),s(fun(X6,fun(X6,X6)),X41),s(fun(X6,X6),X42))) = s('type.ring.ring'(X6),'const.ring.ring_R|48|_fupd_2'(s(fun(X6,X6),X40),s('type.ring.ring'(X6),'const.ring.recordtype|46|ring_5'(s(X6,X25),s(X6,X7),s(fun(X6,fun(X6,X6)),X19),s(fun(X6,fun(X6,X6)),X41),s(fun(X6,X6),X42)))))),
  file('thm.ring.ring_|49||49|','thm.ring.ring_R|48|_fupd')).
fof(f311,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,X2))),s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,X3))))))) = s('type.num.num','const.ring.ring_size_2'(s(fun(X0,'type.num.num'),X1),s('type.ring.ring'(X0),'const.ring.recordtype|46|ring_5'(s(X0,X2),s(X0,X3),s(fun(X0,fun(X0,X0)),X4),s(fun(X0,fun(X0,X0)),X5),s(fun(X0,X0),X6)))))) )),
  inference(cnf_transformation,[],[f166])).
fof(f166,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,X2))),s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,X3))))))) = s('type.num.num','const.ring.ring_size_2'(s(fun(X0,'type.num.num'),X1),s('type.ring.ring'(X0),'const.ring.recordtype|46|ring_5'(s(X0,X2),s(X0,X3),s(fun(X0,fun(X0,X0)),X4),s(fun(X0,fun(X0,X0)),X5),s(fun(X0,X0),X6)))))),
  inference(flattening,[],[f165])).
fof(f165,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,X2))),s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,X3))))))) = s('type.num.num','const.ring.ring_size_2'(s(fun(X0,'type.num.num'),X1),s('type.ring.ring'(X0),'const.ring.recordtype|46|ring_5'(s(X0,X2),s(X0,X3),s(fun(X0,fun(X0,X0)),X4),s(fun(X0,fun(X0,X0)),X5),s(fun(X0,X0),X6)))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X6] : ! [X19,X7,X8,X9,X10,X11] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',happ(s(fun(X6,'type.num.num'),X19),s(X6,X7))),s('type.num.num',happ(s(fun(X6,'type.num.num'),X19),s(X6,X8))))))) = s('type.num.num','const.ring.ring_size_2'(s(fun(X6,'type.num.num'),X19),s('type.ring.ring'(X6),'const.ring.recordtype|46|ring_5'(s(X6,X7),s(X6,X8),s(fun(X6,fun(X6,X6)),X9),s(fun(X6,fun(X6,X6)),X10),s(fun(X6,X6),X11)))))),
  file('thm.ring.ring_|49||49|','thm.ring.ring_size_def')).
fof(f310,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(fun(X0,fun(X0,X0)),X4) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),'const.ring.recordtype|46|ring_5'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4),s(fun(X0,X0),X5)))))) )),
  inference(cnf_transformation,[],[f164])).
fof(f164,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(fun(X0,fun(X0,X0)),X4) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),'const.ring.recordtype|46|ring_5'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4),s(fun(X0,X0),X5)))))),
  inference(flattening,[],[f163])).
fof(f163,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5] : s(fun(X0,fun(X0,X0)),X4) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),'const.ring.recordtype|46|ring_5'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4),s(fun(X0,X0),X5)))))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X6] : ! [X25,X7,X19,X41,X42] : s(fun(X6,fun(X6,X6)),X41) = s(fun(X6,fun(X6,X6)),'const.ring.ring_RM_1'(s('type.ring.ring'(X6),'const.ring.recordtype|46|ring_5'(s(X6,X25),s(X6,X7),s(fun(X6,fun(X6,X6)),X19),s(fun(X6,fun(X6,X6)),X41),s(fun(X6,X6),X42)))))),
  file('thm.ring.ring_|49||49|','thm.ring.ring_RM')).
fof(f308,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(fun(X0,X0),X5) = s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),'const.ring.recordtype|46|ring_5'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4),s(fun(X0,X0),X5)))))) )),
  inference(cnf_transformation,[],[f160])).
fof(f160,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(fun(X0,X0),X5) = s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),'const.ring.recordtype|46|ring_5'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4),s(fun(X0,X0),X5)))))),
  inference(flattening,[],[f159])).
fof(f159,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5] : s(fun(X0,X0),X5) = s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),'const.ring.recordtype|46|ring_5'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4),s(fun(X0,X0),X5)))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X6] : ! [X25,X7,X19,X41,X42] : s(fun(X6,X6),X42) = s(fun(X6,X6),'const.ring.ring_RN_1'(s('type.ring.ring'(X6),'const.ring.recordtype|46|ring_5'(s(X6,X25),s(X6,X7),s(fun(X6,fun(X6,X6)),X19),s(fun(X6,fun(X6,X6)),X41),s(fun(X6,X6),X42)))))),
  file('thm.ring.ring_|49||49|','thm.ring.ring_RN')).
fof(f307,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,X2) = s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),'const.ring.recordtype|46|ring_5'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4),s(fun(X0,X0),X5)))))) )),
  inference(cnf_transformation,[],[f158])).
fof(f158,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,X2) = s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),'const.ring.recordtype|46|ring_5'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4),s(fun(X0,X0),X5)))))),
  inference(flattening,[],[f157])).
fof(f157,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5] : s(X0,X2) = s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),'const.ring.recordtype|46|ring_5'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4),s(fun(X0,X0),X5)))))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X6] : ! [X25,X7,X19,X41,X42] : s(X6,X7) = s(X6,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X6),'const.ring.recordtype|46|ring_5'(s(X6,X25),s(X6,X7),s(fun(X6,fun(X6,X6)),X19),s(fun(X6,fun(X6,X6)),X41),s(fun(X6,X6),X42)))))),
  file('thm.ring.ring_|49||49|','thm.ring.ring_R|49|')).
fof(f306,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,X1) = s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),'const.ring.recordtype|46|ring_5'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4),s(fun(X0,X0),X5)))))) )),
  inference(cnf_transformation,[],[f156])).
fof(f156,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,X1) = s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),'const.ring.recordtype|46|ring_5'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4),s(fun(X0,X0),X5)))))),
  inference(flattening,[],[f155])).
fof(f155,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5] : s(X0,X1) = s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),'const.ring.recordtype|46|ring_5'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4),s(fun(X0,X0),X5)))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X6] : ! [X25,X7,X19,X41,X42] : s(X6,X25) = s(X6,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X6),'const.ring.recordtype|46|ring_5'(s(X6,X25),s(X6,X7),s(fun(X6,fun(X6,X6)),X19),s(fun(X6,fun(X6,X6)),X41),s(fun(X6,X6),X42)))))),
  file('thm.ring.ring_|49||49|','thm.ring.ring_R|48|')).
fof(f305,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))) )),
  inference(cnf_transformation,[],[f154])).
fof(f154,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))),
  inference(flattening,[],[f153])).
fof(f153,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),happ(s(fun(fun('type.num.num',X0),fun('type.num.num',X0)),happ(s(fun(X0,fun(fun('type.num.num',X0),fun('type.num.num',X0))),'const.ind_type.FCONS_0'),s(X0,X1))),s(fun('type.num.num',X0),X2))),s('type.num.num',X3)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X6] : ! [X27,X28,X29] : s(X6,'const.ind_type.FCONS_3'(s(X6,X27),s(fun('type.num.num',X6),X28),s('type.num.num',X29))) = s(X6,happ(s(fun('type.num.num',X6),happ(s(fun(fun('type.num.num',X6),fun('type.num.num',X6)),happ(s(fun(X6,fun(fun('type.num.num',X6),fun('type.num.num',X6))),'const.ind_type.FCONS_0'),s(X6,X27))),s(fun('type.num.num',X6),X28))),s('type.num.num',X29)))),
  file('thm.ring.ring_|49||49|','ARITY_0_3_const.ind_type.FCONS')).
fof(f304,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.ind_type.recspace'(X0),'const.ind_type.CONSTR_3'(s('type.num.num',X1),s(X0,X2),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3))) = s('type.ind_type.recspace'(X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2))),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3)))) )),
  inference(cnf_transformation,[],[f152])).
fof(f152,plain,(
  ! [X0,X1,X2,X3] : s('type.ind_type.recspace'(X0),'const.ind_type.CONSTR_3'(s('type.num.num',X1),s(X0,X2),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3))) = s('type.ind_type.recspace'(X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2))),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3)))),
  inference(flattening,[],[f151])).
fof(f151,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.ind_type.recspace'(X0),'const.ind_type.CONSTR_3'(s('type.num.num',X1),s(X0,X2),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3))) = s('type.ind_type.recspace'(X0),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0)),happ(s(fun(X0,fun(fun('type.num.num','type.ind_type.recspace'(X0)),'type.ind_type.recspace'(X0))),'const.ind_type.CONSTR_1'(s('type.num.num',X1))),s(X0,X2))),s(fun('type.num.num','type.ind_type.recspace'(X0)),X3)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X6] : ! [X27,X28,X29] : s('type.ind_type.recspace'(X6),'const.ind_type.CONSTR_3'(s('type.num.num',X27),s(X6,X28),s(fun('type.num.num','type.ind_type.recspace'(X6)),X29))) = s('type.ind_type.recspace'(X6),happ(s(fun(fun('type.num.num','type.ind_type.recspace'(X6)),'type.ind_type.recspace'(X6)),happ(s(fun(X6,fun(fun('type.num.num','type.ind_type.recspace'(X6)),'type.ind_type.recspace'(X6))),'const.ind_type.CONSTR_1'(s('type.num.num',X27))),s(X6,X28))),s(fun('type.num.num','type.ind_type.recspace'(X6)),X29)))),
  file('thm.ring.ring_|49||49|','ARITY_1_3_const.ind_type.CONSTR')).
fof(f303,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ring.ring_RM_3'(s('type.ring.ring'(X0),X1),s(X0,X2),s(X0,X3))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))),s(X0,X2))),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f150])).
fof(f150,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.ring.ring_RM_3'(s('type.ring.ring'(X0),X1),s(X0,X2),s(X0,X3))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))),s(X0,X2))),s(X0,X3)))),
  inference(flattening,[],[f149])).
fof(f149,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,'const.ring.ring_RM_3'(s('type.ring.ring'(X0),X1),s(X0,X2),s(X0,X3))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))),s(X0,X2))),s(X0,X3)))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X6] : ! [X27,X28,X29] : s(X6,'const.ring.ring_RM_3'(s('type.ring.ring'(X6),X27),s(X6,X28),s(X6,X29))) = s(X6,happ(s(fun(X6,X6),happ(s(fun(X6,fun(X6,X6)),'const.ring.ring_RM_1'(s('type.ring.ring'(X6),X27))),s(X6,X28))),s(X6,X29)))),
  file('thm.ring.ring_|49||49|','ARITY_1_3_const.ring.ring_RM')).
fof(f302,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ring.ring_RP_3'(s('type.ring.ring'(X0),X1),s(X0,X2),s(X0,X3))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))),s(X0,X2))),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f148])).
fof(f148,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.ring.ring_RP_3'(s('type.ring.ring'(X0),X1),s(X0,X2),s(X0,X3))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))),s(X0,X2))),s(X0,X3)))),
  inference(flattening,[],[f147])).
fof(f147,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,'const.ring.ring_RP_3'(s('type.ring.ring'(X0),X1),s(X0,X2),s(X0,X3))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))),s(X0,X2))),s(X0,X3)))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X6] : ! [X27,X28,X29] : s(X6,'const.ring.ring_RP_3'(s('type.ring.ring'(X6),X27),s(X6,X28),s(X6,X29))) = s(X6,happ(s(fun(X6,X6),happ(s(fun(X6,fun(X6,X6)),'const.ring.ring_RP_1'(s('type.ring.ring'(X6),X27))),s(X6,X28))),s(X6,X29)))),
  file('thm.ring.ring_|49||49|','ARITY_1_3_const.ring.ring_RP')).
fof(f301,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))) )),
  inference(cnf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0,X1,X2,X3] : s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))),
  inference(flattening,[],[f145])).
fof(f145,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X3))) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6] : ! [X25,X19,X26] : s(X6,happ(s(fun('type.num.num',X6),X19),s('type.num.num',X26))) = s(X6,'const.ind_type.FCONS_3'(s(X6,X25),s(fun('type.num.num',X6),X19),s('type.num.num','const.num.SUC_1'(s('type.num.num',X26)))))),
  file('thm.ring.ring_|49||49|','thm.ind_type.FCONS|split|1')).
fof(f289,plain,(
  ( ! [X2,X0,X1] : (s('type.ring.ring'(X0),X1) != s('type.ring.ring'(X0),X2) | s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X2)))) )),
  inference(cnf_transformation,[],[f219])).
fof(f290,plain,(
  ( ! [X2,X0,X1] : (s('type.ring.ring'(X0),X1) != s('type.ring.ring'(X0),X2) | s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X2)))) )),
  inference(cnf_transformation,[],[f219])).
fof(f292,plain,(
  ( ! [X2,X0,X1] : (s('type.ring.ring'(X0),X1) != s('type.ring.ring'(X0),X2) | s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X2)))) )),
  inference(cnf_transformation,[],[f219])).
fof(f293,plain,(
  ( ! [X2,X0,X1] : (s('type.ring.ring'(X0),X1) != s('type.ring.ring'(X0),X2) | s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X2)))) )),
  inference(cnf_transformation,[],[f219])).
fof(f294,plain,(
  ( ! [X2,X0,X1] : (s('type.ring.ring'(X0),X1) = s('type.ring.ring'(X0),X2) | s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X1))) != s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X2))) | s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))) != s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X2))) | s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))) != s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X2))) | s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))) != s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X2))) | s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))) != s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X2)))) )),
  inference(cnf_transformation,[],[f219])).
fof(f288,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK11(X2,X1,X0))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f217])).
fof(f286,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.ring.ring'(X0),X1))))) = s(fun(X0,fun(X0,X0)),happ(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f138])).
fof(f138,plain,(
  ! [X0,X1,X2] : s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.ring.ring'(X0),X1))))) = s(fun(X0,fun(X0,X0)),happ(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f137])).
fof(f137,plain,(
  ! [X0] : ! [X1,X2] : s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.ring.ring'(X0),X1))))) = s(fun(X0,fun(X0,X0)),happ(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X6] : ! [X22,X19] : s(fun(X6,fun(X6,X6)),'const.ring.ring_RP_1'(s('type.ring.ring'(X6),'const.ring.ring_RP_fupd_2'(s(fun(fun(X6,fun(X6,X6)),fun(X6,fun(X6,X6))),X19),s('type.ring.ring'(X6),X22))))) = s(fun(X6,fun(X6,X6)),happ(s(fun(fun(X6,fun(X6,X6)),fun(X6,fun(X6,X6))),X19),s(fun(X6,fun(X6,X6)),'const.ring.ring_RP_1'(s('type.ring.ring'(X6),X22)))))),
  file('thm.ring.ring_|49||49|','thm.ring.ring_4|48|_prove_ltb')).
fof(f285,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.ring.ring'(X0),X1))))) = s(fun(X0,fun(X0,X0)),happ(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f136])).
fof(f136,plain,(
  ! [X0,X1,X2] : s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.ring.ring'(X0),X1))))) = s(fun(X0,fun(X0,X0)),happ(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f135])).
fof(f135,plain,(
  ! [X0] : ! [X1,X2] : s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.ring.ring'(X0),X1))))) = s(fun(X0,fun(X0,X0)),happ(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f49])).
fof(f49,axiom,(
  ! [X6] : ! [X22,X19] : s(fun(X6,fun(X6,X6)),'const.ring.ring_RM_1'(s('type.ring.ring'(X6),'const.ring.ring_RM_fupd_2'(s(fun(fun(X6,fun(X6,X6)),fun(X6,fun(X6,X6))),X19),s('type.ring.ring'(X6),X22))))) = s(fun(X6,fun(X6,X6)),happ(s(fun(fun(X6,fun(X6,X6)),fun(X6,fun(X6,X6))),X19),s(fun(X6,fun(X6,X6)),'const.ring.ring_RM_1'(s('type.ring.ring'(X6),X22)))))),
  file('thm.ring.ring_|49||49|','thm.ring.ring_4|49|_prove_ltb')).
fof(f284,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),X2),s('type.ring.ring'(X0),X1))))) = s(fun(X0,X0),happ(s(fun(fun(X0,X0),fun(X0,X0)),X2),s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f134])).
fof(f134,plain,(
  ! [X0,X1,X2] : s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),X2),s('type.ring.ring'(X0),X1))))) = s(fun(X0,X0),happ(s(fun(fun(X0,X0),fun(X0,X0)),X2),s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f133])).
fof(f133,plain,(
  ! [X0] : ! [X1,X2] : s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),X2),s('type.ring.ring'(X0),X1))))) = s(fun(X0,X0),happ(s(fun(fun(X0,X0),fun(X0,X0)),X2),s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f55])).
fof(f55,axiom,(
  ! [X6] : ! [X22,X19] : s(fun(X6,X6),'const.ring.ring_RN_1'(s('type.ring.ring'(X6),'const.ring.ring_RN_fupd_2'(s(fun(fun(X6,X6),fun(X6,X6)),X19),s('type.ring.ring'(X6),X22))))) = s(fun(X6,X6),happ(s(fun(fun(X6,X6),fun(X6,X6)),X19),s(fun(X6,X6),'const.ring.ring_RN_1'(s('type.ring.ring'(X6),X22)))))),
  file('thm.ring.ring_|49||49|','thm.ring.ring_42_prove_ltb')).
fof(f283,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f132])).
fof(f132,plain,(
  ! [X0,X1,X2] : s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f131])).
fof(f131,plain,(
  ! [X0] : ! [X1,X2] : s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f52])).
fof(f52,axiom,(
  ! [X6] : ! [X22,X19] : s(fun(X6,fun(X6,X6)),'const.ring.ring_RM_1'(s('type.ring.ring'(X6),X22))) = s(fun(X6,fun(X6,X6)),'const.ring.ring_RM_1'(s('type.ring.ring'(X6),'const.ring.ring_RP_fupd_2'(s(fun(fun(X6,fun(X6,X6)),fun(X6,fun(X6,X6))),X19),s('type.ring.ring'(X6),X22)))))),
  file('thm.ring.ring_|49||49|','thm.ring.ring_32_prove_ltb')).
fof(f282,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f130])).
fof(f130,plain,(
  ! [X0,X1,X2] : s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f129])).
fof(f129,plain,(
  ! [X0] : ! [X1,X2] : s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f62])).
fof(f62,axiom,(
  ! [X6] : ! [X22,X19] : s(fun(X6,fun(X6,X6)),'const.ring.ring_RP_1'(s('type.ring.ring'(X6),X22))) = s(fun(X6,fun(X6,X6)),'const.ring.ring_RP_1'(s('type.ring.ring'(X6),'const.ring.ring_RM_fupd_2'(s(fun(fun(X6,fun(X6,X6)),fun(X6,fun(X6,X6))),X19),s('type.ring.ring'(X6),X22)))))),
  file('thm.ring.ring_|49||49|','thm.ring.ring_28_prove_ltb')).
fof(f281,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X0,X1,X2] : s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f127])).
fof(f127,plain,(
  ! [X0] : ! [X1,X2] : s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f50])).
fof(f50,axiom,(
  ! [X6] : ! [X22,X19] : s(fun(X6,X6),'const.ring.ring_RN_1'(s('type.ring.ring'(X6),X22))) = s(fun(X6,X6),'const.ring.ring_RN_1'(s('type.ring.ring'(X6),'const.ring.ring_RM_fupd_2'(s(fun(fun(X6,fun(X6,X6)),fun(X6,fun(X6,X6))),X19),s('type.ring.ring'(X6),X22)))))),
  file('thm.ring.ring_|49||49|','thm.ring.ring_37_prove_ltb')).
fof(f280,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0,X1,X2] : s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f125])).
fof(f125,plain,(
  ! [X0] : ! [X1,X2] : s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f48])).
fof(f48,axiom,(
  ! [X6] : ! [X22,X19] : s(fun(X6,X6),'const.ring.ring_RN_1'(s('type.ring.ring'(X6),X22))) = s(fun(X6,X6),'const.ring.ring_RN_1'(s('type.ring.ring'(X6),'const.ring.ring_RP_fupd_2'(s(fun(fun(X6,fun(X6,X6)),fun(X6,fun(X6,X6))),X19),s('type.ring.ring'(X6),X22)))))),
  file('thm.ring.ring_|49||49|','thm.ring.ring_36_prove_ltb')).
fof(f279,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),X2),s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f124])).
fof(f124,plain,(
  ! [X0,X1,X2] : s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),X2),s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f123])).
fof(f123,plain,(
  ! [X0] : ! [X1,X2] : s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),X2),s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f53])).
fof(f53,axiom,(
  ! [X6] : ! [X22,X19] : s(fun(X6,fun(X6,X6)),'const.ring.ring_RM_1'(s('type.ring.ring'(X6),X22))) = s(fun(X6,fun(X6,X6)),'const.ring.ring_RM_1'(s('type.ring.ring'(X6),'const.ring.ring_RN_fupd_2'(s(fun(fun(X6,X6),fun(X6,X6)),X19),s('type.ring.ring'(X6),X22)))))),
  file('thm.ring.ring_|49||49|','thm.ring.ring_33_prove_ltb')).
fof(f278,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),X2),s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0,X1,X2] : s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),X2),s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f121])).
fof(f121,plain,(
  ! [X0] : ! [X1,X2] : s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),X2),s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f66])).
fof(f66,axiom,(
  ! [X6] : ! [X22,X19] : s(fun(X6,fun(X6,X6)),'const.ring.ring_RP_1'(s('type.ring.ring'(X6),X22))) = s(fun(X6,fun(X6,X6)),'const.ring.ring_RP_1'(s('type.ring.ring'(X6),'const.ring.ring_RN_fupd_2'(s(fun(fun(X6,X6),fun(X6,X6)),X19),s('type.ring.ring'(X6),X22)))))),
  file('thm.ring.ring_|49||49|','thm.ring.ring_29_prove_ltb')).
fof(f277,plain,(
  ( ! [X2,X0,X1] : (s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0,X1,X2] : s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f119])).
fof(f119,plain,(
  ! [X0] : ! [X1,X2] : s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f68])).
fof(f68,axiom,(
  ! [X6] : ! [X22,X19] : s(X6,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X6),X22))) = s(X6,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X6),'const.ring.ring_RM_fupd_2'(s(fun(fun(X6,fun(X6,X6)),fun(X6,fun(X6,X6))),X19),s('type.ring.ring'(X6),X22)))))),
  file('thm.ring.ring_|49||49|','thm.ring.ring_2|48|_prove_ltb')).
fof(f276,plain,(
  ( ! [X2,X0,X1] : (s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0,X1,X2] : s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f117])).
fof(f117,plain,(
  ! [X0] : ! [X1,X2] : s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f63])).
fof(f63,axiom,(
  ! [X6] : ! [X22,X19] : s(X6,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X6),X22))) = s(X6,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X6),'const.ring.ring_RM_fupd_2'(s(fun(fun(X6,fun(X6,X6)),fun(X6,fun(X6,X6))),X19),s('type.ring.ring'(X6),X22)))))),
  file('thm.ring.ring_|49||49|','thm.ring.ring_24_prove_ltb')).
fof(f275,plain,(
  ( ! [X2,X0,X1] : (s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f116])).
fof(f116,plain,(
  ! [X0,X1,X2] : s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f115])).
fof(f115,plain,(
  ! [X0] : ! [X1,X2] : s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f61])).
fof(f61,axiom,(
  ! [X6] : ! [X22,X19] : s(X6,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X6),X22))) = s(X6,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X6),'const.ring.ring_RP_fupd_2'(s(fun(fun(X6,fun(X6,X6)),fun(X6,fun(X6,X6))),X19),s('type.ring.ring'(X6),X22)))))),
  file('thm.ring.ring_|49||49|','thm.ring.ring_23_prove_ltb')).
fof(f274,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f114])).
fof(f114,plain,(
  ! [X0,X1,X2] : s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f113])).
fof(f113,plain,(
  ! [X0] : ! [X1,X2] : s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f59])).
fof(f59,axiom,(
  ! [X6] : ! [X22,X19] : s(fun(X6,fun(X6,X6)),'const.ring.ring_RP_1'(s('type.ring.ring'(X6),X22))) = s(fun(X6,fun(X6,X6)),'const.ring.ring_RP_1'(s('type.ring.ring'(X6),'const.ring.ring_R|49|_fupd_2'(s(fun(X6,X6),X19),s('type.ring.ring'(X6),X22)))))),
  file('thm.ring.ring_|49||49|','thm.ring.ring_27_prove_ltb')).
fof(f273,plain,(
  ( ! [X2,X0,X1] : (s(X0,happ(s(fun(X0,X0),X2),s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))))) = s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f112])).
fof(f112,plain,(
  ! [X0,X1,X2] : s(X0,happ(s(fun(X0,X0),X2),s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))))) = s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f111])).
fof(f111,plain,(
  ! [X0] : ! [X1,X2] : s(X0,happ(s(fun(X0,X0),X2),s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))))) = s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X6] : ! [X22,X19] : s(X6,happ(s(fun(X6,X6),X19),s(X6,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X6),X22))))) = s(X6,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X6),'const.ring.ring_R|48|_fupd_2'(s(fun(X6,X6),X19),s('type.ring.ring'(X6),X22)))))),
  file('thm.ring.ring_|49||49|','thm.ring.ring_38_prove_ltb')).
fof(f272,plain,(
  ( ! [X2,X0,X1] : (s(X0,happ(s(fun(X0,X0),X2),s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))))) = s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f110])).
fof(f110,plain,(
  ! [X0,X1,X2] : s(X0,happ(s(fun(X0,X0),X2),s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))))) = s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f109])).
fof(f109,plain,(
  ! [X0] : ! [X1,X2] : s(X0,happ(s(fun(X0,X0),X2),s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))))) = s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X6] : ! [X22,X19] : s(X6,happ(s(fun(X6,X6),X19),s(X6,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X6),X22))))) = s(X6,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X6),'const.ring.ring_R|49|_fupd_2'(s(fun(X6,X6),X19),s('type.ring.ring'(X6),X22)))))),
  file('thm.ring.ring_|49||49|','thm.ring.ring_39_prove_ltb')).
fof(f271,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f108])).
fof(f108,plain,(
  ! [X0,X1,X2] : s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f107])).
fof(f107,plain,(
  ! [X0] : ! [X1,X2] : s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f54])).
fof(f54,axiom,(
  ! [X6] : ! [X22,X19] : s(fun(X6,fun(X6,X6)),'const.ring.ring_RM_1'(s('type.ring.ring'(X6),X22))) = s(fun(X6,fun(X6,X6)),'const.ring.ring_RM_1'(s('type.ring.ring'(X6),'const.ring.ring_R|49|_fupd_2'(s(fun(X6,X6),X19),s('type.ring.ring'(X6),X22)))))),
  file('thm.ring.ring_|49||49|','thm.ring.ring_3|49|_prove_ltb')).
fof(f270,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f106])).
fof(f106,plain,(
  ! [X0,X1,X2] : s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f105])).
fof(f105,plain,(
  ! [X0] : ! [X1,X2] : s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f58])).
fof(f58,axiom,(
  ! [X6] : ! [X22,X19] : s(fun(X6,fun(X6,X6)),'const.ring.ring_RP_1'(s('type.ring.ring'(X6),X22))) = s(fun(X6,fun(X6,X6)),'const.ring.ring_RP_1'(s('type.ring.ring'(X6),'const.ring.ring_R|48|_fupd_2'(s(fun(X6,X6),X19),s('type.ring.ring'(X6),X22)))))),
  file('thm.ring.ring_|49||49|','thm.ring.ring_26_prove_ltb')).
fof(f269,plain,(
  ( ! [X2,X0,X1] : (s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f104])).
fof(f104,plain,(
  ! [X0,X1,X2] : s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f103])).
fof(f103,plain,(
  ! [X0] : ! [X1,X2] : s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f65])).
fof(f65,axiom,(
  ! [X6] : ! [X22,X19] : s(X6,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X6),X22))) = s(X6,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X6),'const.ring.ring_RP_fupd_2'(s(fun(fun(X6,fun(X6,X6)),fun(X6,fun(X6,X6))),X19),s('type.ring.ring'(X6),X22)))))),
  file('thm.ring.ring_|49||49|','thm.ring.ring_|49|9_prove_ltb')).
fof(f268,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0,X1,X2] : s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f101])).
fof(f101,plain,(
  ! [X0] : ! [X1,X2] : s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X6] : ! [X22,X19] : s(fun(X6,fun(X6,X6)),'const.ring.ring_RM_1'(s('type.ring.ring'(X6),X22))) = s(fun(X6,fun(X6,X6)),'const.ring.ring_RM_1'(s('type.ring.ring'(X6),'const.ring.ring_R|48|_fupd_2'(s(fun(X6,X6),X19),s('type.ring.ring'(X6),X22)))))),
  file('thm.ring.ring_|49||49|','thm.ring.ring_3|48|_prove_ltb')).
fof(f267,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f100])).
fof(f100,plain,(
  ! [X0,X1,X2] : s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f99])).
fof(f99,plain,(
  ! [X0] : ! [X1,X2] : s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f56])).
fof(f56,axiom,(
  ! [X6] : ! [X22,X19] : s(fun(X6,X6),'const.ring.ring_RN_1'(s('type.ring.ring'(X6),X22))) = s(fun(X6,X6),'const.ring.ring_RN_1'(s('type.ring.ring'(X6),'const.ring.ring_R|48|_fupd_2'(s(fun(X6,X6),X19),s('type.ring.ring'(X6),X22)))))),
  file('thm.ring.ring_|49||49|','thm.ring.ring_34_prove_ltb')).
fof(f266,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f98])).
fof(f98,plain,(
  ! [X0,X1,X2] : s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f97])).
fof(f97,plain,(
  ! [X0] : ! [X1,X2] : s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f51])).
fof(f51,axiom,(
  ! [X6] : ! [X22,X19] : s(fun(X6,X6),'const.ring.ring_RN_1'(s('type.ring.ring'(X6),X22))) = s(fun(X6,X6),'const.ring.ring_RN_1'(s('type.ring.ring'(X6),'const.ring.ring_R|49|_fupd_2'(s(fun(X6,X6),X19),s('type.ring.ring'(X6),X22)))))),
  file('thm.ring.ring_|49||49|','thm.ring.ring_35_prove_ltb')).
fof(f265,plain,(
  ( ! [X2,X0,X1] : (s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),X2),s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f96])).
fof(f96,plain,(
  ! [X0,X1,X2] : s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),X2),s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f95])).
fof(f95,plain,(
  ! [X0] : ! [X1,X2] : s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),X2),s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f60])).
fof(f60,axiom,(
  ! [X6] : ! [X22,X19] : s(X6,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X6),X22))) = s(X6,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X6),'const.ring.ring_RN_fupd_2'(s(fun(fun(X6,X6),fun(X6,X6)),X19),s('type.ring.ring'(X6),X22)))))),
  file('thm.ring.ring_|49||49|','thm.ring.ring_25_prove_ltb')).
fof(f264,plain,(
  ( ! [X2,X0,X1] : (s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),X2),s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f94])).
fof(f94,plain,(
  ! [X0,X1,X2] : s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),X2),s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f93])).
fof(f93,plain,(
  ! [X0] : ! [X1,X2] : s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),X2),s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f67])).
fof(f67,axiom,(
  ! [X6] : ! [X22,X19] : s(X6,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X6),X22))) = s(X6,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X6),'const.ring.ring_RN_fupd_2'(s(fun(fun(X6,X6),fun(X6,X6)),X19),s('type.ring.ring'(X6),X22)))))),
  file('thm.ring.ring_|49||49|','thm.ring.ring_2|49|_prove_ltb')).
fof(f263,plain,(
  ( ! [X2,X0,X1] : (s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f92])).
fof(f92,plain,(
  ! [X0,X1,X2] : s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f91])).
fof(f91,plain,(
  ! [X0] : ! [X1,X2] : s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f69])).
fof(f69,axiom,(
  ! [X6] : ! [X22,X19] : s(X6,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X6),X22))) = s(X6,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X6),'const.ring.ring_R|49|_fupd_2'(s(fun(X6,X6),X19),s('type.ring.ring'(X6),X22)))))),
  file('thm.ring.ring_|49||49|','thm.ring.ring_|49|8_prove_ltb')).
fof(f262,plain,(
  ( ! [X2,X0,X1] : (s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f90])).
fof(f90,plain,(
  ! [X0,X1,X2] : s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f89])).
fof(f89,plain,(
  ! [X0] : ! [X1,X2] : s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f64])).
fof(f64,axiom,(
  ! [X6] : ! [X22,X19] : s(X6,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X6),X22))) = s(X6,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X6),'const.ring.ring_R|48|_fupd_2'(s(fun(X6,X6),X19),s('type.ring.ring'(X6),X22)))))),
  file('thm.ring.ring_|49||49|','thm.ring.ring_22_prove_ltb')).
fof(f261,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f88])).
fof(f88,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(flattening,[],[f87])).
fof(f87,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X6] : ! [X25,X19] : s(X6,X25) = s(X6,'const.ind_type.FCONS_3'(s(X6,X25),s(fun('type.num.num',X6),X19),s('type.num.num','const.num.|48|_0')))),
  file('thm.ring.ring_|49||49|','thm.ind_type.FCONS|split|0')).
fof(f236,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f202])).
fof(f202,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X31] : (s(bool,'T_0') = s(bool,X31) <=> p(s(bool,X31)))),
  file('thm.ring.ring_|49||49|','thm.bool.EQ_CLAUSES|split|1')).
fof(f237,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f202])).
fof(f235,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.ring.ring_|49||49|',aHL_BOOL_CASES)).
fof(f232,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.ring.ring_|49||49|',aHL_TRUTH)).
fof(f9917,plain,(
  $false | $spl0),
  inference(global_subsumption,[],[f231,f230,f229,f228,f227,f226,f232,f235,f237,f236,f261,f262,f263,f264,f265,f266,f267,f268,f269,f270,f271,f272,f273,f274,f275,f276,f277,f278,f279,f280,f281,f282,f283,f284,f285,f286,f288,f294,f293,f292,f291,f290,f289,f301,f302,f303,f304,f305,f306,f307,f308,f309,f310,f311,f312,f313,f314,f315,f316,f322,f321,f320,f319,f318,f317,f326,f323,f327,f328,f329,f330,f331,f332,f333,f358,f389,f392,f426,f563,f765,f766,f767,f770,f786,f796,f797,f798,f801,f902,f1154,f1155,f1156,f1165,f1167,f1243,f1263,f1264,f1265,f1266,f1267,f1268,f1251,f1270,f1492,f1652,f1783,f1886,f1887,f1888,f1922,f1981,f2149,f2156,f2289,f2359,f2360,f2361,f2367,f2382,f3174,f5008,f5531,f5999,f6033,f6083,f6364,f6417,f6490,f7335,f7345,f7347,f7329,f7352,f7214,f7217,f7356,f7223,f7224,f7225,f7226,f7227,f7328,f7231,f7312,f7886,f7911,f6253,f6204,f6202,f5875,f5828,f5826,f2365,f2364,f1817,f1815,f800,f769,f7084,f7031,f7030,f4351,f4316,f4314,f2363,f2362,f1577,f1575,f799,f768,f1333,f1372,f2901,f7569,f2318,f2319,f2566,f2567,f7079,f6248,f6120,f5978,f5871,f5620,f2962,f2954,f2953,f2952,f2951,f2950,f2949,f2948,f2947,f2946,f1765,f1202,f1201,f1195,f1193,f1190,f1189,f6258,f6257,f5879,f5878,f1854,f1853,f1174,f631,f7092,f7091,f4355,f4354,f1611,f1610,f1173,f740,f5530,f5007,f3090,f3091,f6987,f6959,f5772,f5759,f2131,f1242,f1205,f1176,f1175,f1172,f1171,f1170,f561,f560,f559,f558,f287,f5623,f5483,f5481,f5626,f4964,f4962,f1269,f1277,f1314,f3086,f3084,f3042,f3040,f7924,f7884,f7882,f7881,f7880,f7879,f7878,f7877,f7876,f7875,f7874,f7873,f7872,f7871,f7870,f7869,f7847,f7846,f7845,f7844,f7843,f7842,f7841,f7840,f7839,f7838,f7837,f7836,f7835,f7834,f7833,f7832,f7353,f7249,f7247,f7246,f7244,f7243,f7242,f7240,f7239,f7238,f7237,f7236,f7235,f7234,f7212,f7210,f7209,f7208,f7207,f7206,f7205,f7204,f7203,f7202,f7201,f7200,f7199,f7198,f7197,f7082,f7080,f7074,f7073,f7072,f7069,f7029,f6939,f6938,f6937,f6936,f6932,f6363,f6300,f6251,f6249,f6243,f6242,f6241,f6238,f5919,f5874,f5872,f5866,f5865,f5864,f5861,f5775,f5774,f5773,f5771,f5770,f5722,f5719,f5716,f5567,f5042,f4855,f4691,f4561,f4379,f4350,f4345,f3173,f3172,f3168,f3166,f3165,f3164,f3163,f3162,f3161,f3160,f3159,f3158,f3157,f3116,f3085,f3079,f3078,f3077,f3074,f2976,f2965,f2964,f2963,f2961,f2960,f2619,f2618,f2613,f2612,f2611,f2515,f2488,f2341,f2190,f2188,f2187,f2186,f2185,f2155,f2154,f2153,f1964,f1940,f1907,f1848,f1847,f1846,f1843,f1819,f1785,f1784,f1782,f1781,f1740,f1739,f1684,f1683,f1682,f1681,f1680,f1679,f1678,f1677,f1606,f1605,f1602,f1579,f1554,f1553,f1479,f1478,f1477,f1476,f1475,f1474,f1473,f1459,f1458,f1411,f1410,f1409,f1408,f1407,f1406,f1405,f1244,f1203,f1200,f1192,f1191,f1169,f735,f718,f699,f686,f685,f649,f648,f647,f646,f645,f612,f534,f490,f450,f391,f241,f240,f234,f8641,f336,f9874,f9875,f9880,f9881,f9883,f9884,f9886,f9887,f9900,f9903,f9906,f9910,f9914,f9891])).
fof(f9892,plain,(
  $false | ($spl0 | $spl5)),
  inference(subsumption_resolution,[],[f9891,f341])).
fof(f9878,plain,(
  $false | ($spl0 | $spl5)),
  inference(subsumption_resolution,[],[f9877,f341])).
fof(f9737,plain,(
  $false | ($spl0 | $spl5 | $spl6)),
  inference(subsumption_resolution,[],[f341,f9701])).
fof(f9701,plain,(
  s(sK0,sK2) = s(sK0,sK7) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f9700,f307])).
fof(f9700,plain,(
  s(sK0,sK7) = s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))))) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f9652,f307])).
fof(f9652,plain,(
  s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))))) = s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK6),s(sK0,sK7),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK9),s(fun(sK0,sK0),sK10))))) | ($spl0 | $spl6)),
  inference(unit_resulting_resolution,[],[f9535,f290])).
fof(f9535,plain,(
  s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))) = s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK6),s(sK0,sK7),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK9),s(fun(sK0,sK0),sK10))) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f336,f342])).
fof(f9712,plain,(
  $false | ($spl0 | $spl3 | $spl6)),
  inference(subsumption_resolution,[],[f9711,f339])).
fof(f9711,plain,(
  s(sK0,sK1) = s(sK0,sK6) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f9710,f306])).
fof(f9710,plain,(
  s(sK0,sK6) = s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))))) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f9648,f306])).
fof(f9648,plain,(
  s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))))) = s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK6),s(sK0,sK7),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK9),s(fun(sK0,sK0),sK10))))) | ($spl0 | $spl6)),
  inference(unit_resulting_resolution,[],[f9535,f289])).
fof(f9699,plain,(
  $false | ($spl0 | $spl3 | $spl6)),
  inference(subsumption_resolution,[],[f9698,f339])).
fof(f9698,plain,(
  s(sK0,sK1) = s(sK0,sK6) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f9697,f306])).
fof(f9697,plain,(
  s(sK0,sK6) = s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))))) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f9653,f306])).
fof(f9653,plain,(
  s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))))) = s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK6),s(sK0,sK7),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK9),s(fun(sK0,sK0),sK10))))) | ($spl0 | $spl6)),
  inference(unit_resulting_resolution,[],[f9535,f289])).
fof(f9534,plain,(
  $false | ($spl0 | $spl3 | $spl6 | $spl10)),
  inference(subsumption_resolution,[],[f339,f9500])).
fof(f9500,plain,(
  s(sK0,sK1) = s(sK0,sK6) | ($spl0 | $spl6 | $spl10)),
  inference(forward_demodulation,[],[f9499,f306])).
fof(f9499,plain,(
  s(sK0,sK6) = s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))))) | ($spl0 | $spl6 | $spl10)),
  inference(forward_demodulation,[],[f9456,f306])).
fof(f9456,plain,(
  s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))))) = s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK6),s(sK0,sK7),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK9),s(fun(sK0,sK0),sK5))))) | ($spl0 | $spl6 | $spl10)),
  inference(unit_resulting_resolution,[],[f355,f289])).
fof(f355,plain,(
  s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))) = s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK6),s(sK0,sK7),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK9),s(fun(sK0,sK0),sK5))) | ($spl0 | $spl6 | $spl10)),
  inference(forward_demodulation,[],[f354,f342])).
fof(f354,plain,(
  s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))) = s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK6),s(sK0,sK7),s(fun(sK0,fun(sK0,sK0)),sK8),s(fun(sK0,fun(sK0,sK0)),sK9),s(fun(sK0,sK0),sK5))) | ($spl0 | $spl10)),
  inference(forward_demodulation,[],[f336,f346])).
fof(f9518,plain,(
  $false | ($spl0 | $spl6 | $spl9 | $spl10)),
  inference(subsumption_resolution,[],[f9517,f345])).
fof(f9517,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK4) = s(fun(sK0,fun(sK0,sK0)),sK9) | ($spl0 | $spl6 | $spl10)),
  inference(forward_demodulation,[],[f9516,f310])).
fof(f9516,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK9) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))))) | ($spl0 | $spl6 | $spl10)),
  inference(forward_demodulation,[],[f9447,f310])).
fof(f9447,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK6),s(sK0,sK7),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK9),s(fun(sK0,sK0),sK5))))) | ($spl0 | $spl6 | $spl10)),
  inference(unit_resulting_resolution,[],[f355,f292])).
fof(f9509,plain,(
  $false | ($spl0 | $spl6 | $spl9 | $spl10)),
  inference(subsumption_resolution,[],[f9508,f345])).
fof(f9508,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK4) = s(fun(sK0,fun(sK0,sK0)),sK9) | ($spl0 | $spl6 | $spl10)),
  inference(forward_demodulation,[],[f9507,f310])).
fof(f9507,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK9) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))))) | ($spl0 | $spl6 | $spl10)),
  inference(forward_demodulation,[],[f9452,f310])).
fof(f9452,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK6),s(sK0,sK7),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK9),s(fun(sK0,sK0),sK5))))) | ($spl0 | $spl6 | $spl10)),
  inference(unit_resulting_resolution,[],[f355,f292])).
fof(f9045,plain,(
  $false | ($spl0 | $spl2 | $spl6 | $spl9 | $spl10)),
  inference(subsumption_resolution,[],[f9044,f345])).
fof(f9044,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK4) = s(fun(sK0,fun(sK0,sK0)),sK9) | ($spl0 | $spl2 | $spl6 | $spl10)),
  inference(forward_demodulation,[],[f9043,f310])).
fof(f9043,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK9) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))))) | ($spl0 | $spl2 | $spl6 | $spl10)),
  inference(forward_demodulation,[],[f8979,f310])).
fof(f8979,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK7),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK9),s(fun(sK0,sK0),sK5))))) | ($spl0 | $spl2 | $spl6 | $spl10)),
  inference(unit_resulting_resolution,[],[f8698,f292])).
fof(f8698,plain,(
  s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))) = s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK7),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK9),s(fun(sK0,sK0),sK5))) | ($spl0 | $spl2 | $spl6 | $spl10)),
  inference(forward_demodulation,[],[f355,f338])).
fof(f9037,plain,(
  $false | ($spl0 | $spl2 | $spl6 | $spl9 | $spl10)),
  inference(subsumption_resolution,[],[f9036,f345])).
fof(f9036,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK4) = s(fun(sK0,fun(sK0,sK0)),sK9) | ($spl0 | $spl2 | $spl6 | $spl10)),
  inference(forward_demodulation,[],[f9035,f310])).
fof(f9035,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK9) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))))) | ($spl0 | $spl2 | $spl6 | $spl10)),
  inference(forward_demodulation,[],[f8984,f310])).
fof(f8984,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK7),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK9),s(fun(sK0,sK0),sK5))))) | ($spl0 | $spl2 | $spl6 | $spl10)),
  inference(unit_resulting_resolution,[],[f8698,f292])).
fof(f8665,plain,(
  $false | ($spl0 | $spl2 | $spl4 | $spl6 | $spl9 | $spl10)),
  inference(subsumption_resolution,[],[f8664,f345])).
fof(f8664,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK4) = s(fun(sK0,fun(sK0,sK0)),sK9) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10)),
  inference(forward_demodulation,[],[f8663,f310])).
fof(f8663,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK9) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10)),
  inference(forward_demodulation,[],[f8593,f310])).
fof(f8593,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK9),s(fun(sK0,sK0),sK5))))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10)),
  inference(unit_resulting_resolution,[],[f357,f292])).
fof(f357,plain,(
  s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))) = s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK9),s(fun(sK0,sK0),sK5))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10)),
  inference(forward_demodulation,[],[f356,f338])).
fof(f356,plain,(
  s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))) = s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK6),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK9),s(fun(sK0,sK0),sK5))) | ($spl0 | $spl4 | $spl6 | $spl10)),
  inference(forward_demodulation,[],[f355,f340])).
fof(f8662,plain,(
  $false | ($spl0 | $spl2 | $spl4 | $spl6 | $spl9 | $spl10)),
  inference(subsumption_resolution,[],[f8661,f345])).
fof(f8661,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK4) = s(fun(sK0,fun(sK0,sK0)),sK9) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10)),
  inference(forward_demodulation,[],[f8660,f310])).
fof(f8660,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK9) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10)),
  inference(forward_demodulation,[],[f8609,f310])).
fof(f8609,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK9),s(fun(sK0,sK0),sK5))))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10)),
  inference(unit_resulting_resolution,[],[f357,f292])).
fof(f8268,plain,(
  $false | ($spl22 | $spl24)),
  inference(global_subsumption,[],[f8267,f231,f230,f229,f228,f227,f226,f232,f235,f237,f236,f261,f262,f263,f264,f265,f266,f267,f268,f269,f270,f271,f272,f273,f274,f275,f276,f277,f278,f279,f280,f281,f282,f283,f284,f285,f286,f288,f294,f293,f292,f291,f290,f289,f301,f302,f303,f304,f305,f306,f307,f308,f309,f310,f311,f312,f313,f314,f315,f316,f322,f321,f320,f319,f318,f317,f326,f323,f327,f328,f329,f330,f331,f332,f333,f358,f389,f392,f426,f563,f765,f766,f767,f770,f786,f796,f797,f798,f801,f902,f1154,f1155,f1156,f1165,f1167,f1243,f1263,f1264,f1265,f1266,f1267,f1268,f1251,f1270,f1492,f1652,f1783,f1886,f1887,f1888,f1922,f1981,f2149,f2156,f2289,f2359,f2360,f2361,f2367,f2382,f3174,f2365,f2364,f1817,f1815,f800,f769,f2363,f2362,f1577,f1575,f799,f768,f1333,f1372,f2901,f631,f1174,f1853,f1854,f740,f1173,f1610,f1611,f3090,f3091,f2131,f1242,f1205,f1176,f1175,f1172,f1171,f1170,f561,f560,f559,f558,f287,f3086,f3042,f3040,f1269,f1277,f1314,f3168,f3116,f3085,f3079,f3078,f3077,f3074,f2976,f2619,f2618,f2613,f2612,f2611,f2515,f2488,f2341,f2190,f2188,f2187,f2186,f2185,f2155,f2154,f2153,f1940,f1848,f1847,f1846,f1843,f1819,f1785,f1784,f1782,f1781,f1740,f1739,f1684,f1683,f1682,f1681,f1680,f1679,f1678,f1677,f1606,f1605,f1602,f1579,f1554,f1553,f1479,f1478,f1477,f1476,f1475,f1474,f1473,f1459,f1458,f1411,f1410,f1409,f1408,f1407,f1406,f1405,f1244,f1200,f1169,f735,f718,f699,f686,f685,f649,f648,f647,f646,f645,f612,f534,f490,f450,f391,f241,f240,f234,f2318,f2319,f2566,f2567,f3173,f3172,f3166,f3165,f3164,f3163,f3162,f3161,f3160,f3159,f3158,f3157,f3084,f2965,f2964,f2963,f2962,f2961,f2960,f2954,f2953,f2952,f2951,f2950,f2949,f2948,f2947,f2946,f1964,f1907,f1765,f1203,f1202,f1201,f1195,f1193,f1192,f1191,f1190,f1189,f4379,f4314,f4316,f4345,f4350,f4351,f4354,f4355,f4561,f4691,f4855,f5042,f4962,f4964,f5007,f5008,f5567,f5481,f5483,f5530,f5531,f5620,f5623,f5626,f5716,f5719,f5722,f5770,f5771,f5772,f5773,f5774,f5775,f5759,f5919,f5826,f5828,f5861,f5864,f5865,f5866,f5871,f5872,f5874,f5875,f5878,f5879,f5978,f5999,f6033,f6083,f6120,f6300,f6202,f6204,f6238,f6241,f6242,f6243,f6248,f6249,f6251,f6253,f6257,f6258,f6363,f6364,f6417,f6490,f6932,f6936,f6937,f6938,f6939,f6959,f6987,f7029,f7030,f7031,f7069,f7072,f7073,f7074,f7079,f7080,f7082,f7084,f7091,f7092,f7335,f7345,f7347,f7329,f7352,f7353,f7197,f7198,f7199,f7200,f7201,f7202,f7203,f7204,f7205,f7206,f7207,f7208,f7209,f7210,f7212,f7214,f7217,f7356,f7223,f7224,f7225,f7226,f7227,f7328,f7231,f7234,f7235,f7236,f7237,f7238,f7239,f7240,f7242,f7243,f7244,f7246,f7247,f7249,f7312,f7569,f7832,f7833,f7834,f7835,f7836,f7837,f7838,f7839,f7840,f7841,f7842,f7843,f7844,f7845,f7846,f7847,f7886,f7911,f7924,f7869,f7870,f7871,f7872,f7873,f7874,f7875,f7876,f7877,f7878,f7879,f7880,f7881,f7882,f7884,f7343,f8260])).
fof(f8260,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl22),
  inference(subsumption_resolution,[],[f7926,f7343])).
fof(f7926,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl22),
  inference(backward_demodulation,[],[f7343,f240])).
fof(f8267,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,X1)) ) | ($spl22 | $spl24)),
  inference(subsumption_resolution,[],[f8107,f7343])).
fof(f8107,plain,(
  ( ! [X2,X3,X1] : (s(bool,'T_0') != s(bool,'const.combin.K_2'(s(bool,X1),s(X2,X3))) | s(bool,'F_0') = s(bool,X1)) ) | ($spl22 | $spl24)),
  inference(backward_demodulation,[],[f7343,f7344])).
fof(f8264,plain,(
  $false | $spl22),
  inference(subsumption_resolution,[],[f8001,f7343])).
fof(f8001,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.combin.K_2'(s(bool,'T_0'),s(X0,X1)))) ) | $spl22),
  inference(backward_demodulation,[],[f7343,f2341])).
fof(f8261,plain,(
  $false | $spl22),
  inference(global_subsumption,[],[f231,f230,f229,f228,f227,f226,f232,f235,f237,f236,f261,f262,f263,f264,f265,f266,f267,f268,f269,f270,f271,f272,f273,f274,f275,f276,f277,f278,f279,f280,f281,f282,f283,f284,f285,f286,f288,f294,f293,f292,f291,f290,f289,f301,f302,f303,f304,f305,f306,f307,f308,f309,f310,f311,f312,f313,f314,f315,f316,f322,f321,f320,f319,f318,f317,f326,f323,f327,f328,f329,f330,f331,f332,f333,f358,f389,f392,f426,f563,f765,f766,f767,f770,f786,f796,f797,f798,f801,f902,f1154,f1155,f1156,f1165,f1167,f1243,f1263,f1264,f1265,f1266,f1267,f1268,f1251,f1270,f1492,f1652,f1783,f1886,f1887,f1888,f1922,f1981,f2149,f2156,f2289,f2359,f2360,f2361,f2367,f2382,f3174,f2365,f2364,f1817,f1815,f800,f769,f2363,f2362,f1577,f1575,f799,f768,f1333,f1372,f2901,f631,f1174,f1853,f1854,f740,f1173,f1610,f1611,f3090,f3091,f2131,f1242,f1205,f1176,f1175,f1172,f1171,f1170,f561,f560,f559,f558,f287,f3086,f3042,f3040,f1269,f1277,f1314,f3168,f3116,f3085,f3079,f3078,f3077,f3074,f2976,f2619,f2618,f2613,f2612,f2611,f2515,f2488,f2341,f2190,f2188,f2187,f2186,f2185,f2155,f2154,f2153,f1940,f1848,f1847,f1846,f1843,f1819,f1785,f1784,f1782,f1781,f1740,f1739,f1684,f1683,f1682,f1681,f1680,f1679,f1678,f1677,f1606,f1605,f1602,f1579,f1554,f1553,f1479,f1478,f1477,f1476,f1475,f1474,f1473,f1459,f1458,f1411,f1410,f1409,f1408,f1407,f1406,f1405,f1244,f1200,f1169,f735,f718,f699,f686,f685,f649,f648,f647,f646,f645,f612,f534,f490,f450,f391,f241,f240,f234,f2318,f2319,f2566,f2567,f3173,f3172,f3166,f3165,f3164,f3163,f3162,f3161,f3160,f3159,f3158,f3157,f3084,f2965,f2964,f2963,f2962,f2961,f2960,f2954,f2953,f2952,f2951,f2950,f2949,f2948,f2947,f2946,f1964,f1907,f1765,f1203,f1202,f1201,f1195,f1193,f1192,f1191,f1190,f1189,f4379,f4314,f4316,f4345,f4350,f4351,f4354,f4355,f4561,f4691,f4855,f5042,f4962,f4964,f5007,f5008,f5567,f5481,f5483,f5530,f5531,f5620,f5623,f5626,f5716,f5719,f5722,f5770,f5771,f5772,f5773,f5774,f5775,f5759,f5919,f5826,f5828,f5861,f5864,f5865,f5866,f5871,f5872,f5874,f5875,f5878,f5879,f5978,f5999,f6033,f6083,f6120,f6300,f6202,f6204,f6238,f6241,f6242,f6243,f6248,f6249,f6251,f6253,f6257,f6258,f6363,f6364,f6417,f6490,f6932,f6936,f6937,f6938,f6939,f6959,f6987,f7029,f7030,f7031,f7069,f7072,f7073,f7074,f7079,f7080,f7082,f7084,f7091,f7092,f7335,f7345,f7347,f7329,f7352,f7353,f7197,f7198,f7199,f7200,f7201,f7202,f7203,f7204,f7205,f7206,f7207,f7208,f7209,f7210,f7212,f7214,f7217,f7356,f7223,f7224,f7225,f7226,f7227,f7328,f7231,f7234,f7235,f7236,f7237,f7238,f7239,f7240,f7242,f7243,f7244,f7246,f7247,f7249,f7312,f7569,f7832,f7833,f7834,f7835,f7836,f7837,f7838,f7839,f7840,f7841,f7842,f7843,f7844,f7845,f7846,f7847,f7886,f7911,f7924,f7869,f7870,f7871,f7872,f7873,f7874,f7875,f7876,f7877,f7878,f7879,f7880,f7881,f7882,f7884,f7343,f8260])).
fof(f8259,plain,(
  $false | $spl22),
  inference(subsumption_resolution,[],[f7925,f232])).
fof(f7925,plain,(
  ~p(s(bool,'T_0')) | $spl22),
  inference(backward_demodulation,[],[f7343,f234])).
fof(f7343,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl22),
  inference(cnf_transformation,[],[f7343_D])).
fof(f7343_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f7910,plain,(
  ( ! [X17,X15,X16] : (s(bool,X17) != s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X15,X16))) | s(bool,'F_0') = s(bool,X17)) ) | $spl46),
  inference(cnf_transformation,[],[f7910_D])).
fof(f7910_D,plain,(
  ( ! [X17,X15,X16] : (s(bool,X17) != s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X15,X16))) | s(bool,'F_0') = s(bool,X17)) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f7911,plain,(
  ( ! [X21,X19,X20,X18] : (s(bool,X21) != s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X19,X20))) | s(bool,'T_0') = s(bool,X18) | s(bool,'F_0') = s(bool,X21)) )),
  inference(subsumption_resolution,[],[f7854,f2289])).
fof(f7854,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,X21) != s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X19,X20))) | s(bool,'T_0') = s(bool,X18) | s(bool,'F_0') = s(bool,X21) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X18),s(X22,X23)))) )),
  inference(superposition,[],[f7312,f2618])).
fof(f7886,plain,(
  ( ! [X14,X17,X15,X16] : (s(bool,X17) != s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X15,X16))) | s(bool,'T_0') = s(bool,X14) | s(bool,'F_0') = s(bool,X17)) )),
  inference(duplicate_literal_removal,[],[f7853])).
fof(f7853,plain,(
  ( ! [X14,X17,X15,X16] : (s(bool,X17) != s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X15,X16))) | s(bool,'T_0') = s(bool,X14) | s(bool,'F_0') = s(bool,X17) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f7312,f2515])).
fof(f7365,plain,(
  ( ! [X107,X105,X106] : (s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,X105),s(X106,X107))) | s(bool,'F_0') = s(bool,X105)) ) | $spl44),
  inference(cnf_transformation,[],[f7365_D])).
fof(f7365_D,plain,(
  ( ! [X107,X105,X106] : (s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,X105),s(X106,X107))) | s(bool,'F_0') = s(bool,X105)) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f7355,plain,(
  ( ! [X6,X7,X5] : (s(bool,X5) != s(bool,'const.combin.K_2'(s(bool,'T_0'),s(X6,X7))) | s(bool,'T_0') = s(bool,X5)) ) | $spl30),
  inference(cnf_transformation,[],[f7355_D])).
fof(f7355_D,plain,(
  ( ! [X6,X7,X5] : (s(bool,X5) != s(bool,'const.combin.K_2'(s(bool,'T_0'),s(X6,X7))) | s(bool,'T_0') = s(bool,X5)) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f7231,plain,(
  ( ! [X109,X107,X105,X110,X108,X106] : (s(bool,X108) != s(bool,'const.combin.K_2'(s(bool,'T_0'),s(X109,X110))) | s(bool,'T_0') = s(bool,X108) | s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,X105),s(X106,X107))) | s(bool,'F_0') = s(bool,X105)) )),
  inference(superposition,[],[f686,f391])).
fof(f7364,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,X0),s(X1,X2)))) ) | $spl42),
  inference(cnf_transformation,[],[f7364_D])).
fof(f7364_D,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,X0),s(X1,X2)))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f7328,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,X0),s(X1,X2))) | s(bool,X3) != s(bool,'const.combin.K_2'(s(bool,'T_0'),s(X4,X5))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(condensation,[],[f7229])).
fof(f7229,plain,(
  ( ! [X94,X92,X97,X95,X93,X91,X98,X96] : (s(bool,X94) != s(bool,'const.combin.K_2'(s(bool,'T_0'),s(X95,X96))) | s(bool,'T_0') = s(bool,X94) | s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,X91),s(X92,X93))) | s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,X91),s(X97,X98)))) )),
  inference(superposition,[],[f686,f649])).
fof(f7361,plain,(
  ( ! [X76,X81,X77,X75,X82] : (s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,X75),s(X76,X77))) | s(bool,X75) = s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X81,X82)))) ) | $spl40),
  inference(cnf_transformation,[],[f7361_D])).
fof(f7361_D,plain,(
  ( ! [X76,X81,X77,X75,X82] : (s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,X75),s(X76,X77))) | s(bool,X75) = s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X81,X82)))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f7227,plain,(
  ( ! [X80,X78,X76,X81,X79,X77,X75,X82] : (s(bool,X78) != s(bool,'const.combin.K_2'(s(bool,'T_0'),s(X79,X80))) | s(bool,'T_0') = s(bool,X78) | s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,X75),s(X76,X77))) | s(bool,X75) = s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X81,X82)))) )),
  inference(superposition,[],[f686,f2618])).
fof(f7360,plain,(
  ( ! [X70,X71,X69] : (s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X69,bool),'const.pair.,_2'(s(X69,X70),s(bool,X71))))) | s(bool,'F_0') = s(bool,X71)) ) | $spl38),
  inference(cnf_transformation,[],[f7360_D])).
fof(f7360_D,plain,(
  ( ! [X70,X71,X69] : (s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X69,bool),'const.pair.,_2'(s(X69,X70),s(bool,X71))))) | s(bool,'F_0') = s(bool,X71)) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f7226,plain,(
  ( ! [X70,X74,X72,X71,X69,X73] : (s(bool,X72) != s(bool,'const.combin.K_2'(s(bool,'T_0'),s(X73,X74))) | s(bool,'T_0') = s(bool,X72) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X69,bool),'const.pair.,_2'(s(X69,X70),s(bool,X71))))) | s(bool,'F_0') = s(bool,X71)) )),
  inference(superposition,[],[f686,f534])).
fof(f7359,plain,(
  ( ! [X68,X64,X63] : (s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X63,bool),X64))) | s(X63,X68) = s(X63,sK11(X68,X64,X63))) ) | $spl36),
  inference(cnf_transformation,[],[f7359_D])).
fof(f7359_D,plain,(
  ( ! [X68,X64,X63] : (s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X63,bool),X64))) | s(X63,X68) = s(X63,sK11(X68,X64,X63))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f7225,plain,(
  ( ! [X68,X66,X64,X67,X65,X63] : (s(bool,X65) != s(bool,'const.combin.K_2'(s(bool,'T_0'),s(X66,X67))) | s(bool,'T_0') = s(bool,X65) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X63,bool),X64))) | s(X63,X68) = s(X63,sK11(X68,X64,X63))) )),
  inference(superposition,[],[f686,f563])).
fof(f7358,plain,(
  ( ! [X59,X58] : (s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X58),s(fun('type.num.num',bool),X59),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X58)) ) | $spl34),
  inference(cnf_transformation,[],[f7358_D])).
fof(f7358_D,plain,(
  ( ! [X59,X58] : (s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X58),s(fun('type.num.num',bool),X59),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X58)) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f7224,plain,(
  ( ! [X61,X59,X62,X60,X58] : (s(bool,X60) != s(bool,'const.combin.K_2'(s(bool,'T_0'),s(X61,X62))) | s(bool,'T_0') = s(bool,X60) | s(bool,'F_0') = s(bool,'const.ind_type.FCONS_3'(s(bool,X58),s(fun('type.num.num',bool),X59),s('type.num.num','const.num.|48|_0'))) | s(bool,'F_0') = s(bool,X58)) )),
  inference(superposition,[],[f686,f450])).
fof(f7357,plain,(
  ( ! [X54,X52,X53] : (s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X52),'const.pair.,_2'(s(bool,X53),s(X52,X54))))) | s(bool,'F_0') = s(bool,X53)) ) | $spl32),
  inference(cnf_transformation,[],[f7357_D])).
fof(f7357_D,plain,(
  ( ! [X54,X52,X53] : (s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X52),'const.pair.,_2'(s(bool,X53),s(X52,X54))))) | s(bool,'F_0') = s(bool,X53)) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f7223,plain,(
  ( ! [X57,X54,X52,X56,X55,X53] : (s(bool,X55) != s(bool,'const.combin.K_2'(s(bool,'T_0'),s(X56,X57))) | s(bool,'T_0') = s(bool,X55) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X52),'const.pair.,_2'(s(bool,X53),s(X52,X54))))) | s(bool,'F_0') = s(bool,X53)) )),
  inference(superposition,[],[f686,f490])).
fof(f7354,plain,(
  ( ! [X4,X8] : (s(bool,'F_0') = s(bool,X4) | s(bool,'T_0') = s(bool,X8) | s(bool,X4) = s(bool,X8)) ) | $spl28),
  inference(cnf_transformation,[],[f7354_D])).
fof(f7354_D,plain,(
  ( ! [X4,X8] : (s(bool,'F_0') = s(bool,X4) | s(bool,'T_0') = s(bool,X8) | s(bool,X4) = s(bool,X8)) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f7217,plain,(
  ( ! [X23,X21,X19,X22,X20] : (s(bool,X20) != s(bool,'const.combin.K_2'(s(bool,'T_0'),s(X21,X22))) | s(bool,'T_0') = s(bool,X20) | s(bool,'F_0') = s(bool,X19) | s(bool,X19) = s(bool,X23) | s(bool,'T_0') = s(bool,X23)) )),
  inference(superposition,[],[f686,f389])).
fof(f7214,plain,(
  ( ! [X6,X4,X8,X7,X5] : (s(bool,X5) != s(bool,'const.combin.K_2'(s(bool,'T_0'),s(X6,X7))) | s(bool,'T_0') = s(bool,X5) | s(bool,'F_0') = s(bool,X4) | s(bool,X4) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f686,f389])).
fof(f7344,plain,(
  ( ! [X2,X3,X1] : (s(bool,'F_0') != s(bool,'const.combin.K_2'(s(bool,X1),s(X2,X3))) | s(bool,'F_0') = s(bool,X1)) ) | $spl24),
  inference(cnf_transformation,[],[f7344_D])).
fof(f7344_D,plain,(
  ( ! [X2,X3,X1] : (s(bool,'F_0') != s(bool,'const.combin.K_2'(s(bool,X1),s(X2,X3))) | s(bool,'F_0') = s(bool,X1)) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f7352,plain,(
  ( ! [X101,X97,X102,X100] : (s(bool,'F_0') != s(bool,'const.combin.K_2'(s(bool,X100),s(X101,X102))) | s(bool,'F_0') = s(bool,X100) | s(bool,'T_0') = s(bool,X97)) )),
  inference(subsumption_resolution,[],[f7193,f2289])).
fof(f7193,plain,(
  ( ! [X101,X99,X97,X102,X100,X98] : (s(bool,'F_0') != s(bool,'const.combin.K_2'(s(bool,X100),s(X101,X102))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X97),s(X98,X99))) | s(bool,'F_0') = s(bool,X100) | s(bool,'T_0') = s(bool,X97)) )),
  inference(superposition,[],[f686,f1940])).
fof(f7349,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X0),s(X1,X2)))) ) | $spl26),
  inference(cnf_transformation,[],[f7349_D])).
fof(f7349_D,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X0),s(X1,X2)))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f7329,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X0),s(X1,X2))) | s(bool,'F_0') != s(bool,'const.combin.K_2'(s(bool,X3),s(X4,X5))) | s(bool,'F_0') = s(bool,X3)) )),
  inference(condensation,[],[f7191])).
fof(f7191,plain,(
  ( ! [X88,X87,X85,X83,X81,X86,X84,X82] : (s(bool,'F_0') != s(bool,'const.combin.K_2'(s(bool,X84),s(X85,X86))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X81),s(X82,X83))) | s(bool,'F_0') = s(bool,X84) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X81),s(X87,X88)))) )),
  inference(superposition,[],[f686,f649])).
fof(f7347,plain,(
  ( ! [X30,X35,X33,X34] : (s(bool,'F_0') != s(bool,'const.combin.K_2'(s(bool,X33),s(X34,X35))) | s(bool,'T_0') = s(bool,X30) | s(bool,'F_0') = s(bool,X33)) )),
  inference(forward_demodulation,[],[f7346,f328])).
fof(f7346,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X31,X32))) != s(bool,'const.combin.K_2'(s(bool,X33),s(X34,X35))) | s(bool,'T_0') = s(bool,X30) | s(bool,'F_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f7185,f2289])).
fof(f7185,plain,(
  ( ! [X30,X37,X35,X33,X31,X36,X34,X32] : (s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X31,X32))) != s(bool,'const.combin.K_2'(s(bool,X33),s(X34,X35))) | s(bool,'T_0') = s(bool,X30) | s(bool,'F_0') = s(bool,X33) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X30),s(X36,X37)))) )),
  inference(superposition,[],[f686,f2618])).
fof(f7345,plain,(
  ( ! [X28,X24,X29,X27] : (s(bool,'F_0') != s(bool,'const.combin.K_2'(s(bool,X27),s(X28,X29))) | s(bool,'T_0') = s(bool,X24) | s(bool,'F_0') = s(bool,X27)) )),
  inference(forward_demodulation,[],[f7330,f328])).
fof(f7330,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X25,X26))) != s(bool,'const.combin.K_2'(s(bool,X27),s(X28,X29))) | s(bool,'T_0') = s(bool,X24) | s(bool,'F_0') = s(bool,X27)) )),
  inference(duplicate_literal_removal,[],[f7184])).
fof(f7184,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X25,X26))) != s(bool,'const.combin.K_2'(s(bool,X27),s(X28,X29))) | s(bool,'T_0') = s(bool,X24) | s(bool,'F_0') = s(bool,X27) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f686,f2515])).
fof(f7335,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') != s(bool,'const.combin.K_2'(s(bool,X1),s(X2,X3))) | s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f7179])).
fof(f7179,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') != s(bool,'const.combin.K_2'(s(bool,X1),s(X2,X3))) | s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f686,f235])).
fof(f3447,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f3184,f232])).
fof(f3184,plain,(
  ~p(s(bool,'T_0')) | $spl20),
  inference(unit_resulting_resolution,[],[f358,f3176,f326])).
fof(f3446,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f3186,f232])).
fof(f3186,plain,(
  ~p(s(bool,'T_0')) | $spl20),
  inference(unit_resulting_resolution,[],[f2341,f3176,f326])).
fof(f3445,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f3188,f3176])).
fof(f3188,plain,(
  ~p(s(bool,'F_0')) | $spl20),
  inference(unit_resulting_resolution,[],[f358,f3176,f326])).
fof(f3444,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f3190,f3176])).
fof(f3190,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X0,X1))))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f2341,f3176,f326])).
fof(f3443,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f3194,f232])).
fof(f3194,plain,(
  ~p(s(bool,'T_0')) | $spl20),
  inference(unit_resulting_resolution,[],[f358,f3176,f326])).
fof(f3442,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f3196,f232])).
fof(f3196,plain,(
  ~p(s(bool,'T_0')) | $spl20),
  inference(unit_resulting_resolution,[],[f2341,f3176,f326])).
fof(f3441,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f3198,f3176])).
fof(f3198,plain,(
  ~p(s(bool,'F_0')) | $spl20),
  inference(unit_resulting_resolution,[],[f358,f3176,f326])).
fof(f3440,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f3200,f3176])).
fof(f3200,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X0,X1))))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f2341,f3176,f326])).
fof(f3437,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f3362,f3245])).
fof(f3245,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl20),
  inference(unit_resulting_resolution,[],[f2341,f2341,f3176,f1154])).
fof(f3362,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.combin.K_2'(s(bool,'T_0'),s(X0,X1)))) ) | $spl20),
  inference(backward_demodulation,[],[f3245,f2341])).
fof(f3434,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f3433,f3176])).
fof(f3433,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl20),
  inference(subsumption_resolution,[],[f3288,f3245])).
fof(f3288,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl20),
  inference(backward_demodulation,[],[f3245,f240])).
fof(f3432,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f3287,f232])).
fof(f3287,plain,(
  ~p(s(bool,'T_0')) | $spl20),
  inference(backward_demodulation,[],[f3245,f234])).
fof(f3246,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f358,f2341,f2341,f3176,f1154])).
fof(f3240,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f358,f358,f358,f3176,f1154])).
fof(f3229,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f358,f358,f3176,f685])).
fof(f3223,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f328,f2341,f3176,f392])).
fof(f3221,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f329,f358,f3176,f392])).
fof(f3201,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f3176,f2341,f3176,f326])).
fof(f3199,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f3176,f358,f3176,f326])).
fof(f3197,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f232,f2341,f3176,f326])).
fof(f3195,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f232,f358,f3176,f326])).
fof(f3191,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f3176,f2341,f3176,f326])).
fof(f3189,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f3176,f358,f3176,f326])).
fof(f3187,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f232,f2341,f3176,f326])).
fof(f3185,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f232,f358,f3176,f326])).
fof(f3181,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f329,f3176,f240])).
fof(f3179,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f358,f3176,f237])).
fof(f3177,plain,(
  $false | $spl20),
  inference(unit_resulting_resolution,[],[f234,f3176])).
fof(f3175,plain,(
  ( ! [X61,X64,X62] : (s(bool,'F_0') = s(bool,sK11('const.combin.K_2'(s(bool,'F_0'),s(X61,X62)),X64,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X64)))) ) | $spl18),
  inference(cnf_transformation,[],[f3175_D])).
fof(f3175_D,plain,(
  ( ! [X61,X64,X62] : (s(bool,'F_0') = s(bool,sK11('const.combin.K_2'(s(bool,'F_0'),s(X61,X62)),X64,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X64)))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f3176,plain,(
  ( ! [X63] : (p(s(bool,X63))) ) | $spl20),
  inference(cnf_transformation,[],[f3176_D])).
fof(f3176_D,plain,(
  ( ! [X63] : (p(s(bool,X63))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f3174,plain,(
  ( ! [X61,X64,X62,X63] : (p(s(bool,X63)) | s(bool,'F_0') = s(bool,sK11('const.combin.K_2'(s(bool,'F_0'),s(X61,X62)),X64,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X64)))) )),
  inference(subsumption_resolution,[],[f3156,f236])).
fof(f3156,plain,(
  ( ! [X61,X64,X62,X63] : (p(s(bool,X63)) | s(bool,'F_0') = s(bool,sK11('const.combin.K_2'(s(bool,'F_0'),s(X61,X62)),X64,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X64))) | s(bool,'T_0') = s(bool,X63)) )),
  inference(superposition,[],[f1192,f2515])).
fof(f2112,plain,(
  ( ! [X12,X17,X16] : (s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,X12),s(X16,X17))) | s(bool,'T_0') = s(bool,X12)) ) | $spl12),
  inference(cnf_transformation,[],[f2112_D])).
fof(f2112_D,plain,(
  ( ! [X12,X17,X16] : (s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,X12),s(X16,X17))) | s(bool,'T_0') = s(bool,X12)) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f2114,plain,(
  ( ! [X26,X25] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X25,X26)))) ) | $spl16),
  inference(cnf_transformation,[],[f2114_D])).
fof(f2114_D,plain,(
  ( ! [X26,X25] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X25,X26)))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f1888,plain,(
  ( ! [X28,X26,X24,X27,X25] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X25,X26))) | s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,X24),s(X27,X28))) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f649,f235])).
fof(f2113,plain,(
  ( ! [X14,X15,X13] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X13),s(X14,X15))) | s(bool,'T_0') = s(bool,X13)) ) | $spl14),
  inference(cnf_transformation,[],[f2113_D])).
fof(f2113_D,plain,(
  ( ! [X14,X15,X13] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X13),s(X14,X15))) | s(bool,'T_0') = s(bool,X13)) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f1887,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X19),s(X20,X21))) | s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,X18),s(X22,X23))) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f649,f389])).
fof(f1886,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X13),s(X14,X15))) | s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,X12),s(X16,X17))) | s(bool,'T_0') = s(bool,X12) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f649,f389])).
fof(f353,plain,(
  $false | ($spl1 | $spl2 | $spl4 | $spl6 | $spl8 | $spl10)),
  inference(trivial_inequality_removal,[],[f352])).
fof(f352,plain,(
  s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))) != s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))) | ($spl1 | $spl2 | $spl4 | $spl6 | $spl8 | $spl10)),
  inference(forward_demodulation,[],[f351,f338])).
fof(f351,plain,(
  s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))) != s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK6),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))) | ($spl1 | $spl4 | $spl6 | $spl8 | $spl10)),
  inference(forward_demodulation,[],[f350,f340])).
fof(f350,plain,(
  s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))) != s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK6),s(sK0,sK7),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))) | ($spl1 | $spl6 | $spl8 | $spl10)),
  inference(forward_demodulation,[],[f349,f342])).
fof(f349,plain,(
  s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))) != s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK6),s(sK0,sK7),s(fun(sK0,fun(sK0,sK0)),sK8),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))) | ($spl1 | $spl8 | $spl10)),
  inference(forward_demodulation,[],[f348,f344])).
fof(f348,plain,(
  s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))) != s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK6),s(sK0,sK7),s(fun(sK0,fun(sK0,sK0)),sK8),s(fun(sK0,fun(sK0,sK0)),sK9),s(fun(sK0,sK0),sK5))) | ($spl1 | $spl10)),
  inference(forward_demodulation,[],[f337,f346])).
fof(f336,plain,(
  s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))) = s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK6),s(sK0,sK7),s(fun(sK0,fun(sK0,sK0)),sK8),s(fun(sK0,fun(sK0,sK0)),sK9),s(fun(sK0,sK0),sK10))) | $spl0),
  inference(cnf_transformation,[],[f336_D])).
fof(f336_D,plain,(
  s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))) = s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK6),s(sK0,sK7),s(fun(sK0,fun(sK0,sK0)),sK8),s(fun(sK0,fun(sK0,sK0)),sK9),s(fun(sK0,sK0),sK10))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f338,plain,(
  s(sK0,sK1) = s(sK0,sK6) | $spl2),
  inference(cnf_transformation,[],[f338_D])).
fof(f338_D,plain,(
  s(sK0,sK1) = s(sK0,sK6) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f226,plain,(
  s(sK0,sK1) = s(sK0,sK6) | s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))) = s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK6),s(sK0,sK7),s(fun(sK0,fun(sK0,sK0)),sK8),s(fun(sK0,fun(sK0,sK0)),sK9),s(fun(sK0,sK0),sK10)))),
  inference(cnf_transformation,[],[f201])).
fof(f201,plain,(
  (s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))) = s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK6),s(sK0,sK7),s(fun(sK0,fun(sK0,sK0)),sK8),s(fun(sK0,fun(sK0,sK0)),sK9),s(fun(sK0,sK0),sK10))) | (s(sK0,sK1) = s(sK0,sK6) & s(sK0,sK2) = s(sK0,sK7) & s(fun(sK0,fun(sK0,sK0)),sK3) = s(fun(sK0,fun(sK0,sK0)),sK8) & s(fun(sK0,fun(sK0,sK0)),sK4) = s(fun(sK0,fun(sK0,sK0)),sK9) & s(fun(sK0,sK0),sK5) = s(fun(sK0,sK0),sK10))) & (s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))) != s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK6),s(sK0,sK7),s(fun(sK0,fun(sK0,sK0)),sK8),s(fun(sK0,fun(sK0,sK0)),sK9),s(fun(sK0,sK0),sK10))) | s(sK0,sK1) != s(sK0,sK6) | s(sK0,sK2) != s(sK0,sK7) | s(fun(sK0,fun(sK0,sK0)),sK3) != s(fun(sK0,fun(sK0,sK0)),sK8) | s(fun(sK0,fun(sK0,sK0)),sK4) != s(fun(sK0,fun(sK0,sK0)),sK9) | s(fun(sK0,sK0),sK5) != s(fun(sK0,sK0),sK10))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK9,sK10])],[f200])).
fof(f200,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : ((s('type.ring.ring'(X0),'const.ring.recordtype|46|ring_5'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4),s(fun(X0,X0),X5))) = s('type.ring.ring'(X0),'const.ring.recordtype|46|ring_5'(s(X0,X6),s(X0,X7),s(fun(X0,fun(X0,X0)),X8),s(fun(X0,fun(X0,X0)),X9),s(fun(X0,X0),X10))) | (s(X0,X1) = s(X0,X6) & s(X0,X2) = s(X0,X7) & s(fun(X0,fun(X0,X0)),X3) = s(fun(X0,fun(X0,X0)),X8) & s(fun(X0,fun(X0,X0)),X4) = s(fun(X0,fun(X0,X0)),X9) & s(fun(X0,X0),X5) = s(fun(X0,X0),X10))) & (s('type.ring.ring'(X0),'const.ring.recordtype|46|ring_5'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4),s(fun(X0,X0),X5))) != s('type.ring.ring'(X0),'const.ring.recordtype|46|ring_5'(s(X0,X6),s(X0,X7),s(fun(X0,fun(X0,X0)),X8),s(fun(X0,fun(X0,X0)),X9),s(fun(X0,X0),X10))) | s(X0,X1) != s(X0,X6) | s(X0,X2) != s(X0,X7) | s(fun(X0,fun(X0,X0)),X3) != s(fun(X0,fun(X0,X0)),X8) | s(fun(X0,fun(X0,X0)),X4) != s(fun(X0,fun(X0,X0)),X9) | s(fun(X0,X0),X5) != s(fun(X0,X0),X10)))),
  inference(flattening,[],[f199])).
fof(f199,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : ((s('type.ring.ring'(X0),'const.ring.recordtype|46|ring_5'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4),s(fun(X0,X0),X5))) = s('type.ring.ring'(X0),'const.ring.recordtype|46|ring_5'(s(X0,X6),s(X0,X7),s(fun(X0,fun(X0,X0)),X8),s(fun(X0,fun(X0,X0)),X9),s(fun(X0,X0),X10))) | (s(X0,X1) = s(X0,X6) & s(X0,X2) = s(X0,X7) & s(fun(X0,fun(X0,X0)),X3) = s(fun(X0,fun(X0,X0)),X8) & s(fun(X0,fun(X0,X0)),X4) = s(fun(X0,fun(X0,X0)),X9) & s(fun(X0,X0),X5) = s(fun(X0,X0),X10))) & (s('type.ring.ring'(X0),'const.ring.recordtype|46|ring_5'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4),s(fun(X0,X0),X5))) != s('type.ring.ring'(X0),'const.ring.recordtype|46|ring_5'(s(X0,X6),s(X0,X7),s(fun(X0,fun(X0,X0)),X8),s(fun(X0,fun(X0,X0)),X9),s(fun(X0,X0),X10))) | (s(X0,X1) != s(X0,X6) | s(X0,X2) != s(X0,X7) | s(fun(X0,fun(X0,X0)),X3) != s(fun(X0,fun(X0,X0)),X8) | s(fun(X0,fun(X0,X0)),X4) != s(fun(X0,fun(X0,X0)),X9) | s(fun(X0,X0),X5) != s(fun(X0,X0),X10))))),
  inference(nnf_transformation,[],[f194])).
fof(f194,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : (s('type.ring.ring'(X0),'const.ring.recordtype|46|ring_5'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4),s(fun(X0,X0),X5))) = s('type.ring.ring'(X0),'const.ring.recordtype|46|ring_5'(s(X0,X6),s(X0,X7),s(fun(X0,fun(X0,X0)),X8),s(fun(X0,fun(X0,X0)),X9),s(fun(X0,X0),X10))) <~> (s(X0,X1) = s(X0,X6) & s(X0,X2) = s(X0,X7) & s(fun(X0,fun(X0,X0)),X3) = s(fun(X0,fun(X0,X0)),X8) & s(fun(X0,fun(X0,X0)),X4) = s(fun(X0,fun(X0,X0)),X9) & s(fun(X0,X0),X5) = s(fun(X0,X0),X10)))),
  inference(ennf_transformation,[],[f72])).
fof(f72,plain,(
  ~! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : (s('type.ring.ring'(X0),'const.ring.recordtype|46|ring_5'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4),s(fun(X0,X0),X5))) = s('type.ring.ring'(X0),'const.ring.recordtype|46|ring_5'(s(X0,X6),s(X0,X7),s(fun(X0,fun(X0,X0)),X8),s(fun(X0,fun(X0,X0)),X9),s(fun(X0,X0),X10))) <=> (s(X0,X1) = s(X0,X6) & s(X0,X2) = s(X0,X7) & s(fun(X0,fun(X0,X0)),X3) = s(fun(X0,fun(X0,X0)),X8) & s(fun(X0,fun(X0,X0)),X4) = s(fun(X0,fun(X0,X0)),X9) & s(fun(X0,X0),X5) = s(fun(X0,X0),X10)))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ~! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : (s('type.ring.ring'(X0),'const.ring.recordtype|46|ring_5'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4),s(fun(X0,X0),X5))) = s('type.ring.ring'(X0),'const.ring.recordtype|46|ring_5'(s(X0,X6),s(X0,X7),s(fun(X0,fun(X0,X0)),X8),s(fun(X0,fun(X0,X0)),X9),s(fun(X0,X0),X10))) <=> (s(X0,X1) = s(X0,X6) & s(X0,X2) = s(X0,X7) & s(fun(X0,fun(X0,X0)),X3) = s(fun(X0,fun(X0,X0)),X8) & s(fun(X0,fun(X0,X0)),X4) = s(fun(X0,fun(X0,X0)),X9) & s(fun(X0,X0),X5) = s(fun(X0,X0),X10)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9,X10,X11,X12,X13,X14,X15,X16] : (s('type.ring.ring'(X6),'const.ring.recordtype|46|ring_5'(s(X6,X7),s(X6,X8),s(fun(X6,fun(X6,X6)),X9),s(fun(X6,fun(X6,X6)),X10),s(fun(X6,X6),X11))) = s('type.ring.ring'(X6),'const.ring.recordtype|46|ring_5'(s(X6,X12),s(X6,X13),s(fun(X6,fun(X6,X6)),X14),s(fun(X6,fun(X6,X6)),X15),s(fun(X6,X6),X16))) <=> (s(X6,X7) = s(X6,X12) & s(X6,X8) = s(X6,X13) & s(fun(X6,fun(X6,X6)),X9) = s(fun(X6,fun(X6,X6)),X14) & s(fun(X6,fun(X6,X6)),X10) = s(fun(X6,fun(X6,X6)),X15) & s(fun(X6,X6),X11) = s(fun(X6,X6),X16)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9,X10,X11,X12,X13,X14,X15,X16] : (s('type.ring.ring'(X6),'const.ring.recordtype|46|ring_5'(s(X6,X7),s(X6,X8),s(fun(X6,fun(X6,X6)),X9),s(fun(X6,fun(X6,X6)),X10),s(fun(X6,X6),X11))) = s('type.ring.ring'(X6),'const.ring.recordtype|46|ring_5'(s(X6,X12),s(X6,X13),s(fun(X6,fun(X6,X6)),X14),s(fun(X6,fun(X6,X6)),X15),s(fun(X6,X6),X16))) <=> (s(X6,X7) = s(X6,X12) & s(X6,X8) = s(X6,X13) & s(fun(X6,fun(X6,X6)),X9) = s(fun(X6,fun(X6,X6)),X14) & s(fun(X6,fun(X6,X6)),X10) = s(fun(X6,fun(X6,X6)),X15) & s(fun(X6,X6),X11) = s(fun(X6,X6),X16)))),
  file('thm.ring.ring_|49||49|','thm.ring.ring_|49||49|')).
fof(f340,plain,(
  s(sK0,sK2) = s(sK0,sK7) | $spl4),
  inference(cnf_transformation,[],[f340_D])).
fof(f340_D,plain,(
  s(sK0,sK2) = s(sK0,sK7) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f227,plain,(
  s(sK0,sK2) = s(sK0,sK7) | s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))) = s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK6),s(sK0,sK7),s(fun(sK0,fun(sK0,sK0)),sK8),s(fun(sK0,fun(sK0,sK0)),sK9),s(fun(sK0,sK0),sK10)))),
  inference(cnf_transformation,[],[f201])).
fof(f342,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK3) = s(fun(sK0,fun(sK0,sK0)),sK8) | $spl6),
  inference(cnf_transformation,[],[f342_D])).
fof(f342_D,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK3) = s(fun(sK0,fun(sK0,sK0)),sK8) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f228,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK3) = s(fun(sK0,fun(sK0,sK0)),sK8) | s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))) = s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK6),s(sK0,sK7),s(fun(sK0,fun(sK0,sK0)),sK8),s(fun(sK0,fun(sK0,sK0)),sK9),s(fun(sK0,sK0),sK10)))),
  inference(cnf_transformation,[],[f201])).
fof(f344,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK4) = s(fun(sK0,fun(sK0,sK0)),sK9) | $spl8),
  inference(cnf_transformation,[],[f344_D])).
fof(f344_D,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK4) = s(fun(sK0,fun(sK0,sK0)),sK9) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f229,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK4) = s(fun(sK0,fun(sK0,sK0)),sK9) | s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))) = s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK6),s(sK0,sK7),s(fun(sK0,fun(sK0,sK0)),sK8),s(fun(sK0,fun(sK0,sK0)),sK9),s(fun(sK0,sK0),sK10)))),
  inference(cnf_transformation,[],[f201])).
fof(f346,plain,(
  s(fun(sK0,sK0),sK5) = s(fun(sK0,sK0),sK10) | $spl10),
  inference(cnf_transformation,[],[f346_D])).
fof(f346_D,plain,(
  s(fun(sK0,sK0),sK5) = s(fun(sK0,sK0),sK10) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f230,plain,(
  s(fun(sK0,sK0),sK5) = s(fun(sK0,sK0),sK10) | s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))) = s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK6),s(sK0,sK7),s(fun(sK0,fun(sK0,sK0)),sK8),s(fun(sK0,fun(sK0,sK0)),sK9),s(fun(sK0,sK0),sK10)))),
  inference(cnf_transformation,[],[f201])).
fof(f337,plain,(
  s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))) != s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK6),s(sK0,sK7),s(fun(sK0,fun(sK0,sK0)),sK8),s(fun(sK0,fun(sK0,sK0)),sK9),s(fun(sK0,sK0),sK10))) | $spl1),
  inference(cnf_transformation,[],[f337_D])).
fof(f337_D,plain,(
  s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))) != s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK6),s(sK0,sK7),s(fun(sK0,fun(sK0,sK0)),sK8),s(fun(sK0,fun(sK0,sK0)),sK9),s(fun(sK0,sK0),sK10))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f339,plain,(
  s(sK0,sK1) != s(sK0,sK6) | $spl3),
  inference(cnf_transformation,[],[f339_D])).
fof(f339_D,plain,(
  s(sK0,sK1) != s(sK0,sK6) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f341,plain,(
  s(sK0,sK2) != s(sK0,sK7) | $spl5),
  inference(cnf_transformation,[],[f341_D])).
fof(f341_D,plain,(
  s(sK0,sK2) != s(sK0,sK7) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f343,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK3) != s(fun(sK0,fun(sK0,sK0)),sK8) | $spl7),
  inference(cnf_transformation,[],[f343_D])).
fof(f343_D,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK3) != s(fun(sK0,fun(sK0,sK0)),sK8) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f345,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK4) != s(fun(sK0,fun(sK0,sK0)),sK9) | $spl9),
  inference(cnf_transformation,[],[f345_D])).
fof(f345_D,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK4) != s(fun(sK0,fun(sK0,sK0)),sK9) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f347,plain,(
  s(fun(sK0,sK0),sK5) != s(fun(sK0,sK0),sK10) | $spl11),
  inference(cnf_transformation,[],[f347_D])).
fof(f347_D,plain,(
  s(fun(sK0,sK0),sK5) != s(fun(sK0,sK0),sK10) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f231,plain,(
  s(fun(sK0,sK0),sK5) != s(fun(sK0,sK0),sK10) | s(fun(sK0,fun(sK0,sK0)),sK4) != s(fun(sK0,fun(sK0,sK0)),sK9) | s(fun(sK0,fun(sK0,sK0)),sK3) != s(fun(sK0,fun(sK0,sK0)),sK8) | s(sK0,sK2) != s(sK0,sK7) | s(sK0,sK1) != s(sK0,sK6) | s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK1),s(sK0,sK2),s(fun(sK0,fun(sK0,sK0)),sK3),s(fun(sK0,fun(sK0,sK0)),sK4),s(fun(sK0,sK0),sK5))) != s('type.ring.ring'(sK0),'const.ring.recordtype|46|ring_5'(s(sK0,sK6),s(sK0,sK7),s(fun(sK0,fun(sK0,sK0)),sK8),s(fun(sK0,fun(sK0,sK0)),sK9),s(fun(sK0,sK0),sK10)))),
  inference(cnf_transformation,[],[f201])).
% SZS output end Proof for thm.ring
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 113
% Passive clauses: 1085
% Generated clauses: 10167
% Final active clauses: 55
% Final passive clauses: 431
% Input formulas: 69
% Initial clauses: 108
% 
% Duplicate literals: 267
% Trivial inequalities: 75
% Fw subsumption resolutions: 320
% Fw demodulations: 1276
% Bw demodulations: 432
% Condensations: 635
% Global subsumptions: 4
% 
% Simple tautologies: 99
% Equational tautologies: 96
% Forward subsumptions: 6087
% Fw demodulations to eq. taut.: 150
% Bw demodulations to eq. taut.: 119
% 
% Unit resulting resolution: 573
% Factoring: 1
% Forward superposition: 3440
% Backward superposition: 2285
% Self superposition: 133
% Equality factoring: 35
% Equality resolution: 16
% 
% Split clauses: 26
% Split components: 56
% SAT solver clauses: 1190
% SAT solver unit clauses: 88
% SAT solver binary clauses: 437
% 
% Sat splits: 26
% Sat splitting refutations: 47
% 
% Memory used [KB]: 7164
% Time elapsed: 0.419 s
% ------------------------------
% ------------------------------
% Success in time 1.72 s
