Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.lebesgue
% (89845)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1186
% Passive clauses: 9914
% Generated clauses: 29001
% Final active clauses: 1012
% Final passive clauses: 7999
% Input formulas: 61
% Initial clauses: 221
% 
% Duplicate literals: 518
% Trivial inequalities: 299
% Fw subsumption resolutions: 2637
% Fw demodulations: 1628
% Bw demodulations: 160
% 
% Simple tautologies: 264
% Equational tautologies: 64
% Forward subsumptions: 12578
% 
% Binary resolution: 9843
% Factoring: 923
% Forward superposition: 5093
% Backward superposition: 7101
% Self superposition: 63
% Equality factoring: 26
% Equality resolution: 36
% 
% Split clauses: 1033
% Split components: 2067
% SAT solver clauses: 1653
% SAT solver unit clauses: 132
% SAT solver binary clauses: 366
% 
% Sat splits: 1033
% Sat splitting refutations: 16
% 
% TWLsolver clauses: 1650
% TWLsolver calls for satisfiability: 678
% 
% Memory used [KB]: 14583
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.lebesgue
% (89868)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 569
% Passive clauses: 2560
% Generated clauses: 15423
% Final active clauses: 544
% Final passive clauses: 1946
% Input formulas: 61
% Initial clauses: 221
% 
% Duplicate literals: 314
% Trivial inequalities: 188
% Fw subsumption resolutions: 1346
% Fw demodulations: 1201
% Bw demodulations: 66
% Condensations: 352
% Global subsumptions: 14
% 
% Simple tautologies: 306
% Equational tautologies: 21
% Forward subsumptions: 5929
% 
% Binary resolution: 3036
% Factoring: 269
% Forward superposition: 2472
% Backward superposition: 3207
% Self superposition: 59
% Equality factoring: 30
% Equality resolution: 27
% 
% Split clauses: 143
% Split components: 286
% SAT solver clauses: 3353
% SAT solver unit clauses: 255
% SAT solver binary clauses: 103
% SAT solver learnt clauses: 28
% SAT solver learnt literals: 18
% 
% Sat splits: 143
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 3297
% TWLsolver calls for satisfiability: 2360
% 
% Memory used [KB]: 15095
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.lebesgue
% (89881)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 2463
% Passive clauses: 17798
% Generated clauses: 155061
% Final active clauses: 2250
% Final passive clauses: 5086
% Input formulas: 61
% Initial clauses: 221
% Discarded non-redundant clauses: 12777
% 
% Duplicate literals: 4165
% Trivial inequalities: 87
% Fw subsumption resolutions: 3571
% Fw demodulations: 32318
% Bw demodulations: 1063
% Condensations: 8159
% Global subsumptions: 1883
% 
% Simple tautologies: 135
% Equational tautologies: 103
% Forward subsumptions: 82070
% Fw demodulations to eq. taut.: 36
% Bw demodulations to eq. taut.: 16
% 
% Unit resulting resolution: 31940
% Factoring: 2
% Forward superposition: 48829
% Backward superposition: 21713
% Self superposition: 152
% Equality factoring: 32
% Equality resolution: 21
% 
% Split clauses: 131
% Split components: 267
% SAT solver clauses: 24068
% SAT solver unit clauses: 9232
% SAT solver binary clauses: 5595
% 
% Sat splits: 131
% Sat splitting refutations: 4123
% 
% Memory used [KB]: 216201
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.lebesgue
% (90162)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 2821
% Passive clauses: 43568
% Generated clauses: 254925
% Final active clauses: 2524
% Final passive clauses: 38822
% Input formulas: 61
% Initial clauses: 192
% 
% Selected by SInE selection: 55
% SInE iterations: 1
% 
% Duplicate literals: 5336
% Trivial inequalities: 1146
% Fw subsumption resolutions: 20169
% Fw demodulations: 3072
% Bw demodulations: 260
% 
% Simple tautologies: 3502
% Equational tautologies: 554
% Forward subsumptions: 172101
% 
% Binary resolution: 206004
% Factoring: 140
% Forward superposition: 10836
% Backward superposition: 6656
% Self superposition: 51
% Equality factoring: 24
% Equality resolution: 101
% 
% Split clauses: 5215
% Split components: 10922
% SAT solver clauses: 5367
% SAT solver unit clauses: 43
% SAT solver binary clauses: 667
% 
% Sat splits: 5215
% Sat splitting refutations: 49
% 
% Memory used [KB]: 59103
% Time elapsed: 4.0000 s
% ------------------------------
% ------------------------------
% remaining time: 142 next slice time: 6
lrs+1002_2:3_br=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:sas=minisat:stl=30:sac=on:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none:sp=reverse_arity:urr=on_4 on thm.lebesgue
% (90304)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 497
% Passive clauses: 5523
% Generated clauses: 18768
% Final active clauses: 481
% Final passive clauses: 4784
% Input formulas: 61
% Initial clauses: 221
% 
% Duplicate literals: 263
% Trivial inequalities: 75
% Fw subsumption resolutions: 1542
% Fw demodulations: 1747
% Bw demodulations: 198
% Global subsumptions: 18
% 
% Simple tautologies: 280
% Equational tautologies: 165
% Forward subsumptions: 9102
% 
% Unit resulting resolution: 5246
% Factoring: 2
% Forward superposition: 4992
% Backward superposition: 4600
% Self superposition: 106
% Equality factoring: 2
% Equality resolution: 20
% 
% Split clauses: 5
% Split components: 10
% SAT solver clauses: 5601
% SAT solver unit clauses: 1521
% SAT solver binary clauses: 1153
% 
% Sat splits: 5
% Sat splitting refutations: 53
% 
% Memory used [KB]: 16502
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 135 next slice time: 15
lrs+1011_12_bs=on:bsr=unit_only:cond=on:gs=on:gsaa=from_current:gsssp=full:nwc=1.1:sas=minisat:stl=60:sos=all:sac=on:sdd=off:sser=off:sfr=on:ssfp=100000:ssfq=1.2:ssnc=none:sp=reverse_arity:updr=off_14 on thm.lebesgue
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.lebesgue
% SZS output start Proof for thm.lebesgue
fof(f971,plain,(
  $false),
  inference(subsumption_resolution,[],[f970,f536])).
fof(f536,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f360])).
fof(f360,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f230])).
fof(f230,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f95])).
fof(f95,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f94])).
fof(f94,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X6] : ! [X14] : (s(X6,X14) = s(X6,X14) <=> p(s(bool,'T_0')))),
  file('thm.lebesgue.integral_cmul_indicator','thm.bool.REFL_CLAUSE')).
fof(f970,plain,(
  ~p(s(bool,'T_0'))),
  inference(forward_demodulation,[],[f969,f544])).
fof(f544,plain,(
  s(bool,'T_0') = s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK17,bool),'type.pair.prod'(fun(fun(sK17,bool),bool),fun(fun(sK17,bool),'type.realax.real'))),sK18)))),
  inference(resolution,[],[f310,f332])).
fof(f332,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f214])).
fof(f214,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X11] : (s(bool,'T_0') = s(bool,X11) <=> p(s(bool,X11)))),
  file('thm.lebesgue.integral_cmul_indicator','thm.bool.EQ_CLAUSES|split|0')).
fof(f310,plain,(
  p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK17,bool),'type.pair.prod'(fun(fun(sK17,bool),bool),fun(fun(sK17,bool),'type.realax.real'))),sK18))))),
  inference(cnf_transformation,[],[f211])).
fof(f211,plain,(
  p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK17,bool),'type.pair.prod'(fun(fun(sK17,bool),bool),fun(fun(sK17,bool),'type.realax.real'))),sK18)))) & p(s(bool,'const.bool.IN_2'(s(fun(sK17,bool),sK19),s(fun(fun(sK17,bool),bool),'const.measure.measurable_sets_1'(s('type.pair.prod'(fun(sK17,bool),'type.pair.prod'(fun(fun(sK17,bool),bool),fun(fun(sK17,bool),'type.realax.real'))),sK18)))))) & s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK20),s('type.realax.real','const.measure.measure_2'(s('type.pair.prod'(fun(sK17,bool),'type.pair.prod'(fun(fun(sK17,bool),bool),fun(fun(sK17,bool),'type.realax.real'))),sK18),s(fun(sK17,bool),sK19))))))) != s('type.extreal.extreal','const.lebesgue.integral_2'(s('type.pair.prod'(fun(sK17,bool),'type.pair.prod'(fun(fun(sK17,bool),bool),fun(fun(sK17,bool),'type.realax.real'))),sK18),s(fun(sK17,'type.extreal.extreal'),'_dst_x0x1_2'(s(fun('type.extreal.extreal','type.extreal.extreal'),'const.extreal.extreal_mul_1'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK20))))),s(fun(sK17,'type.extreal.extreal'),'const.measure.indicator_fn_1'(s(fun(sK17,bool),sK19)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK17,sK18,sK19,sK20])],[f153])).
fof(f153,plain,(
  ? [X0,X1,X2,X3] : (p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.measure.measurable_sets_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))))) & s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X3),s('type.realax.real','const.measure.measure_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,bool),X2))))))) != s('type.extreal.extreal','const.lebesgue.integral_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,'type.extreal.extreal'),'_dst_x0x1_2'(s(fun('type.extreal.extreal','type.extreal.extreal'),'const.extreal.extreal_mul_1'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))))),s(fun(X0,'type.extreal.extreal'),'const.measure.indicator_fn_1'(s(fun(X0,bool),X2))))))))),
  inference(flattening,[],[f152])).
fof(f152,plain,(
  ? [X0,X1,X2,X3] : ((p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.measure.measurable_sets_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1))))))) & s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X3),s('type.realax.real','const.measure.measure_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,bool),X2))))))) != s('type.extreal.extreal','const.lebesgue.integral_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,'type.extreal.extreal'),'_dst_x0x1_2'(s(fun('type.extreal.extreal','type.extreal.extreal'),'const.extreal.extreal_mul_1'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))))),s(fun(X0,'type.extreal.extreal'),'const.measure.indicator_fn_1'(s(fun(X0,bool),X2))))))))),
  inference(ennf_transformation,[],[f64])).
fof(f64,plain,(
  ~! [X0,X1,X2,X3] : ((p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.measure.measurable_sets_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1))))))) => s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X3),s('type.realax.real','const.measure.measure_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,bool),X2))))))) = s('type.extreal.extreal','const.lebesgue.integral_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,'type.extreal.extreal'),'_dst_x0x1_2'(s(fun('type.extreal.extreal','type.extreal.extreal'),'const.extreal.extreal_mul_1'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))))),s(fun(X0,'type.extreal.extreal'),'const.measure.indicator_fn_1'(s(fun(X0,bool),X2))))))))),
  inference(flattening,[],[f63])).
fof(f63,plain,(
  ~! [X0] : ! [X1,X2,X3] : ((p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.measure.measurable_sets_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1))))))) => s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X3),s('type.realax.real','const.measure.measure_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,bool),X2))))))) = s('type.extreal.extreal','const.lebesgue.integral_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,'type.extreal.extreal'),'_dst_x0x1_2'(s(fun('type.extreal.extreal','type.extreal.extreal'),'const.extreal.extreal_mul_1'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))))),s(fun(X0,'type.extreal.extreal'),'const.measure.indicator_fn_1'(s(fun(X0,bool),X2))))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : ((p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X6,bool),'type.pair.prod'(fun(fun(X6,bool),bool),fun(fun(X6,bool),'type.realax.real'))),X7)))) & p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),X8),s(fun(fun(X6,bool),bool),'const.measure.measurable_sets_1'(s('type.pair.prod'(fun(X6,bool),'type.pair.prod'(fun(fun(X6,bool),bool),fun(fun(X6,bool),'type.realax.real'))),X7))))))) => s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X9),s('type.realax.real','const.measure.measure_2'(s('type.pair.prod'(fun(X6,bool),'type.pair.prod'(fun(fun(X6,bool),bool),fun(fun(X6,bool),'type.realax.real'))),X7),s(fun(X6,bool),X8))))))) = s('type.extreal.extreal','const.lebesgue.integral_2'(s('type.pair.prod'(fun(X6,bool),'type.pair.prod'(fun(fun(X6,bool),bool),fun(fun(X6,bool),'type.realax.real'))),X7),s(fun(X6,'type.extreal.extreal'),'_dst_x0x1_2'(s(fun('type.extreal.extreal','type.extreal.extreal'),'const.extreal.extreal_mul_1'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X9))))),s(fun(X6,'type.extreal.extreal'),'const.measure.indicator_fn_1'(s(fun(X6,bool),X8))))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : ((p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X6,bool),'type.pair.prod'(fun(fun(X6,bool),bool),fun(fun(X6,bool),'type.realax.real'))),X7)))) & p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),X8),s(fun(fun(X6,bool),bool),'const.measure.measurable_sets_1'(s('type.pair.prod'(fun(X6,bool),'type.pair.prod'(fun(fun(X6,bool),bool),fun(fun(X6,bool),'type.realax.real'))),X7))))))) => s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X9),s('type.realax.real','const.measure.measure_2'(s('type.pair.prod'(fun(X6,bool),'type.pair.prod'(fun(fun(X6,bool),bool),fun(fun(X6,bool),'type.realax.real'))),X7),s(fun(X6,bool),X8))))))) = s('type.extreal.extreal','const.lebesgue.integral_2'(s('type.pair.prod'(fun(X6,bool),'type.pair.prod'(fun(fun(X6,bool),bool),fun(fun(X6,bool),'type.realax.real'))),X7),s(fun(X6,'type.extreal.extreal'),'_dst_x0x1_2'(s(fun('type.extreal.extreal','type.extreal.extreal'),'const.extreal.extreal_mul_1'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X9))))),s(fun(X6,'type.extreal.extreal'),'const.measure.indicator_fn_1'(s(fun(X6,bool),X8))))))))),
  file('thm.lebesgue.integral_cmul_indicator','thm.lebesgue.integral_cmul_indicator')).
fof(f969,plain,(
  ~p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK17,bool),'type.pair.prod'(fun(fun(sK17,bool),bool),fun(fun(sK17,bool),'type.realax.real'))),sK18))))),
  inference(subsumption_resolution,[],[f968,f381])).
fof(f381,plain,(
  ( ! [X0,X1] : (s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X0))))) = s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))))) )),
  inference(cnf_transformation,[],[f110])).
fof(f110,plain,(
  ! [X0,X1] : s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X0))))) = s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X18,X14] : s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X14),s('type.realax.real',X18))))) = s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X14))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X18)))))),
  file('thm.lebesgue.integral_cmul_indicator','thm.extreal.extreal_mul_def|split|8')).
fof(f968,plain,(
  s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK20),s('type.realax.real','const.measure.measure_2'(s('type.pair.prod'(fun(sK17,bool),'type.pair.prod'(fun(fun(sK17,bool),bool),fun(fun(sK17,bool),'type.realax.real'))),sK18),s(fun(sK17,bool),sK19))))))) != s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK20))),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.measure.measure_2'(s('type.pair.prod'(fun(sK17,bool),'type.pair.prod'(fun(fun(sK17,bool),bool),fun(fun(sK17,bool),'type.realax.real'))),sK18),s(fun(sK17,bool),sK19))))))) | ~p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK17,bool),'type.pair.prod'(fun(fun(sK17,bool),bool),fun(fun(sK17,bool),'type.realax.real'))),sK18))))),
  inference(forward_demodulation,[],[f967,f833])).
fof(f833,plain,(
  s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.measure.measure_2'(s('type.pair.prod'(fun(sK17,bool),'type.pair.prod'(fun(fun(sK17,bool),bool),fun(fun(sK17,bool),'type.realax.real'))),sK18),s(fun(sK17,bool),sK19))))) = s('type.extreal.extreal','const.lebesgue.integral_2'(s('type.pair.prod'(fun(sK17,bool),'type.pair.prod'(fun(fun(sK17,bool),bool),fun(fun(sK17,bool),'type.realax.real'))),sK18),s(fun(sK17,'type.extreal.extreal'),'const.measure.indicator_fn_1'(s(fun(sK17,bool),sK19)))))),
  inference(subsumption_resolution,[],[f832,f313])).
fof(f313,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f8])).
fof(f8,axiom,(
  p(s(bool,'T_0'))),
  file('thm.lebesgue.integral_cmul_indicator','thm.bool.TRUTH')).
fof(f832,plain,(
  ~p(s(bool,'T_0')) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.measure.measure_2'(s('type.pair.prod'(fun(sK17,bool),'type.pair.prod'(fun(fun(sK17,bool),bool),fun(fun(sK17,bool),'type.realax.real'))),sK18),s(fun(sK17,bool),sK19))))) = s('type.extreal.extreal','const.lebesgue.integral_2'(s('type.pair.prod'(fun(sK17,bool),'type.pair.prod'(fun(fun(sK17,bool),bool),fun(fun(sK17,bool),'type.realax.real'))),sK18),s(fun(sK17,'type.extreal.extreal'),'const.measure.indicator_fn_1'(s(fun(sK17,bool),sK19)))))),
  inference(forward_demodulation,[],[f699,f544])).
fof(f699,plain,(
  s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.measure.measure_2'(s('type.pair.prod'(fun(sK17,bool),'type.pair.prod'(fun(fun(sK17,bool),bool),fun(fun(sK17,bool),'type.realax.real'))),sK18),s(fun(sK17,bool),sK19))))) = s('type.extreal.extreal','const.lebesgue.integral_2'(s('type.pair.prod'(fun(sK17,bool),'type.pair.prod'(fun(fun(sK17,bool),bool),fun(fun(sK17,bool),'type.realax.real'))),sK18),s(fun(sK17,'type.extreal.extreal'),'const.measure.indicator_fn_1'(s(fun(sK17,bool),sK19))))) | ~p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK17,bool),'type.pair.prod'(fun(fun(sK17,bool),bool),fun(fun(sK17,bool),'type.realax.real'))),sK18))))),
  inference(resolution,[],[f311,f366])).
fof(f366,plain,(
  ( ! [X2,X0,X1] : (s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.measure.measure_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,bool),X2))))) = s('type.extreal.extreal','const.lebesgue.integral_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,'type.extreal.extreal'),'const.measure.indicator_fn_1'(s(fun(X0,bool),X2))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.measure.measurable_sets_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))))) | ~p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1))))) )),
  inference(cnf_transformation,[],[f167])).
fof(f167,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.measure.measurable_sets_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.measure.measure_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,bool),X2))))) = s('type.extreal.extreal','const.lebesgue.integral_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,'type.extreal.extreal'),'const.measure.indicator_fn_1'(s(fun(X0,bool),X2))))))),
  inference(flattening,[],[f166])).
fof(f166,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.measure.measurable_sets_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.measure.measure_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,bool),X2))))) = s('type.extreal.extreal','const.lebesgue.integral_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,'type.extreal.extreal'),'const.measure.indicator_fn_1'(s(fun(X0,bool),X2))))))),
  inference(ennf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0,X1,X2] : ((p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.measure.measurable_sets_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1))))))) => s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.measure.measure_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,bool),X2))))) = s('type.extreal.extreal','const.lebesgue.integral_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,'type.extreal.extreal'),'const.measure.indicator_fn_1'(s(fun(X0,bool),X2))))))),
  inference(flattening,[],[f101])).
fof(f101,plain,(
  ! [X0] : ! [X1,X2] : ((p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.measure.measurable_sets_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1))))))) => s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.measure.measure_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,bool),X2))))) = s('type.extreal.extreal','const.lebesgue.integral_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,'type.extreal.extreal'),'const.measure.indicator_fn_1'(s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X6] : ! [X7,X8] : ((p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X6,bool),'type.pair.prod'(fun(fun(X6,bool),bool),fun(fun(X6,bool),'type.realax.real'))),X7)))) & p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),X8),s(fun(fun(X6,bool),bool),'const.measure.measurable_sets_1'(s('type.pair.prod'(fun(X6,bool),'type.pair.prod'(fun(fun(X6,bool),bool),fun(fun(X6,bool),'type.realax.real'))),X7))))))) => s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.measure.measure_2'(s('type.pair.prod'(fun(X6,bool),'type.pair.prod'(fun(fun(X6,bool),bool),fun(fun(X6,bool),'type.realax.real'))),X7),s(fun(X6,bool),X8))))) = s('type.extreal.extreal','const.lebesgue.integral_2'(s('type.pair.prod'(fun(X6,bool),'type.pair.prod'(fun(fun(X6,bool),bool),fun(fun(X6,bool),'type.realax.real'))),X7),s(fun(X6,'type.extreal.extreal'),'const.measure.indicator_fn_1'(s(fun(X6,bool),X8))))))),
  file('thm.lebesgue.integral_cmul_indicator','thm.lebesgue.integral_indicator')).
fof(f311,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK17,bool),sK19),s(fun(fun(sK17,bool),bool),'const.measure.measurable_sets_1'(s('type.pair.prod'(fun(sK17,bool),'type.pair.prod'(fun(fun(sK17,bool),bool),fun(fun(sK17,bool),'type.realax.real'))),sK18))))))),
  inference(cnf_transformation,[],[f211])).
fof(f967,plain,(
  s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK20),s('type.realax.real','const.measure.measure_2'(s('type.pair.prod'(fun(sK17,bool),'type.pair.prod'(fun(fun(sK17,bool),bool),fun(fun(sK17,bool),'type.realax.real'))),sK18),s(fun(sK17,bool),sK19))))))) != s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK20))),s('type.extreal.extreal','const.lebesgue.integral_2'(s('type.pair.prod'(fun(sK17,bool),'type.pair.prod'(fun(fun(sK17,bool),bool),fun(fun(sK17,bool),'type.realax.real'))),sK18),s(fun(sK17,'type.extreal.extreal'),'const.measure.indicator_fn_1'(s(fun(sK17,bool),sK19))))))) | ~p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK17,bool),'type.pair.prod'(fun(fun(sK17,bool),bool),fun(fun(sK17,bool),'type.realax.real'))),sK18))))),
  inference(forward_demodulation,[],[f966,f380])).
fof(f380,plain,(
  ( ! [X0,X1] : (s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))) = s('type.extreal.extreal',happ(s(fun('type.extreal.extreal','type.extreal.extreal'),'const.extreal.extreal_mul_1'(s('type.extreal.extreal',X0))),s('type.extreal.extreal',X1)))) )),
  inference(cnf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0,X1] : s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))) = s('type.extreal.extreal',happ(s(fun('type.extreal.extreal','type.extreal.extreal'),'const.extreal.extreal_mul_1'(s('type.extreal.extreal',X0))),s('type.extreal.extreal',X1)))),
  inference(rectify,[],[f61])).
fof(f61,axiom,(
  ! [X24,X25] : s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X24),s('type.extreal.extreal',X25))) = s('type.extreal.extreal',happ(s(fun('type.extreal.extreal','type.extreal.extreal'),'const.extreal.extreal_mul_1'(s('type.extreal.extreal',X24))),s('type.extreal.extreal',X25)))),
  file('thm.lebesgue.integral_cmul_indicator','ARITY_1_2_const.extreal.extreal_mul')).
fof(f966,plain,(
  s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK20),s('type.realax.real','const.measure.measure_2'(s('type.pair.prod'(fun(sK17,bool),'type.pair.prod'(fun(fun(sK17,bool),bool),fun(fun(sK17,bool),'type.realax.real'))),sK18),s(fun(sK17,bool),sK19))))))) != s('type.extreal.extreal',happ(s(fun('type.extreal.extreal','type.extreal.extreal'),'const.extreal.extreal_mul_1'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK20))))),s('type.extreal.extreal','const.lebesgue.integral_2'(s('type.pair.prod'(fun(sK17,bool),'type.pair.prod'(fun(fun(sK17,bool),bool),fun(fun(sK17,bool),'type.realax.real'))),sK18),s(fun(sK17,'type.extreal.extreal'),'const.measure.indicator_fn_1'(s(fun(sK17,bool),sK19))))))) | ~p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK17,bool),'type.pair.prod'(fun(fun(sK17,bool),bool),fun(fun(sK17,bool),'type.realax.real'))),sK18))))),
  inference(subsumption_resolution,[],[f965,f835])).
fof(f835,plain,(
  p(s(bool,'const.lebesgue.integrable_2'(s('type.pair.prod'(fun(sK17,bool),'type.pair.prod'(fun(fun(sK17,bool),bool),fun(fun(sK17,bool),'type.realax.real'))),sK18),s(fun(sK17,'type.extreal.extreal'),'const.measure.indicator_fn_1'(s(fun(sK17,bool),sK19))))))),
  inference(subsumption_resolution,[],[f834,f313])).
fof(f834,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.lebesgue.integrable_2'(s('type.pair.prod'(fun(sK17,bool),'type.pair.prod'(fun(fun(sK17,bool),bool),fun(fun(sK17,bool),'type.realax.real'))),sK18),s(fun(sK17,'type.extreal.extreal'),'const.measure.indicator_fn_1'(s(fun(sK17,bool),sK19))))))),
  inference(forward_demodulation,[],[f700,f544])).
fof(f700,plain,(
  p(s(bool,'const.lebesgue.integrable_2'(s('type.pair.prod'(fun(sK17,bool),'type.pair.prod'(fun(fun(sK17,bool),bool),fun(fun(sK17,bool),'type.realax.real'))),sK18),s(fun(sK17,'type.extreal.extreal'),'const.measure.indicator_fn_1'(s(fun(sK17,bool),sK19)))))) | ~p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK17,bool),'type.pair.prod'(fun(fun(sK17,bool),bool),fun(fun(sK17,bool),'type.realax.real'))),sK18))))),
  inference(resolution,[],[f311,f365])).
fof(f365,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.lebesgue.integrable_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,'type.extreal.extreal'),'const.measure.indicator_fn_1'(s(fun(X0,bool),X2)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.measure.measurable_sets_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))))) | ~p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1))))) )),
  inference(cnf_transformation,[],[f165])).
fof(f165,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.measure.measurable_sets_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))))) | p(s(bool,'const.lebesgue.integrable_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,'type.extreal.extreal'),'const.measure.indicator_fn_1'(s(fun(X0,bool),X2)))))))),
  inference(flattening,[],[f164])).
fof(f164,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.measure.measurable_sets_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1))))))) | p(s(bool,'const.lebesgue.integrable_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,'type.extreal.extreal'),'const.measure.indicator_fn_1'(s(fun(X0,bool),X2)))))))),
  inference(ennf_transformation,[],[f100])).
fof(f100,plain,(
  ! [X0,X1,X2] : ((p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.measure.measurable_sets_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1))))))) => p(s(bool,'const.lebesgue.integrable_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,'type.extreal.extreal'),'const.measure.indicator_fn_1'(s(fun(X0,bool),X2)))))))),
  inference(flattening,[],[f99])).
fof(f99,plain,(
  ! [X0] : ! [X1,X2] : ((p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.measure.measurable_sets_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1))))))) => p(s(bool,'const.lebesgue.integrable_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,'type.extreal.extreal'),'const.measure.indicator_fn_1'(s(fun(X0,bool),X2)))))))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X6] : ! [X7,X8] : ((p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X6,bool),'type.pair.prod'(fun(fun(X6,bool),bool),fun(fun(X6,bool),'type.realax.real'))),X7)))) & p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),X8),s(fun(fun(X6,bool),bool),'const.measure.measurable_sets_1'(s('type.pair.prod'(fun(X6,bool),'type.pair.prod'(fun(fun(X6,bool),bool),fun(fun(X6,bool),'type.realax.real'))),X7))))))) => p(s(bool,'const.lebesgue.integrable_2'(s('type.pair.prod'(fun(X6,bool),'type.pair.prod'(fun(fun(X6,bool),bool),fun(fun(X6,bool),'type.realax.real'))),X7),s(fun(X6,'type.extreal.extreal'),'const.measure.indicator_fn_1'(s(fun(X6,bool),X8)))))))),
  file('thm.lebesgue.integral_cmul_indicator','thm.lebesgue.integrable_indicator')).
fof(f965,plain,(
  s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK20),s('type.realax.real','const.measure.measure_2'(s('type.pair.prod'(fun(sK17,bool),'type.pair.prod'(fun(fun(sK17,bool),bool),fun(fun(sK17,bool),'type.realax.real'))),sK18),s(fun(sK17,bool),sK19))))))) != s('type.extreal.extreal',happ(s(fun('type.extreal.extreal','type.extreal.extreal'),'const.extreal.extreal_mul_1'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK20))))),s('type.extreal.extreal','const.lebesgue.integral_2'(s('type.pair.prod'(fun(sK17,bool),'type.pair.prod'(fun(fun(sK17,bool),bool),fun(fun(sK17,bool),'type.realax.real'))),sK18),s(fun(sK17,'type.extreal.extreal'),'const.measure.indicator_fn_1'(s(fun(sK17,bool),sK19))))))) | ~p(s(bool,'const.lebesgue.integrable_2'(s('type.pair.prod'(fun(sK17,bool),'type.pair.prod'(fun(fun(sK17,bool),bool),fun(fun(sK17,bool),'type.realax.real'))),sK18),s(fun(sK17,'type.extreal.extreal'),'const.measure.indicator_fn_1'(s(fun(sK17,bool),sK19)))))) | ~p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(sK17,bool),'type.pair.prod'(fun(fun(sK17,bool),bool),fun(fun(sK17,bool),'type.realax.real'))),sK18))))),
  inference(superposition,[],[f312,f367])).
fof(f367,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.extreal.extreal',happ(s(fun('type.extreal.extreal','type.extreal.extreal'),'const.extreal.extreal_mul_1'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))))),s('type.extreal.extreal','const.lebesgue.integral_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,'type.extreal.extreal'),X2))))) = s('type.extreal.extreal','const.lebesgue.integral_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,'type.extreal.extreal'),'_dst_x0x1_2'(s(fun('type.extreal.extreal','type.extreal.extreal'),'const.extreal.extreal_mul_1'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))))),s(fun(X0,'type.extreal.extreal'),X2))))) | ~p(s(bool,'const.lebesgue.integrable_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,'type.extreal.extreal'),X2)))) | ~p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1))))) )),
  inference(cnf_transformation,[],[f169])).
fof(f169,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))) | ~p(s(bool,'const.lebesgue.integrable_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,'type.extreal.extreal'),X2)))) | s('type.extreal.extreal',happ(s(fun('type.extreal.extreal','type.extreal.extreal'),'const.extreal.extreal_mul_1'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))))),s('type.extreal.extreal','const.lebesgue.integral_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,'type.extreal.extreal'),X2))))) = s('type.extreal.extreal','const.lebesgue.integral_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,'type.extreal.extreal'),'_dst_x0x1_2'(s(fun('type.extreal.extreal','type.extreal.extreal'),'const.extreal.extreal_mul_1'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))))),s(fun(X0,'type.extreal.extreal'),X2))))))),
  inference(flattening,[],[f168])).
fof(f168,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))) | ~p(s(bool,'const.lebesgue.integrable_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,'type.extreal.extreal'),X2))))) | s('type.extreal.extreal',happ(s(fun('type.extreal.extreal','type.extreal.extreal'),'const.extreal.extreal_mul_1'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))))),s('type.extreal.extreal','const.lebesgue.integral_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,'type.extreal.extreal'),X2))))) = s('type.extreal.extreal','const.lebesgue.integral_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,'type.extreal.extreal'),'_dst_x0x1_2'(s(fun('type.extreal.extreal','type.extreal.extreal'),'const.extreal.extreal_mul_1'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))))),s(fun(X0,'type.extreal.extreal'),X2))))))),
  inference(ennf_transformation,[],[f104])).
fof(f104,plain,(
  ! [X0,X1,X2,X3] : ((p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))) & p(s(bool,'const.lebesgue.integrable_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,'type.extreal.extreal'),X2))))) => s('type.extreal.extreal',happ(s(fun('type.extreal.extreal','type.extreal.extreal'),'const.extreal.extreal_mul_1'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))))),s('type.extreal.extreal','const.lebesgue.integral_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,'type.extreal.extreal'),X2))))) = s('type.extreal.extreal','const.lebesgue.integral_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,'type.extreal.extreal'),'_dst_x0x1_2'(s(fun('type.extreal.extreal','type.extreal.extreal'),'const.extreal.extreal_mul_1'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))))),s(fun(X0,'type.extreal.extreal'),X2))))))),
  inference(flattening,[],[f103])).
fof(f103,plain,(
  ! [X0] : ! [X1,X2,X3] : ((p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1)))) & p(s(bool,'const.lebesgue.integrable_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,'type.extreal.extreal'),X2))))) => s('type.extreal.extreal',happ(s(fun('type.extreal.extreal','type.extreal.extreal'),'const.extreal.extreal_mul_1'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))))),s('type.extreal.extreal','const.lebesgue.integral_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,'type.extreal.extreal'),X2))))) = s('type.extreal.extreal','const.lebesgue.integral_2'(s('type.pair.prod'(fun(X0,bool),'type.pair.prod'(fun(fun(X0,bool),bool),fun(fun(X0,bool),'type.realax.real'))),X1),s(fun(X0,'type.extreal.extreal'),'_dst_x0x1_2'(s(fun('type.extreal.extreal','type.extreal.extreal'),'const.extreal.extreal_mul_1'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))))),s(fun(X0,'type.extreal.extreal'),X2))))))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X6] : ! [X7,X16,X9] : ((p(s(bool,'const.measure.measure_space_1'(s('type.pair.prod'(fun(X6,bool),'type.pair.prod'(fun(fun(X6,bool),bool),fun(fun(X6,bool),'type.realax.real'))),X7)))) & p(s(bool,'const.lebesgue.integrable_2'(s('type.pair.prod'(fun(X6,bool),'type.pair.prod'(fun(fun(X6,bool),bool),fun(fun(X6,bool),'type.realax.real'))),X7),s(fun(X6,'type.extreal.extreal'),X16))))) => s('type.extreal.extreal',happ(s(fun('type.extreal.extreal','type.extreal.extreal'),'const.extreal.extreal_mul_1'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X9))))),s('type.extreal.extreal','const.lebesgue.integral_2'(s('type.pair.prod'(fun(X6,bool),'type.pair.prod'(fun(fun(X6,bool),bool),fun(fun(X6,bool),'type.realax.real'))),X7),s(fun(X6,'type.extreal.extreal'),X16))))) = s('type.extreal.extreal','const.lebesgue.integral_2'(s('type.pair.prod'(fun(X6,bool),'type.pair.prod'(fun(fun(X6,bool),bool),fun(fun(X6,bool),'type.realax.real'))),X7),s(fun(X6,'type.extreal.extreal'),'_dst_x0x1_2'(s(fun('type.extreal.extreal','type.extreal.extreal'),'const.extreal.extreal_mul_1'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X9))))),s(fun(X6,'type.extreal.extreal'),X16))))))),
  file('thm.lebesgue.integral_cmul_indicator','thm.lebesgue.integral_cmul')).
fof(f312,plain,(
  s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK20),s('type.realax.real','const.measure.measure_2'(s('type.pair.prod'(fun(sK17,bool),'type.pair.prod'(fun(fun(sK17,bool),bool),fun(fun(sK17,bool),'type.realax.real'))),sK18),s(fun(sK17,bool),sK19))))))) != s('type.extreal.extreal','const.lebesgue.integral_2'(s('type.pair.prod'(fun(sK17,bool),'type.pair.prod'(fun(fun(sK17,bool),bool),fun(fun(sK17,bool),'type.realax.real'))),sK18),s(fun(sK17,'type.extreal.extreal'),'_dst_x0x1_2'(s(fun('type.extreal.extreal','type.extreal.extreal'),'const.extreal.extreal_mul_1'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK20))))),s(fun(sK17,'type.extreal.extreal'),'const.measure.indicator_fn_1'(s(fun(sK17,bool),sK19)))))))),
  inference(cnf_transformation,[],[f211])).
% SZS output end Proof for thm.lebesgue
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 197
% Passive clauses: 209
% Generated clauses: 433
% Final active clauses: 194
% Final passive clauses: 180
% Input formulas: 61
% Initial clauses: 221
% 
% Duplicate literals: 13
% Trivial inequalities: 3
% Fw subsumption resolutions: 47
% Bw subsumption resolutions: 2
% Fw demodulations: 30
% Bw demodulations: 2
% Condensations: 11
% 
% Simple tautologies: 50
% Equational tautologies: 2
% Forward subsumptions: 113
% Backward subsumptions: 13
% 
% Binary resolution: 269
% Forward superposition: 63
% 
% SAT solver clauses: 211
% SAT solver unit clauses: 43
% SAT solver binary clauses: 67
% 
% Memory used [KB]: 2046
% Time elapsed: 0.025 s
% ------------------------------
% ------------------------------
% Success in time 16.447 s
