Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.integer
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.integer
% SZS output start Proof for thm.integer
fof(f27695,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f279,f27620,f8656,f27621,f17920,f17918,f17912,f283,f503,f501,f490,f504,f27582,f8655,f284,f499,f497,f495,f27563,f280,f8619,f292,f507,f27551,f282,f502,f27644,f27671,f281,f498,f500,f494,f27669,f496])).
fof(f496,plain,(
  s('type.integer.int',sK10) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK11))) | $spl2),
  inference(cnf_transformation,[],[f496_D])).
fof(f496_D,plain,(
  s('type.integer.int',sK10) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK11))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f27669,plain,(
  s('type.integer.int',sK10) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK11))) | ($spl0 | $spl7)),
  inference(subsumption_resolution,[],[f27668,f501])).
fof(f27668,plain,(
  s('type.integer.int',sK10) = s('type.integer.int',sK11) | s('type.integer.int',sK10) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK11))) | $spl0),
  inference(forward_demodulation,[],[f27649,f494])).
fof(f27649,plain,(
  s('type.integer.int',sK10) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK11))) | s('type.integer.int',sK10) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK10))) | $spl0),
  inference(superposition,[],[f6570,f494])).
fof(f6570,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))) )),
  inference(equality_resolution,[],[f372])).
fof(f372,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X1) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))) | s('type.integer.int',X0) = s('type.integer.int',X1) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))) )),
  inference(cnf_transformation,[],[f228])).
fof(f228,plain,(
  ! [X0,X1] : ((s('type.integer.int',X1) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))) | (s('type.integer.int',X0) = s('type.integer.int',X1) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1))))) | (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1)))))) & (((s('type.integer.int',X0) != s('type.integer.int',X1) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1))))) & (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1)))))) | s('type.integer.int',X1) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))),
  inference(flattening,[],[f227])).
fof(f227,plain,(
  ! [X0,X1] : ((s('type.integer.int',X1) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))) | ((s('type.integer.int',X0) = s('type.integer.int',X1) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1))))) | (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1))))))) & (((s('type.integer.int',X0) != s('type.integer.int',X1) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1))))) & (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1)))))) | s('type.integer.int',X1) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))),
  inference(nnf_transformation,[],[f127])).
fof(f127,plain,(
  ! [X0,X1] : (s('type.integer.int',X1) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))) <=> ((s('type.integer.int',X0) = s('type.integer.int',X1) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1))))) | (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1)))))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6,X7] : (s('type.integer.int',X7) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X6))) <=> ((s('type.integer.int',X6) = s('type.integer.int',X7) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X7))))) | (s('type.integer.int',X6) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X7))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X7)))))))),
  file('thm.integer.INT_ABS_EQ|split|1','thm.integer.INT_ABS_EQ|split|0')).
fof(f494,plain,(
  s('type.integer.int',sK11) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK10))) | $spl0),
  inference(cnf_transformation,[],[f494_D])).
fof(f494_D,plain,(
  s('type.integer.int',sK11) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK10))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f500,plain,(
  s('type.integer.int',sK10) = s('type.integer.int',sK11) | $spl6),
  inference(cnf_transformation,[],[f500_D])).
fof(f500_D,plain,(
  s('type.integer.int',sK10) = s('type.integer.int',sK11) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f498,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK11)))) | $spl4),
  inference(cnf_transformation,[],[f498_D])).
fof(f498_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK11)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f281,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK11)))) | s('type.integer.int',sK10) = s('type.integer.int',sK11) | s('type.integer.int',sK11) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK10)))),
  inference(cnf_transformation,[],[f195])).
fof(f195,plain,(
  (s('type.integer.int',sK11) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK10))) | (s('type.integer.int',sK10) = s('type.integer.int',sK11) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK11))))) | (s('type.integer.int',sK10) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK11))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK11)))))) & (s('type.integer.int',sK11) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK10))) | ((s('type.integer.int',sK10) != s('type.integer.int',sK11) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK11))))) & (s('type.integer.int',sK10) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK11))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK11)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK10,sK11])],[f194])).
fof(f194,plain,(
  ? [X0,X1] : ((s('type.integer.int',X1) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))) | (s('type.integer.int',X0) = s('type.integer.int',X1) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1))))) | (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1)))))) & (s('type.integer.int',X1) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))) | ((s('type.integer.int',X0) != s('type.integer.int',X1) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1))))) & (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1))))))))),
  inference(flattening,[],[f193])).
fof(f193,plain,(
  ? [X0,X1] : ((s('type.integer.int',X1) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))) | ((s('type.integer.int',X0) = s('type.integer.int',X1) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1))))) | (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1))))))) & (s('type.integer.int',X1) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))) | ((s('type.integer.int',X0) != s('type.integer.int',X1) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1))))) & (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1))))))))),
  inference(nnf_transformation,[],[f157])).
fof(f157,plain,(
  ? [X0,X1] : (s('type.integer.int',X1) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))) <~> ((s('type.integer.int',X0) = s('type.integer.int',X1) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1))))) | (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1)))))))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ~! [X0,X1] : (s('type.integer.int',X1) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))) <=> ((s('type.integer.int',X0) = s('type.integer.int',X1) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1))))) | (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (s('type.integer.int',X7) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X6))) <=> ((s('type.integer.int',X6) = s('type.integer.int',X7) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X7))))) | (s('type.integer.int',X6) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X7))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X7)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (s('type.integer.int',X7) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X6))) <=> ((s('type.integer.int',X6) = s('type.integer.int',X7) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X7))))) | (s('type.integer.int',X6) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X7))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X7)))))))),
  file('thm.integer.INT_ABS_EQ|split|1','thm.integer.INT_ABS_EQ|split|1')).
fof(f27671,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK11)))) | ($spl0 | $spl7)),
  inference(subsumption_resolution,[],[f27670,f501])).
fof(f27670,plain,(
  s('type.integer.int',sK10) = s('type.integer.int',sK11) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK11)))) | $spl0),
  inference(forward_demodulation,[],[f27650,f494])).
fof(f27650,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK11)))) | s('type.integer.int',sK10) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK10))) | $spl0),
  inference(superposition,[],[f7531,f494])).
fof(f7531,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))) )),
  inference(equality_resolution,[],[f374])).
fof(f374,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X1) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))) | s('type.integer.int',X0) = s('type.integer.int',X1) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f228])).
fof(f27644,plain,(
  $false | ($spl0 | $spl6 | $spl294)),
  inference(trivial_inequality_removal,[],[f27643])).
fof(f27643,plain,(
  s('type.integer.int',sK10) != s('type.integer.int',sK10) | ($spl0 | $spl6 | $spl294)),
  inference(superposition,[],[f27621,f17921])).
fof(f17921,plain,(
  s('type.integer.int',sK10) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK10))) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f494,f500])).
fof(f502,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK11)))) | $spl8),
  inference(cnf_transformation,[],[f502_D])).
fof(f502_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK11)))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f282,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK11)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK11)))) | s('type.integer.int',sK11) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK10)))),
  inference(cnf_transformation,[],[f195])).
fof(f27551,plain,(
  $false | ($spl0 | $spl3 | $spl6 | $spl13 | $spl237)),
  inference(subsumption_resolution,[],[f27550,f8653])).
fof(f8653,plain,(
  s('type.integer.int',sK10) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK10))) | ($spl3 | $spl6)),
  inference(forward_demodulation,[],[f497,f500])).
fof(f27550,plain,(
  s('type.integer.int',sK10) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK10))) | ($spl0 | $spl6 | $spl13 | $spl237)),
  inference(forward_demodulation,[],[f27549,f17921])).
fof(f27549,plain,(
  s('type.integer.int',sK10) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK10))))) | ($spl0 | $spl6 | $spl13 | $spl237)),
  inference(forward_demodulation,[],[f27548,f500])).
fof(f27548,plain,(
  s('type.integer.int',sK11) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK11))))) | ($spl0 | $spl6 | $spl13 | $spl237)),
  inference(subsumption_resolution,[],[f27547,f507])).
fof(f27547,plain,(
  p(s(bool,'F_0')) | s('type.integer.int',sK11) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK11))))) | ($spl0 | $spl6 | $spl237)),
  inference(forward_demodulation,[],[f27546,f17961])).
fof(f17961,plain,(
  s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK10))) | $spl237),
  inference(resolution,[],[f17918,f320])).
fof(f320,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f200])).
fof(f200,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f98])).
fof(f98,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f97])).
fof(f97,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X8] : (s(bool,'F_0') = s(bool,X8) <=> ~p(s(bool,X8)))),
  file('thm.integer.INT_ABS_EQ|split|1','thm.bool.EQ_CLAUSES|split|3')).
fof(f27546,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK10)))) | s('type.integer.int',sK11) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK11))))) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f27535,f17921])).
fof(f27535,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK10)))))) | s('type.integer.int',sK11) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK11))))) | $spl6),
  inference(superposition,[],[f8035,f500])).
fof(f8035,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))) )),
  inference(equality_resolution,[],[f373])).
fof(f373,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X1) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1)))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))) )),
  inference(cnf_transformation,[],[f228])).
fof(f507,plain,(
  ~p(s(bool,'F_0')) | $spl13),
  inference(cnf_transformation,[],[f507_D])).
fof(f507_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f292,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f74])).
fof(f74,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.integer.INT_ABS_EQ|split|1',aHL_FALSITY)).
fof(f8619,plain,(
  $false | ($spl1 | $spl2 | $spl4 | $spl10)),
  inference(subsumption_resolution,[],[f8618,f495])).
fof(f8618,plain,(
  s('type.integer.int',sK11) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK10))) | ($spl2 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f8617,f496])).
fof(f8617,plain,(
  s('type.integer.int',sK11) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK11))))) | ($spl2 | $spl4 | $spl10)),
  inference(trivial_inequality_removal,[],[f8615])).
fof(f8615,plain,(
  s('type.integer.int',sK10) != s('type.integer.int',sK10) | s('type.integer.int',sK11) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK11))))) | ($spl2 | $spl4 | $spl10)),
  inference(superposition,[],[f8100,f496])).
fof(f8100,plain,(
  ( ! [X0] : (s('type.integer.int',sK10) != s('type.integer.int',X0) | s('type.integer.int',sK11) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))) ) | ($spl2 | $spl4 | $spl10)),
  inference(subsumption_resolution,[],[f8099,f504])).
fof(f8099,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s('type.integer.int',sK10) != s('type.integer.int',X0) | s('type.integer.int',sK11) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))) ) | ($spl2 | $spl4)),
  inference(forward_demodulation,[],[f8094,f948])).
fof(f948,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK11))) | $spl4),
  inference(resolution,[],[f498,f316])).
fof(f316,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f198])).
fof(f198,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f95])).
fof(f95,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X8] : (s(bool,'T_0') = s(bool,X8) <=> p(s(bool,X8)))),
  file('thm.integer.INT_ABS_EQ|split|1','thm.bool.EQ_CLAUSES|split|1')).
fof(f8094,plain,(
  ( ! [X0] : (s('type.integer.int',sK10) != s('type.integer.int',X0) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK11)))) | s('type.integer.int',sK11) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))) ) | $spl2),
  inference(superposition,[],[f377,f496])).
fof(f377,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1)))) | s('type.integer.int',X1) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))) )),
  inference(cnf_transformation,[],[f228])).
fof(f280,plain,(
  s('type.integer.int',sK10) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK11))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK11)))) | s('type.integer.int',sK11) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK10)))),
  inference(cnf_transformation,[],[f195])).
fof(f27563,plain,(
  $false | ($spl0 | $spl3 | $spl4 | $spl6 | $spl10 | $spl13 | $spl237)),
  inference(subsumption_resolution,[],[f27562,f8653])).
fof(f27562,plain,(
  s('type.integer.int',sK10) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK10))) | ($spl0 | $spl4 | $spl6 | $spl10 | $spl13 | $spl237)),
  inference(forward_demodulation,[],[f27561,f17921])).
fof(f27561,plain,(
  s('type.integer.int',sK10) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK10))))) | ($spl0 | $spl4 | $spl6 | $spl10 | $spl13 | $spl237)),
  inference(forward_demodulation,[],[f27560,f18079])).
fof(f18079,plain,(
  s('type.integer.int',sK10) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK10))))) | ($spl0 | $spl4 | $spl6 | $spl10)),
  inference(equality_resolution,[],[f17956])).
fof(f17956,plain,(
  ( ! [X14] : (s('type.integer.int',X14) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK10))) | s('type.integer.int',sK10) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X14)))) ) | ($spl0 | $spl4 | $spl6 | $spl10)),
  inference(forward_demodulation,[],[f17955,f17921])).
fof(f17955,plain,(
  ( ! [X14] : (s('type.integer.int',X14) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK10))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK10))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X14)))) ) | ($spl0 | $spl4 | $spl6 | $spl10)),
  inference(subsumption_resolution,[],[f17954,f504])).
fof(f17954,plain,(
  ( ! [X14] : (~p(s(bool,'T_0')) | s('type.integer.int',X14) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK10))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK10))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X14)))) ) | ($spl0 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f17953,f8622])).
fof(f8622,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK10))) | ($spl4 | $spl6)),
  inference(backward_demodulation,[],[f500,f948])).
fof(f17953,plain,(
  ( ! [X14] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK10)))) | s('type.integer.int',X14) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK10))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK10))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X14)))) ) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f17937,f17921])).
fof(f17937,plain,(
  ( ! [X14] : (s('type.integer.int',X14) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK10))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK10)))))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK10))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X14)))) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f377,f17921])).
fof(f27560,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK10))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK10))))))))) | ($spl0 | $spl4 | $spl6 | $spl10 | $spl13 | $spl237)),
  inference(subsumption_resolution,[],[f27559,f507])).
fof(f27559,plain,(
  p(s(bool,'F_0')) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK10))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK10))))))))) | ($spl0 | $spl4 | $spl6 | $spl10 | $spl237)),
  inference(forward_demodulation,[],[f27558,f17961])).
fof(f27558,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK10)))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK10))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK10))))))))) | ($spl0 | $spl4 | $spl6 | $spl10)),
  inference(forward_demodulation,[],[f27538,f17921])).
fof(f27538,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK10)))))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK10))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK10))))))))) | ($spl0 | $spl4 | $spl6 | $spl10)),
  inference(superposition,[],[f8035,f18079])).
fof(f495,plain,(
  s('type.integer.int',sK11) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK10))) | $spl1),
  inference(cnf_transformation,[],[f495_D])).
fof(f495_D,plain,(
  s('type.integer.int',sK11) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK10))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f497,plain,(
  s('type.integer.int',sK10) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK11))) | $spl3),
  inference(cnf_transformation,[],[f497_D])).
fof(f497_D,plain,(
  s('type.integer.int',sK10) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK11))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f499,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK11)))) | $spl5),
  inference(cnf_transformation,[],[f499_D])).
fof(f499_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK11)))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f284,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK11)))) | s('type.integer.int',sK10) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK11))) | s('type.integer.int',sK11) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK10)))),
  inference(cnf_transformation,[],[f195])).
fof(f8655,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK10)))) | $spl149),
  inference(cnf_transformation,[],[f8655_D])).
fof(f8655_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK10)))) <=> ~$spl149),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl149])])).
fof(f27582,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK10)))) | ($spl5 | $spl6)),
  inference(forward_demodulation,[],[f499,f500])).
fof(f504,plain,(
  p(s(bool,'T_0')) | $spl10),
  inference(cnf_transformation,[],[f504_D])).
fof(f504_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f490,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f356])).
fof(f356,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f223])).
fof(f223,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f116])).
fof(f116,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f115])).
fof(f115,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f59])).
fof(f59,axiom,(
  ! [X15] : ! [X16] : (s(X15,X16) = s(X15,X16) <=> p(s(bool,'T_0')))),
  file('thm.integer.INT_ABS_EQ|split|1','thm.bool.REFL_CLAUSE')).
fof(f501,plain,(
  s('type.integer.int',sK10) != s('type.integer.int',sK11) | $spl7),
  inference(cnf_transformation,[],[f501_D])).
fof(f501_D,plain,(
  s('type.integer.int',sK10) != s('type.integer.int',sK11) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f503,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK11)))) | $spl9),
  inference(cnf_transformation,[],[f503_D])).
fof(f503_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK11)))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f283,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK11)))) | s('type.integer.int',sK10) != s('type.integer.int',sK11) | s('type.integer.int',sK11) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK10)))),
  inference(cnf_transformation,[],[f195])).
fof(f17912,plain,(
  $false | ($spl1 | $spl6 | $spl8 | $spl10)),
  inference(subsumption_resolution,[],[f17911,f8621])).
fof(f8621,plain,(
  s('type.integer.int',sK10) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK10))) | ($spl1 | $spl6)),
  inference(backward_demodulation,[],[f500,f495])).
fof(f17911,plain,(
  s('type.integer.int',sK10) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK10))) | ($spl6 | $spl8 | $spl10)),
  inference(forward_demodulation,[],[f17910,f500])).
fof(f17910,plain,(
  s('type.integer.int',sK11) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK11))) | ($spl6 | $spl8 | $spl10)),
  inference(subsumption_resolution,[],[f17909,f504])).
fof(f17909,plain,(
  ~p(s(bool,'T_0')) | s('type.integer.int',sK11) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK11))) | ($spl6 | $spl8)),
  inference(forward_demodulation,[],[f17901,f8665])).
fof(f8665,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK10))) | ($spl6 | $spl8)),
  inference(resolution,[],[f8661,f316])).
fof(f8661,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK10)))) | ($spl6 | $spl8)),
  inference(forward_demodulation,[],[f502,f500])).
fof(f17901,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK10)))) | s('type.integer.int',sK11) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK11))) | $spl6),
  inference(superposition,[],[f7537,f500])).
fof(f7537,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))) )),
  inference(equality_resolution,[],[f376])).
fof(f376,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X0) != s('type.integer.int',X1) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1)))) | s('type.integer.int',X1) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))) )),
  inference(cnf_transformation,[],[f228])).
fof(f17918,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK10)))) | $spl237),
  inference(cnf_transformation,[],[f17918_D])).
fof(f17918_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK10)))) <=> ~$spl237),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl237])])).
fof(f17920,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK10)))) | ($spl6 | $spl9)),
  inference(forward_demodulation,[],[f503,f500])).
fof(f27621,plain,(
  ( ! [X7] : (s('type.integer.int',sK10) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X7)))) ) | $spl294),
  inference(cnf_transformation,[],[f27621_D])).
fof(f27621_D,plain,(
  ( ! [X7] : (s('type.integer.int',sK10) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X7)))) ) <=> ~$spl294),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl294])])).
fof(f8656,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK10)))) | $spl148),
  inference(cnf_transformation,[],[f8656_D])).
fof(f8656_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK10)))) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f27620,plain,(
  ( ! [X7] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK10)))) | s('type.integer.int',sK10) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X7)))) ) | ($spl2 | $spl6 | $spl13 | $spl237)),
  inference(forward_demodulation,[],[f27619,f27583])).
fof(f27583,plain,(
  s('type.integer.int',sK10) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK10))) | ($spl2 | $spl6)),
  inference(forward_demodulation,[],[f496,f500])).
fof(f27619,plain,(
  ( ! [X7] : (s('type.integer.int',sK10) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X7))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK10))))))) ) | ($spl2 | $spl6 | $spl13 | $spl237)),
  inference(subsumption_resolution,[],[f27618,f507])).
fof(f27618,plain,(
  ( ! [X7] : (p(s(bool,'F_0')) | s('type.integer.int',sK10) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X7))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK10))))))) ) | ($spl2 | $spl6 | $spl237)),
  inference(forward_demodulation,[],[f27617,f17961])).
fof(f27617,plain,(
  ( ! [X7] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK10)))) | s('type.integer.int',sK10) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X7))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK10))))))) ) | ($spl2 | $spl6)),
  inference(forward_demodulation,[],[f27591,f27583])).
fof(f27591,plain,(
  ( ! [X7] : (s('type.integer.int',sK10) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X7))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK10)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK10))))))) ) | ($spl2 | $spl6)),
  inference(superposition,[],[f375,f27583])).
fof(f375,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X1) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f228])).
fof(f279,plain,(
  s('type.integer.int',sK10) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK11))) | s('type.integer.int',sK10) = s('type.integer.int',sK11) | s('type.integer.int',sK11) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK10)))),
  inference(cnf_transformation,[],[f195])).
% SZS output end Proof for thm.integer
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 773
% Passive clauses: 8340
% Generated clauses: 27423
% Final active clauses: 699
% Final passive clauses: 7001
% Input formulas: 69
% Initial clauses: 209
% 
% Duplicate literals: 295
% Trivial inequalities: 27
% Fw subsumption resolutions: 2157
% Fw demodulations: 1956
% Bw demodulations: 33
% 
% Simple tautologies: 238
% Equational tautologies: 58
% Forward subsumptions: 14122
% Fw demodulations to eq. taut.: 3
% 
% Binary resolution: 9028
% Factoring: 1043
% Forward superposition: 7790
% Backward superposition: 4610
% Self superposition: 56
% Equality factoring: 26
% Equality resolution: 16
% 
% Split clauses: 290
% Split components: 588
% SAT solver clauses: 594
% SAT solver unit clauses: 24
% SAT solver binary clauses: 118
% SAT solver learnt clauses: 7
% SAT solver learnt literals: 10
% 
% Sat splits: 290
% Sat splitting refutations: 9
% 
% TWLsolver clauses: 587
% TWLsolver calls for satisfiability: 380
% 
% Memory used [KB]: 13304
% Time elapsed: 0.589 s
% ------------------------------
% ------------------------------
% Success in time 0.6 s
