Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.integer
% (35880)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1394
% Passive clauses: 12613
% Generated clauses: 42900
% Final active clauses: 1219
% Final passive clauses: 10577
% Input formulas: 261
% Initial clauses: 522
% 
% Duplicate literals: 225
% Trivial inequalities: 89
% Fw subsumption resolutions: 2416
% Fw demodulations: 5577
% Bw demodulations: 171
% 
% Simple tautologies: 400
% Equational tautologies: 36
% Forward subsumptions: 21052
% Fw demodulations to eq. taut.: 4
% 
% Binary resolution: 11105
% Factoring: 819
% Forward superposition: 12255
% Backward superposition: 9298
% Self superposition: 49
% Equality factoring: 2
% Equality resolution: 23
% 
% Split clauses: 643
% Split components: 1288
% SAT solver clauses: 880
% SAT solver unit clauses: 15
% SAT solver binary clauses: 88
% 
% Sat splits: 643
% Sat splitting refutations: 10
% 
% TWLsolver clauses: 875
% TWLsolver calls for satisfiability: 538
% 
% Memory used [KB]: 20084
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.integer
% (35937)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit

% Active clauses: 603
% Passive clauses: 3426
% Generated clauses: 20469
% Final active clauses: 559
% Final passive clauses: 2528
% Input formulas: 261
% Initial clauses: 522
% 
% Duplicate literals: 97
% Trivial inequalities: 44
% Fw subsumption resolutions: 834
% Fw demodulations: 3109
% Bw demodulations: 321
% Condensations: 56
% Global subsumptions: 24
% 
% Simple tautologies: 325
% Equational tautologies: 9
% Forward subsumptions: 8599
% Fw demodulations to eq. taut.: 14
% 
% Binary resolution: 2563
% Factoring: 98
% Forward superposition: 5346
% Backward superposition: 3842
% Self superposition: 61
% Equality factoring: 2
% Equality resolution: 12
% 
% Split clauses: 34
% Split components: 69
% SAT solver clauses: 4382
% SAT solver unit clauses: 627
% SAT solver binary clauses: 168
% SAT solver learnt clauses: 48
% SAT solver learnt literals: 12
% 
% Sat splits: 34
% Sat splitting refutations: 12
% 
% TWLsolver clauses: 4286
% TWLsolver calls for satisfiability: 3484
% 
% Memory used [KB]: 18038
% Time elapsed: 0.718 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.integer
% (35967)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1616
% Passive clauses: 19289
% Generated clauses: 141552
% Final active clauses: 1074
% Final passive clauses: 4734
% Input formulas: 261
% Initial clauses: 522
% Discarded non-redundant clauses: 9795
% 
% Duplicate literals: 113
% Trivial inequalities: 140
% Fw subsumption resolutions: 2424
% Fw demodulations: 45655
% Bw demodulations: 7675
% Condensations: 60
% Global subsumptions: 11
% 
% Simple tautologies: 352
% Equational tautologies: 68
% Forward subsumptions: 67673
% Fw demodulations to eq. taut.: 270
% Bw demodulations to eq. taut.: 57
% 
% Unit resulting resolution: 47662
% Factoring: 9
% Forward superposition: 16854
% Backward superposition: 12929
% Self superposition: 131
% Equality factoring: 2
% Equality resolution: 21
% 
% Split clauses: 42
% Split components: 85
% SAT solver clauses: 21457
% SAT solver unit clauses: 4117
% SAT solver binary clauses: 8574
% 
% Sat splits: 42
% Sat splitting refutations: 2071
% 
% Memory used [KB]: 125115
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 21 next slice time: 21
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.integer
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.integer
% SZS output start Proof for thm.integer
fof(f62860,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f602,f906,f601,f600,f915,f918,f1047,f1048,f1187,f1189,f1332,f1424,f1474,f1476,f1612,f1898,f2037,f2042,f1993,f2049,f2056,f2061,f2067,f2069,f2072,f2077,f2080,f2083,f2086,f2111,f2114,f2232,f2324,f2366,f2367,f2369,f2516,f2664,f2672,f2623,f2673,f2678,f2681,f2685,f2687,f2688,f2689,f2690,f2691,f2694,f2782,f2784,f3197,f3203,f3216,f3225,f3289,f3293,f3295,f3297,f3401,f3544,f3636,f3679,f3681,f3825,f3971,f3979,f3932,f1167,f3980,f3985,f3988,f3992,f3994,f3995,f3996,f3997,f3998,f4001,f4089,f4091,f4244,f4245,f4971,f4972,f4986,f4822,f4988,f4687,f4989,f4990,f4991,f4992,f4993,f4994,f4995,f4996,f4997,f4998,f4999,f5289,f5291,f5449,f5584,f5587,f5544,f5588,f5593,f5596,f5600,f5602,f5603,f5604,f5605,f5606,f5609,f5721,f5724,f6029,f6191,f6337,f6343,f6345,f6350,f6353,f6357,f6359,f6360,f6361,f6362,f6363,f6366,f6432,f6787,f2487,f6790,f6924,f7125,f7879,f7767,f8208,f8363,f8251,f8387,f8400,f8971,f9336,f9217,f9683,f9564,f10362,f8089,f905,f10391,f10392,f10396,f7539,f8364,f7405,f6561,f9684,f6427,f6182,f6032,f7743,f10406,f7608,f7536,f10407,f6034,f7880,f5899,f9098,f10409,f8059,f8209,f7925,f10411,f10412,f6782,f10419,f10420,f10569,f10568,f10563,f10542,f10539,f10538,f10537,f10536,f10532,f10531,f10515,f10650,f10651,f10652,f10653,f10654,f10655,f10656,f10657,f10658,f10659,f10660,f10661,f5716,f7613,f10395,f7754,f10676,f10677,f10678,f10679,f8080,f10711,f11008,f11076,f11234,f11236,f11238,f11240,f11256,f11258,f11385,f11637,f11638,f11639,f11387,f11517,f11707,f12025,f12026,f12792,f13209,f13575,f13581,f13428,f13582,f13584,f13588,f13591,f12938,f13592,f13593,f13594,f13595,f13596,f13597,f13600,f13601,f13602,f13603,f13604,f13611,f13612,f13619,f13620,f13627,f13628,f13636,f13767,f13768,f13774,f13910,f13912,f14151,f14972,f14853,f14987,f15300,f15301,f15311,f15312,f15613,f16623,f16624,f16625,f16626,f16627,f16628,f16629,f16630,f16178,f14032,f11211,f16648,f16649,f14844,f16666,f16667,f16668,f16669,f1029,f16966,f16980,f16984,f17143,f17165,f1032,f17437,f17407,f17503,f17638,f17646,f17649,f17655,f17656,f17657,f17717,f1451,f17719,f2894,f1594,f18100,f18415,f18417,f1298,f18675,f18838,f19030,f18911,f19070,f19071,f19199,f1596,f19202,f19839,f19346,f19852,f19853,f19854,f20041,f20187,f2662,f20198,f20201,f20207,f20208,f20209,f2050,f20532,f20696,f20697,f20698,f20708,f20709,f20710,f20975,f8388,f20987,f8401,f20996,f21627,f21900,f21788,f22187,f22188,f22189,f22190,f22191,f22210,f22211,f22447,f22448,f22464,f22465,f1302,f22473,f22608,f22631,f22633,f22652,f22653,f22669,f22670,f23721,f1164,f13330,f24280,f24296,f24418,f24440,f24556,f1300,f24727,f24728,f25294,f25175,f25589,f25607,f25740,f25742,f25735,f25744,f25757,f25874,f26347,f26481,f26483,f26476,f26478,f26498,f26921,f26924,f26928,f26910,f26912,f27127,f27130,f27134,f27116,f27118,f28294,f28427,f1592,f28455,f28461,f28465,f28469,f28600,f28610,f28621,f28622,f23566,f23259,f28623,f22971,f22802,f28624,f18414,f1462,f17162,f1030,f24295,f25172,f5837,f28628,f28626,f28619,f1028,f28617,f28615,f26496,f22632,f22630,f28637,f28640,f28644,f28482,f28646,f29018,f29019,f29020,f29021,f29022,f29023,f29024,f29025,f29026,f29267,f29271,f29274,f29275,f29279,f3205,f29280,f29282,f29289,f29293,f29295,f29300,f29305,f29307,f29308,f29309,f29310,f29311,f29312,f29313,f29314,f29315,f29316,f29317,f29318,f29319,f29320,f29321,f29322,f29323,f29324,f29325,f29326,f29327,f29328,f25759,f29330,f29337,f29339,f13427,f29342,f29344,f29340,f29345,f29338,f29348,f3201,f29350,f29352,f29618,f29749,f25737,f29753,f29755,f2062,f29779,f29781,f30287,f30296,f30541,f30543,f30551,f29904,f29269,f25876,f30556,f30558,f30561,f31171,f31330,f31454,f31254,f31249,f31156,f30560,f30557,f30555,f30366,f31506,f31619,f31832,f31863,f31868,f19194,f907,f32198,f32199,f32200,f32201,f32202,f22673,f32203,f32204,f23075,f23849,f32205,f2488,f32206,f25167,f32207,f10393,f32208,f32210,f32212,f32214,f32216,f32218,f32220,f32222,f32224,f32226,f32209,f32227,f32211,f32228,f32213,f32229,f32215,f32230,f32217,f32231,f32219,f32232,f32221,f32233,f32223,f32234,f32225,f28287,f1031,f27598,f32235,f27617,f32236,f27635,f32237,f27655,f32238,f27674,f32239,f27692,f32240,f27711,f32241,f27731,f32242,f27751,f32243,f22977,f908,f32252,f6927,f1721,f10397,f10398,f10401,f28434,f32259,f32261,f29299,f32263,f29296,f29294,f32265,f28454,f28460,f28459,f28464,f28457,f28468,f28467,f28463,f29284,f22679,f23855,f10675,f2043,f32285,f32287,f32293,f32298,f32301,f32303,f32305,f32315,f32317,f32320,f27454,f32327,f27434,f32328,f27414,f32329,f27395,f32330,f27376,f32331,f27357,f32332,f27338,f32333,f27319,f32334,f27300,f32335,f23953,f1027,f22769,f22399,f7126,f22237,f6925,f3221,f28288,f4364,f32344,f32345,f32351,f1454,f32472,f32443,f1461,f32486,f32487,f10394,f32494,f32495,f32502,f32503,f1463,f32521,f32656,f32659,f32616,f32660,f32665,f2058,f32668,f32672,f32674,f32675,f32676,f32677,f32678,f32681,f32690,f1595,f32692,f1452,f1597,f32826,f1299,f32828,f1305,f1303,f32948,f1724,f1722,f33136,f33143,f33159,f33166,f33179,f1460,f33515,f33517,f3199,f33519,f33521,f33526,f33534,f33536,f3218,f33539,f3222,f33447,f3220,f33604,f2045,f33607,f909,f3291,f33610,f33613,f3290,f28387,f1034,f32257,f33979,f34264,f34267,f35304,f35305,f36189,f6903,f36191,f36735,f6904,f36493,f37150,f37588,f37195,f37592,f37596,f37600,f37603,f37202,f37604,f37606,f37206,f37607,f37695,f38110,f38112,f38495,f38117,f39455,f39496,f39055,f39498,f39502,f39506,f37602,f39509,f39511,f37597,f40016,f41132,f41133,f41134,f41136,f41173,f41175,f41185,f41188,f41190,f41192,f41205,f41208,f41210,f41212,f41215,f41217,f41383,f910,f10400,f41529,f41659,f41621,f2044,f36190,f41624,f41625,f1166,f41664,f41665,f41634,f41667,f2048,f37590,f42131,f42135,f42136,f42138,f42140,f42141,f44438,f44495,f44500,f44511,f44515,f44519,f44501,f44522,f44524,f44512,f37605,f44527,f44529,f44516,f37609,f44557,f44629,f45638,f45641,f42132,f45642,f45643,f42137,f45644,f42139,f45645,f45646,f47533,f47845,f48615,f47896,f47899,f48619,f47900,f48620,f47901,f48617,f48621,f48623,f48622,f48624,f48626,f48625,f48634,f48636,f48635,f48789,f48795,f28634,f28632,f48797,f26613,f50885,f51737,f2079,f50938,f50941,f51741,f50942,f51742,f50943,f51739,f51743,f51745,f51744,f51746,f51748,f51747,f51759,f51761,f51760,f52374,f35355,f53512,f2051,f52427,f37593,f52430,f53517,f52431,f53518,f52432,f37601,f53515,f53519,f53521,f37598,f53520,f53522,f53524,f37608,f53523,f53541,f45640,f53543,f42134,f53542,f54124,f54130,f54126,f54134,f54128,f54137,f54136,f54132,f53855,f54144,f53856,f54145,f54143,f53857,f54146,f53860,f54147,f54142,f54148,f54293,f55015,f13585,f55027,f13583,f55218,f55219,f55229,f55230,f4363,f55393,f4987,f55274,f9338,f5839,f55413,f55414,f55428,f55429,f3198,f55477,f28629,f55486,f5840,f55617,f2046,f28630,f55651,f55652,f55666,f55667,f29332,f2097,f56143,f2661,f1165,f56573,f56691,f26615,f57034,f57349,f2033,f57483,f1723,f2039,f57646,f25593,f25591,f58219,f58898,f7881,f58786,f7884,f59181,f4362,f8210,f59339,f8365,f59227,f8368,f59472,f1593,f10405,f59353,f1301,f7883,f59768,f59770,f59772,f59775,f60203,f60204,f60206,f60208,f60210,f60212,f60213,f60214,f60220,f60769,f2486,f9685,f60650,f2031,f8367,f61530,f61534,f61522,f61524,f26922,f61865,f2063,f61869,f2071,f61857,f2076,f61859,f27128,f2082,f2074,f62302,f62326,f62312,f62313,f62314,f62315,f62316,f62317,f62327,f62325,f62431,f40369,f62455,f44439,f62441,f37557,f62442,f35353,f62443,f53501,f62444,f51729,f62445,f48607,f62446,f35356,f62456,f39456,f62454,f62849])).
fof(f62849,plain,(
  $false | ($spl0 | $spl15)),
  inference(subsumption_resolution,[],[f62725,f1165])).
fof(f62725,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK12)))) | $spl0),
  inference(resolution,[],[f1893,f840])).
fof(f840,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f581])).
fof(f581,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int',X1)))) | (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0)))))) & (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int',X1))))))),
  inference(flattening,[],[f580])).
fof(f580,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int',X1)))) | (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0)))))) & ((~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int',X1))))))),
  inference(nnf_transformation,[],[f362])).
fof(f362,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int',X1)))) <=> (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0))))))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X6,X7] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X6))),s('type.integer.int',X7)))) <=> (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X6),s('type.integer.int',X7)))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X7))),s('type.integer.int',X6))))))),
  file('thm.integer.INT_ABS_LE|split|3','thm.integer.INT_ABS_LE|split|0')).
fof(f1893,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl0),
  inference(resolution,[],[f905,f838])).
fof(f838,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))))))) )),
  inference(cnf_transformation,[],[f579])).
fof(f579,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))))))) & ((~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) & ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))))))),
  inference(flattening,[],[f578])).
fof(f578,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))) | (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))))) & ((~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) & ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))))))),
  inference(nnf_transformation,[],[f361])).
fof(f361,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))) <=> (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))))))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6,X7] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X6)))))) <=> (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X7),s('type.integer.int',X6)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X6),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X7))))))))),
  file('thm.integer.INT_ABS_LE|split|3','thm.integer.INT_ABS_LE|split|1')).
fof(f62454,plain,(
  ( ! [X1] : (sP11('const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X1)) ) | $spl454),
  inference(cnf_transformation,[],[f62454_D])).
fof(f62454_D,plain,(
  ( ! [X1] : (sP11('const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X1)) ) <=> ~$spl454),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl454])])).
fof(f39456,plain,(
  ( ! [X462] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X462)))) ) | $spl366),
  inference(cnf_transformation,[],[f39456_D])).
fof(f39456_D,plain,(
  ( ! [X462] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X462)))) ) <=> ~$spl366),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl366])])).
fof(f62456,plain,(
  ( ! [X37,X36] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X36))) | sP11('const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X37)) )),
  inference(forward_demodulation,[],[f62447,f892])).
fof(f892,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))) )),
  inference(cnf_transformation,[],[f397])).
fof(f397,plain,(
  ! [X0] : s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))),
  inference(rectify,[],[f48])).
fof(f48,axiom,(
  ! [X6] : s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X6))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X6)))))),
  file('thm.integer.INT_ABS_LE|split|3','thm.integer.INT_ABS_ABS')).
fof(f62447,plain,(
  ( ! [X37,X36] : (sP11('const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X37) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X36)))))) )),
  inference(superposition,[],[f34485,f890])).
fof(f890,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))) )),
  inference(cnf_transformation,[],[f596])).
fof(f596,plain,(
  ! [X0] : ((s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))),
  inference(nnf_transformation,[],[f396])).
fof(f396,plain,(
  ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))) <=> s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(rectify,[],[f51])).
fof(f51,axiom,(
  ! [X6] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X6))) <=> s('type.integer.int',X6) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  file('thm.integer.INT_ABS_LE|split|3','thm.integer.INT_ABS_EQ|48|')).
fof(f34485,plain,(
  ( ! [X206,X205] : (sP11('const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X205)))),X206)) )),
  inference(resolution,[],[f854,f729])).
fof(f729,plain,(
  ( ! [X0,X1] : (sP11(X0,X1) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f522])).
fof(f522,plain,(
  ! [X0,X1] : ((~sP11(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP11(X0,X1)))),
  inference(flattening,[],[f521])).
fof(f521,plain,(
  ! [X0,X1] : ((~sP11(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP11(X0,X1)))),
  inference(nnf_transformation,[],[f460])).
fof(f460,plain,(
  ! [X0,X1] : (sP11(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f854,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))))) )),
  inference(cnf_transformation,[],[f373])).
fof(f373,plain,(
  ! [X0] : p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X6] : p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X6))))))),
  file('thm.integer.INT_ABS_LE|split|3','thm.integer.INT_ABS_POS')).
fof(f35356,plain,(
  ( ! [X12] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.ABS_1'(s('type.integer.int',X12)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.ABS_1'(s('type.integer.int',X12))))) ) | $spl336),
  inference(cnf_transformation,[],[f35356_D])).
fof(f35356_D,plain,(
  ( ! [X12] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.ABS_1'(s('type.integer.int',X12)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.ABS_1'(s('type.integer.int',X12))))) ) <=> ~$spl336),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl336])])).
fof(f62446,plain,(
  ( ! [X35,X34] : (sP11('const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X35) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.ABS_1'(s('type.integer.int',X34)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.ABS_1'(s('type.integer.int',X34))))) )),
  inference(superposition,[],[f34485,f880])).
fof(f880,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31(X0)) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30(X0))) )),
  inference(cnf_transformation,[],[f595])).
fof(f595,plain,(
  ! [X0] : ((s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30(X0)))) & s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30(X0))) | (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31(X0)))))) & s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31(X0))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK30,sK31])],[f386])).
fof(f386,plain,(
  ! [X0] : (? [X2] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X2))) & s('type.num.num','const.num.|48|_0') != s('type.num.num',X2)) | ? [X1] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))))) & s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(flattening,[],[f385])).
fof(f385,plain,(
  ! [X0] : (? [X2] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X2))) & ~s('type.num.num','const.num.|48|_0') = s('type.num.num',X2)) | ? [X1] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))))) & ~s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X6] : (? [X10] : (s('type.integer.int',X6) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X10))) & ~s('type.num.num','const.num.|48|_0') = s('type.num.num',X10)) | ? [X10] : (s('type.integer.int',X6) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X10))))) & ~s('type.num.num','const.num.|48|_0') = s('type.num.num',X10)) | s('type.integer.int',X6) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  file('thm.integer.INT_ABS_LE|split|3','thm.integer.INT_NUM_CASES')).
fof(f48607,plain,(
  ( ! [X480] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.ABS_1'(s('type.integer.int',X480)))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X480))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.ABS_1'(s('type.integer.int',X480))))))) ) | $spl390),
  inference(cnf_transformation,[],[f48607_D])).
fof(f48607_D,plain,(
  ( ! [X480] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.ABS_1'(s('type.integer.int',X480)))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X480))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.ABS_1'(s('type.integer.int',X480))))))) ) <=> ~$spl390),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl390])])).
fof(f62445,plain,(
  ( ! [X33,X32] : (sP11('const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X33) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.ABS_1'(s('type.integer.int',X32)))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X32))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.ABS_1'(s('type.integer.int',X32))))))) )),
  inference(superposition,[],[f34485,f879])).
fof(f879,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31(X0)) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30(X0))))) )),
  inference(cnf_transformation,[],[f595])).
fof(f51729,plain,(
  ( ! [X487] : (s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X487))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.ABS_1'(s('type.integer.int',X487)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.ABS_1'(s('type.integer.int',X487))))) ) | $spl404),
  inference(cnf_transformation,[],[f51729_D])).
fof(f51729_D,plain,(
  ( ! [X487] : (s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X487))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.ABS_1'(s('type.integer.int',X487)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.ABS_1'(s('type.integer.int',X487))))) ) <=> ~$spl404),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl404])])).
fof(f62444,plain,(
  ( ! [X30,X31] : (sP11('const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X31) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X30))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.ABS_1'(s('type.integer.int',X30)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.ABS_1'(s('type.integer.int',X30))))) )),
  inference(superposition,[],[f34485,f878])).
fof(f878,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31(X0)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30(X0))) )),
  inference(cnf_transformation,[],[f595])).
fof(f53501,plain,(
  ( ! [X488] : (s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X488))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.ABS_1'(s('type.integer.int',X488)))))))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X488))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.ABS_1'(s('type.integer.int',X488))))))) ) | $spl418),
  inference(cnf_transformation,[],[f53501_D])).
fof(f53501_D,plain,(
  ( ! [X488] : (s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X488))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.ABS_1'(s('type.integer.int',X488)))))))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X488))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.ABS_1'(s('type.integer.int',X488))))))) ) <=> ~$spl418),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl418])])).
fof(f62443,plain,(
  ( ! [X28,X29] : (sP11('const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X29) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X28))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.ABS_1'(s('type.integer.int',X28)))))))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X28))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.ABS_1'(s('type.integer.int',X28))))))) )),
  inference(superposition,[],[f34485,f877])).
fof(f877,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31(X0)))))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30(X0))))) )),
  inference(cnf_transformation,[],[f595])).
fof(f35353,plain,(
  ( ! [X11] : (s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X11))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X11)))))) ) | $spl332),
  inference(cnf_transformation,[],[f35353_D])).
fof(f35353_D,plain,(
  ( ! [X11] : (s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X11))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X11)))))) ) <=> ~$spl332),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl332])])).
fof(f62442,plain,(
  ( ! [X26,X27] : (sP11('const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X27) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X26))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X26)))))) )),
  inference(superposition,[],[f34485,f875])).
fof(f875,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))) )),
  inference(cnf_transformation,[],[f594])).
fof(f594,plain,(
  ! [X0] : ((s('type.integer.int',X0) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))),
  inference(nnf_transformation,[],[f384])).
fof(f384,plain,(
  ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) <=> s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(rectify,[],[f172])).
fof(f172,axiom,(
  ! [X6] : (s('type.integer.int',X6) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X6))) <=> s('type.integer.int',X6) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  file('thm.integer.INT_ABS_LE|split|3','thm.integer.INT_NEG_SAME_EQ')).
fof(f37557,plain,(
  ( ! [X458] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X458)))))) ) | $spl340),
  inference(cnf_transformation,[],[f37557_D])).
fof(f37557_D,plain,(
  ( ! [X458] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X458)))))) ) <=> ~$spl340),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl340])])).
fof(f62441,plain,(
  ( ! [X24,X25] : (sP11('const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X25) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X24)))))) )),
  inference(superposition,[],[f34485,f873])).
fof(f873,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))) )),
  inference(cnf_transformation,[],[f593])).
fof(f593,plain,(
  ! [X0] : ((s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))),
  inference(nnf_transformation,[],[f383])).
fof(f383,plain,(
  ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) <=> s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(rectify,[],[f145])).
fof(f145,axiom,(
  ! [X9] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9))) <=> s('type.integer.int',X9) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  file('thm.integer.INT_ABS_LE|split|3','thm.integer.INT_NEG_EQ|48|')).
fof(f44439,plain,(
  ( ! [X469] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X469))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl380),
  inference(cnf_transformation,[],[f44439_D])).
fof(f44439_D,plain,(
  ( ! [X469] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X469))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) <=> ~$spl380),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl380])])).
fof(f62455,plain,(
  ( ! [X17,X16] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X16))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | sP11('const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X17)) )),
  inference(forward_demodulation,[],[f62438,f892])).
fof(f62438,plain,(
  ( ! [X17,X16] : (sP11('const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X17) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X16))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f34485,f845])).
fof(f845,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(cnf_transformation,[],[f582])).
fof(f582,plain,(
  ! [X0] : ((~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))),
  inference(nnf_transformation,[],[f366])).
fof(f366,plain,(
  ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X6] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X6))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> s('type.integer.int',X6) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  file('thm.integer.INT_ABS_LE|split|3','thm.integer.INT_ABS_LE|48|')).
fof(f40369,plain,(
  ( ! [X49] : (s('type.integer.int',X49) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl370),
  inference(cnf_transformation,[],[f40369_D])).
fof(f40369_D,plain,(
  ( ! [X49] : (s('type.integer.int',X49) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl370),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl370])])).
fof(f62431,plain,(
  ( ! [X0,X1] : (sP11('const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X1) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f34485,f891])).
fof(f891,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f596])).
fof(f62325,plain,(
  ( ! [X1] : (sP11(X1,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl452),
  inference(cnf_transformation,[],[f62325_D])).
fof(f62325_D,plain,(
  ( ! [X1] : (sP11(X1,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) <=> ~$spl452),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl452])])).
fof(f62327,plain,(
  ( ! [X37,X36] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X36))) | sP11(X37,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(forward_demodulation,[],[f62318,f892])).
fof(f62318,plain,(
  ( ! [X37,X36] : (sP11(X37,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X36)))))) )),
  inference(superposition,[],[f34484,f890])).
fof(f34484,plain,(
  ( ! [X204,X203] : (sP11(X203,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X204)))))) )),
  inference(resolution,[],[f854,f728])).
fof(f728,plain,(
  ( ! [X0,X1] : (sP11(X0,X1) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f522])).
fof(f62317,plain,(
  ( ! [X35,X34] : (sP11(X35,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.ABS_1'(s('type.integer.int',X34)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.ABS_1'(s('type.integer.int',X34))))) )),
  inference(superposition,[],[f34484,f880])).
fof(f62316,plain,(
  ( ! [X33,X32] : (sP11(X33,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.ABS_1'(s('type.integer.int',X32)))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X32))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.ABS_1'(s('type.integer.int',X32))))))) )),
  inference(superposition,[],[f34484,f879])).
fof(f62315,plain,(
  ( ! [X30,X31] : (sP11(X31,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X30))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.ABS_1'(s('type.integer.int',X30)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.ABS_1'(s('type.integer.int',X30))))) )),
  inference(superposition,[],[f34484,f878])).
fof(f62314,plain,(
  ( ! [X28,X29] : (sP11(X29,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X28))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.ABS_1'(s('type.integer.int',X28)))))))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X28))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.ABS_1'(s('type.integer.int',X28))))))) )),
  inference(superposition,[],[f34484,f877])).
fof(f62313,plain,(
  ( ! [X26,X27] : (sP11(X27,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X26))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X26)))))) )),
  inference(superposition,[],[f34484,f875])).
fof(f62312,plain,(
  ( ! [X24,X25] : (sP11(X25,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X24)))))) )),
  inference(superposition,[],[f34484,f873])).
fof(f62326,plain,(
  ( ! [X17,X16] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X16))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | sP11(X17,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(forward_demodulation,[],[f62309,f892])).
fof(f62309,plain,(
  ( ! [X17,X16] : (sP11(X17,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X16))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f34484,f845])).
fof(f62302,plain,(
  ( ! [X0,X1] : (sP11(X1,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f34484,f891])).
fof(f2074,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl64),
  inference(cnf_transformation,[],[f2074_D])).
fof(f2074_D,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f2082,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl71),
  inference(cnf_transformation,[],[f2082_D])).
fof(f2082_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) <=> ~$spl71),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl71])])).
fof(f27128,plain,(
  ( ! [X10] : (sP11('const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X10)) ) | $spl176),
  inference(cnf_transformation,[],[f27128_D])).
fof(f27128_D,plain,(
  ( ! [X10] : (sP11('const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X10)) ) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f61859,plain,(
  ( ! [X19] : (sP11('const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X19) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))))) ) | $spl0),
  inference(superposition,[],[f1970,f879])).
fof(f1970,plain,(
  ( ! [X128] : (sP11('const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))),X128)) ) | $spl0),
  inference(resolution,[],[f905,f729])).
fof(f2076,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | $spl66),
  inference(cnf_transformation,[],[f2076_D])).
fof(f2076_D,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f61857,plain,(
  ( ! [X17] : (sP11('const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X17) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))))) ) | $spl0),
  inference(superposition,[],[f1970,f877])).
fof(f2071,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) | $spl63),
  inference(cnf_transformation,[],[f2071_D])).
fof(f2071_D,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) <=> ~$spl63),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl63])])).
fof(f61869,plain,(
  ( ! [X16] : (s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) | sP11('const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X16)) ) | $spl0),
  inference(forward_demodulation,[],[f61856,f886])).
fof(f886,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))) )),
  inference(cnf_transformation,[],[f392])).
fof(f392,plain,(
  ! [X0] : s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X9] : s('type.integer.int',X9) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9)))))),
  file('thm.integer.INT_ABS_LE|split|3','thm.integer.INT_NEGNEG')).
fof(f61856,plain,(
  ( ! [X16] : (sP11('const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X16) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) ) | $spl0),
  inference(superposition,[],[f1970,f875])).
fof(f2063,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl61),
  inference(cnf_transformation,[],[f2063_D])).
fof(f2063_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl61),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl61])])).
fof(f61865,plain,(
  ( ! [X10] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | sP11('const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X10)) ) | $spl0),
  inference(forward_demodulation,[],[f61864,f892])).
fof(f61864,plain,(
  ( ! [X10] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | sP11('const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X10)) ) | $spl0),
  inference(forward_demodulation,[],[f61852,f885])).
fof(f885,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))) )),
  inference(cnf_transformation,[],[f391])).
fof(f391,plain,(
  ! [X0] : s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X6] : s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X6))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X6)))))),
  file('thm.integer.INT_ABS_LE|split|3','thm.integer.INT_ABS_NEG')).
fof(f61852,plain,(
  ( ! [X10] : (sP11('const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X10) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl0),
  inference(superposition,[],[f1970,f845])).
fof(f26922,plain,(
  ( ! [X10] : (sP11(X10,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl174),
  inference(cnf_transformation,[],[f26922_D])).
fof(f26922_D,plain,(
  ( ! [X10] : (sP11(X10,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f61524,plain,(
  ( ! [X19] : (sP11(X19,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))))) ) | $spl0),
  inference(superposition,[],[f1969,f879])).
fof(f1969,plain,(
  ( ! [X127] : (sP11(X127,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) ) | $spl0),
  inference(resolution,[],[f905,f728])).
fof(f61522,plain,(
  ( ! [X17] : (sP11(X17,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))))) ) | $spl0),
  inference(superposition,[],[f1969,f877])).
fof(f61534,plain,(
  ( ! [X16] : (s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) | sP11(X16,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl0),
  inference(forward_demodulation,[],[f61521,f886])).
fof(f61521,plain,(
  ( ! [X16] : (sP11(X16,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) ) | $spl0),
  inference(superposition,[],[f1969,f875])).
fof(f61530,plain,(
  ( ! [X10] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | sP11(X10,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl0),
  inference(forward_demodulation,[],[f61529,f892])).
fof(f61529,plain,(
  ( ! [X10] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | sP11(X10,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl0),
  inference(forward_demodulation,[],[f61517,f885])).
fof(f61517,plain,(
  ( ! [X10] : (sP11(X10,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl0),
  inference(superposition,[],[f1969,f845])).
fof(f8367,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl123),
  inference(cnf_transformation,[],[f8367_D])).
fof(f8367_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) <=> ~$spl123),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl123])])).
fof(f2031,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl44),
  inference(cnf_transformation,[],[f2031_D])).
fof(f2031_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f60650,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl16),
  inference(resolution,[],[f5718,f861])).
fof(f861,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))))) )),
  inference(cnf_transformation,[],[f589])).
fof(f589,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))))))) & ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) & ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))))))),
  inference(flattening,[],[f588])).
fof(f588,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))) | (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))))) & ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) & ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))))))),
  inference(nnf_transformation,[],[f376])).
fof(f376,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))) <=> (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))))))))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X6,X7] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X6)))))) <=> (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X7),s('type.integer.int',X6)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X6),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X7))))))))),
  file('thm.integer.INT_ABS_LE|split|3','thm.integer.INT_ABS_LT|split|1')).
fof(f5718,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | $spl16),
  inference(resolution,[],[f1166,f708])).
fof(f708,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) )),
  inference(cnf_transformation,[],[f513])).
fof(f513,plain,(
  ! [X0,X1] : ((p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) & (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))),
  inference(nnf_transformation,[],[f284])).
fof(f284,plain,(
  ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) <=> p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(flattening,[],[f283])).
fof(f283,plain,(
  ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) <=> p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X9,X11] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X9),s('type.integer.int',X11)))) <=> p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X11),s('type.integer.int',X9)))))),
  file('thm.integer.INT_ABS_LE|split|3','thm.integer.INT_NOT_LT')).
fof(f9685,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl130),
  inference(cnf_transformation,[],[f9685_D])).
fof(f9685_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f2486,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int',sK13)))) | $spl75),
  inference(cnf_transformation,[],[f2486_D])).
fof(f2486_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int',sK13)))) <=> ~$spl75),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl75])])).
fof(f60769,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl16),
  inference(forward_demodulation,[],[f60648,f892])).
fof(f60648,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))),s('type.integer.int',sK13)))) | $spl16),
  inference(resolution,[],[f5718,f855])).
fof(f855,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f585])).
fof(f585,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) & ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0)))) & ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))),s('type.integer.int',X1))))))),
  inference(flattening,[],[f584])).
fof(f584,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))),s('type.integer.int',X1)))) | (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))) & ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0)))) & ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))),s('type.integer.int',X1))))))),
  inference(nnf_transformation,[],[f374])).
fof(f374,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))),s('type.integer.int',X1)))) <=> (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6,X7] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X6))))),s('type.integer.int',X7)))) <=> (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X7))),s('type.integer.int',X6)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X6),s('type.integer.int',X7))))))),
  file('thm.integer.INT_ABS_LE|split|3','thm.integer.INT_ABS_LT|split|2')).
fof(f60220,plain,(
  $false | ($spl88 | $spl90)),
  inference(subsumption_resolution,[],[f60219,f3218])).
fof(f60219,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl90),
  inference(forward_demodulation,[],[f60218,f887])).
fof(f887,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(cnf_transformation,[],[f129])).
fof(f129,axiom,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.integer.INT_ABS_LE|split|3','thm.integer.INT_NEG_|48|')).
fof(f60218,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl90),
  inference(subsumption_resolution,[],[f60217,f889])).
fof(f889,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(cnf_transformation,[],[f395])).
fof(f395,plain,(
  ! [X0] : ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(flattening,[],[f394])).
fof(f394,plain,(
  ! [X0] : ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X6] : ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X6))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  file('thm.integer.INT_ABS_LE|split|3','thm.integer.INT_ABS_LT|48|')).
fof(f60217,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl90),
  inference(forward_demodulation,[],[f60064,f883])).
fof(f883,plain,(
  ( ! [X0] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))) )),
  inference(cnf_transformation,[],[f389])).
fof(f389,plain,(
  ! [X0] : s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))),
  inference(rectify,[],[f55])).
fof(f55,axiom,(
  ! [X9] : s(bool,'const.integer.int_lt_2'(s('type.integer.int',X9),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9)))))),
  file('thm.integer.INT_ABS_LE|split|3','thm.integer.INT_NEG_GT|48|')).
fof(f60064,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl90),
  inference(resolution,[],[f3220,f860])).
fof(f860,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))))))) )),
  inference(cnf_transformation,[],[f587])).
fof(f587,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))))) | (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))) & (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))))))))),
  inference(flattening,[],[f586])).
fof(f586,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))))) | (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))) & ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))))))))),
  inference(nnf_transformation,[],[f375])).
fof(f375,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))))) <=> (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0))))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X6,X7] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X6)))))))) <=> (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X6),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X7)))))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X7),s('type.integer.int',X6))))))),
  file('thm.integer.INT_ABS_LE|split|3','thm.integer.INT_ABS_LT|split|3')).
fof(f60214,plain,(
  $false | ($spl88 | $spl90)),
  inference(subsumption_resolution,[],[f60062,f3218])).
fof(f60062,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl90),
  inference(resolution,[],[f3220,f766])).
fof(f766,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f428])).
fof(f428,plain,(
  ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(ennf_transformation,[],[f312])).
fof(f312,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) => ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(flattening,[],[f311])).
fof(f311,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) => ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(rectify,[],[f171])).
fof(f171,axiom,(
  ! [X9,X11] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X9),s('type.integer.int',X11)))) => ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X11),s('type.integer.int',X9)))))),
  file('thm.integer.INT_ABS_LE|split|3','thm.integer.INT_LT_GT')).
fof(f60213,plain,(
  $false | ($spl88 | $spl90)),
  inference(subsumption_resolution,[],[f60061,f3218])).
fof(f60061,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl90),
  inference(resolution,[],[f3220,f766])).
fof(f60212,plain,(
  $false | ($spl53 | $spl90)),
  inference(subsumption_resolution,[],[f60211,f2046])).
fof(f60211,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | $spl90),
  inference(forward_demodulation,[],[f60060,f852])).
fof(f852,plain,(
  ( ! [X0] : (s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int',X0)))) )),
  inference(cnf_transformation,[],[f371])).
fof(f371,plain,(
  ! [X0] : s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int',X0)))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X9] : s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X9))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9))),s('type.integer.int',X9)))),
  file('thm.integer.INT_ABS_LE|split|3','thm.integer.INT_LE_NEGL')).
fof(f60060,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | $spl90),
  inference(resolution,[],[f3220,f750])).
fof(f750,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f533])).
fof(f533,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s('type.integer.int',X0) = s('type.integer.int',X1)) & ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & s('type.integer.int',X0) != s('type.integer.int',X1)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))),
  inference(flattening,[],[f532])).
fof(f532,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s('type.integer.int',X0) = s('type.integer.int',X1))) & ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & s('type.integer.int',X0) != s('type.integer.int',X1)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))),
  inference(nnf_transformation,[],[f300])).
fof(f300,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) <=> (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s('type.integer.int',X0) = s('type.integer.int',X1)))),
  inference(rectify,[],[f50])).
fof(f50,axiom,(
  ! [X9,X11] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X9),s('type.integer.int',X11)))) <=> (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X9),s('type.integer.int',X11)))) | s('type.integer.int',X9) = s('type.integer.int',X11)))),
  file('thm.integer.INT_ABS_LE|split|3','thm.integer.INT_LE_LT')).
fof(f60210,plain,(
  $false | ($spl82 | $spl90)),
  inference(subsumption_resolution,[],[f60209,f3198])).
fof(f60209,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl90),
  inference(forward_demodulation,[],[f60059,f853])).
fof(f853,plain,(
  ( ! [X0] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))) )),
  inference(cnf_transformation,[],[f372])).
fof(f372,plain,(
  ! [X0] : s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X9] : s(bool,'const.integer.int_le_2'(s('type.integer.int',X9),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',X9),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9)))))),
  file('thm.integer.INT_ABS_LE|split|3','thm.integer.INT_LE_NEGR')).
fof(f60059,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl90),
  inference(resolution,[],[f3220,f747])).
fof(f747,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f531])).
fof(f531,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) & (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))),
  inference(nnf_transformation,[],[f299])).
fof(f299,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) <=> ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(flattening,[],[f298])).
fof(f298,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) <=> ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X9,X11] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X9),s('type.integer.int',X11)))) <=> ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X11),s('type.integer.int',X9)))))),
  file('thm.integer.INT_ABS_LE|split|3','thm.integer.int_le')).
fof(f60208,plain,(
  $false | ($spl82 | $spl90)),
  inference(subsumption_resolution,[],[f60207,f3198])).
fof(f60207,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl90),
  inference(forward_demodulation,[],[f60058,f853])).
fof(f60058,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl90),
  inference(resolution,[],[f3220,f710])).
fof(f710,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) )),
  inference(cnf_transformation,[],[f514])).
fof(f514,plain,(
  ! [X0,X1] : ((p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) & (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))),
  inference(nnf_transformation,[],[f286])).
fof(f286,plain,(
  ! [X0,X1] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) <=> p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(flattening,[],[f285])).
fof(f285,plain,(
  ! [X0,X1] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) <=> p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(rectify,[],[f144])).
fof(f144,axiom,(
  ! [X9,X11] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X9),s('type.integer.int',X11)))) <=> p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X11),s('type.integer.int',X9)))))),
  file('thm.integer.INT_ABS_LE|split|3','thm.integer.INT_NOT_LE')).
fof(f60206,plain,(
  $false | ($spl82 | $spl90)),
  inference(subsumption_resolution,[],[f60205,f3198])).
fof(f60205,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl90),
  inference(forward_demodulation,[],[f60057,f853])).
fof(f60057,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl90),
  inference(resolution,[],[f3220,f708])).
fof(f60204,plain,(
  $false | ($spl88 | $spl90)),
  inference(subsumption_resolution,[],[f60056,f3218])).
fof(f60056,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl90),
  inference(resolution,[],[f3220,f698])).
fof(f698,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f410])).
fof(f410,plain,(
  ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(ennf_transformation,[],[f278])).
fof(f278,plain,(
  ! [X0,X1] : ~(p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(rectify,[],[f212])).
fof(f212,axiom,(
  ! [X9,X11] : ~(p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X9),s('type.integer.int',X11)))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X11),s('type.integer.int',X9)))))),
  file('thm.integer.INT_ABS_LE|split|3','thm.integer.INT_LT_ANTISYM')).
fof(f60203,plain,(
  $false | ($spl88 | $spl90)),
  inference(subsumption_resolution,[],[f60055,f3218])).
fof(f60055,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl90),
  inference(resolution,[],[f3220,f698])).
fof(f59775,plain,(
  $false | ($spl82 | $spl89 | $spl93)),
  inference(subsumption_resolution,[],[f59774,f3198])).
fof(f59774,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | ($spl89 | $spl93)),
  inference(forward_demodulation,[],[f59773,f853])).
fof(f59773,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl89 | $spl93)),
  inference(subsumption_resolution,[],[f59619,f3221])).
fof(f59619,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl89),
  inference(resolution,[],[f3217,f749])).
fof(f749,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X0) = s('type.integer.int',X1) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f533])).
fof(f3217,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl89),
  inference(cnf_transformation,[],[f3217_D])).
fof(f3217_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl89),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl89])])).
fof(f59772,plain,(
  $false | ($spl53 | $spl89)),
  inference(subsumption_resolution,[],[f59771,f2046])).
fof(f59771,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | $spl89),
  inference(forward_demodulation,[],[f59618,f852])).
fof(f59618,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | $spl89),
  inference(resolution,[],[f3217,f748])).
fof(f748,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) )),
  inference(cnf_transformation,[],[f531])).
fof(f59770,plain,(
  $false | ($spl53 | $spl89)),
  inference(subsumption_resolution,[],[f59769,f2046])).
fof(f59769,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | $spl89),
  inference(forward_demodulation,[],[f59617,f852])).
fof(f59617,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | $spl89),
  inference(resolution,[],[f3217,f709])).
fof(f709,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f514])).
fof(f59768,plain,(
  $false | ($spl53 | $spl89)),
  inference(subsumption_resolution,[],[f59767,f2046])).
fof(f59767,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | $spl89),
  inference(forward_demodulation,[],[f59616,f852])).
fof(f59616,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | $spl89),
  inference(resolution,[],[f3217,f707])).
fof(f707,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f513])).
fof(f7883,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl117),
  inference(cnf_transformation,[],[f7883_D])).
fof(f7883_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) <=> ~$spl117),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl117])])).
fof(f1301,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl20),
  inference(cnf_transformation,[],[f1301_D])).
fof(f1301_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f59353,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl81),
  inference(resolution,[],[f2893,f861])).
fof(f2893,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) | $spl81),
  inference(cnf_transformation,[],[f2893_D])).
fof(f2893_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) <=> ~$spl81),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl81])])).
fof(f10405,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl136),
  inference(cnf_transformation,[],[f10405_D])).
fof(f10405_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f1593,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK12)))) | $spl35),
  inference(cnf_transformation,[],[f1593_D])).
fof(f1593_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK12)))) <=> ~$spl35),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl35])])).
fof(f59472,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK12)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl81),
  inference(forward_demodulation,[],[f59351,f892])).
fof(f59351,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))),s('type.integer.int',sK12)))) | $spl81),
  inference(resolution,[],[f2893,f855])).
fof(f8368,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl122),
  inference(cnf_transformation,[],[f8368_D])).
fof(f8368_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f59227,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK12)))) | $spl76),
  inference(resolution,[],[f2487,f866])).
fof(f866,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f591])).
fof(f591,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int',X1)))) | (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0)))))) & (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int',X1))))))),
  inference(flattening,[],[f590])).
fof(f590,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int',X1)))) | (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0)))))) & ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int',X1))))))),
  inference(nnf_transformation,[],[f377])).
fof(f377,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int',X1)))) <=> (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0))))))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X6,X7] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X6))),s('type.integer.int',X7)))) <=> (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X6),s('type.integer.int',X7)))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X7))),s('type.integer.int',X6))))))),
  file('thm.integer.INT_ABS_LE|split|3','thm.integer.INT_ABS_LT|split|0')).
fof(f8365,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl121),
  inference(cnf_transformation,[],[f8365_D])).
fof(f8365_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) <=> ~$spl121),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl121])])).
fof(f59339,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl76),
  inference(forward_demodulation,[],[f59225,f892])).
fof(f59225,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl76),
  inference(resolution,[],[f2487,f860])).
fof(f8210,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl118),
  inference(cnf_transformation,[],[f8210_D])).
fof(f8210_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f4362,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int',sK13)))) | $spl97),
  inference(cnf_transformation,[],[f4362_D])).
fof(f4362_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int',sK13)))) <=> ~$spl97),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl97])])).
fof(f59181,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl43),
  inference(forward_demodulation,[],[f59061,f892])).
fof(f59061,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))),s('type.integer.int',sK13)))) | $spl43),
  inference(resolution,[],[f1723,f833])).
fof(f833,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f577])).
fof(f577,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) & ((~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0)))) & ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))),s('type.integer.int',X1))))))),
  inference(flattening,[],[f576])).
fof(f576,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))),s('type.integer.int',X1)))) | (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))) & ((~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0)))) & ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))),s('type.integer.int',X1))))))),
  inference(nnf_transformation,[],[f360])).
fof(f360,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))),s('type.integer.int',X1)))) <=> (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X6,X7] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X6))))),s('type.integer.int',X7)))) <=> (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X7))),s('type.integer.int',X6)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X6),s('type.integer.int',X7))))))),
  file('thm.integer.INT_ABS_LE|split|3','thm.integer.INT_ABS_LE|split|2')).
fof(f7884,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl116),
  inference(cnf_transformation,[],[f7884_D])).
fof(f7884_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f58786,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int',sK13)))) | $spl38),
  inference(resolution,[],[f1596,f866])).
fof(f7881,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl115),
  inference(cnf_transformation,[],[f7881_D])).
fof(f7881_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) <=> ~$spl115),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl115])])).
fof(f58898,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl38),
  inference(forward_demodulation,[],[f58784,f892])).
fof(f58784,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl38),
  inference(resolution,[],[f1596,f860])).
fof(f58219,plain,(
  $false | $spl354),
  inference(equality_resolution,[],[f37602])).
fof(f25591,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl152),
  inference(cnf_transformation,[],[f25591_D])).
fof(f25591_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f25593,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int',sK12)))) | $spl155),
  inference(cnf_transformation,[],[f25593_D])).
fof(f25593_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int',sK12)))) <=> ~$spl155),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl155])])).
fof(f57646,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int',sK12)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl237),
  inference(forward_demodulation,[],[f57528,f892])).
fof(f57528,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))),s('type.integer.int',sK12)))) | $spl237),
  inference(resolution,[],[f29755,f833])).
fof(f2039,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl48),
  inference(cnf_transformation,[],[f2039_D])).
fof(f2039_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f1723,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | $spl43),
  inference(cnf_transformation,[],[f1723_D])).
fof(f1723_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) <=> ~$spl43),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl43])])).
fof(f57483,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(forward_demodulation,[],[f57352,f886])).
fof(f57352,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))),s('type.integer.int',sK13)))) | $spl0),
  inference(resolution,[],[f905,f841])).
fof(f841,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f581])).
fof(f2033,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) | $spl46),
  inference(cnf_transformation,[],[f2033_D])).
fof(f2033_D,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f57349,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(resolution,[],[f905,f749])).
fof(f57034,plain,(
  $false | ($spl0 | $spl15 | $spl234)),
  inference(subsumption_resolution,[],[f56733,f10215])).
fof(f10215,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | $spl15),
  inference(resolution,[],[f1165,f835])).
fof(f835,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f577])).
fof(f56733,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | ($spl0 | $spl234)),
  inference(backward_demodulation,[],[f29749,f905])).
fof(f26615,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl172),
  inference(cnf_transformation,[],[f26615_D])).
fof(f26615_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f56691,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) | $spl450),
  inference(cnf_transformation,[],[f56691_D])).
fof(f56691_D,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) <=> ~$spl450),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl450])])).
fof(f56573,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl56),
  inference(resolution,[],[f2051,f749])).
fof(f1165,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK12)))) | $spl15),
  inference(cnf_transformation,[],[f1165_D])).
fof(f1165_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK12)))) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f2661,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl78),
  inference(cnf_transformation,[],[f2661_D])).
fof(f2661_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f56143,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK12)))) | $spl2),
  inference(resolution,[],[f1456,f841])).
fof(f1456,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl2),
  inference(forward_demodulation,[],[f1334,f885])).
fof(f1334,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))))) | $spl2),
  inference(resolution,[],[f907,f837])).
fof(f837,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) )),
  inference(cnf_transformation,[],[f579])).
fof(f2097,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl72),
  inference(cnf_transformation,[],[f2097_D])).
fof(f2097_D,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f29332,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl226),
  inference(cnf_transformation,[],[f29332_D])).
fof(f29332_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) <=> ~$spl226),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl226])])).
fof(f55667,plain,(
  ( ! [X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | p(s(bool,X1))) ) | $spl53),
  inference(subsumption_resolution,[],[f55656,f790])).
fof(f790,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f558])).
fof(f558,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f557])).
fof(f557,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f328])).
fof(f328,plain,(
  ! [X0] : ((p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X12] : ((p(s(bool,'F_0')) | p(s(bool,X12))) <=> p(s(bool,X12)))),
  file('thm.integer.INT_ABS_LE|split|3','thm.bool.OR_CLAUSES|split|2')).
fof(f55656,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | p(s(bool,X1))) ) | $spl53),
  inference(resolution,[],[f25689,f727])).
fof(f727,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP11(X0,X1)) )),
  inference(cnf_transformation,[],[f522])).
fof(f25689,plain,(
  ( ! [X128] : (sP11('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))),X128)) ) | $spl53),
  inference(resolution,[],[f6930,f729])).
fof(f6930,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl53),
  inference(resolution,[],[f2046,f707])).
fof(f55666,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | p(s(bool,X0))) ) | $spl53),
  inference(subsumption_resolution,[],[f55655,f790])).
fof(f55655,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | p(s(bool,X0))) ) | $spl53),
  inference(resolution,[],[f25689,f731])).
fof(f731,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP11(X0,X1)) )),
  inference(cnf_transformation,[],[f524])).
fof(f524,plain,(
  ! [X0,X1] : ((~sP11(X0,X1) | (p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP11(X0,X1)))),
  inference(flattening,[],[f523])).
fof(f523,plain,(
  ! [X0,X1] : ((~sP11(X0,X1) | ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0')))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP11(X0,X1)))),
  inference(nnf_transformation,[],[f461])).
fof(f461,plain,(
  ! [X0,X1] : (sP11(X0,X1) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(definition_folding,[],[f416,f460])).
fof(f416,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f415])).
fof(f415,plain,(
  ! [X0,X1] : (((p(s(bool,X1)) | p(s(bool,X0))) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | (p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(ennf_transformation,[],[f292])).
fof(f292,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(flattening,[],[f291])).
fof(f291,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(rectify,[],[f62])).
fof(f62,axiom,(
  ! [X21,X8] : ((~(p(s(bool,X8)) | p(s(bool,X21))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X8)) => p(s(bool,'F_0'))) => (~p(s(bool,X21)) => p(s(bool,'F_0')))))),
  file('thm.integer.INT_ABS_LE|split|3','thm.sat.sat_9_prove_ltb')).
fof(f55652,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))))))) ) | $spl53),
  inference(subsumption_resolution,[],[f55641,f790])).
fof(f55641,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))))))) ) | $spl53),
  inference(resolution,[],[f25688,f727])).
fof(f25688,plain,(
  ( ! [X127] : (sP11(X127,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) ) | $spl53),
  inference(resolution,[],[f6930,f728])).
fof(f55651,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))))))) ) | $spl53),
  inference(subsumption_resolution,[],[f55640,f790])).
fof(f55640,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))))))) ) | $spl53),
  inference(resolution,[],[f25688,f731])).
fof(f28630,plain,(
  ( ! [X0] : (sP11('const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)),X0)) ) | $spl190),
  inference(cnf_transformation,[],[f28630_D])).
fof(f28630_D,plain,(
  ( ! [X0] : (sP11('const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)),X0)) ) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f2046,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | $spl53),
  inference(cnf_transformation,[],[f2046_D])).
fof(f2046_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) <=> ~$spl53),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl53])])).
fof(f55617,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | sP11('const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)),X0)) ) | $spl38),
  inference(forward_demodulation,[],[f55612,f852])).
fof(f55612,plain,(
  ( ! [X0] : (sP11('const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)),X0) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12))))) ) | $spl38),
  inference(superposition,[],[f7836,f848])).
fof(f848,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))))) )),
  inference(cnf_transformation,[],[f583])).
fof(f583,plain,(
  ! [X0] : ((s('type.integer.int',X0) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))))) & (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))),
  inference(nnf_transformation,[],[f367])).
fof(f367,plain,(
  ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))) <=> p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X6] : (s('type.integer.int',X6) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X6))) <=> p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X6)))))),
  file('thm.integer.INT_ABS_LE|split|3','thm.integer.INT_ABS_EQ_ID')).
fof(f7836,plain,(
  ( ! [X128] : (sP11('const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))),X128)) ) | $spl38),
  inference(resolution,[],[f1596,f729])).
fof(f5840,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl104),
  inference(cnf_transformation,[],[f5840_D])).
fof(f5840_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f55486,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int',sK13)))) | $spl16),
  inference(resolution,[],[f1166,f841])).
fof(f28629,plain,(
  ( ! [X0] : (sP11(X0,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) | $spl188),
  inference(cnf_transformation,[],[f28629_D])).
fof(f28629_D,plain,(
  ( ! [X0] : (sP11(X0,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f55477,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | sP11(X0,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) | $spl38),
  inference(forward_demodulation,[],[f55472,f852])).
fof(f55472,plain,(
  ( ! [X0] : (sP11(X0,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12))))) ) | $spl38),
  inference(superposition,[],[f7835,f848])).
fof(f7835,plain,(
  ( ! [X127] : (sP11(X127,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) ) | $spl38),
  inference(resolution,[],[f1596,f728])).
fof(f3198,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl82),
  inference(cnf_transformation,[],[f3198_D])).
fof(f3198_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f55429,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))))))) ) | $spl222),
  inference(subsumption_resolution,[],[f55418,f790])).
fof(f55418,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))))))) ) | $spl222),
  inference(resolution,[],[f29328,f727])).
fof(f55428,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))))))) ) | $spl222),
  inference(subsumption_resolution,[],[f55417,f790])).
fof(f55417,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))))))) ) | $spl222),
  inference(resolution,[],[f29328,f731])).
fof(f55414,plain,(
  ( ! [X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | p(s(bool,X1))) ) | $spl220),
  inference(subsumption_resolution,[],[f55403,f790])).
fof(f55403,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | p(s(bool,X1))) ) | $spl220),
  inference(resolution,[],[f29326,f727])).
fof(f55413,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | p(s(bool,X0))) ) | $spl220),
  inference(subsumption_resolution,[],[f55402,f790])).
fof(f55402,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | p(s(bool,X0))) ) | $spl220),
  inference(resolution,[],[f29326,f731])).
fof(f5839,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl105),
  inference(cnf_transformation,[],[f5839_D])).
fof(f5839_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) <=> ~$spl105),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl105])])).
fof(f9338,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl128),
  inference(cnf_transformation,[],[f9338_D])).
fof(f9338_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f55274,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl13),
  inference(resolution,[],[f1033,f836])).
fof(f836,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))))) )),
  inference(cnf_transformation,[],[f579])).
fof(f1033,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) | $spl13),
  inference(cnf_transformation,[],[f1033_D])).
fof(f1033_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f4987,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl100),
  inference(cnf_transformation,[],[f4987_D])).
fof(f4987_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f55393,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK12)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl13),
  inference(forward_demodulation,[],[f55273,f892])).
fof(f55273,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))),s('type.integer.int',sK12)))) | $spl13),
  inference(resolution,[],[f1033,f833])).
fof(f4363,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl98),
  inference(cnf_transformation,[],[f4363_D])).
fof(f4363_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f55230,plain,(
  ( ! [X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | p(s(bool,X1))) ) | $spl2),
  inference(subsumption_resolution,[],[f55221,f790])).
fof(f55221,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | p(s(bool,X1))) ) | $spl2),
  inference(resolution,[],[f6859,f727])).
fof(f6859,plain,(
  ( ! [X128] : (sP11('const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))),X128)) ) | $spl2),
  inference(resolution,[],[f1456,f729])).
fof(f55229,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | p(s(bool,X0))) ) | $spl2),
  inference(subsumption_resolution,[],[f55220,f790])).
fof(f55220,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | p(s(bool,X0))) ) | $spl2),
  inference(resolution,[],[f6859,f731])).
fof(f55219,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))) ) | $spl2),
  inference(subsumption_resolution,[],[f55210,f790])).
fof(f55210,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))) ) | $spl2),
  inference(resolution,[],[f6858,f727])).
fof(f6858,plain,(
  ( ! [X127] : (sP11(X127,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) ) | $spl2),
  inference(resolution,[],[f1456,f728])).
fof(f55218,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))) ) | $spl2),
  inference(subsumption_resolution,[],[f55209,f790])).
fof(f55209,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))) ) | $spl2),
  inference(resolution,[],[f6858,f731])).
fof(f13583,plain,(
  ( ! [X0] : (sP11('const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)),X0)) ) | $spl144),
  inference(cnf_transformation,[],[f13583_D])).
fof(f13583_D,plain,(
  ( ! [X0] : (sP11('const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)),X0)) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f55027,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | sP11('const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)),X0)) ) | $spl16),
  inference(forward_demodulation,[],[f55022,f852])).
fof(f55022,plain,(
  ( ! [X0] : (sP11('const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)),X0) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12))))) ) | $spl16),
  inference(superposition,[],[f5793,f848])).
fof(f5793,plain,(
  ( ! [X128] : (sP11('const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))),X128)) ) | $spl16),
  inference(resolution,[],[f1166,f729])).
fof(f13585,plain,(
  ( ! [X0] : (sP11(X0,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) | $spl146),
  inference(cnf_transformation,[],[f13585_D])).
fof(f13585_D,plain,(
  ( ! [X0] : (sP11(X0,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f55015,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | sP11(X0,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) | $spl16),
  inference(forward_demodulation,[],[f55010,f852])).
fof(f55010,plain,(
  ( ! [X0] : (sP11(X0,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12))))) ) | $spl16),
  inference(superposition,[],[f5792,f848])).
fof(f5792,plain,(
  ( ! [X127] : (sP11(X127,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) ) | $spl16),
  inference(resolution,[],[f1166,f728])).
fof(f54293,plain,(
  $false | ($spl86 | $spl95)),
  inference(subsumption_resolution,[],[f54292,f34268])).
fof(f34268,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))) )),
  inference(forward_demodulation,[],[f827,f844])).
fof(f844,plain,(
  ( ! [X0,X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f365])).
fof(f365,plain,(
  ! [X0,X1] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f178])).
fof(f178,axiom,(
  ! [X20,X10] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X20),s('type.num.num',X10))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X20))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X10)))))),
  file('thm.integer.INT_ABS_LE|split|3','thm.integer.INT_LE')).
fof(f827,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))))) )),
  inference(cnf_transformation,[],[f355])).
fof(f355,plain,(
  ! [X0] : p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))))),
  inference(rectify,[],[f158])).
fof(f158,axiom,(
  ! [X10] : p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X10))))))),
  file('thm.integer.INT_ABS_LE|split|3','thm.integer.INT_POS')).
fof(f54292,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK28(sK13))))) | ($spl86 | $spl95)),
  inference(forward_demodulation,[],[f54291,f844])).
fof(f54291,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK28(sK13))))))) | ($spl86 | $spl95)),
  inference(subsumption_resolution,[],[f53987,f3291])).
fof(f53987,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK28(sK13))))))) | $spl86),
  inference(superposition,[],[f848,f3205])).
fof(f54148,plain,(
  $false | ($spl86 | $spl95)),
  inference(subsumption_resolution,[],[f53867,f3291])).
fof(f53867,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) | $spl86),
  inference(superposition,[],[f893,f3205])).
fof(f893,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f398])).
fof(f398,plain,(
  ! [X0] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X10] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X10))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X10)))))),
  file('thm.integer.INT_ABS_LE|split|3','thm.integer.INT_ABS_NUM')).
fof(f54142,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK28(sK13)) | $spl440),
  inference(cnf_transformation,[],[f54142_D])).
fof(f54142_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK28(sK13)) <=> ~$spl440),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl440])])).
fof(f54147,plain,(
  ( ! [X5] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))))),s('type.integer.int',sK12))))) ) | $spl448),
  inference(cnf_transformation,[],[f54147_D])).
fof(f54147_D,plain,(
  ( ! [X5] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))))),s('type.integer.int',sK12))))) ) <=> ~$spl448),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl448])])).
fof(f53860,plain,(
  ( ! [X5] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X5))))),s('type.integer.int',sK12)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK28(sK13))) ) | $spl86),
  inference(superposition,[],[f746,f3205])).
fof(f746,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f530])).
fof(f530,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)) & ((s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))))))))),
  inference(flattening,[],[f529])).
fof(f529,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))) | (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1))) & ((s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))))))))),
  inference(nnf_transformation,[],[f297])).
fof(f297,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))) <=> (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)))),
  inference(flattening,[],[f296])).
fof(f296,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))) <=> (~s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | ~s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)))),
  inference(rectify,[],[f134])).
fof(f134,axiom,(
  ! [X10,X20] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X10))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X20)))))) <=> (~s('type.num.num','const.num.|48|_0') = s('type.num.num',X10) | ~s('type.num.num','const.num.|48|_0') = s('type.num.num',X20)))),
  file('thm.integer.INT_ABS_LE|split|3','thm.integer.INT_LT_CALCULATE|split|2')).
fof(f54146,plain,(
  ( ! [X2] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X2))))))) ) | $spl446),
  inference(cnf_transformation,[],[f54146_D])).
fof(f54146_D,plain,(
  ( ! [X2] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X2))))))) ) <=> ~$spl446),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl446])])).
fof(f53857,plain,(
  ( ! [X2] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X2)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK28(sK13))) ) | $spl86),
  inference(superposition,[],[f745,f3205])).
fof(f745,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f530])).
fof(f54143,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK28(sK13)) | $spl441),
  inference(cnf_transformation,[],[f54143_D])).
fof(f54143_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK28(sK13)) <=> ~$spl441),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl441])])).
fof(f54145,plain,(
  ( ! [X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))))),s('type.integer.int',sK12)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)) ) | $spl444),
  inference(cnf_transformation,[],[f54145_D])).
fof(f54145_D,plain,(
  ( ! [X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))))),s('type.integer.int',sK12)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)) ) <=> ~$spl444),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl444])])).
fof(f53856,plain,(
  ( ! [X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))))),s('type.integer.int',sK12)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK28(sK13)) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)) ) | $spl86),
  inference(superposition,[],[f744,f3205])).
fof(f744,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X1) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))))))) )),
  inference(cnf_transformation,[],[f530])).
fof(f54144,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0)) ) | $spl442),
  inference(cnf_transformation,[],[f54144_D])).
fof(f54144_D,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0)) ) <=> ~$spl442),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl442])])).
fof(f53855,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK28(sK13))) ) | $spl86),
  inference(superposition,[],[f744,f3205])).
fof(f54132,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_of_num_1'(s('type.num.num',sK28(sK13))))) | $spl437),
  inference(cnf_transformation,[],[f54132_D])).
fof(f54132_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_of_num_1'(s('type.num.num',sK28(sK13))))) <=> ~$spl437),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl437])])).
fof(f54136,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_of_num_1'(s('type.num.num',sK28(sK13))))) | $spl439),
  inference(cnf_transformation,[],[f54136_D])).
fof(f54136_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_of_num_1'(s('type.num.num',sK28(sK13))))) <=> ~$spl439),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl439])])).
fof(f54137,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_of_num_1'(s('type.num.num',sK28(sK13))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_of_num_1'(s('type.num.num',sK28(sK13))))) | ($spl86 | $spl93)),
  inference(subsumption_resolution,[],[f53852,f3221])).
fof(f53852,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_of_num_1'(s('type.num.num',sK28(sK13))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_of_num_1'(s('type.num.num',sK28(sK13))))) | $spl86),
  inference(superposition,[],[f3205,f880])).
fof(f54128,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_of_num_1'(s('type.num.num',sK28(sK13))))))) | $spl434),
  inference(cnf_transformation,[],[f54128_D])).
fof(f54128_D,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_of_num_1'(s('type.num.num',sK28(sK13))))))) <=> ~$spl434),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl434])])).
fof(f54134,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_of_num_1'(s('type.num.num',sK28(sK13))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_of_num_1'(s('type.num.num',sK28(sK13))))) | ($spl86 | $spl93)),
  inference(forward_demodulation,[],[f54133,f3205])).
fof(f54133,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_of_num_1'(s('type.num.num',sK28(sK13))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK28(sK13)))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_of_num_1'(s('type.num.num',sK28(sK13))))))) | ($spl86 | $spl93)),
  inference(subsumption_resolution,[],[f53851,f3221])).
fof(f53851,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_of_num_1'(s('type.num.num',sK28(sK13))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK28(sK13)))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_of_num_1'(s('type.num.num',sK28(sK13))))))) | $spl86),
  inference(superposition,[],[f3205,f879])).
fof(f54126,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_of_num_1'(s('type.num.num',sK28(sK13))))))))) | $spl432),
  inference(cnf_transformation,[],[f54126_D])).
fof(f54126_D,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_of_num_1'(s('type.num.num',sK28(sK13))))))))) <=> ~$spl432),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl432])])).
fof(f54130,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_of_num_1'(s('type.num.num',sK28(sK13))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_of_num_1'(s('type.num.num',sK28(sK13))))) | ($spl86 | $spl93)),
  inference(forward_demodulation,[],[f54129,f3205])).
fof(f54129,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK28(sK13)))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_of_num_1'(s('type.num.num',sK28(sK13))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_of_num_1'(s('type.num.num',sK28(sK13))))) | ($spl86 | $spl93)),
  inference(subsumption_resolution,[],[f53850,f3221])).
fof(f53850,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK28(sK13)))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_of_num_1'(s('type.num.num',sK28(sK13))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_of_num_1'(s('type.num.num',sK28(sK13))))) | $spl86),
  inference(superposition,[],[f3205,f878])).
fof(f54124,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_of_num_1'(s('type.num.num',sK28(sK13))))))) | s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_of_num_1'(s('type.num.num',sK28(sK13))))))))) | ($spl86 | $spl93)),
  inference(forward_demodulation,[],[f54123,f3205])).
fof(f54123,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_of_num_1'(s('type.num.num',sK28(sK13))))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK28(sK13)))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_of_num_1'(s('type.num.num',sK28(sK13))))))) | ($spl86 | $spl93)),
  inference(forward_demodulation,[],[f54122,f3205])).
fof(f54122,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK28(sK13)))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_of_num_1'(s('type.num.num',sK28(sK13))))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK28(sK13)))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_of_num_1'(s('type.num.num',sK28(sK13))))))) | ($spl86 | $spl93)),
  inference(subsumption_resolution,[],[f53849,f3221])).
fof(f53849,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK28(sK13)))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_of_num_1'(s('type.num.num',sK28(sK13))))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK28(sK13)))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_of_num_1'(s('type.num.num',sK28(sK13))))))) | $spl86),
  inference(superposition,[],[f3205,f877])).
fof(f53542,plain,(
  ( ! [X219] : (s('type.integer.int',X219) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31(X219)))))) | s('type.integer.int',X219) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30(X219))))) ) | $spl430),
  inference(cnf_transformation,[],[f53542_D])).
fof(f53542_D,plain,(
  ( ! [X219] : (s('type.integer.int',X219) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31(X219)))))) | s('type.integer.int',X219) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30(X219))))) ) <=> ~$spl430),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl430])])).
fof(f42134,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK28('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl374),
  inference(cnf_transformation,[],[f42134_D])).
fof(f42134_D,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK28('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl374),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl374])])).
fof(f53543,plain,(
  ( ! [X220] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK28('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X220) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31(X220)))))) | s('type.integer.int',X220) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30(X220))))) )),
  inference(subsumption_resolution,[],[f52627,f828])).
fof(f828,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X0))))) )),
  inference(cnf_transformation,[],[f356])).
fof(f356,plain,(
  ! [X0] : p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X0))))),
  inference(rectify,[],[f211])).
fof(f211,axiom,(
  ! [X9] : p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X9),s('type.integer.int',X9))))),
  file('thm.integer.INT_ABS_LE|split|3','thm.integer.INT_LE_REFL')).
fof(f52627,plain,(
  ( ! [X220] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X220),s('type.integer.int',X220)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK28('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X220) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31(X220)))))) | s('type.integer.int',X220) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30(X220))))) )),
  inference(superposition,[],[f821,f877])).
fof(f821,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK28(X0)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))))) )),
  inference(cnf_transformation,[],[f572])).
fof(f572,plain,(
  ! [X0] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK28(X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK28])],[f440])).
fof(f440,plain,(
  ! [X0] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) | ? [X1] : s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))))),
  inference(ennf_transformation,[],[f347])).
fof(f347,plain,(
  ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) => ? [X1] : s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))))),
  inference(rectify,[],[f181])).
fof(f181,axiom,(
  ! [X16] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X16)))) => ? [X10] : s('type.integer.int',X16) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X10))))),
  file('thm.integer.INT_ABS_LE|split|3','thm.integer.NUM_POSINT_EXISTS')).
fof(f45640,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl388),
  inference(cnf_transformation,[],[f45640_D])).
fof(f45640_D,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) <=> ~$spl388),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl388])])).
fof(f53541,plain,(
  ( ! [X219] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | s('type.integer.int',X219) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31(X219)))))) | s('type.integer.int',X219) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30(X219))))) )),
  inference(subsumption_resolution,[],[f52626,f828])).
fof(f52626,plain,(
  ( ! [X219] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X219),s('type.integer.int',X219)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | s('type.integer.int',X219) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31(X219)))))) | s('type.integer.int',X219) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30(X219))))) )),
  inference(superposition,[],[f820,f877])).
fof(f820,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27(X0)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(cnf_transformation,[],[f571])).
fof(f571,plain,(
  ! [X0] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27(X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK27])],[f439])).
fof(f439,plain,(
  ! [X0] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ? [X1] : s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))))))),
  inference(ennf_transformation,[],[f346])).
fof(f346,plain,(
  ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) => ? [X1] : s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))))))),
  inference(rectify,[],[f130])).
fof(f130,axiom,(
  ! [X16] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X16),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) => ? [X10] : s('type.integer.int',X16) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X10))))))),
  file('thm.integer.INT_ABS_LE|split|3','thm.integer.NUM_NEGINT_EXISTS')).
fof(f53523,plain,(
  ( ! [X596] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X596) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X596))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_of_num_1'(s('type.num.num',X596)))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X596))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_of_num_1'(s('type.num.num',X596))))))))) ) | $spl428),
  inference(cnf_transformation,[],[f53523_D])).
fof(f53523_D,plain,(
  ( ! [X596] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X596) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X596))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_of_num_1'(s('type.num.num',X596)))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X596))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_of_num_1'(s('type.num.num',X596))))))))) ) <=> ~$spl428),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl428])])).
fof(f37608,plain,(
  ( ! [X552] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X552))))))) ) | $spl362),
  inference(cnf_transformation,[],[f37608_D])).
fof(f37608_D,plain,(
  ( ! [X552] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X552))))))) ) <=> ~$spl362),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl362])])).
fof(f53524,plain,(
  ( ! [X602,X603] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X603)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X602) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X602))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_of_num_1'(s('type.num.num',X602)))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X602))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_of_num_1'(s('type.num.num',X602))))))) )),
  inference(forward_demodulation,[],[f52440,f882])).
fof(f882,plain,(
  ( ! [X0] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(cnf_transformation,[],[f388])).
fof(f388,plain,(
  ! [X0] : s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X9] : s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X9))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.integer.INT_ABS_LE|split|3','thm.integer.INT_NEG_LT|48|')).
fof(f52440,plain,(
  ( ! [X602,X603] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X603))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X602) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X602))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_of_num_1'(s('type.num.num',X602)))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X602))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_of_num_1'(s('type.num.num',X602))))))) )),
  inference(superposition,[],[f746,f877])).
fof(f53522,plain,(
  ( ! [X596,X597] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X597)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X596) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X596))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_of_num_1'(s('type.num.num',X596)))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X596))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_of_num_1'(s('type.num.num',X596))))))) )),
  inference(forward_demodulation,[],[f52437,f887])).
fof(f52437,plain,(
  ( ! [X596,X597] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X597)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X596) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X596))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_of_num_1'(s('type.num.num',X596)))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X596))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_of_num_1'(s('type.num.num',X596))))))) )),
  inference(superposition,[],[f745,f877])).
fof(f53520,plain,(
  ( ! [X592] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X592) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X592))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_of_num_1'(s('type.num.num',X592)))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X592))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_of_num_1'(s('type.num.num',X592))))))))) ) | $spl426),
  inference(cnf_transformation,[],[f53520_D])).
fof(f53520_D,plain,(
  ( ! [X592] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X592) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X592))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_of_num_1'(s('type.num.num',X592)))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X592))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_of_num_1'(s('type.num.num',X592))))))))) ) <=> ~$spl426),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl426])])).
fof(f37598,plain,(
  ( ! [X542] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X542)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X542)) ) | $spl350),
  inference(cnf_transformation,[],[f37598_D])).
fof(f37598_D,plain,(
  ( ! [X542] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X542)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X542)) ) <=> ~$spl350),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl350])])).
fof(f53521,plain,(
  ( ! [X594,X595] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X595)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X594) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X595) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X594))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_of_num_1'(s('type.num.num',X594)))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X594))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_of_num_1'(s('type.num.num',X594))))))) )),
  inference(forward_demodulation,[],[f52436,f882])).
fof(f52436,plain,(
  ( ! [X594,X595] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X595))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X594) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X595) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X594))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_of_num_1'(s('type.num.num',X594)))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X594))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_of_num_1'(s('type.num.num',X594))))))) )),
  inference(superposition,[],[f744,f877])).
fof(f53519,plain,(
  ( ! [X592,X593] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X593)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X593) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X592) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X592))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_of_num_1'(s('type.num.num',X592)))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X592))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_of_num_1'(s('type.num.num',X592))))))) )),
  inference(forward_demodulation,[],[f52435,f887])).
fof(f52435,plain,(
  ( ! [X592,X593] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X593)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X593) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X592) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X592))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_of_num_1'(s('type.num.num',X592)))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X592))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_of_num_1'(s('type.num.num',X592))))))) )),
  inference(superposition,[],[f744,f877])).
fof(f53515,plain,(
  ( ! [X578] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X578))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X578)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X578))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X578))))))))) ) | $spl420),
  inference(cnf_transformation,[],[f53515_D])).
fof(f53515_D,plain,(
  ( ! [X578] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X578))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X578)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X578))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X578))))))))) ) <=> ~$spl420),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl420])])).
fof(f37601,plain,(
  ( ! [X546] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X546)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X546)) ) | $spl352),
  inference(cnf_transformation,[],[f37601_D])).
fof(f37601_D,plain,(
  ( ! [X546] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X546)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X546)) ) <=> ~$spl352),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl352])])).
fof(f52432,plain,(
  ( ! [X588,X589] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X589)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X589) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X588))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X588)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X588))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X588))))))))) )),
  inference(superposition,[],[f746,f877])).
fof(f53518,plain,(
  ( ! [X586] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X586) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X586))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X586)))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X586))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X586))))))))))) ) | $spl424),
  inference(cnf_transformation,[],[f53518_D])).
fof(f53518_D,plain,(
  ( ! [X586] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X586) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X586))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X586)))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X586))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X586))))))))))) ) <=> ~$spl424),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl424])])).
fof(f52431,plain,(
  ( ! [X586,X587] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X587)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X586) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X586))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X586)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X586))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X586))))))))) )),
  inference(superposition,[],[f745,f877])).
fof(f53517,plain,(
  ( ! [X584] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X584) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X584))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X584)))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X584))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X584))))))))))) ) | $spl422),
  inference(cnf_transformation,[],[f53517_D])).
fof(f53517_D,plain,(
  ( ! [X584] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X584) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X584))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X584)))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X584))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X584))))))))))) ) <=> ~$spl422),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl422])])).
fof(f52430,plain,(
  ( ! [X584,X585] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X585)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X585) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X584) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X584))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X584)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X584))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X584))))))))) )),
  inference(superposition,[],[f744,f877])).
fof(f37593,plain,(
  ( ! [X536] : (s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X536))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl346),
  inference(cnf_transformation,[],[f37593_D])).
fof(f37593_D,plain,(
  ( ! [X536] : (s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X536))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) <=> ~$spl346),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl346])])).
fof(f52427,plain,(
  ( ! [X578,X579] : (s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X579))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X578))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X578)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X578))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X578))))))))) )),
  inference(superposition,[],[f872,f877])).
fof(f872,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))))) )),
  inference(cnf_transformation,[],[f382])).
fof(f382,plain,(
  ! [X0,X1] : s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))))),
  inference(rectify,[],[f133])).
fof(f133,axiom,(
  ! [X10,X20] : s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X10))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X20)))))))),
  file('thm.integer.INT_ABS_LE|split|3','thm.integer.INT_LT_CALCULATE|split|3')).
fof(f2051,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl56),
  inference(cnf_transformation,[],[f2051_D])).
fof(f2051_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f53512,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(forward_demodulation,[],[f52425,f853])).
fof(f52425,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(superposition,[],[f905,f877])).
fof(f35355,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl335),
  inference(cnf_transformation,[],[f35355_D])).
fof(f35355_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl335),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl335])])).
fof(f52374,plain,(
  ( ! [X488] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X488))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.ABS_1'(s('type.integer.int',X488)))))))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X488))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.ABS_1'(s('type.integer.int',X488))))))) )),
  inference(superposition,[],[f889,f877])).
fof(f51760,plain,(
  ( ! [X217] : (s('type.integer.int',X217) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31(X217)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30(X217))) ) | $spl416),
  inference(cnf_transformation,[],[f51760_D])).
fof(f51760_D,plain,(
  ( ! [X217] : (s('type.integer.int',X217) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31(X217)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30(X217))) ) <=> ~$spl416),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl416])])).
fof(f51761,plain,(
  ( ! [X218] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK28('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X218) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31(X218)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30(X218))) )),
  inference(subsumption_resolution,[],[f51134,f828])).
fof(f51134,plain,(
  ( ! [X218] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X218),s('type.integer.int',X218)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK28('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X218) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31(X218)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30(X218))) )),
  inference(superposition,[],[f821,f878])).
fof(f51759,plain,(
  ( ! [X217] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | s('type.integer.int',X217) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31(X217)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30(X217))) )),
  inference(subsumption_resolution,[],[f51133,f828])).
fof(f51133,plain,(
  ( ! [X217] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X217),s('type.integer.int',X217)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | s('type.integer.int',X217) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31(X217)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30(X217))) )),
  inference(superposition,[],[f820,f878])).
fof(f51747,plain,(
  ( ! [X594] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X594) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_of_num_1'(s('type.num.num',X594)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X594))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_of_num_1'(s('type.num.num',X594))))))))) ) | $spl414),
  inference(cnf_transformation,[],[f51747_D])).
fof(f51747_D,plain,(
  ( ! [X594] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X594) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_of_num_1'(s('type.num.num',X594)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X594))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_of_num_1'(s('type.num.num',X594))))))))) ) <=> ~$spl414),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl414])])).
fof(f51748,plain,(
  ( ! [X600,X601] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X601)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X600) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X600))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_of_num_1'(s('type.num.num',X600)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_of_num_1'(s('type.num.num',X600))))) )),
  inference(forward_demodulation,[],[f50950,f882])).
fof(f50950,plain,(
  ( ! [X600,X601] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X601))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X600) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X600))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_of_num_1'(s('type.num.num',X600)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_of_num_1'(s('type.num.num',X600))))) )),
  inference(superposition,[],[f746,f878])).
fof(f51746,plain,(
  ( ! [X594,X595] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X595)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X594) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X594))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_of_num_1'(s('type.num.num',X594)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_of_num_1'(s('type.num.num',X594))))) )),
  inference(forward_demodulation,[],[f50947,f887])).
fof(f50947,plain,(
  ( ! [X594,X595] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X595)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X594) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X594))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_of_num_1'(s('type.num.num',X594)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_of_num_1'(s('type.num.num',X594))))) )),
  inference(superposition,[],[f745,f878])).
fof(f51744,plain,(
  ( ! [X590] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X590) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_of_num_1'(s('type.num.num',X590)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X590))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_of_num_1'(s('type.num.num',X590))))))))) ) | $spl412),
  inference(cnf_transformation,[],[f51744_D])).
fof(f51744_D,plain,(
  ( ! [X590] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X590) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_of_num_1'(s('type.num.num',X590)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X590))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_of_num_1'(s('type.num.num',X590))))))))) ) <=> ~$spl412),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl412])])).
fof(f51745,plain,(
  ( ! [X592,X593] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X593)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X592) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X593) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X592))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_of_num_1'(s('type.num.num',X592)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_of_num_1'(s('type.num.num',X592))))) )),
  inference(forward_demodulation,[],[f50946,f882])).
fof(f50946,plain,(
  ( ! [X592,X593] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X593))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X592) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X593) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X592))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_of_num_1'(s('type.num.num',X592)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_of_num_1'(s('type.num.num',X592))))) )),
  inference(superposition,[],[f744,f878])).
fof(f51743,plain,(
  ( ! [X590,X591] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X591)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X591) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X590) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X590))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_of_num_1'(s('type.num.num',X590)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_of_num_1'(s('type.num.num',X590))))) )),
  inference(forward_demodulation,[],[f50945,f887])).
fof(f50945,plain,(
  ( ! [X590,X591] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X591)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X591) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X590) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X590))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_of_num_1'(s('type.num.num',X590)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_of_num_1'(s('type.num.num',X590))))) )),
  inference(superposition,[],[f744,f878])).
fof(f51739,plain,(
  ( ! [X577] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X577))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X577)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X577))))))) ) | $spl406),
  inference(cnf_transformation,[],[f51739_D])).
fof(f51739_D,plain,(
  ( ! [X577] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X577))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X577)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X577))))))) ) <=> ~$spl406),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl406])])).
fof(f50943,plain,(
  ( ! [X588,X587] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X588)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X588) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X587))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X587)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X587))))))) )),
  inference(superposition,[],[f746,f878])).
fof(f51742,plain,(
  ( ! [X585] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X585) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X585)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X585))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X585))))))))))) ) | $spl410),
  inference(cnf_transformation,[],[f51742_D])).
fof(f51742_D,plain,(
  ( ! [X585] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X585) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X585)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X585))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X585))))))))))) ) <=> ~$spl410),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl410])])).
fof(f50942,plain,(
  ( ! [X586,X585] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X586)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X585) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X585))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X585)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X585))))))) )),
  inference(superposition,[],[f745,f878])).
fof(f51741,plain,(
  ( ! [X583] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X583) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X583)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X583))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X583))))))))))) ) | $spl408),
  inference(cnf_transformation,[],[f51741_D])).
fof(f51741_D,plain,(
  ( ! [X583] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X583) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X583)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X583))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X583))))))))))) ) <=> ~$spl408),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl408])])).
fof(f50941,plain,(
  ( ! [X584,X583] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X584)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X584) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X583) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X583))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X583)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X583))))))) )),
  inference(superposition,[],[f744,f878])).
fof(f50938,plain,(
  ( ! [X578,X577] : (s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X578))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X577))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X577)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X577))))))) )),
  inference(superposition,[],[f872,f878])).
fof(f2079,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl69),
  inference(cnf_transformation,[],[f2079_D])).
fof(f2079_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) <=> ~$spl69),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl69])])).
fof(f51737,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl0),
  inference(forward_demodulation,[],[f50936,f853])).
fof(f50936,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl0),
  inference(superposition,[],[f905,f878])).
fof(f50885,plain,(
  ( ! [X487] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X487))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.ABS_1'(s('type.integer.int',X487)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.ABS_1'(s('type.integer.int',X487))))) )),
  inference(superposition,[],[f889,f878])).
fof(f26613,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl171),
  inference(cnf_transformation,[],[f26613_D])).
fof(f26613_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) <=> ~$spl171),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl171])])).
fof(f48797,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl167),
  inference(resolution,[],[f26476,f861])).
fof(f28632,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK13)))) | $spl193),
  inference(cnf_transformation,[],[f28632_D])).
fof(f28632_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK13)))) <=> ~$spl193),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl193])])).
fof(f28634,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | $spl194),
  inference(cnf_transformation,[],[f28634_D])).
fof(f28634_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f48795,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK13)))) | $spl167),
  inference(resolution,[],[f26476,f855])).
fof(f48789,plain,(
  $false | ($spl7 | $spl80)),
  inference(subsumption_resolution,[],[f24428,f1027])).
fof(f24428,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | $spl80),
  inference(resolution,[],[f2894,f864])).
fof(f864,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f591])).
fof(f48635,plain,(
  ( ! [X208] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31(X208)) | s('type.integer.int',X208) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30(X208))))) ) | $spl402),
  inference(cnf_transformation,[],[f48635_D])).
fof(f48635_D,plain,(
  ( ! [X208] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31(X208)) | s('type.integer.int',X208) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30(X208))))) ) <=> ~$spl402),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl402])])).
fof(f48636,plain,(
  ( ! [X209] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK28('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31(X209)) | s('type.integer.int',X209) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30(X209))))) )),
  inference(subsumption_resolution,[],[f48080,f828])).
fof(f48080,plain,(
  ( ! [X209] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X209),s('type.integer.int',X209)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK28('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31(X209)) | s('type.integer.int',X209) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30(X209))))) )),
  inference(superposition,[],[f821,f879])).
fof(f48634,plain,(
  ( ! [X208] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31(X208)) | s('type.integer.int',X208) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30(X208))))) )),
  inference(subsumption_resolution,[],[f48079,f828])).
fof(f48079,plain,(
  ( ! [X208] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X208),s('type.integer.int',X208)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31(X208)) | s('type.integer.int',X208) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30(X208))))) )),
  inference(superposition,[],[f820,f879])).
fof(f48625,plain,(
  ( ! [X584] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X584) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X584))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_of_num_1'(s('type.num.num',X584)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_of_num_1'(s('type.num.num',X584))))) ) | $spl400),
  inference(cnf_transformation,[],[f48625_D])).
fof(f48625_D,plain,(
  ( ! [X584] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X584) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X584))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_of_num_1'(s('type.num.num',X584)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_of_num_1'(s('type.num.num',X584))))) ) <=> ~$spl400),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl400])])).
fof(f48626,plain,(
  ( ! [X590,X591] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X591)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X590) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_of_num_1'(s('type.num.num',X590)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X590))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_of_num_1'(s('type.num.num',X590))))))) )),
  inference(forward_demodulation,[],[f47908,f882])).
fof(f47908,plain,(
  ( ! [X590,X591] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X591))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X590) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_of_num_1'(s('type.num.num',X590)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X590))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_of_num_1'(s('type.num.num',X590))))))) )),
  inference(superposition,[],[f746,f879])).
fof(f48624,plain,(
  ( ! [X584,X585] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X585)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X584) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_of_num_1'(s('type.num.num',X584)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X584))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_of_num_1'(s('type.num.num',X584))))))) )),
  inference(forward_demodulation,[],[f47905,f887])).
fof(f47905,plain,(
  ( ! [X584,X585] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X585)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X584) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_of_num_1'(s('type.num.num',X584)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X584))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_of_num_1'(s('type.num.num',X584))))))) )),
  inference(superposition,[],[f745,f879])).
fof(f48622,plain,(
  ( ! [X580] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X580) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X580))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_of_num_1'(s('type.num.num',X580)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_of_num_1'(s('type.num.num',X580))))) ) | $spl398),
  inference(cnf_transformation,[],[f48622_D])).
fof(f48622_D,plain,(
  ( ! [X580] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X580) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X580))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_of_num_1'(s('type.num.num',X580)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_of_num_1'(s('type.num.num',X580))))) ) <=> ~$spl398),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl398])])).
fof(f48623,plain,(
  ( ! [X582,X583] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X583)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X582) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X583) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_of_num_1'(s('type.num.num',X582)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X582))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_of_num_1'(s('type.num.num',X582))))))) )),
  inference(forward_demodulation,[],[f47904,f882])).
fof(f47904,plain,(
  ( ! [X582,X583] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X583))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X582) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X583) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_of_num_1'(s('type.num.num',X582)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X582))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_of_num_1'(s('type.num.num',X582))))))) )),
  inference(superposition,[],[f744,f879])).
fof(f48621,plain,(
  ( ! [X580,X581] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X581)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X581) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X580) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_of_num_1'(s('type.num.num',X580)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X580))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_of_num_1'(s('type.num.num',X580))))))) )),
  inference(forward_demodulation,[],[f47903,f887])).
fof(f47903,plain,(
  ( ! [X580,X581] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X581)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X581) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X580) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_of_num_1'(s('type.num.num',X580)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X580))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_of_num_1'(s('type.num.num',X580))))))) )),
  inference(superposition,[],[f744,f879])).
fof(f48617,plain,(
  ( ! [X567] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X567)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X567))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X567))))))))) ) | $spl392),
  inference(cnf_transformation,[],[f48617_D])).
fof(f48617_D,plain,(
  ( ! [X567] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X567)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X567))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X567))))))))) ) <=> ~$spl392),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl392])])).
fof(f47901,plain,(
  ( ! [X578,X577] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X578)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X578) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X577)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X577))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X577))))))))) )),
  inference(superposition,[],[f746,f879])).
fof(f48620,plain,(
  ( ! [X575] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X575) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X575))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X575)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X575))))))) ) | $spl396),
  inference(cnf_transformation,[],[f48620_D])).
fof(f48620_D,plain,(
  ( ! [X575] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X575) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X575))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X575)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X575))))))) ) <=> ~$spl396),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl396])])).
fof(f47900,plain,(
  ( ! [X575,X576] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X576)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X575) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X575)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X575))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X575))))))))) )),
  inference(superposition,[],[f745,f879])).
fof(f48619,plain,(
  ( ! [X573] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X573) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X573))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X573)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X573))))))) ) | $spl394),
  inference(cnf_transformation,[],[f48619_D])).
fof(f48619_D,plain,(
  ( ! [X573] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X573) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X573))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X573)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X573))))))) ) <=> ~$spl394),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl394])])).
fof(f47899,plain,(
  ( ! [X573,X574] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X574)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X574) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X573) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X573)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X573))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X573))))))))) )),
  inference(superposition,[],[f744,f879])).
fof(f47896,plain,(
  ( ! [X567,X568] : (s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X568))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X567)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X567))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X567))))))))) )),
  inference(superposition,[],[f872,f879])).
fof(f48615,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(forward_demodulation,[],[f47894,f853])).
fof(f47894,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(superposition,[],[f905,f879])).
fof(f47845,plain,(
  ( ! [X480] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.ABS_1'(s('type.integer.int',X480)))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X480))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.ABS_1'(s('type.integer.int',X480))))))) )),
  inference(superposition,[],[f889,f879])).
fof(f47533,plain,(
  $false | ($spl11 | $spl140)),
  inference(subsumption_resolution,[],[f10675,f4692])).
fof(f4692,plain,(
  s('type.integer.int',sK13) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))) | $spl11),
  inference(resolution,[],[f1032,f751])).
fof(f751,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s('type.integer.int',X0) != s('type.integer.int',X1)) )),
  inference(cnf_transformation,[],[f533])).
fof(f45646,plain,(
  ( ! [X36] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | s('type.integer.int',X36) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(subsumption_resolution,[],[f45277,f828])).
fof(f45277,plain,(
  ( ! [X36] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X36))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X36)))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | s('type.integer.int',X36) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f820,f891])).
fof(f45645,plain,(
  ( ! [X35] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X35)))) )),
  inference(subsumption_resolution,[],[f45276,f828])).
fof(f45276,plain,(
  ( ! [X35] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X35),s('type.integer.int',X35)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X35)))) )),
  inference(superposition,[],[f820,f890])).
fof(f42139,plain,(
  ( ! [X31] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31(X31)) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30(X31))) ) | $spl378),
  inference(cnf_transformation,[],[f42139_D])).
fof(f42139_D,plain,(
  ( ! [X31] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31(X31)) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30(X31))) ) <=> ~$spl378),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl378])])).
fof(f45644,plain,(
  ( ! [X34] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31(X34)) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30(X34))) )),
  inference(subsumption_resolution,[],[f45275,f828])).
fof(f45275,plain,(
  ( ! [X34] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X34),s('type.integer.int',X34)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31(X34)) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30(X34))) )),
  inference(superposition,[],[f820,f880])).
fof(f42137,plain,(
  ( ! [X30] : (s('type.integer.int',X30) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X30)))) ) | $spl376),
  inference(cnf_transformation,[],[f42137_D])).
fof(f42137_D,plain,(
  ( ! [X30] : (s('type.integer.int',X30) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X30)))) ) <=> ~$spl376),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl376])])).
fof(f45643,plain,(
  ( ! [X33] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | s('type.integer.int',X33) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X33)))) )),
  inference(subsumption_resolution,[],[f45274,f828])).
fof(f45274,plain,(
  ( ! [X33] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X33),s('type.integer.int',X33)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | s('type.integer.int',X33) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X33)))) )),
  inference(superposition,[],[f820,f875])).
fof(f45642,plain,(
  ( ! [X32] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(subsumption_resolution,[],[f45273,f828])).
fof(f45273,plain,(
  ( ! [X32] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X32))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X32)))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f820,f874])).
fof(f874,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f593])).
fof(f42132,plain,(
  ( ! [X28] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X28)))) ) | $spl372),
  inference(cnf_transformation,[],[f42132_D])).
fof(f42132_D,plain,(
  ( ! [X28] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X28)))) ) <=> ~$spl372),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl372])])).
fof(f45641,plain,(
  ( ! [X31] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X31)))) )),
  inference(subsumption_resolution,[],[f45272,f828])).
fof(f45272,plain,(
  ( ! [X31] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X31),s('type.integer.int',X31)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X31)))) )),
  inference(superposition,[],[f820,f873])).
fof(f45638,plain,(
  ( ! [X30] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X30))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(subsumption_resolution,[],[f45271,f828])).
fof(f45271,plain,(
  ( ! [X30] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X30),s('type.integer.int',X30)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X30))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f820,f845])).
fof(f44629,plain,(
  $false | $spl382),
  inference(resolution,[],[f44501,f34268])).
fof(f44557,plain,(
  ( ! [X190] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK28('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X190))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(subsumption_resolution,[],[f44122,f828])).
fof(f44122,plain,(
  ( ! [X190] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X190),s('type.integer.int',X190)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK28('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X190))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f821,f845])).
fof(f37609,plain,(
  ( ! [X558] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X558))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl364),
  inference(cnf_transformation,[],[f37609_D])).
fof(f37609_D,plain,(
  ( ! [X558] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X558))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) <=> ~$spl364),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl364])])).
fof(f44516,plain,(
  ( ! [X560] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X560),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X560)) ) | $spl386),
  inference(cnf_transformation,[],[f44516_D])).
fof(f44516_D,plain,(
  ( ! [X560] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X560),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X560)) ) <=> ~$spl386),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl386])])).
fof(f44529,plain,(
  ( ! [X575,X574] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X574),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X575))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X574)) )),
  inference(forward_demodulation,[],[f44528,f844])).
fof(f44528,plain,(
  ( ! [X575,X574] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X574))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X575))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X574)) )),
  inference(forward_demodulation,[],[f43970,f893])).
fof(f43970,plain,(
  ( ! [X575,X574] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X575))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X574) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X574))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f746,f845])).
fof(f44527,plain,(
  ( ! [X569,X568] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X568),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X569)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X568)) )),
  inference(forward_demodulation,[],[f44526,f844])).
fof(f44526,plain,(
  ( ! [X569,X568] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X568))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X569)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X568)) )),
  inference(forward_demodulation,[],[f44525,f893])).
fof(f44525,plain,(
  ( ! [X569,X568] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X569)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X568) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X568))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(forward_demodulation,[],[f43967,f887])).
fof(f43967,plain,(
  ( ! [X569,X568] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X569)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X568) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X568))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f745,f845])).
fof(f37605,plain,(
  ( ! [X550] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X550))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X550)) ) | $spl358),
  inference(cnf_transformation,[],[f37605_D])).
fof(f37605_D,plain,(
  ( ! [X550] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X550))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X550)) ) <=> ~$spl358),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl358])])).
fof(f44512,plain,(
  ( ! [X558] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X558),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X558)) ) | $spl384),
  inference(cnf_transformation,[],[f44512_D])).
fof(f44512_D,plain,(
  ( ! [X558] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X558),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X558)) ) <=> ~$spl384),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl384])])).
fof(f44524,plain,(
  ( ! [X567,X566] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X566),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X567))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X566) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X567)) )),
  inference(forward_demodulation,[],[f44523,f844])).
fof(f44523,plain,(
  ( ! [X567,X566] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X566))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X567))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X566) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X567)) )),
  inference(forward_demodulation,[],[f43966,f893])).
fof(f43966,plain,(
  ( ! [X567,X566] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X567))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X566) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X567) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X566))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f744,f845])).
fof(f44522,plain,(
  ( ! [X565,X564] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X564),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X565)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X565) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X564)) )),
  inference(forward_demodulation,[],[f44521,f844])).
fof(f44521,plain,(
  ( ! [X565,X564] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X564))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X565)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X565) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X564)) )),
  inference(forward_demodulation,[],[f44520,f893])).
fof(f44520,plain,(
  ( ! [X565,X564] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X565)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X565) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X564) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X564))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(forward_demodulation,[],[f43965,f887])).
fof(f43965,plain,(
  ( ! [X565,X564] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X565)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X565) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X564) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X564))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f744,f845])).
fof(f44501,plain,(
  ( ! [X552] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X552),s('type.num.num','const.num.|48|_0'))))) ) | $spl382),
  inference(cnf_transformation,[],[f44501_D])).
fof(f44501_D,plain,(
  ( ! [X552] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X552),s('type.num.num','const.num.|48|_0'))))) ) <=> ~$spl382),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl382])])).
fof(f44519,plain,(
  ( ! [X563,X562] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X562),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X563)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X563)) )),
  inference(forward_demodulation,[],[f44518,f844])).
fof(f44518,plain,(
  ( ! [X563,X562] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X562))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X563)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X563)) )),
  inference(forward_demodulation,[],[f44517,f893])).
fof(f44517,plain,(
  ( ! [X563,X562] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X562))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X563)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X563)) )),
  inference(forward_demodulation,[],[f43964,f885])).
fof(f43964,plain,(
  ( ! [X563,X562] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X563)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X563) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X562))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f746,f845])).
fof(f44515,plain,(
  ( ! [X561,X560] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X560),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X561)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X560)) )),
  inference(forward_demodulation,[],[f44514,f844])).
fof(f44514,plain,(
  ( ! [X561,X560] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X560))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X561)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X560)) )),
  inference(forward_demodulation,[],[f44513,f893])).
fof(f44513,plain,(
  ( ! [X561,X560] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X560))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X561)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X560)) )),
  inference(forward_demodulation,[],[f43963,f885])).
fof(f43963,plain,(
  ( ! [X561,X560] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X561)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X560) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X560))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f745,f845])).
fof(f44511,plain,(
  ( ! [X559,X558] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X558),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X559)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X559) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X558)) )),
  inference(forward_demodulation,[],[f44510,f844])).
fof(f44510,plain,(
  ( ! [X559,X558] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X558))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X559)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X559) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X558)) )),
  inference(forward_demodulation,[],[f44509,f893])).
fof(f44509,plain,(
  ( ! [X559,X558] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X558))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X559)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X559) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X558)) )),
  inference(forward_demodulation,[],[f43962,f885])).
fof(f43962,plain,(
  ( ! [X559,X558] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X559)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X559) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X558) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X558))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f744,f845])).
fof(f44500,plain,(
  ( ! [X553,X552] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X552),s('type.num.num','const.num.|48|_0')))) | s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X553))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(forward_demodulation,[],[f44499,f844])).
fof(f44499,plain,(
  ( ! [X553,X552] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X552))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X553))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(forward_demodulation,[],[f44498,f893])).
fof(f44498,plain,(
  ( ! [X553,X552] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X552))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X553))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(forward_demodulation,[],[f43959,f885])).
fof(f43959,plain,(
  ( ! [X553,X552] : (s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X553))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X552))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f872,f845])).
fof(f44495,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl0),
  inference(forward_demodulation,[],[f44494,f892])).
fof(f44494,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl0),
  inference(forward_demodulation,[],[f44493,f885])).
fof(f44493,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl0),
  inference(forward_demodulation,[],[f43957,f853])).
fof(f43957,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl0),
  inference(superposition,[],[f905,f845])).
fof(f44438,plain,(
  ( ! [X469] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X469))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(forward_demodulation,[],[f43912,f892])).
fof(f43912,plain,(
  ( ! [X469] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X469))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f889,f845])).
fof(f42141,plain,(
  ( ! [X33] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK28('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X33) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(subsumption_resolution,[],[f41838,f828])).
fof(f41838,plain,(
  ( ! [X33] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X33))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X33)))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK28('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X33) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f821,f891])).
fof(f42140,plain,(
  ( ! [X32] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK28('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X32)))) )),
  inference(subsumption_resolution,[],[f41837,f828])).
fof(f41837,plain,(
  ( ! [X32] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X32),s('type.integer.int',X32)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK28('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X32)))) )),
  inference(superposition,[],[f821,f890])).
fof(f42138,plain,(
  ( ! [X31] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK28('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31(X31)) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30(X31))) )),
  inference(subsumption_resolution,[],[f41836,f828])).
fof(f41836,plain,(
  ( ! [X31] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X31),s('type.integer.int',X31)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK28('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31(X31)) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30(X31))) )),
  inference(superposition,[],[f821,f880])).
fof(f42136,plain,(
  ( ! [X30] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK28('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X30) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X30)))) )),
  inference(subsumption_resolution,[],[f41835,f828])).
fof(f41835,plain,(
  ( ! [X30] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X30),s('type.integer.int',X30)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK28('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X30) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X30)))) )),
  inference(superposition,[],[f821,f875])).
fof(f42135,plain,(
  ( ! [X29] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK28('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X29) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(subsumption_resolution,[],[f41834,f828])).
fof(f41834,plain,(
  ( ! [X29] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X29))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X29)))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK28('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X29) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f821,f874])).
fof(f42131,plain,(
  ( ! [X28] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK28('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X28)))) )),
  inference(subsumption_resolution,[],[f41833,f828])).
fof(f41833,plain,(
  ( ! [X28] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X28),s('type.integer.int',X28)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK28('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X28)))) )),
  inference(superposition,[],[f821,f873])).
fof(f37590,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl342),
  inference(cnf_transformation,[],[f37590_D])).
fof(f37590_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl342),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl342])])).
fof(f2048,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl55),
  inference(cnf_transformation,[],[f2048_D])).
fof(f2048_D,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl55),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl55])])).
fof(f41667,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl0),
  inference(forward_demodulation,[],[f41666,f892])).
fof(f41666,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl0),
  inference(forward_demodulation,[],[f41635,f885])).
fof(f41635,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))) | $spl0),
  inference(superposition,[],[f905,f890])).
fof(f41634,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl0),
  inference(superposition,[],[f905,f880])).
fof(f41665,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl0),
  inference(forward_demodulation,[],[f41633,f886])).
fof(f41633,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))) | $spl0),
  inference(superposition,[],[f905,f875])).
fof(f41664,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl0),
  inference(forward_demodulation,[],[f41632,f886])).
fof(f41632,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))) | $spl0),
  inference(superposition,[],[f905,f873])).
fof(f1166,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl16),
  inference(cnf_transformation,[],[f1166_D])).
fof(f1166_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f41625,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl0),
  inference(superposition,[],[f905,f876])).
fof(f876,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f594])).
fof(f41624,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl0),
  inference(superposition,[],[f905,f874])).
fof(f36190,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | $spl339),
  inference(cnf_transformation,[],[f36190_D])).
fof(f36190_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) <=> ~$spl339),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl339])])).
fof(f2044,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl50),
  inference(cnf_transformation,[],[f2044_D])).
fof(f2044_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f41621,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | $spl0),
  inference(superposition,[],[f905,f848])).
fof(f41659,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(forward_demodulation,[],[f41532,f886])).
fof(f41532,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))),s('type.integer.int',sK13)))) | $spl0),
  inference(resolution,[],[f905,f841])).
fof(f41529,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(resolution,[],[f905,f749])).
fof(f10400,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl134),
  inference(cnf_transformation,[],[f10400_D])).
fof(f10400_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f910,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | $spl5),
  inference(cnf_transformation,[],[f910_D])).
fof(f910_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f41383,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl2),
  inference(forward_demodulation,[],[f41264,f886])).
fof(f41264,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK12)))) | $spl2),
  inference(resolution,[],[f907,f841])).
fof(f41217,plain,(
  $false | ($spl7 | $spl28 | $spl38 | $spl157)),
  inference(subsumption_resolution,[],[f41216,f25735])).
fof(f41216,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)))) | ($spl7 | $spl28 | $spl38)),
  inference(forward_demodulation,[],[f13892,f1454])).
fof(f13892,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | ($spl7 | $spl38)),
  inference(subsumption_resolution,[],[f13776,f1596])).
fof(f13776,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl7),
  inference(resolution,[],[f1027,f861])).
fof(f41215,plain,(
  $false | ($spl28 | $spl41)),
  inference(subsumption_resolution,[],[f41214,f828])).
fof(f41214,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)))) | ($spl28 | $spl41)),
  inference(forward_demodulation,[],[f1722,f1454])).
fof(f41212,plain,(
  $false | ($spl28 | $spl41 | $spl157)),
  inference(subsumption_resolution,[],[f41211,f25735])).
fof(f41211,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)))) | ($spl28 | $spl41)),
  inference(forward_demodulation,[],[f7926,f1454])).
fof(f7926,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl41),
  inference(resolution,[],[f1722,f707])).
fof(f41210,plain,(
  $false | ($spl28 | $spl41 | $spl157)),
  inference(subsumption_resolution,[],[f41209,f25735])).
fof(f41209,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)))) | ($spl28 | $spl41)),
  inference(forward_demodulation,[],[f7927,f1454])).
fof(f7927,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl41),
  inference(resolution,[],[f1722,f709])).
fof(f41208,plain,(
  $false | ($spl28 | $spl41 | $spl157)),
  inference(subsumption_resolution,[],[f41207,f25735])).
fof(f41207,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)))) | ($spl28 | $spl41)),
  inference(forward_demodulation,[],[f7928,f1454])).
fof(f7928,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl41),
  inference(resolution,[],[f1722,f748])).
fof(f41205,plain,(
  $false | ($spl28 | $spl41)),
  inference(subsumption_resolution,[],[f7930,f1454])).
fof(f7930,plain,(
  s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) | $spl41),
  inference(resolution,[],[f1722,f751])).
fof(f41192,plain,(
  $false | ($spl28 | $spl41 | $spl157)),
  inference(subsumption_resolution,[],[f41191,f25735])).
fof(f41191,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)))) | ($spl28 | $spl41)),
  inference(forward_demodulation,[],[f21908,f1454])).
fof(f21908,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl41),
  inference(resolution,[],[f1722,f707])).
fof(f41190,plain,(
  $false | ($spl28 | $spl41 | $spl157)),
  inference(subsumption_resolution,[],[f41189,f25735])).
fof(f41189,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)))) | ($spl28 | $spl41)),
  inference(forward_demodulation,[],[f21909,f1454])).
fof(f21909,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl41),
  inference(resolution,[],[f1722,f709])).
fof(f41188,plain,(
  $false | ($spl28 | $spl41 | $spl157)),
  inference(subsumption_resolution,[],[f41187,f25735])).
fof(f41187,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)))) | ($spl28 | $spl41)),
  inference(forward_demodulation,[],[f21910,f1454])).
fof(f21910,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl41),
  inference(resolution,[],[f1722,f748])).
fof(f41185,plain,(
  $false | ($spl28 | $spl41)),
  inference(subsumption_resolution,[],[f21912,f1454])).
fof(f21912,plain,(
  s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) | $spl41),
  inference(resolution,[],[f1722,f751])).
fof(f41175,plain,(
  $false | ($spl5 | $spl28 | $spl43)),
  inference(subsumption_resolution,[],[f41174,f828])).
fof(f41174,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)))) | ($spl5 | $spl28 | $spl43)),
  inference(forward_demodulation,[],[f14141,f1454])).
fof(f14141,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) | ($spl5 | $spl43)),
  inference(subsumption_resolution,[],[f14032,f1723])).
fof(f41173,plain,(
  $false | ($spl0 | $spl5 | $spl28 | $spl43)),
  inference(subsumption_resolution,[],[f41172,f1045])).
fof(f1045,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | $spl5),
  inference(resolution,[],[f910,f767])).
fof(f767,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f313])).
fof(f313,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(rectify,[],[f146])).
fof(f146,axiom,(
  ! [X9,X11] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X9),s('type.integer.int',X11)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X11),s('type.integer.int',X9)))))),
  file('thm.integer.INT_ABS_LE|split|3','thm.integer.INT_LE_TOTAL')).
fof(f41172,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | ($spl0 | $spl28 | $spl43)),
  inference(forward_demodulation,[],[f28609,f1454])).
fof(f28609,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | ($spl0 | $spl43)),
  inference(forward_demodulation,[],[f28608,f852])).
fof(f28608,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | ($spl0 | $spl43)),
  inference(subsumption_resolution,[],[f28577,f1723])).
fof(f28577,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | $spl0),
  inference(superposition,[],[f1893,f848])).
fof(f41136,plain,(
  $false | ($spl0 | $spl7 | $spl9 | $spl28 | $spl102)),
  inference(subsumption_resolution,[],[f41135,f1027])).
fof(f41135,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | ($spl0 | $spl9 | $spl28 | $spl102)),
  inference(subsumption_resolution,[],[f41020,f1029])).
fof(f41020,plain,(
  s('type.integer.int',sK12) = s('type.integer.int',sK13) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | ($spl0 | $spl28 | $spl102)),
  inference(resolution,[],[f40728,f749])).
fof(f40728,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | ($spl0 | $spl28 | $spl102)),
  inference(forward_demodulation,[],[f40494,f1454])).
fof(f40494,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | ($spl0 | $spl102)),
  inference(backward_demodulation,[],[f5837,f905])).
fof(f41134,plain,(
  $false | ($spl0 | $spl5 | $spl28 | $spl102)),
  inference(subsumption_resolution,[],[f41019,f1040])).
fof(f1040,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | $spl5),
  inference(resolution,[],[f910,f707])).
fof(f41019,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | ($spl0 | $spl28 | $spl102)),
  inference(resolution,[],[f40728,f747])).
fof(f41133,plain,(
  $false | ($spl0 | $spl5 | $spl28 | $spl102)),
  inference(subsumption_resolution,[],[f41018,f1040])).
fof(f41018,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | ($spl0 | $spl28 | $spl102)),
  inference(resolution,[],[f40728,f710])).
fof(f41132,plain,(
  $false | ($spl0 | $spl5 | $spl28 | $spl102)),
  inference(subsumption_resolution,[],[f41017,f1040])).
fof(f41017,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | ($spl0 | $spl28 | $spl102)),
  inference(resolution,[],[f40728,f708])).
fof(f40016,plain,(
  ( ! [X49] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X49) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f889,f891])).
fof(f37597,plain,(
  ( ! [X541] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X541)) ) | $spl348),
  inference(cnf_transformation,[],[f37597_D])).
fof(f37597_D,plain,(
  ( ! [X541] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X541)) ) <=> ~$spl348),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl348])])).
fof(f39511,plain,(
  ( ! [X555,X556] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X556))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X555) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X556)) )),
  inference(subsumption_resolution,[],[f39510,f898])).
fof(f898,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))) | s('type.num.num',X0) != s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f599])).
fof(f599,plain,(
  ! [X0,X1] : ((s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) & (s('type.num.num',X0) != s('type.num.num',X1) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))),
  inference(nnf_transformation,[],[f403])).
fof(f403,plain,(
  ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))) <=> s('type.num.num',X0) = s('type.num.num',X1))),
  inference(rectify,[],[f177])).
fof(f177,axiom,(
  ! [X20,X10] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X10))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X20))) <=> s('type.num.num',X10) = s('type.num.num',X20))),
  file('thm.integer.INT_ABS_LE|split|3','thm.integer.INT_INJ')).
fof(f39510,plain,(
  ( ! [X555,X556] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X555))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X556))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X555) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X556)) )),
  inference(forward_demodulation,[],[f39062,f893])).
fof(f39062,plain,(
  ( ! [X555,X556] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X556))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X555) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X556) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X555)))))) )),
  inference(superposition,[],[f744,f890])).
fof(f39509,plain,(
  ( ! [X553,X554] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X554)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X554) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X553)) )),
  inference(subsumption_resolution,[],[f39508,f898])).
fof(f39508,plain,(
  ( ! [X553,X554] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X553))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X554)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X554) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X553)) )),
  inference(forward_demodulation,[],[f39507,f893])).
fof(f39507,plain,(
  ( ! [X553,X554] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X554)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X554) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X553) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X553)))))) )),
  inference(forward_demodulation,[],[f39061,f887])).
fof(f39061,plain,(
  ( ! [X553,X554] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X554)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X554) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X553) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X553)))))) )),
  inference(superposition,[],[f744,f890])).
fof(f37602,plain,(
  ( ! [X545] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X545)))) ) | $spl354),
  inference(cnf_transformation,[],[f37602_D])).
fof(f37602_D,plain,(
  ( ! [X545] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X545)))) ) <=> ~$spl354),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl354])])).
fof(f39506,plain,(
  ( ! [X551,X552] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X551))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X552)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X552)) )),
  inference(forward_demodulation,[],[f39505,f893])).
fof(f39505,plain,(
  ( ! [X551,X552] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X551))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X552)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X552)) )),
  inference(forward_demodulation,[],[f39060,f885])).
fof(f39060,plain,(
  ( ! [X551,X552] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X552)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X552) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X551)))))))) )),
  inference(superposition,[],[f746,f890])).
fof(f39502,plain,(
  ( ! [X547,X548] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X548)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X548) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X547)) )),
  inference(subsumption_resolution,[],[f39501,f898])).
fof(f39501,plain,(
  ( ! [X547,X548] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X547))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X548)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X548) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X547)) )),
  inference(forward_demodulation,[],[f39500,f893])).
fof(f39500,plain,(
  ( ! [X547,X548] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X547))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X548)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X548) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X547)) )),
  inference(forward_demodulation,[],[f39058,f885])).
fof(f39058,plain,(
  ( ! [X547,X548] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X548)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X548) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X547) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X547)))))))) )),
  inference(superposition,[],[f744,f890])).
fof(f39498,plain,(
  ( ! [X541] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X541)))))))) ) | $spl368),
  inference(cnf_transformation,[],[f39498_D])).
fof(f39498_D,plain,(
  ( ! [X541] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X541)))))))) ) <=> ~$spl368),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl368])])).
fof(f39055,plain,(
  ( ! [X542,X541] : (s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X542))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X541)))))))) )),
  inference(superposition,[],[f872,f890])).
fof(f39496,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl0),
  inference(forward_demodulation,[],[f39495,f892])).
fof(f39495,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl0),
  inference(forward_demodulation,[],[f39053,f885])).
fof(f39053,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))) | $spl0),
  inference(superposition,[],[f905,f890])).
fof(f39455,plain,(
  ( ! [X462] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X462))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(forward_demodulation,[],[f39009,f892])).
fof(f39009,plain,(
  ( ! [X462] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X462)))))) )),
  inference(superposition,[],[f889,f890])).
fof(f38117,plain,(
  ( ! [X83,X84] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X84)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X84) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X83)))) )),
  inference(superposition,[],[f746,f874])).
fof(f38495,plain,(
  ( ! [X80,X79] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X80)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X80) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X79)))) )),
  inference(subsumption_resolution,[],[f38115,f897])).
fof(f897,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) = s('type.num.num',X1) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f599])).
fof(f38115,plain,(
  ( ! [X80,X79] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X80)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X80) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X79) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X79)))) )),
  inference(superposition,[],[f744,f874])).
fof(f38112,plain,(
  ( ! [X74,X73] : (s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X74))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X73)))) )),
  inference(superposition,[],[f872,f874])).
fof(f38110,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl0),
  inference(superposition,[],[f905,f874])).
fof(f37695,plain,(
  $false | $spl348),
  inference(equality_resolution,[],[f37597])).
fof(f37607,plain,(
  ( ! [X551] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X551) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X551)))))) ) | $spl360),
  inference(cnf_transformation,[],[f37607_D])).
fof(f37607_D,plain,(
  ( ! [X551] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X551) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X551)))))) ) <=> ~$spl360),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl360])])).
fof(f37206,plain,(
  ( ! [X557,X558] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X558))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X557) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X557)))))) )),
  inference(superposition,[],[f746,f873])).
fof(f37606,plain,(
  ( ! [X551,X552] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X552)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X551) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X551)))))) )),
  inference(forward_demodulation,[],[f37203,f887])).
fof(f37203,plain,(
  ( ! [X551,X552] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X552)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X551) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X551)))))) )),
  inference(superposition,[],[f745,f873])).
fof(f37604,plain,(
  ( ! [X547] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X547) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X547)))))) ) | $spl356),
  inference(cnf_transformation,[],[f37604_D])).
fof(f37604_D,plain,(
  ( ! [X547] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X547) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X547)))))) ) <=> ~$spl356),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl356])])).
fof(f37202,plain,(
  ( ! [X549,X550] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X550))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X549) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X550) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X549)))))) )),
  inference(superposition,[],[f744,f873])).
fof(f37603,plain,(
  ( ! [X547,X548] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X548)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X548) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X547) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X547)))))) )),
  inference(forward_demodulation,[],[f37201,f887])).
fof(f37201,plain,(
  ( ! [X547,X548] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X548)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X548) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X547) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X547)))))) )),
  inference(superposition,[],[f744,f873])).
fof(f37600,plain,(
  ( ! [X546,X545] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X545))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X546)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X546)) )),
  inference(forward_demodulation,[],[f37200,f886])).
fof(f37200,plain,(
  ( ! [X546,X545] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X546)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X546) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X545)))))))) )),
  inference(superposition,[],[f746,f873])).
fof(f37596,plain,(
  ( ! [X542,X541] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X542)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X542) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X541)) )),
  inference(subsumption_resolution,[],[f37595,f898])).
fof(f37595,plain,(
  ( ! [X542,X541] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X541))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X542)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X542) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X541)) )),
  inference(forward_demodulation,[],[f37198,f886])).
fof(f37198,plain,(
  ( ! [X542,X541] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X542)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X542) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X541) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X541)))))))) )),
  inference(superposition,[],[f744,f873])).
fof(f37592,plain,(
  ( ! [X535] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X535)))))))) ) | $spl344),
  inference(cnf_transformation,[],[f37592_D])).
fof(f37592_D,plain,(
  ( ! [X535] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X535)))))))) ) <=> ~$spl344),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl344])])).
fof(f37195,plain,(
  ( ! [X536,X535] : (s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X536))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X535)))))))) )),
  inference(superposition,[],[f872,f873])).
fof(f37588,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl0),
  inference(forward_demodulation,[],[f37193,f886])).
fof(f37193,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))) | $spl0),
  inference(superposition,[],[f905,f873])).
fof(f37150,plain,(
  ( ! [X458] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X458)))))) )),
  inference(superposition,[],[f889,f873])).
fof(f36493,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | $spl0),
  inference(superposition,[],[f905,f848])).
fof(f6904,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) | $spl106),
  inference(cnf_transformation,[],[f6904_D])).
fof(f6904_D,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f36735,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | $spl28),
  inference(forward_demodulation,[],[f36734,f1454])).
fof(f36734,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | $spl28),
  inference(forward_demodulation,[],[f36439,f885])).
fof(f36439,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))))) | $spl28),
  inference(superposition,[],[f848,f1454])).
fof(f36191,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | $spl338),
  inference(cnf_transformation,[],[f36191_D])).
fof(f36191_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) <=> ~$spl338),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl338])])).
fof(f6903,plain,(
  s('type.integer.int',sK12) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) | $spl107),
  inference(cnf_transformation,[],[f6903_D])).
fof(f6903_D,plain,(
  s('type.integer.int',sK12) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) <=> ~$spl107),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl107])])).
fof(f36189,plain,(
  s('type.integer.int',sK12) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | $spl28),
  inference(forward_demodulation,[],[f36188,f1454])).
fof(f36188,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | $spl28),
  inference(forward_demodulation,[],[f36106,f885])).
fof(f36106,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))))) | $spl28),
  inference(superposition,[],[f847,f1454])).
fof(f847,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) | s('type.integer.int',X0) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))) )),
  inference(cnf_transformation,[],[f583])).
fof(f35305,plain,(
  ( ! [X12] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.ABS_1'(s('type.integer.int',X12)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.ABS_1'(s('type.integer.int',X12))))) )),
  inference(superposition,[],[f889,f880])).
fof(f35304,plain,(
  ( ! [X11] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X11))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X11)))))) )),
  inference(superposition,[],[f889,f875])).
fof(f34267,plain,(
  $false | ($spl9 | $spl28 | $spl30)),
  inference(subsumption_resolution,[],[f34266,f1029])).
fof(f34266,plain,(
  s('type.integer.int',sK12) = s('type.integer.int',sK13) | ($spl28 | $spl30)),
  inference(forward_demodulation,[],[f34265,f1454])).
fof(f34265,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) | $spl30),
  inference(forward_demodulation,[],[f33969,f1460])).
fof(f33969,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl30),
  inference(trivial_inequality_removal,[],[f33950])).
fof(f33950,plain,(
  s('type.integer.int',sK13) != s('type.integer.int',sK13) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl30),
  inference(superposition,[],[f876,f1460])).
fof(f34264,plain,(
  $false | ($spl9 | $spl28 | $spl30)),
  inference(subsumption_resolution,[],[f34263,f1029])).
fof(f34263,plain,(
  s('type.integer.int',sK12) = s('type.integer.int',sK13) | ($spl28 | $spl30)),
  inference(forward_demodulation,[],[f33970,f1454])).
fof(f33970,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) | $spl30),
  inference(trivial_inequality_removal,[],[f33949])).
fof(f33949,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl30),
  inference(superposition,[],[f876,f1460])).
fof(f33979,plain,(
  $false | ($spl9 | $spl28 | $spl30)),
  inference(subsumption_resolution,[],[f33978,f1029])).
fof(f33978,plain,(
  s('type.integer.int',sK12) = s('type.integer.int',sK13) | ($spl28 | $spl30)),
  inference(forward_demodulation,[],[f33736,f1454])).
fof(f33736,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) | $spl30),
  inference(backward_demodulation,[],[f1460,f887])).
fof(f32257,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl277),
  inference(cnf_transformation,[],[f32257_D])).
fof(f32257_D,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl277),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl277])])).
fof(f1034,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) | $spl12),
  inference(cnf_transformation,[],[f1034_D])).
fof(f1034_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f28387,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl14),
  inference(superposition,[],[f1164,f876])).
fof(f3290,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) | $spl94),
  inference(cnf_transformation,[],[f3290_D])).
fof(f3290_D,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f33613,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | s('type.integer.int',sK12) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) | $spl28),
  inference(forward_demodulation,[],[f33612,f1454])).
fof(f33612,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.integer.int',sK12) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) | $spl28),
  inference(forward_demodulation,[],[f33611,f853])).
fof(f33611,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',sK12) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) | $spl28),
  inference(forward_demodulation,[],[f33332,f850])).
fof(f850,plain,(
  ( ! [X0] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))) )),
  inference(cnf_transformation,[],[f369])).
fof(f369,plain,(
  ! [X0] : s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X9] : s(bool,'const.integer.int_le_2'(s('type.integer.int',X9),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9)))))),
  file('thm.integer.INT_ABS_LE|split|3','thm.integer.INT_NEG_GE|48|')).
fof(f33332,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl28),
  inference(superposition,[],[f848,f1454])).
fof(f33610,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | $spl28),
  inference(forward_demodulation,[],[f33609,f1454])).
fof(f33609,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | $spl28),
  inference(forward_demodulation,[],[f33608,f885])).
fof(f33608,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))))) | $spl28),
  inference(forward_demodulation,[],[f33331,f852])).
fof(f33331,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))))) | $spl28),
  inference(superposition,[],[f848,f1454])).
fof(f3291,plain,(
  s('type.integer.int',sK12) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) | $spl95),
  inference(cnf_transformation,[],[f3291_D])).
fof(f3291_D,plain,(
  s('type.integer.int',sK12) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) <=> ~$spl95),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl95])])).
fof(f909,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | $spl4),
  inference(cnf_transformation,[],[f909_D])).
fof(f909_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f33607,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | s('type.integer.int',sK12) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) | $spl28),
  inference(forward_demodulation,[],[f33606,f1454])).
fof(f33606,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.integer.int',sK12) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) | $spl28),
  inference(forward_demodulation,[],[f33605,f853])).
fof(f33605,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',sK12) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) | $spl28),
  inference(forward_demodulation,[],[f33330,f850])).
fof(f33330,plain,(
  s('type.integer.int',sK12) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl28),
  inference(superposition,[],[f847,f1454])).
fof(f2045,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | $spl52),
  inference(cnf_transformation,[],[f2045_D])).
fof(f2045_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f33604,plain,(
  s('type.integer.int',sK12) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | $spl28),
  inference(forward_demodulation,[],[f33603,f1454])).
fof(f33603,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | $spl28),
  inference(forward_demodulation,[],[f33602,f885])).
fof(f33602,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))))) | $spl28),
  inference(forward_demodulation,[],[f33329,f852])).
fof(f33329,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))))) | $spl28),
  inference(superposition,[],[f847,f1454])).
fof(f3220,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | $spl90),
  inference(cnf_transformation,[],[f3220_D])).
fof(f3220_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f33447,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | $spl324),
  inference(cnf_transformation,[],[f33447_D])).
fof(f33447_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) <=> ~$spl324),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl324])])).
fof(f3222,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl92),
  inference(cnf_transformation,[],[f3222_D])).
fof(f3222_D,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f33539,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | $spl28),
  inference(forward_demodulation,[],[f33538,f1454])).
fof(f33538,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl28),
  inference(forward_demodulation,[],[f33537,f882])).
fof(f33537,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl28),
  inference(forward_demodulation,[],[f33276,f883])).
fof(f33276,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl28),
  inference(superposition,[],[f825,f1454])).
fof(f825,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f352])).
fof(f352,plain,(
  ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))))),
  inference(rectify,[],[f58])).
fof(f58,axiom,(
  ! [X9] : (s('type.integer.int',X9) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X9)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9)))))))),
  file('thm.integer.INT_ABS_LE|split|3','thm.integer.INT_LT_NEGTOTAL')).
fof(f3218,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl88),
  inference(cnf_transformation,[],[f3218_D])).
fof(f3218_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f33536,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl330),
  inference(cnf_transformation,[],[f33536_D])).
fof(f33536_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl330),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl330])])).
fof(f33534,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl28),
  inference(forward_demodulation,[],[f33533,f1454])).
fof(f33533,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl28),
  inference(forward_demodulation,[],[f33532,f883])).
fof(f33532,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl28),
  inference(forward_demodulation,[],[f33275,f883])).
fof(f33275,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl28),
  inference(superposition,[],[f825,f1454])).
fof(f33526,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl28),
  inference(forward_demodulation,[],[f33525,f1454])).
fof(f33525,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl28),
  inference(forward_demodulation,[],[f33524,f853])).
fof(f33524,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl28),
  inference(forward_demodulation,[],[f33523,f850])).
fof(f33523,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl28),
  inference(forward_demodulation,[],[f33522,f853])).
fof(f33522,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl28),
  inference(forward_demodulation,[],[f33273,f850])).
fof(f33273,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl28),
  inference(superposition,[],[f824,f1454])).
fof(f824,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))))) )),
  inference(cnf_transformation,[],[f351])).
fof(f351,plain,(
  ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X9] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X9)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9)))))))),
  file('thm.integer.INT_ABS_LE|split|3','thm.integer.INT_LE_NEGTOTAL')).
fof(f33521,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK28('const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl328),
  inference(cnf_transformation,[],[f33521_D])).
fof(f33521_D,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK28('const.integer.int_neg_1'(s('type.integer.int',sK13)))))) <=> ~$spl328),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl328])])).
fof(f33519,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK28('const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | $spl28),
  inference(forward_demodulation,[],[f33518,f1454])).
fof(f33518,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK28('const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl28),
  inference(forward_demodulation,[],[f33272,f852])).
fof(f33272,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK28('const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl28),
  inference(superposition,[],[f821,f1454])).
fof(f3199,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl83),
  inference(cnf_transformation,[],[f3199_D])).
fof(f3199_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) <=> ~$spl83),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl83])])).
fof(f33517,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int',sK13)))))))) | $spl326),
  inference(cnf_transformation,[],[f33517_D])).
fof(f33517_D,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int',sK13)))))))) <=> ~$spl326),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl326])])).
fof(f33515,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int',sK13)))))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl28),
  inference(forward_demodulation,[],[f33514,f1454])).
fof(f33514,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int',sK13)))))))) | $spl28),
  inference(forward_demodulation,[],[f33271,f853])).
fof(f33271,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.int_neg_1'(s('type.integer.int',sK13)))))))) | $spl28),
  inference(superposition,[],[f820,f1454])).
fof(f1460,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl30),
  inference(cnf_transformation,[],[f1460_D])).
fof(f1460_D,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f33179,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl28),
  inference(superposition,[],[f825,f1454])).
fof(f33166,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)),X5) | ~sP6(X5,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)),X4) | ~sP8(X4,X5,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) | $spl322),
  inference(cnf_transformation,[],[f33166_D])).
fof(f33166_D,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)),X5) | ~sP6(X5,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)),X4) | ~sP8(X4,X5,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) <=> ~$spl322),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl322])])).
fof(f33159,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)),X5) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | ~sP8(X4,X5,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | ~sP6(X5,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)),X4)) ) | $spl5),
  inference(resolution,[],[f1103,f659])).
fof(f659,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | p(s(bool,X1)) | ~sP8(X0,X2,X1) | ~sP7(X1,X2,X0) | ~sP6(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f490])).
fof(f490,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | (sP6(X2,X1,X0) & sP7(X1,X2,X0) & sP8(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP6(X2,X1,X0) | ~sP7(X1,X2,X0) | ~sP8(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f489])).
fof(f489,plain,(
  ! [X2,X1,X0] : ((~sP9(X2,X1,X0) | (sP6(X0,X1,X2) & sP7(X1,X0,X2) & sP8(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP6(X0,X1,X2) | ~sP7(X1,X0,X2) | ~sP8(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP9(X2,X1,X0)))),
  inference(flattening,[],[f488])).
fof(f488,plain,(
  ! [X2,X1,X0] : ((~sP9(X2,X1,X0) | (sP6(X0,X1,X2) & sP7(X1,X0,X2) & sP8(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP6(X0,X1,X2) | ~sP7(X1,X0,X2) | ~sP8(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP9(X2,X1,X0)))),
  inference(nnf_transformation,[],[f456])).
fof(f456,plain,(
  ! [X2,X1,X0] : (sP9(X2,X1,X0) <=> (sP6(X0,X1,X2) & sP7(X1,X0,X2) & sP8(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f1103,plain,(
  ( ! [X107,X106] : (sP7('const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)),X106,X107)) ) | $spl5),
  inference(resolution,[],[f910,f668])).
fof(f668,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f496])).
fof(f496,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f495])).
fof(f495,plain,(
  ! [X1,X0,X2] : ((~sP7(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP7(X1,X0,X2)))),
  inference(flattening,[],[f494])).
fof(f494,plain,(
  ! [X1,X0,X2] : ((~sP7(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP7(X1,X0,X2)))),
  inference(nnf_transformation,[],[f454])).
fof(f454,plain,(
  ! [X1,X0,X2] : (sP7(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f33143,plain,(
  ( ! [X6,X7] : (sP9('const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)),X6,X7) | ~sP6(X7,X6,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | ~sP7(X6,X7,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) | $spl318),
  inference(cnf_transformation,[],[f33143_D])).
fof(f33143_D,plain,(
  ( ! [X6,X7] : (sP9('const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)),X6,X7) | ~sP6(X7,X6,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | ~sP7(X6,X7,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) <=> ~$spl318),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl318])])).
fof(f33136,plain,(
  ( ! [X6,X7] : (sP9('const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)),X6,X7) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | ~sP7(X6,X7,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | ~sP6(X7,X6,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) | $spl5),
  inference(resolution,[],[f1100,f660])).
fof(f660,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | p(s(bool,X0)) | ~sP8(X0,X2,X1) | ~sP7(X1,X2,X0) | ~sP6(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f490])).
fof(f1100,plain,(
  ( ! [X101,X100] : (sP8('const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)),X100,X101)) ) | $spl5),
  inference(resolution,[],[f910,f664])).
fof(f664,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f493])).
fof(f493,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f492])).
fof(f492,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f491])).
fof(f491,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f455])).
fof(f455,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f1722,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) | $spl41),
  inference(cnf_transformation,[],[f1722_D])).
fof(f1722_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) <=> ~$spl41),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl41])])).
fof(f1724,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | $spl42),
  inference(cnf_transformation,[],[f1724_D])).
fof(f1724_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f32948,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) | $spl5),
  inference(resolution,[],[f1045,f841])).
fof(f1303,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) | $spl23),
  inference(cnf_transformation,[],[f1303_D])).
fof(f1303_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f1305,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | $spl24),
  inference(cnf_transformation,[],[f1305_D])).
fof(f1305_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f32828,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) | $spl5),
  inference(resolution,[],[f1040,f866])).
fof(f1299,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl19),
  inference(cnf_transformation,[],[f1299_D])).
fof(f1299_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f32826,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl5),
  inference(resolution,[],[f1040,f860])).
fof(f1597,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl39),
  inference(cnf_transformation,[],[f1597_D])).
fof(f1597_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) <=> ~$spl39),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl39])])).
fof(f1452,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl26),
  inference(cnf_transformation,[],[f1452_D])).
fof(f1452_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f32692,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl7),
  inference(resolution,[],[f1027,f861])).
fof(f1595,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | $spl36),
  inference(cnf_transformation,[],[f1595_D])).
fof(f1595_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f32690,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK12)))) | $spl7),
  inference(resolution,[],[f1027,f855])).
fof(f32681,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl0),
  inference(forward_demodulation,[],[f32680,f892])).
fof(f32680,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl0),
  inference(forward_demodulation,[],[f32679,f885])).
fof(f32679,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))) | $spl0),
  inference(forward_demodulation,[],[f32632,f853])).
fof(f32632,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))) | $spl0),
  inference(superposition,[],[f905,f890])).
fof(f32678,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl0),
  inference(forward_demodulation,[],[f32631,f853])).
fof(f32631,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl0),
  inference(superposition,[],[f905,f880])).
fof(f32677,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(forward_demodulation,[],[f32630,f853])).
fof(f32630,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(superposition,[],[f905,f879])).
fof(f32676,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl0),
  inference(forward_demodulation,[],[f32629,f853])).
fof(f32629,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl0),
  inference(superposition,[],[f905,f878])).
fof(f32675,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(forward_demodulation,[],[f32628,f853])).
fof(f32628,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(superposition,[],[f905,f877])).
fof(f32674,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl0),
  inference(forward_demodulation,[],[f32673,f886])).
fof(f32673,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))) | $spl0),
  inference(forward_demodulation,[],[f32627,f853])).
fof(f32627,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))) | $spl0),
  inference(superposition,[],[f905,f875])).
fof(f32672,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl0),
  inference(forward_demodulation,[],[f32671,f886])).
fof(f32671,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))) | $spl0),
  inference(forward_demodulation,[],[f32626,f853])).
fof(f32626,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))) | $spl0),
  inference(superposition,[],[f905,f873])).
fof(f32668,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl0),
  inference(forward_demodulation,[],[f32667,f892])).
fof(f32667,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl0),
  inference(forward_demodulation,[],[f32666,f885])).
fof(f32666,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl0),
  inference(forward_demodulation,[],[f32623,f853])).
fof(f32623,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl0),
  inference(superposition,[],[f905,f845])).
fof(f2058,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl58),
  inference(cnf_transformation,[],[f2058_D])).
fof(f2058_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f32665,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl0),
  inference(subsumption_resolution,[],[f32664,f889])).
fof(f32664,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl0),
  inference(forward_demodulation,[],[f32663,f883])).
fof(f32663,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(forward_demodulation,[],[f32662,f882])).
fof(f32662,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(forward_demodulation,[],[f32661,f883])).
fof(f32661,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(forward_demodulation,[],[f32622,f853])).
fof(f32622,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(superposition,[],[f905,f825])).
fof(f32660,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl0),
  inference(forward_demodulation,[],[f32617,f853])).
fof(f32617,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl0),
  inference(superposition,[],[f905,f874])).
fof(f32616,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl0),
  inference(superposition,[],[f905,f876])).
fof(f32659,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl0),
  inference(forward_demodulation,[],[f32613,f852])).
fof(f32613,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | $spl0),
  inference(superposition,[],[f905,f848])).
fof(f32656,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(forward_demodulation,[],[f32524,f886])).
fof(f32524,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))),s('type.integer.int',sK13)))) | $spl0),
  inference(resolution,[],[f905,f841])).
fof(f32521,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(resolution,[],[f905,f749])).
fof(f1463,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | $spl32),
  inference(cnf_transformation,[],[f1463_D])).
fof(f1463_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f32503,plain,(
  ( ! [X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | p(s(bool,X1))) ) | $spl5),
  inference(subsumption_resolution,[],[f32497,f790])).
fof(f32497,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | p(s(bool,X1))) ) | $spl5),
  inference(resolution,[],[f1681,f727])).
fof(f1681,plain,(
  ( ! [X128] : (sP11('const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)),X128)) ) | $spl5),
  inference(resolution,[],[f1045,f729])).
fof(f32502,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | p(s(bool,X0))) ) | $spl5),
  inference(subsumption_resolution,[],[f32496,f790])).
fof(f32496,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | p(s(bool,X0))) ) | $spl5),
  inference(resolution,[],[f1681,f731])).
fof(f32495,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))) ) | $spl5),
  inference(subsumption_resolution,[],[f32489,f790])).
fof(f32489,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))) ) | $spl5),
  inference(resolution,[],[f1680,f727])).
fof(f1680,plain,(
  ( ! [X127] : (sP11(X127,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) | $spl5),
  inference(resolution,[],[f1045,f728])).
fof(f32494,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))) ) | $spl5),
  inference(subsumption_resolution,[],[f32488,f790])).
fof(f32488,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))) ) | $spl5),
  inference(resolution,[],[f1680,f731])).
fof(f10394,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | $spl132),
  inference(cnf_transformation,[],[f10394_D])).
fof(f10394_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f32487,plain,(
  ( ! [X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | p(s(bool,X1))) ) | $spl5),
  inference(subsumption_resolution,[],[f32481,f790])).
fof(f32481,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | p(s(bool,X1))) ) | $spl5),
  inference(resolution,[],[f1258,f727])).
fof(f1258,plain,(
  ( ! [X128] : (sP11('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)),X128)) ) | $spl5),
  inference(resolution,[],[f1040,f729])).
fof(f32486,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | p(s(bool,X0))) ) | $spl5),
  inference(subsumption_resolution,[],[f32480,f790])).
fof(f32480,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | p(s(bool,X0))) ) | $spl5),
  inference(resolution,[],[f1258,f731])).
fof(f1461,plain,(
  s('type.integer.int',sK13) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl31),
  inference(cnf_transformation,[],[f1461_D])).
fof(f1461_D,plain,(
  s('type.integer.int',sK13) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl31),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl31])])).
fof(f32443,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | s('type.integer.int',sK13) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl2),
  inference(superposition,[],[f907,f876])).
fof(f32472,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl2),
  inference(forward_demodulation,[],[f32354,f886])).
fof(f32354,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK12)))) | $spl2),
  inference(resolution,[],[f907,f841])).
fof(f1454,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) | $spl28),
  inference(cnf_transformation,[],[f1454_D])).
fof(f1454_D,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f32351,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl2),
  inference(resolution,[],[f907,f749])).
fof(f32345,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))) ) | $spl5),
  inference(subsumption_resolution,[],[f32339,f790])).
fof(f32339,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))) ) | $spl5),
  inference(resolution,[],[f1257,f727])).
fof(f1257,plain,(
  ( ! [X127] : (sP11(X127,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) | $spl5),
  inference(resolution,[],[f1040,f728])).
fof(f32344,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))) ) | $spl5),
  inference(subsumption_resolution,[],[f32338,f790])).
fof(f32338,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))) ) | $spl5),
  inference(resolution,[],[f1257,f731])).
fof(f4364,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl99),
  inference(cnf_transformation,[],[f4364_D])).
fof(f4364_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) <=> ~$spl99),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl99])])).
fof(f28288,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl14),
  inference(resolution,[],[f1164,f841])).
fof(f3221,plain,(
  s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl93),
  inference(cnf_transformation,[],[f3221_D])).
fof(f3221_D,plain,(
  s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl93),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl93])])).
fof(f6925,plain,(
  ( ! [X0] : (sP11(X0,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) | $spl108),
  inference(cnf_transformation,[],[f6925_D])).
fof(f6925_D,plain,(
  ( ! [X0] : (sP11(X0,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f22237,plain,(
  ( ! [X0] : (sP11(X0,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl10),
  inference(superposition,[],[f11588,f876])).
fof(f11588,plain,(
  ( ! [X127] : (sP11(X127,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) ) | $spl10),
  inference(resolution,[],[f1031,f728])).
fof(f7126,plain,(
  ( ! [X0] : (sP11('const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)),X0)) ) | $spl112),
  inference(cnf_transformation,[],[f7126_D])).
fof(f7126_D,plain,(
  ( ! [X0] : (sP11('const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)),X0)) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f22399,plain,(
  ( ! [X0] : (sP11('const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)),X0) | s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl10),
  inference(superposition,[],[f11589,f876])).
fof(f11589,plain,(
  ( ! [X128] : (sP11('const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)),X128)) ) | $spl10),
  inference(resolution,[],[f1031,f729])).
fof(f22769,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl10),
  inference(superposition,[],[f1031,f876])).
fof(f1027,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | $spl7),
  inference(cnf_transformation,[],[f1027_D])).
fof(f1027_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f23953,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl19),
  inference(superposition,[],[f1299,f876])).
fof(f32335,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)),X1)) ) | $spl302),
  inference(cnf_transformation,[],[f32335_D])).
fof(f32335_D,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)),X1)) ) <=> ~$spl302),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl302])])).
fof(f27300,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)),X1) | s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl10),
  inference(superposition,[],[f11528,f876])).
fof(f11528,plain,(
  ( ! [X15,X16] : (sP0(X15,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)),X16)) ) | $spl10),
  inference(resolution,[],[f1031,f612])).
fof(f612,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f470])).
fof(f470,plain,(
  ! [X0,X1,X2] : ((~sP0(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP0(X0,X1,X2)))),
  inference(rectify,[],[f469])).
fof(f469,plain,(
  ! [X2,X0,X1] : ((~sP0(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP0(X2,X0,X1)))),
  inference(flattening,[],[f468])).
fof(f468,plain,(
  ! [X2,X0,X1] : ((~sP0(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP0(X2,X0,X1)))),
  inference(nnf_transformation,[],[f444])).
fof(f444,plain,(
  ! [X2,X0,X1] : (sP0(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f32334,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) | $spl300),
  inference(cnf_transformation,[],[f32334_D])).
fof(f32334_D,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) <=> ~$spl300),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl300])])).
fof(f27319,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl10),
  inference(superposition,[],[f11541,f876])).
fof(f11541,plain,(
  ( ! [X41,X42] : (sP2(X41,X42,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) ) | $spl10),
  inference(resolution,[],[f1031,f628])).
fof(f628,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f478])).
fof(f478,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(rectify,[],[f477])).
fof(f477,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(flattening,[],[f476])).
fof(f476,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(nnf_transformation,[],[f447])).
fof(f447,plain,(
  ! [X2,X0,X1] : (sP2(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f32333,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)),X1)) ) | $spl298),
  inference(cnf_transformation,[],[f32333_D])).
fof(f32333_D,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)),X1)) ) <=> ~$spl298),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl298])])).
fof(f27338,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)),X1) | s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl10),
  inference(superposition,[],[f11542,f876])).
fof(f11542,plain,(
  ( ! [X43,X44] : (sP2(X43,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)),X44)) ) | $spl10),
  inference(resolution,[],[f1031,f629])).
fof(f629,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f478])).
fof(f32332,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) | $spl296),
  inference(cnf_transformation,[],[f32332_D])).
fof(f32332_D,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) <=> ~$spl296),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl296])])).
fof(f27357,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl10),
  inference(superposition,[],[f11558,f876])).
fof(f11558,plain,(
  ( ! [X76,X75] : (sP4(X75,X76,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) ) | $spl10),
  inference(resolution,[],[f1031,f645])).
fof(f645,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f485])).
fof(f485,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(flattening,[],[f484])).
fof(f484,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(nnf_transformation,[],[f450])).
fof(f450,plain,(
  ! [X0,X1,X2] : (sP4(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f32331,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) | $spl294),
  inference(cnf_transformation,[],[f32331_D])).
fof(f32331_D,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) <=> ~$spl294),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl294])])).
fof(f27376,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl10),
  inference(superposition,[],[f11571,f876])).
fof(f11571,plain,(
  ( ! [X101,X102] : (sP8(X101,X102,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) ) | $spl10),
  inference(resolution,[],[f1031,f662])).
fof(f662,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f493])).
fof(f32330,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) | $spl292),
  inference(cnf_transformation,[],[f32330_D])).
fof(f32330_D,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) <=> ~$spl292),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl292])])).
fof(f27395,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl10),
  inference(superposition,[],[f11574,f876])).
fof(f11574,plain,(
  ( ! [X107,X108] : (sP7(X107,X108,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) ) | $spl10),
  inference(resolution,[],[f1031,f666])).
fof(f666,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f496])).
fof(f32329,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) | $spl290),
  inference(cnf_transformation,[],[f32329_D])).
fof(f32329_D,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) <=> ~$spl290),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl290])])).
fof(f27414,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl10),
  inference(superposition,[],[f11575,f876])).
fof(f11575,plain,(
  ( ! [X109,X110] : (sP6(X109,X110,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) ) | $spl10),
  inference(resolution,[],[f1031,f670])).
fof(f670,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f498])).
fof(f498,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(flattening,[],[f497])).
fof(f497,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(nnf_transformation,[],[f453])).
fof(f453,plain,(
  ! [X0,X1,X2] : (sP6(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f32328,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)),X1)) ) | $spl288),
  inference(cnf_transformation,[],[f32328_D])).
fof(f32328_D,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)),X1)) ) <=> ~$spl288),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl288])])).
fof(f27434,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)),X1) | s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl10),
  inference(superposition,[],[f11576,f876])).
fof(f11576,plain,(
  ( ! [X111,X112] : (sP6(X111,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)),X112)) ) | $spl10),
  inference(resolution,[],[f1031,f671])).
fof(f671,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f498])).
fof(f32327,plain,(
  ( ! [X0,X1] : (sP6('const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)),X0,X1)) ) | $spl286),
  inference(cnf_transformation,[],[f32327_D])).
fof(f32327_D,plain,(
  ( ! [X0,X1] : (sP6('const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)),X0,X1)) ) <=> ~$spl286),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl286])])).
fof(f27454,plain,(
  ( ! [X0,X1] : (sP6('const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)),X0,X1) | s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl10),
  inference(superposition,[],[f11577,f876])).
fof(f11577,plain,(
  ( ! [X114,X113] : (sP6('const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)),X113,X114)) ) | $spl10),
  inference(resolution,[],[f1031,f672])).
fof(f672,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f498])).
fof(f32320,plain,(
  $false | ($spl53 | $spl94)),
  inference(subsumption_resolution,[],[f32319,f828])).
fof(f32319,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)))) | ($spl53 | $spl94)),
  inference(forward_demodulation,[],[f6928,f3290])).
fof(f6928,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | $spl53),
  inference(resolution,[],[f2046,f840])).
fof(f32317,plain,(
  $false | ($spl53 | $spl94 | $spl157)),
  inference(subsumption_resolution,[],[f32316,f25735])).
fof(f32316,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)))) | ($spl53 | $spl94)),
  inference(forward_demodulation,[],[f25738,f3290])).
fof(f25738,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl53),
  inference(forward_demodulation,[],[f25617,f871])).
fof(f871,plain,(
  ( ! [X0,X1] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))) )),
  inference(cnf_transformation,[],[f381])).
fof(f381,plain,(
  ! [X0,X1] : s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))),
  inference(rectify,[],[f53])).
fof(f53,axiom,(
  ! [X9,X11] : s(bool,'const.integer.int_lt_2'(s('type.integer.int',X11),s('type.integer.int',X9))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X11)))))),
  file('thm.integer.INT_ABS_LE|split|3','thm.integer.INT_LT_NEG')).
fof(f25617,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl53),
  inference(resolution,[],[f6930,f857])).
fof(f857,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f585])).
fof(f32315,plain,(
  $false | ($spl53 | $spl94 | $spl157)),
  inference(subsumption_resolution,[],[f32314,f25735])).
fof(f32314,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)))) | ($spl53 | $spl94)),
  inference(forward_demodulation,[],[f25741,f3290])).
fof(f25741,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl53),
  inference(forward_demodulation,[],[f25619,f885])).
fof(f25619,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))))) | $spl53),
  inference(resolution,[],[f6930,f862])).
fof(f862,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) )),
  inference(cnf_transformation,[],[f589])).
fof(f32305,plain,(
  $false | ($spl53 | $spl94 | $spl157)),
  inference(subsumption_resolution,[],[f32304,f25735])).
fof(f32304,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)))) | ($spl53 | $spl94)),
  inference(forward_demodulation,[],[f25464,f3290])).
fof(f25464,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl53),
  inference(resolution,[],[f6928,f707])).
fof(f32303,plain,(
  $false | ($spl53 | $spl94 | $spl157)),
  inference(subsumption_resolution,[],[f32302,f25735])).
fof(f32302,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)))) | ($spl53 | $spl94)),
  inference(forward_demodulation,[],[f25465,f3290])).
fof(f25465,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl53),
  inference(resolution,[],[f6928,f709])).
fof(f32301,plain,(
  $false | ($spl53 | $spl94 | $spl157)),
  inference(subsumption_resolution,[],[f32300,f25735])).
fof(f32300,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)))) | ($spl53 | $spl94)),
  inference(forward_demodulation,[],[f25466,f3290])).
fof(f25466,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl53),
  inference(resolution,[],[f6928,f748])).
fof(f32298,plain,(
  $false | ($spl53 | $spl94)),
  inference(subsumption_resolution,[],[f25468,f3290])).
fof(f25468,plain,(
  s('type.integer.int',sK12) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) | $spl53),
  inference(resolution,[],[f6928,f751])).
fof(f32293,plain,(
  $false | ($spl53 | $spl94 | $spl157)),
  inference(subsumption_resolution,[],[f32292,f25735])).
fof(f32292,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)))) | ($spl53 | $spl94)),
  inference(forward_demodulation,[],[f25606,f3290])).
fof(f25606,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl53),
  inference(resolution,[],[f6930,f863])).
fof(f863,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))))))) )),
  inference(cnf_transformation,[],[f589])).
fof(f32287,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK28('const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl284),
  inference(cnf_transformation,[],[f32287_D])).
fof(f32287_D,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK28('const.integer.ABS_1'(s('type.integer.int',sK12)))))) <=> ~$spl284),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl284])])).
fof(f32285,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK28('const.integer.ABS_1'(s('type.integer.int',sK12)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | $spl94),
  inference(forward_demodulation,[],[f30360,f3290])).
fof(f30360,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK28('const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl94),
  inference(forward_demodulation,[],[f30126,f852])).
fof(f30126,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK28('const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl94),
  inference(superposition,[],[f821,f3290])).
fof(f2043,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl51),
  inference(cnf_transformation,[],[f2043_D])).
fof(f2043_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) <=> ~$spl51),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl51])])).
fof(f10675,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))) | $spl140),
  inference(cnf_transformation,[],[f10675_D])).
fof(f10675_D,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f23855,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl19),
  inference(resolution,[],[f1299,f749])).
fof(f22679,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | $spl10),
  inference(resolution,[],[f1031,f749])).
fof(f29284,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | $spl14),
  inference(forward_demodulation,[],[f29283,f892])).
fof(f29283,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | $spl14),
  inference(forward_demodulation,[],[f28470,f885])).
fof(f28470,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))) | $spl14),
  inference(forward_demodulation,[],[f28403,f852])).
fof(f28403,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))) | $spl14),
  inference(superposition,[],[f1164,f890])).
fof(f28463,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl185),
  inference(cnf_transformation,[],[f28463_D])).
fof(f28463_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) <=> ~$spl185),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl185])])).
fof(f28467,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl187),
  inference(cnf_transformation,[],[f28467_D])).
fof(f28467_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) <=> ~$spl187),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl187])])).
fof(f28468,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl14),
  inference(forward_demodulation,[],[f28402,f852])).
fof(f28402,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl14),
  inference(superposition,[],[f1164,f880])).
fof(f28457,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl180),
  inference(cnf_transformation,[],[f28457_D])).
fof(f28457_D,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f28464,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl14),
  inference(forward_demodulation,[],[f28401,f852])).
fof(f28401,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl14),
  inference(superposition,[],[f1164,f879])).
fof(f28459,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | $spl182),
  inference(cnf_transformation,[],[f28459_D])).
fof(f28459_D,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f28460,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl14),
  inference(forward_demodulation,[],[f28400,f852])).
fof(f28400,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl14),
  inference(superposition,[],[f1164,f878])).
fof(f28454,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl14),
  inference(forward_demodulation,[],[f28399,f852])).
fof(f28399,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl14),
  inference(superposition,[],[f1164,f877])).
fof(f32265,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) | $spl283),
  inference(cnf_transformation,[],[f32265_D])).
fof(f32265_D,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) <=> ~$spl283),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl283])])).
fof(f29294,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | $spl14),
  inference(forward_demodulation,[],[f28451,f886])).
fof(f28451,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))) | $spl14),
  inference(forward_demodulation,[],[f28398,f852])).
fof(f28398,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))) | $spl14),
  inference(superposition,[],[f1164,f875])).
fof(f29296,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | $spl14),
  inference(forward_demodulation,[],[f28448,f886])).
fof(f28448,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))) | $spl14),
  inference(forward_demodulation,[],[f28397,f852])).
fof(f28397,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))) | $spl14),
  inference(superposition,[],[f1164,f873])).
fof(f32263,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl281),
  inference(cnf_transformation,[],[f32263_D])).
fof(f32263_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl281),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl281])])).
fof(f29299,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | $spl14),
  inference(forward_demodulation,[],[f29298,f892])).
fof(f29298,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | $spl14),
  inference(forward_demodulation,[],[f28442,f885])).
fof(f28442,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl14),
  inference(forward_demodulation,[],[f28394,f852])).
fof(f28394,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl14),
  inference(superposition,[],[f1164,f845])).
fof(f32261,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl278),
  inference(cnf_transformation,[],[f32261_D])).
fof(f32261_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) <=> ~$spl278),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl278])])).
fof(f32259,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | $spl14),
  inference(subsumption_resolution,[],[f32258,f889])).
fof(f32258,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | $spl14),
  inference(forward_demodulation,[],[f29302,f883])).
fof(f29302,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl14),
  inference(forward_demodulation,[],[f29301,f882])).
fof(f29301,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl14),
  inference(forward_demodulation,[],[f28436,f883])).
fof(f28436,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl14),
  inference(forward_demodulation,[],[f28393,f852])).
fof(f28393,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl14),
  inference(superposition,[],[f1164,f825])).
fof(f28434,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl14),
  inference(forward_demodulation,[],[f28388,f852])).
fof(f28388,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl14),
  inference(superposition,[],[f1164,f874])).
fof(f10401,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | $spl16),
  inference(forward_demodulation,[],[f5813,f852])).
fof(f5813,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | $spl16),
  inference(superposition,[],[f1166,f848])).
fof(f10398,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | sP11(X0,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) | $spl16),
  inference(forward_demodulation,[],[f6410,f852])).
fof(f6410,plain,(
  ( ! [X0] : (sP11(X0,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12))))) ) | $spl16),
  inference(superposition,[],[f5792,f848])).
fof(f10397,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | sP11('const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)),X0)) ) | $spl16),
  inference(forward_demodulation,[],[f6421,f852])).
fof(f6421,plain,(
  ( ! [X0] : (sP11('const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)),X0) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12))))) ) | $spl16),
  inference(superposition,[],[f5793,f848])).
fof(f1721,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) | $spl40),
  inference(cnf_transformation,[],[f1721_D])).
fof(f1721_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f6927,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | $spl111),
  inference(cnf_transformation,[],[f6927_D])).
fof(f6927_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) <=> ~$spl111),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl111])])).
fof(f32252,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) | $spl14),
  inference(forward_demodulation,[],[f28384,f852])).
fof(f28384,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | $spl14),
  inference(superposition,[],[f1164,f848])).
fof(f908,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl3),
  inference(cnf_transformation,[],[f908_D])).
fof(f908_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f22977,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl27),
  inference(resolution,[],[f1451,f749])).
fof(f32243,plain,(
  ( ! [X0,X1] : (sP7('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)),X0,X1)) ) | $spl274),
  inference(cnf_transformation,[],[f32243_D])).
fof(f32243_D,plain,(
  ( ! [X0,X1] : (sP7('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)),X0,X1)) ) <=> ~$spl274),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl274])])).
fof(f27751,plain,(
  ( ! [X0,X1] : (sP7('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)),X0,X1) | s('type.integer.int',sK13) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl27),
  inference(superposition,[],[f11936,f876])).
fof(f11936,plain,(
  ( ! [X107,X108] : (sP7('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))),X107,X108)) ) | $spl27),
  inference(resolution,[],[f1451,f668])).
fof(f32242,plain,(
  ( ! [X0,X1] : (sP7(X0,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)),X1)) ) | $spl272),
  inference(cnf_transformation,[],[f32242_D])).
fof(f32242_D,plain,(
  ( ! [X0,X1] : (sP7(X0,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)),X1)) ) <=> ~$spl272),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl272])])).
fof(f27731,plain,(
  ( ! [X0,X1] : (sP7(X0,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)),X1) | s('type.integer.int',sK13) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl27),
  inference(superposition,[],[f11935,f876])).
fof(f11935,plain,(
  ( ! [X105,X106] : (sP7(X105,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))),X106)) ) | $spl27),
  inference(resolution,[],[f1451,f667])).
fof(f667,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f496])).
fof(f32241,plain,(
  ( ! [X0,X1] : (sP8('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)),X0,X1)) ) | $spl270),
  inference(cnf_transformation,[],[f32241_D])).
fof(f32241_D,plain,(
  ( ! [X0,X1] : (sP8('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)),X0,X1)) ) <=> ~$spl270),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl270])])).
fof(f27711,plain,(
  ( ! [X0,X1] : (sP8('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)),X0,X1) | s('type.integer.int',sK13) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl27),
  inference(superposition,[],[f11933,f876])).
fof(f11933,plain,(
  ( ! [X101,X102] : (sP8('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))),X101,X102)) ) | $spl27),
  inference(resolution,[],[f1451,f664])).
fof(f32240,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)),X1)) ) | $spl268),
  inference(cnf_transformation,[],[f32240_D])).
fof(f32240_D,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)),X1)) ) <=> ~$spl268),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl268])])).
fof(f27692,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)),X1) | s('type.integer.int',sK13) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl27),
  inference(superposition,[],[f11932,f876])).
fof(f11932,plain,(
  ( ! [X99,X100] : (sP8(X99,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))),X100)) ) | $spl27),
  inference(resolution,[],[f1451,f663])).
fof(f663,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f493])).
fof(f32239,plain,(
  ( ! [X0,X1] : (sP4('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)),X0,X1)) ) | $spl266),
  inference(cnf_transformation,[],[f32239_D])).
fof(f32239_D,plain,(
  ( ! [X0,X1] : (sP4('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)),X0,X1)) ) <=> ~$spl266),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl266])])).
fof(f27674,plain,(
  ( ! [X0,X1] : (sP4('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)),X0,X1) | s('type.integer.int',sK13) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl27),
  inference(superposition,[],[f11920,f876])).
fof(f11920,plain,(
  ( ! [X76,X75] : (sP4('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))),X75,X76)) ) | $spl27),
  inference(resolution,[],[f1451,f647])).
fof(f647,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f485])).
fof(f32238,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)),X1)) ) | $spl264),
  inference(cnf_transformation,[],[f32238_D])).
fof(f32238_D,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)),X1)) ) <=> ~$spl264),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl264])])).
fof(f27655,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)),X1) | s('type.integer.int',sK13) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl27),
  inference(superposition,[],[f11919,f876])).
fof(f11919,plain,(
  ( ! [X74,X73] : (sP4(X73,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))),X74)) ) | $spl27),
  inference(resolution,[],[f1451,f646])).
fof(f646,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f485])).
fof(f32237,plain,(
  ( ! [X0,X1] : (sP2('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)),X0,X1)) ) | $spl262),
  inference(cnf_transformation,[],[f32237_D])).
fof(f32237_D,plain,(
  ( ! [X0,X1] : (sP2('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)),X0,X1)) ) <=> ~$spl262),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl262])])).
fof(f27635,plain,(
  ( ! [X0,X1] : (sP2('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)),X0,X1) | s('type.integer.int',sK13) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl27),
  inference(superposition,[],[f11906,f876])).
fof(f11906,plain,(
  ( ! [X47,X48] : (sP2('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))),X47,X48)) ) | $spl27),
  inference(resolution,[],[f1451,f630])).
fof(f630,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f478])).
fof(f32236,plain,(
  ( ! [X0,X1] : (sP0('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)),X0,X1)) ) | $spl260),
  inference(cnf_transformation,[],[f32236_D])).
fof(f32236_D,plain,(
  ( ! [X0,X1] : (sP0('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)),X0,X1)) ) <=> ~$spl260),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl260])])).
fof(f27617,plain,(
  ( ! [X0,X1] : (sP0('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)),X0,X1) | s('type.integer.int',sK13) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl27),
  inference(superposition,[],[f11890,f876])).
fof(f11890,plain,(
  ( ! [X15,X16] : (sP0('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))),X15,X16)) ) | $spl27),
  inference(resolution,[],[f1451,f613])).
fof(f613,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f470])).
fof(f32235,plain,(
  ( ! [X0,X1] : (sP0(X0,X1,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) | $spl258),
  inference(cnf_transformation,[],[f32235_D])).
fof(f32235_D,plain,(
  ( ! [X0,X1] : (sP0(X0,X1,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) <=> ~$spl258),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl258])])).
fof(f27598,plain,(
  ( ! [X0,X1] : (sP0(X0,X1,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | s('type.integer.int',sK13) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl27),
  inference(superposition,[],[f11889,f876])).
fof(f11889,plain,(
  ( ! [X14,X13] : (sP0(X13,X14,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) ) | $spl27),
  inference(resolution,[],[f1451,f611])).
fof(f611,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f470])).
fof(f1031,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | $spl10),
  inference(cnf_transformation,[],[f1031_D])).
fof(f1031_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f28287,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | $spl14),
  inference(resolution,[],[f1164,f833])).
fof(f32225,plain,(
  ( ! [X0,X1] : (sP6('const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)),X0,X1)) ) | $spl256),
  inference(cnf_transformation,[],[f32225_D])).
fof(f32225_D,plain,(
  ( ! [X0,X1] : (sP6('const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)),X0,X1)) ) <=> ~$spl256),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl256])])).
fof(f32234,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | sP6('const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)),X0,X1)) ) | $spl12),
  inference(forward_demodulation,[],[f27581,f852])).
fof(f27581,plain,(
  ( ! [X0,X1] : (sP6('const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)),X0,X1) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) ) | $spl12),
  inference(superposition,[],[f11767,f848])).
fof(f11767,plain,(
  ( ! [X114,X113] : (sP6('const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)),X113,X114)) ) | $spl12),
  inference(resolution,[],[f1034,f672])).
fof(f32223,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)),X1)) ) | $spl254),
  inference(cnf_transformation,[],[f32223_D])).
fof(f32223_D,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)),X1)) ) <=> ~$spl254),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl254])])).
fof(f32233,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | sP6(X0,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)),X1)) ) | $spl12),
  inference(forward_demodulation,[],[f27567,f852])).
fof(f27567,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)),X1) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) ) | $spl12),
  inference(superposition,[],[f11766,f848])).
fof(f11766,plain,(
  ( ! [X111,X112] : (sP6(X111,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)),X112)) ) | $spl12),
  inference(resolution,[],[f1034,f671])).
fof(f32221,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) | $spl252),
  inference(cnf_transformation,[],[f32221_D])).
fof(f32221_D,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) <=> ~$spl252),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl252])])).
fof(f32232,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | sP6(X0,X1,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) | $spl12),
  inference(forward_demodulation,[],[f27553,f852])).
fof(f27553,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) ) | $spl12),
  inference(superposition,[],[f11765,f848])).
fof(f11765,plain,(
  ( ! [X109,X110] : (sP6(X109,X110,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) ) | $spl12),
  inference(resolution,[],[f1034,f670])).
fof(f32219,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) | $spl250),
  inference(cnf_transformation,[],[f32219_D])).
fof(f32219_D,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) <=> ~$spl250),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl250])])).
fof(f32231,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | sP7(X0,X1,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) | $spl12),
  inference(forward_demodulation,[],[f27540,f852])).
fof(f27540,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) ) | $spl12),
  inference(superposition,[],[f11764,f848])).
fof(f11764,plain,(
  ( ! [X107,X108] : (sP7(X107,X108,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) ) | $spl12),
  inference(resolution,[],[f1034,f666])).
fof(f32217,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) | $spl248),
  inference(cnf_transformation,[],[f32217_D])).
fof(f32217_D,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) <=> ~$spl248),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl248])])).
fof(f32230,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | sP8(X0,X1,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) | $spl12),
  inference(forward_demodulation,[],[f27527,f852])).
fof(f27527,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) ) | $spl12),
  inference(superposition,[],[f11761,f848])).
fof(f11761,plain,(
  ( ! [X101,X102] : (sP8(X101,X102,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) ) | $spl12),
  inference(resolution,[],[f1034,f662])).
fof(f32215,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) | $spl246),
  inference(cnf_transformation,[],[f32215_D])).
fof(f32215_D,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) <=> ~$spl246),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl246])])).
fof(f32229,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | sP4(X0,X1,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) | $spl12),
  inference(forward_demodulation,[],[f27514,f852])).
fof(f27514,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) ) | $spl12),
  inference(superposition,[],[f11748,f848])).
fof(f11748,plain,(
  ( ! [X76,X75] : (sP4(X75,X76,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) ) | $spl12),
  inference(resolution,[],[f1034,f645])).
fof(f32213,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)),X1)) ) | $spl244),
  inference(cnf_transformation,[],[f32213_D])).
fof(f32213_D,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)),X1)) ) <=> ~$spl244),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl244])])).
fof(f32228,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | sP2(X0,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)),X1)) ) | $spl12),
  inference(forward_demodulation,[],[f27501,f852])).
fof(f27501,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)),X1) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) ) | $spl12),
  inference(superposition,[],[f11732,f848])).
fof(f11732,plain,(
  ( ! [X43,X44] : (sP2(X43,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)),X44)) ) | $spl12),
  inference(resolution,[],[f1034,f629])).
fof(f32211,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) | $spl242),
  inference(cnf_transformation,[],[f32211_D])).
fof(f32211_D,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) <=> ~$spl242),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl242])])).
fof(f32227,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | sP2(X0,X1,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) | $spl12),
  inference(forward_demodulation,[],[f27488,f852])).
fof(f27488,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) ) | $spl12),
  inference(superposition,[],[f11731,f848])).
fof(f11731,plain,(
  ( ! [X41,X42] : (sP2(X41,X42,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) ) | $spl12),
  inference(resolution,[],[f1034,f628])).
fof(f32209,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)),X1)) ) | $spl240),
  inference(cnf_transformation,[],[f32209_D])).
fof(f32209_D,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)),X1)) ) <=> ~$spl240),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl240])])).
fof(f32226,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | sP0(X0,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)),X1)) ) | $spl12),
  inference(forward_demodulation,[],[f27475,f852])).
fof(f27475,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)),X1) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) ) | $spl12),
  inference(superposition,[],[f11718,f848])).
fof(f11718,plain,(
  ( ! [X15,X16] : (sP0(X15,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)),X16)) ) | $spl12),
  inference(resolution,[],[f1034,f612])).
fof(f32224,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | sP6('const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)),X0,X1)) ) | $spl16),
  inference(forward_demodulation,[],[f25457,f852])).
fof(f25457,plain,(
  ( ! [X0,X1] : (sP6('const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)),X0,X1) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12))))) ) | $spl16),
  inference(superposition,[],[f5781,f848])).
fof(f5781,plain,(
  ( ! [X114,X113] : (sP6('const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))),X113,X114)) ) | $spl16),
  inference(resolution,[],[f1166,f672])).
fof(f32222,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | sP6(X0,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)),X1)) ) | $spl16),
  inference(forward_demodulation,[],[f25430,f852])).
fof(f25430,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)),X1) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12))))) ) | $spl16),
  inference(superposition,[],[f5780,f848])).
fof(f5780,plain,(
  ( ! [X111,X112] : (sP6(X111,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))),X112)) ) | $spl16),
  inference(resolution,[],[f1166,f671])).
fof(f32220,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | sP6(X0,X1,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) | $spl16),
  inference(forward_demodulation,[],[f25416,f852])).
fof(f25416,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12))))) ) | $spl16),
  inference(superposition,[],[f5779,f848])).
fof(f5779,plain,(
  ( ! [X109,X110] : (sP6(X109,X110,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) ) | $spl16),
  inference(resolution,[],[f1166,f670])).
fof(f32218,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | sP7(X0,X1,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) | $spl16),
  inference(forward_demodulation,[],[f25403,f852])).
fof(f25403,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12))))) ) | $spl16),
  inference(superposition,[],[f5778,f848])).
fof(f5778,plain,(
  ( ! [X107,X108] : (sP7(X107,X108,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) ) | $spl16),
  inference(resolution,[],[f1166,f666])).
fof(f32216,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | sP8(X0,X1,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) | $spl16),
  inference(forward_demodulation,[],[f25390,f852])).
fof(f25390,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12))))) ) | $spl16),
  inference(superposition,[],[f5775,f848])).
fof(f5775,plain,(
  ( ! [X101,X102] : (sP8(X101,X102,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) ) | $spl16),
  inference(resolution,[],[f1166,f662])).
fof(f32214,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | sP4(X0,X1,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) | $spl16),
  inference(forward_demodulation,[],[f25377,f852])).
fof(f25377,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12))))) ) | $spl16),
  inference(superposition,[],[f5762,f848])).
fof(f5762,plain,(
  ( ! [X76,X75] : (sP4(X75,X76,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) ) | $spl16),
  inference(resolution,[],[f1166,f645])).
fof(f32212,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | sP2(X0,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)),X1)) ) | $spl16),
  inference(forward_demodulation,[],[f25351,f852])).
fof(f25351,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)),X1) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12))))) ) | $spl16),
  inference(superposition,[],[f5746,f848])).
fof(f5746,plain,(
  ( ! [X43,X44] : (sP2(X43,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))),X44)) ) | $spl16),
  inference(resolution,[],[f1166,f629])).
fof(f32210,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | sP2(X0,X1,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) | $spl16),
  inference(forward_demodulation,[],[f25338,f852])).
fof(f25338,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12))))) ) | $spl16),
  inference(superposition,[],[f5745,f848])).
fof(f5745,plain,(
  ( ! [X41,X42] : (sP2(X41,X42,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) ) | $spl16),
  inference(resolution,[],[f1166,f628])).
fof(f32208,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | sP0(X0,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)),X1)) ) | $spl16),
  inference(forward_demodulation,[],[f25325,f852])).
fof(f25325,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)),X1) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12))))) ) | $spl16),
  inference(superposition,[],[f5732,f848])).
fof(f5732,plain,(
  ( ! [X15,X16] : (sP0(X15,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))),X16)) ) | $spl16),
  inference(resolution,[],[f1166,f612])).
fof(f10393,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | $spl133),
  inference(cnf_transformation,[],[f10393_D])).
fof(f10393_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) <=> ~$spl133),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl133])])).
fof(f32207,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | $spl16),
  inference(forward_demodulation,[],[f25268,f852])).
fof(f25268,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | $spl16),
  inference(superposition,[],[f5718,f848])).
fof(f25167,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | $spl16),
  inference(resolution,[],[f5718,f866])).
fof(f32206,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | $spl77),
  inference(forward_demodulation,[],[f24393,f852])).
fof(f24393,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | $spl77),
  inference(superposition,[],[f2488,f848])).
fof(f2488,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl77),
  inference(cnf_transformation,[],[f2488_D])).
fof(f2488_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) <=> ~$spl77),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl77])])).
fof(f32205,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | $spl19),
  inference(forward_demodulation,[],[f23982,f871])).
fof(f23982,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl19),
  inference(forward_demodulation,[],[f23856,f886])).
fof(f23856,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl19),
  inference(resolution,[],[f1299,f855])).
fof(f23849,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl19),
  inference(resolution,[],[f1299,f861])).
fof(f23075,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | s('type.integer.int',sK13) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl27),
  inference(superposition,[],[f1451,f876])).
fof(f32204,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | $spl27),
  inference(forward_demodulation,[],[f23105,f885])).
fof(f23105,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))))) | $spl27),
  inference(forward_demodulation,[],[f22980,f871])).
fof(f22980,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))))) | $spl27),
  inference(resolution,[],[f1451,f861])).
fof(f32203,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | $spl12),
  inference(forward_demodulation,[],[f22905,f852])).
fof(f22905,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | $spl12),
  inference(superposition,[],[f1034,f848])).
fof(f22673,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | $spl10),
  inference(resolution,[],[f1031,f841])).
fof(f32202,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | sP11('const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)),X0)) ) | $spl12),
  inference(forward_demodulation,[],[f22427,f852])).
fof(f22427,plain,(
  ( ! [X0] : (sP11('const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)),X0) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) ) | $spl12),
  inference(superposition,[],[f11779,f848])).
fof(f11779,plain,(
  ( ! [X128] : (sP11('const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)),X128)) ) | $spl12),
  inference(resolution,[],[f1034,f729])).
fof(f32201,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | sP11(X0,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) | $spl12),
  inference(forward_demodulation,[],[f22416,f852])).
fof(f22416,plain,(
  ( ! [X0] : (sP11(X0,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) ) | $spl12),
  inference(superposition,[],[f11778,f848])).
fof(f11778,plain,(
  ( ! [X127] : (sP11(X127,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) ) | $spl12),
  inference(resolution,[],[f1034,f728])).
fof(f32200,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | sP11('const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)),X0)) ) | $spl16),
  inference(forward_demodulation,[],[f20628,f852])).
fof(f20628,plain,(
  ( ! [X0] : (sP11('const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)),X0) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12))))) ) | $spl16),
  inference(superposition,[],[f5793,f848])).
fof(f32199,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | sP11(X0,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) | $spl16),
  inference(forward_demodulation,[],[f20617,f852])).
fof(f20617,plain,(
  ( ! [X0] : (sP11(X0,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12))))) ) | $spl16),
  inference(superposition,[],[f5792,f848])).
fof(f32198,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | $spl16),
  inference(forward_demodulation,[],[f19291,f852])).
fof(f19291,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | $spl16),
  inference(superposition,[],[f1166,f848])).
fof(f907,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl2),
  inference(cnf_transformation,[],[f907_D])).
fof(f907_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f19194,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | $spl16),
  inference(resolution,[],[f1166,f836])).
fof(f31868,plain,(
  $false | ($spl5 | $spl8 | $spl157)),
  inference(subsumption_resolution,[],[f31867,f25735])).
fof(f31867,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)))) | ($spl5 | $spl8)),
  inference(forward_demodulation,[],[f1040,f1030])).
fof(f31863,plain,(
  $false | ($spl5 | $spl8)),
  inference(subsumption_resolution,[],[f1044,f1030])).
fof(f1044,plain,(
  s('type.integer.int',sK12) != s('type.integer.int',sK13) | $spl5),
  inference(resolution,[],[f910,f751])).
fof(f31832,plain,(
  $false | ($spl5 | $spl8)),
  inference(subsumption_resolution,[],[f13633,f1030])).
fof(f13633,plain,(
  s('type.integer.int',sK12) != s('type.integer.int',sK13) | $spl5),
  inference(resolution,[],[f910,f751])).
fof(f31619,plain,(
  $false | ($spl8 | $spl132 | $spl157)),
  inference(subsumption_resolution,[],[f31618,f25735])).
fof(f31618,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)))) | ($spl8 | $spl132)),
  inference(forward_demodulation,[],[f10394,f1030])).
fof(f31506,plain,(
  $false | ($spl1 | $spl8 | $spl28 | $spl94)),
  inference(subsumption_resolution,[],[f30875,f30600])).
fof(f30600,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))) | ($spl8 | $spl28)),
  inference(forward_demodulation,[],[f1454,f1030])).
fof(f30875,plain,(
  s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))) | ($spl1 | $spl8 | $spl94)),
  inference(forward_demodulation,[],[f30874,f1030])).
fof(f30874,plain,(
  s('type.integer.int',sK13) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))) | ($spl1 | $spl94)),
  inference(forward_demodulation,[],[f2521,f3290])).
fof(f2521,plain,(
  s('type.integer.int',sK13) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) | $spl1),
  inference(resolution,[],[f906,f751])).
fof(f30366,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | $spl94),
  inference(forward_demodulation,[],[f30365,f3290])).
fof(f30365,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl94),
  inference(subsumption_resolution,[],[f30130,f859])).
fof(f859,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))))))) )),
  inference(cnf_transformation,[],[f587])).
fof(f30130,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl94),
  inference(superposition,[],[f825,f3290])).
fof(f30555,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl94),
  inference(forward_demodulation,[],[f30554,f3290])).
fof(f30554,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl94),
  inference(forward_demodulation,[],[f30364,f3290])).
fof(f30364,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl94),
  inference(forward_demodulation,[],[f30129,f883])).
fof(f30129,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl94),
  inference(superposition,[],[f825,f3290])).
fof(f30557,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | $spl8),
  inference(forward_demodulation,[],[f29106,f1030])).
fof(f29106,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl8),
  inference(forward_demodulation,[],[f29105,f1030])).
fof(f29105,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl8),
  inference(forward_demodulation,[],[f28866,f883])).
fof(f28866,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl8),
  inference(superposition,[],[f825,f1030])).
fof(f30560,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl8),
  inference(forward_demodulation,[],[f30559,f1030])).
fof(f30559,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl8),
  inference(forward_demodulation,[],[f29104,f1030])).
fof(f29104,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl8),
  inference(forward_demodulation,[],[f28865,f883])).
fof(f28865,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl8),
  inference(superposition,[],[f825,f1030])).
fof(f31156,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ($spl8 | $spl28)),
  inference(duplicate_literal_removal,[],[f30898])).
fof(f30898,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ($spl8 | $spl28)),
  inference(superposition,[],[f825,f30600])).
fof(f31249,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl8 | $spl28)),
  inference(forward_demodulation,[],[f31248,f30600])).
fof(f31248,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ($spl8 | $spl28)),
  inference(forward_demodulation,[],[f31153,f883])).
fof(f31153,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ($spl8 | $spl28)),
  inference(duplicate_literal_removal,[],[f30992])).
fof(f30992,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ($spl8 | $spl28)),
  inference(superposition,[],[f825,f30600])).
fof(f31254,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | ($spl8 | $spl28)),
  inference(forward_demodulation,[],[f31253,f30600])).
fof(f31253,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ($spl8 | $spl28)),
  inference(duplicate_literal_removal,[],[f31252])).
fof(f31252,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ($spl8 | $spl28)),
  inference(forward_demodulation,[],[f31251,f882])).
fof(f31251,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ($spl8 | $spl28)),
  inference(forward_demodulation,[],[f30993,f883])).
fof(f30993,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ($spl8 | $spl28)),
  inference(superposition,[],[f825,f30600])).
fof(f31454,plain,(
  $false | ($spl8 | $spl28 | $spl31)),
  inference(subsumption_resolution,[],[f31453,f16673])).
fof(f16673,plain,(
  s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ($spl8 | $spl31)),
  inference(backward_demodulation,[],[f1030,f1461])).
fof(f31453,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ($spl8 | $spl28)),
  inference(forward_demodulation,[],[f31452,f30600])).
fof(f31452,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ($spl8 | $spl28)),
  inference(subsumption_resolution,[],[f31126,f30600])).
fof(f31126,plain,(
  s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ($spl8 | $spl28)),
  inference(superposition,[],[f875,f30600])).
fof(f31330,plain,(
  $false | ($spl8 | $spl28 | $spl31 | $spl94)),
  inference(subsumption_resolution,[],[f31329,f16673])).
fof(f31329,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ($spl8 | $spl28 | $spl94)),
  inference(forward_demodulation,[],[f31328,f30600])).
fof(f31328,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ($spl8 | $spl28 | $spl94)),
  inference(subsumption_resolution,[],[f31327,f828])).
fof(f31327,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ($spl8 | $spl28 | $spl94)),
  inference(forward_demodulation,[],[f31326,f30600])).
fof(f31326,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ($spl8 | $spl28 | $spl94)),
  inference(forward_demodulation,[],[f31325,f853])).
fof(f31325,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ($spl8 | $spl28 | $spl94)),
  inference(forward_demodulation,[],[f31044,f3290])).
fof(f31044,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ($spl8 | $spl28)),
  inference(superposition,[],[f845,f30600])).
fof(f31171,plain,(
  $false | ($spl8 | $spl28 | $spl31)),
  inference(subsumption_resolution,[],[f31155,f16673])).
fof(f31155,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ($spl8 | $spl28)),
  inference(trivial_inequality_removal,[],[f30926])).
fof(f30926,plain,(
  s('type.integer.int',sK12) != s('type.integer.int',sK12) | s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ($spl8 | $spl28)),
  inference(superposition,[],[f875,f30600])).
fof(f30561,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl8 | $spl31)),
  inference(subsumption_resolution,[],[f30560,f16673])).
fof(f30558,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | ($spl8 | $spl31)),
  inference(subsumption_resolution,[],[f30557,f16673])).
fof(f30556,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl8 | $spl31 | $spl94)),
  inference(subsumption_resolution,[],[f30555,f16673])).
fof(f25876,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | $spl164),
  inference(cnf_transformation,[],[f25876_D])).
fof(f25876_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f29269,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))) | $spl198),
  inference(cnf_transformation,[],[f29269_D])).
fof(f29269_D,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f29904,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | $spl52),
  inference(resolution,[],[f2045,f749])).
fof(f30551,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))) | ($spl8 | $spl27)),
  inference(forward_demodulation,[],[f30550,f1030])).
fof(f30550,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | ($spl8 | $spl27)),
  inference(forward_demodulation,[],[f22977,f1030])).
fof(f30543,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl238),
  inference(cnf_transformation,[],[f30543_D])).
fof(f30543_D,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.ABS_1'(s('type.integer.int',sK12)))))))) <=> ~$spl238),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl238])])).
fof(f30541,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl94),
  inference(forward_demodulation,[],[f30359,f3290])).
fof(f30359,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl94),
  inference(forward_demodulation,[],[f30125,f853])).
fof(f30125,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27('const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl94),
  inference(superposition,[],[f820,f3290])).
fof(f30296,plain,(
  $false | ($spl88 | $spl94)),
  inference(subsumption_resolution,[],[f30067,f3218])).
fof(f30067,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl94),
  inference(superposition,[],[f889,f3290])).
fof(f30287,plain,(
  $false | ($spl0 | $spl3 | $spl8 | $spl94)),
  inference(subsumption_resolution,[],[f30035,f28648])).
fof(f28648,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | ($spl3 | $spl8)),
  inference(backward_demodulation,[],[f1030,f908])).
fof(f30035,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | ($spl0 | $spl8 | $spl94)),
  inference(backward_demodulation,[],[f3290,f16670])).
fof(f16670,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | ($spl0 | $spl8)),
  inference(backward_demodulation,[],[f1030,f905])).
fof(f29781,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | ($spl0 | $spl3 | $spl8)),
  inference(subsumption_resolution,[],[f29780,f28648])).
fof(f29780,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | ($spl0 | $spl8)),
  inference(forward_demodulation,[],[f29725,f853])).
fof(f29725,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | ($spl0 | $spl8)),
  inference(superposition,[],[f16670,f879])).
fof(f29779,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | ($spl0 | $spl3 | $spl8)),
  inference(subsumption_resolution,[],[f29778,f28648])).
fof(f29778,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | ($spl0 | $spl8)),
  inference(forward_demodulation,[],[f29723,f853])).
fof(f29723,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | ($spl0 | $spl8)),
  inference(superposition,[],[f16670,f877])).
fof(f2062,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl60),
  inference(cnf_transformation,[],[f2062_D])).
fof(f2062_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f29755,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | $spl237),
  inference(cnf_transformation,[],[f29755_D])).
fof(f29755_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) <=> ~$spl237),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl237])])).
fof(f29753,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl0 | $spl8)),
  inference(forward_demodulation,[],[f29752,f886])).
fof(f29752,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))),s('type.integer.int',sK12)))) | ($spl0 | $spl8)),
  inference(forward_demodulation,[],[f29621,f853])).
fof(f29621,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))),s('type.integer.int',sK12)))) | ($spl0 | $spl8)),
  inference(resolution,[],[f16670,f841])).
fof(f25737,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl158),
  inference(cnf_transformation,[],[f25737_D])).
fof(f25737_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f29749,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) | $spl234),
  inference(cnf_transformation,[],[f29749_D])).
fof(f29749_D,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) <=> ~$spl234),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl234])])).
fof(f29618,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | ($spl0 | $spl8)),
  inference(resolution,[],[f16670,f749])).
fof(f29352,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | s('type.integer.int',sK12) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) | $spl8),
  inference(forward_demodulation,[],[f29351,f1030])).
fof(f29351,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | s('type.integer.int',sK12) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) | $spl8),
  inference(forward_demodulation,[],[f28920,f852])).
fof(f28920,plain,(
  s('type.integer.int',sK12) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | $spl8),
  inference(superposition,[],[f847,f1030])).
fof(f29350,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | s('type.integer.int',sK12) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) | $spl8),
  inference(forward_demodulation,[],[f29349,f1030])).
fof(f29349,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | s('type.integer.int',sK12) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) | $spl8),
  inference(forward_demodulation,[],[f28922,f852])).
fof(f28922,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | $spl8),
  inference(superposition,[],[f848,f1030])).
fof(f3201,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27(sK13)))))) | $spl84),
  inference(cnf_transformation,[],[f3201_D])).
fof(f3201_D,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27(sK13)))))) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f29348,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27(sK13)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl8),
  inference(forward_demodulation,[],[f29093,f1030])).
fof(f29093,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27(sK13)))))) | $spl8),
  inference(forward_demodulation,[],[f28861,f853])).
fof(f28861,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27(sK13)))))) | $spl8),
  inference(superposition,[],[f820,f1030])).
fof(f29338,plain,(
  ( ! [X0] : (sP11('const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X0)) ) | $spl228),
  inference(cnf_transformation,[],[f29338_D])).
fof(f29338_D,plain,(
  ( ! [X0] : (sP11('const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X0)) ) <=> ~$spl228),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl228])])).
fof(f29345,plain,(
  ( ! [X0] : (sP11('const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X0) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12))))) ) | ($spl8 | $spl16)),
  inference(forward_demodulation,[],[f10397,f1030])).
fof(f29340,plain,(
  ( ! [X0] : (sP11(X0,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)))) ) | $spl230),
  inference(cnf_transformation,[],[f29340_D])).
fof(f29340_D,plain,(
  ( ! [X0] : (sP11(X0,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)))) ) <=> ~$spl230),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl230])])).
fof(f29344,plain,(
  ( ! [X0] : (sP11(X0,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK12))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12))))) ) | ($spl8 | $spl16)),
  inference(forward_demodulation,[],[f10398,f1030])).
fof(f29342,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int',sK12)))) | $spl233),
  inference(cnf_transformation,[],[f29342_D])).
fof(f29342_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int',sK12)))) <=> ~$spl233),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl233])])).
fof(f13427,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int',sK12)))) | $spl69),
  inference(forward_demodulation,[],[f13418,f852])).
fof(f13418,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int',sK12)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | $spl69),
  inference(superposition,[],[f2079,f848])).
fof(f29339,plain,(
  ( ! [X0] : (sP11(X0,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK12))) | s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl10)),
  inference(forward_demodulation,[],[f22237,f1030])).
fof(f29337,plain,(
  ( ! [X0] : (sP11('const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X0) | s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl10)),
  inference(forward_demodulation,[],[f22399,f1030])).
fof(f29330,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl225),
  inference(cnf_transformation,[],[f29330_D])).
fof(f29330_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) <=> ~$spl225),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl225])])).
fof(f25759,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl157),
  inference(resolution,[],[f25735,f861])).
fof(f29328,plain,(
  ( ! [X0] : (sP11(X0,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) ) | $spl222),
  inference(cnf_transformation,[],[f29328_D])).
fof(f29328_D,plain,(
  ( ! [X0] : (sP11(X0,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) ) <=> ~$spl222),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl222])])).
fof(f29327,plain,(
  ( ! [X0] : (sP11(X0,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12))))) ) | ($spl0 | $spl8)),
  inference(forward_demodulation,[],[f26917,f1030])).
fof(f26917,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | sP11(X0,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) ) | $spl0),
  inference(forward_demodulation,[],[f26895,f852])).
fof(f26895,plain,(
  ( ! [X0] : (sP11(X0,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12))))) ) | $spl0),
  inference(superposition,[],[f1969,f848])).
fof(f29326,plain,(
  ( ! [X0] : (sP11('const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))),X0)) ) | $spl220),
  inference(cnf_transformation,[],[f29326_D])).
fof(f29326_D,plain,(
  ( ! [X0] : (sP11('const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))),X0)) ) <=> ~$spl220),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl220])])).
fof(f29325,plain,(
  ( ! [X0] : (sP11('const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))),X0) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12))))) ) | ($spl0 | $spl8)),
  inference(forward_demodulation,[],[f27123,f1030])).
fof(f27123,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | sP11('const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))),X0)) ) | $spl0),
  inference(forward_demodulation,[],[f27101,f852])).
fof(f27101,plain,(
  ( ! [X0] : (sP11('const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))),X0) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12))))) ) | $spl0),
  inference(superposition,[],[f1970,f848])).
fof(f29324,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X1)) ) | $spl218),
  inference(cnf_transformation,[],[f29324_D])).
fof(f29324_D,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X1)) ) <=> ~$spl218),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl218])])).
fof(f29323,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X1) | s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl10)),
  inference(forward_demodulation,[],[f27300,f1030])).
fof(f29322,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)))) ) | $spl216),
  inference(cnf_transformation,[],[f29322_D])).
fof(f29322_D,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)))) ) <=> ~$spl216),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl216])])).
fof(f29321,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK12))) | s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl10)),
  inference(forward_demodulation,[],[f27319,f1030])).
fof(f29320,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X1)) ) | $spl214),
  inference(cnf_transformation,[],[f29320_D])).
fof(f29320_D,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X1)) ) <=> ~$spl214),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl214])])).
fof(f29319,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X1) | s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl10)),
  inference(forward_demodulation,[],[f27338,f1030])).
fof(f29318,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)))) ) | $spl212),
  inference(cnf_transformation,[],[f29318_D])).
fof(f29318_D,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)))) ) <=> ~$spl212),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl212])])).
fof(f29317,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK12))) | s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl10)),
  inference(forward_demodulation,[],[f27357,f1030])).
fof(f29316,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)))) ) | $spl210),
  inference(cnf_transformation,[],[f29316_D])).
fof(f29316_D,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)))) ) <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f29315,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK12))) | s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl10)),
  inference(forward_demodulation,[],[f27376,f1030])).
fof(f29314,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)))) ) | $spl208),
  inference(cnf_transformation,[],[f29314_D])).
fof(f29314_D,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)))) ) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f29313,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK12))) | s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl10)),
  inference(forward_demodulation,[],[f27395,f1030])).
fof(f29312,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)))) ) | $spl206),
  inference(cnf_transformation,[],[f29312_D])).
fof(f29312_D,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)))) ) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f29311,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK12))) | s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl10)),
  inference(forward_demodulation,[],[f27414,f1030])).
fof(f29310,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X1)) ) | $spl204),
  inference(cnf_transformation,[],[f29310_D])).
fof(f29310_D,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X1)) ) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f29309,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X1) | s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl10)),
  inference(forward_demodulation,[],[f27434,f1030])).
fof(f29308,plain,(
  ( ! [X0,X1] : (sP6('const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X0,X1)) ) | $spl202),
  inference(cnf_transformation,[],[f29308_D])).
fof(f29308_D,plain,(
  ( ! [X0,X1] : (sP6('const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X0,X1)) ) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f29307,plain,(
  ( ! [X0,X1] : (sP6('const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)),X0,X1) | s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl8 | $spl10)),
  inference(forward_demodulation,[],[f27454,f1030])).
fof(f29305,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | ($spl8 | $spl14)),
  inference(subsumption_resolution,[],[f29304,f889])).
fof(f29304,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | ($spl8 | $spl14)),
  inference(forward_demodulation,[],[f29303,f883])).
fof(f29303,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | ($spl8 | $spl14)),
  inference(forward_demodulation,[],[f29302,f1030])).
fof(f29300,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | ($spl8 | $spl14)),
  inference(forward_demodulation,[],[f29299,f1030])).
fof(f29295,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | ($spl8 | $spl14)),
  inference(forward_demodulation,[],[f29294,f1030])).
fof(f29293,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | ($spl8 | $spl14)),
  inference(forward_demodulation,[],[f29292,f1030])).
fof(f29292,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | ($spl8 | $spl14)),
  inference(forward_demodulation,[],[f28454,f1030])).
fof(f29289,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | ($spl8 | $spl14)),
  inference(forward_demodulation,[],[f29288,f1030])).
fof(f29288,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | ($spl8 | $spl14)),
  inference(forward_demodulation,[],[f28464,f1030])).
fof(f29282,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl200),
  inference(cnf_transformation,[],[f29282_D])).
fof(f29282_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f29280,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | ($spl0 | $spl8)),
  inference(forward_demodulation,[],[f28625,f1030])).
fof(f28625,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl0),
  inference(forward_demodulation,[],[f28572,f852])).
fof(f28572,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | $spl0),
  inference(superposition,[],[f1893,f848])).
fof(f3205,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK28(sK13)))) | $spl86),
  inference(cnf_transformation,[],[f3205_D])).
fof(f3205_D,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK28(sK13)))) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f29279,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK28(sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | $spl8),
  inference(forward_demodulation,[],[f29096,f1030])).
fof(f29096,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK28(sK13)))) | $spl8),
  inference(forward_demodulation,[],[f28862,f852])).
fof(f28862,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK28(sK13)))) | $spl8),
  inference(superposition,[],[f821,f1030])).
fof(f29275,plain,(
  s('type.integer.int',sK12) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | $spl8),
  inference(forward_demodulation,[],[f29169,f1030])).
fof(f29169,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | s('type.integer.int',sK13) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) | $spl8),
  inference(forward_demodulation,[],[f28919,f852])).
fof(f28919,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | s('type.integer.int',sK13) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) | $spl8),
  inference(superposition,[],[f847,f1030])).
fof(f29274,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | $spl8),
  inference(forward_demodulation,[],[f29172,f1030])).
fof(f29172,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | s('type.integer.int',sK13) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) | $spl8),
  inference(forward_demodulation,[],[f28921,f852])).
fof(f28921,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | s('type.integer.int',sK13) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) | $spl8),
  inference(superposition,[],[f848,f1030])).
fof(f29271,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))) | ($spl8 | $spl10)),
  inference(forward_demodulation,[],[f29270,f1030])).
fof(f29270,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | ($spl8 | $spl10)),
  inference(forward_demodulation,[],[f22679,f1030])).
fof(f29267,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))) | ($spl8 | $spl19)),
  inference(forward_demodulation,[],[f29266,f1030])).
fof(f29266,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | ($spl8 | $spl19)),
  inference(forward_demodulation,[],[f23855,f1030])).
fof(f29026,plain,(
  $false | ($spl8 | $spl51 | $spl82)),
  inference(subsumption_resolution,[],[f28684,f3198])).
fof(f28684,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | ($spl8 | $spl51)),
  inference(backward_demodulation,[],[f1030,f2043])).
fof(f29025,plain,(
  $false | ($spl3 | $spl8 | $spl53)),
  inference(subsumption_resolution,[],[f28671,f2046])).
fof(f28671,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | ($spl3 | $spl8)),
  inference(backward_demodulation,[],[f1030,f1736])).
fof(f1736,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) | $spl3),
  inference(resolution,[],[f908,f767])).
fof(f29024,plain,(
  $false | ($spl3 | $spl8 | $spl53)),
  inference(subsumption_resolution,[],[f28670,f6933])).
fof(f6933,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | $spl53),
  inference(resolution,[],[f2046,f750])).
fof(f28670,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | ($spl3 | $spl8)),
  inference(backward_demodulation,[],[f1030,f1731])).
fof(f1731,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) | $spl3),
  inference(resolution,[],[f908,f707])).
fof(f29023,plain,(
  $false | ($spl8 | $spl36 | $spl53)),
  inference(subsumption_resolution,[],[f28669,f6933])).
fof(f28669,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | ($spl8 | $spl36)),
  inference(backward_demodulation,[],[f1030,f1595])).
fof(f29022,plain,(
  $false | ($spl8 | $spl27 | $spl53)),
  inference(subsumption_resolution,[],[f28666,f6930])).
fof(f28666,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | ($spl8 | $spl27)),
  inference(backward_demodulation,[],[f1030,f1451])).
fof(f29021,plain,(
  $false | ($spl8 | $spl19 | $spl53)),
  inference(subsumption_resolution,[],[f28665,f6930])).
fof(f28665,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | ($spl8 | $spl19)),
  inference(backward_demodulation,[],[f1030,f1299])).
fof(f29020,plain,(
  $false | ($spl8 | $spl12 | $spl53)),
  inference(subsumption_resolution,[],[f28662,f6928])).
fof(f28662,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | ($spl8 | $spl12)),
  inference(backward_demodulation,[],[f1030,f1034])).
fof(f29019,plain,(
  $false | ($spl8 | $spl10 | $spl53)),
  inference(subsumption_resolution,[],[f28661,f2046])).
fof(f28661,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | ($spl8 | $spl10)),
  inference(backward_demodulation,[],[f1030,f1031])).
fof(f29018,plain,(
  $false | ($spl3 | $spl8 | $spl82)),
  inference(subsumption_resolution,[],[f28648,f3198])).
fof(f28646,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) | $spl196),
  inference(cnf_transformation,[],[f28646_D])).
fof(f28646_D,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f28482,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl0),
  inference(resolution,[],[f1893,f749])).
fof(f28644,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl0),
  inference(subsumption_resolution,[],[f28643,f889])).
fof(f28643,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl0),
  inference(forward_demodulation,[],[f28642,f883])).
fof(f28642,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(forward_demodulation,[],[f28641,f882])).
fof(f28641,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(forward_demodulation,[],[f22612,f883])).
fof(f22612,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(forward_demodulation,[],[f22574,f853])).
fof(f22574,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(superposition,[],[f905,f825])).
fof(f28640,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl0),
  inference(forward_demodulation,[],[f28639,f892])).
fof(f28639,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl0),
  inference(forward_demodulation,[],[f22618,f885])).
fof(f22618,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl0),
  inference(forward_demodulation,[],[f22575,f853])).
fof(f22575,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl0),
  inference(superposition,[],[f905,f845])).
fof(f28637,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl0),
  inference(forward_demodulation,[],[f22627,f886])).
fof(f22627,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))) | $spl0),
  inference(forward_demodulation,[],[f22579,f853])).
fof(f22579,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))) | $spl0),
  inference(superposition,[],[f905,f875])).
fof(f22630,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(forward_demodulation,[],[f22580,f853])).
fof(f22580,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(superposition,[],[f905,f877])).
fof(f22632,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(forward_demodulation,[],[f22582,f853])).
fof(f22582,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(superposition,[],[f905,f879])).
fof(f26496,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK13)))) | $spl167),
  inference(resolution,[],[f26476,f855])).
fof(f28615,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | sP11('const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)),X0)) ) | $spl80),
  inference(forward_demodulation,[],[f24578,f852])).
fof(f24578,plain,(
  ( ! [X0] : (sP11('const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)),X0) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) ) | $spl80),
  inference(superposition,[],[f24511,f848])).
fof(f24511,plain,(
  ( ! [X128] : (sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)),X128)) ) | $spl80),
  inference(resolution,[],[f2894,f729])).
fof(f28617,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | sP11(X0,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) | $spl80),
  inference(forward_demodulation,[],[f24567,f852])).
fof(f24567,plain,(
  ( ! [X0] : (sP11(X0,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) ) | $spl80),
  inference(superposition,[],[f24510,f848])).
fof(f24510,plain,(
  ( ! [X127] : (sP11(X127,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) ) | $spl80),
  inference(resolution,[],[f2894,f728])).
fof(f1028,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | $spl6),
  inference(cnf_transformation,[],[f1028_D])).
fof(f1028_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f28619,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | $spl80),
  inference(forward_demodulation,[],[f24529,f852])).
fof(f24529,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | $spl80),
  inference(superposition,[],[f2894,f848])).
fof(f28626,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | $spl0),
  inference(forward_demodulation,[],[f28577,f852])).
fof(f28628,plain,(
  $false | ($spl8 | $spl57 | $spl82)),
  inference(subsumption_resolution,[],[f16675,f3198])).
fof(f16675,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | ($spl8 | $spl57)),
  inference(backward_demodulation,[],[f1030,f2050])).
fof(f5837,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) | $spl102),
  inference(cnf_transformation,[],[f5837_D])).
fof(f5837_D,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f25172,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | $spl16),
  inference(resolution,[],[f5718,f749])).
fof(f24295,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl77),
  inference(resolution,[],[f2488,f749])).
fof(f1030,plain,(
  s('type.integer.int',sK12) = s('type.integer.int',sK13) | $spl8),
  inference(cnf_transformation,[],[f1030_D])).
fof(f1030_D,plain,(
  s('type.integer.int',sK12) = s('type.integer.int',sK13) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f17162,plain,(
  s('type.integer.int',sK12) = s('type.integer.int',sK13) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | $spl4),
  inference(resolution,[],[f909,f749])).
fof(f1462,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | $spl33),
  inference(cnf_transformation,[],[f1462_D])).
fof(f1462_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f18414,plain,(
  s('type.integer.int',sK12) = s('type.integer.int',sK13) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | $spl4),
  inference(resolution,[],[f10705,f749])).
fof(f10705,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | $spl4),
  inference(resolution,[],[f909,f708])).
fof(f28624,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | $spl3),
  inference(forward_demodulation,[],[f22383,f885])).
fof(f22383,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))))) | $spl3),
  inference(forward_demodulation,[],[f22259,f843])).
fof(f843,plain,(
  ( ! [X0,X1] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))) )),
  inference(cnf_transformation,[],[f364])).
fof(f364,plain,(
  ! [X0,X1] : s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X9,X11] : s(bool,'const.integer.int_le_2'(s('type.integer.int',X11),s('type.integer.int',X9))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X11)))))),
  file('thm.integer.INT_ABS_LE|split|3','thm.integer.INT_LE_NEG')).
fof(f22259,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))))) | $spl3),
  inference(resolution,[],[f908,f836])).
fof(f22802,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | $spl10),
  inference(forward_demodulation,[],[f22801,f843])).
fof(f22801,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl10),
  inference(forward_demodulation,[],[f22682,f885])).
fof(f22682,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl10),
  inference(resolution,[],[f1031,f841])).
fof(f22971,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl27),
  inference(resolution,[],[f1451,f861])).
fof(f28623,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | $spl27),
  inference(forward_demodulation,[],[f23103,f871])).
fof(f23103,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl27),
  inference(forward_demodulation,[],[f22978,f886])).
fof(f22978,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK12)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl27),
  inference(resolution,[],[f1451,f855])).
fof(f23259,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | $spl3),
  inference(resolution,[],[f1736,f841])).
fof(f23566,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl51),
  inference(resolution,[],[f2043,f836])).
fof(f28622,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | $spl51),
  inference(forward_demodulation,[],[f23699,f843])).
fof(f23699,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl51),
  inference(forward_demodulation,[],[f23574,f886])).
fof(f23574,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl51),
  inference(resolution,[],[f2043,f833])).
fof(f28621,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | $spl19),
  inference(forward_demodulation,[],[f23986,f885])).
fof(f23986,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))))) | $spl19),
  inference(forward_demodulation,[],[f23858,f871])).
fof(f23858,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))))) | $spl19),
  inference(resolution,[],[f1299,f861])).
fof(f28610,plain,(
  $false | ($spl0 | $spl43 | $spl57)),
  inference(subsumption_resolution,[],[f28609,f7166])).
fof(f7166,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | $spl57),
  inference(resolution,[],[f2050,f767])).
fof(f28600,plain,(
  $false | ($spl0 | $spl3 | $spl33)),
  inference(subsumption_resolution,[],[f28475,f5274])).
fof(f5274,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | ($spl3 | $spl33)),
  inference(forward_demodulation,[],[f5273,f885])).
fof(f5273,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))))) | ($spl3 | $spl33)),
  inference(subsumption_resolution,[],[f5272,f1462])).
fof(f5272,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))))) | $spl3),
  inference(forward_demodulation,[],[f5148,f843])).
fof(f5148,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))))) | $spl3),
  inference(resolution,[],[f908,f836])).
fof(f28475,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl0),
  inference(resolution,[],[f1893,f839])).
fof(f839,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f581])).
fof(f28469,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | ($spl14 | $spl53)),
  inference(subsumption_resolution,[],[f28468,f2046])).
fof(f28465,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | ($spl14 | $spl53)),
  inference(subsumption_resolution,[],[f28464,f2046])).
fof(f28461,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | ($spl14 | $spl53)),
  inference(subsumption_resolution,[],[f28460,f2046])).
fof(f28455,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | ($spl14 | $spl53)),
  inference(subsumption_resolution,[],[f28454,f2046])).
fof(f1592,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK12)))) | $spl34),
  inference(cnf_transformation,[],[f1592_D])).
fof(f1592_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK12)))) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f28427,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) | $spl178),
  inference(cnf_transformation,[],[f28427_D])).
fof(f28427_D,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f28294,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK12)))) | $spl14),
  inference(resolution,[],[f1164,f749])).
fof(f27118,plain,(
  ( ! [X20] : (sP11('const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X20) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))))) ) | $spl0),
  inference(superposition,[],[f1970,f879])).
fof(f27116,plain,(
  ( ! [X18] : (sP11('const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X18) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))))) ) | $spl0),
  inference(superposition,[],[f1970,f877])).
fof(f27134,plain,(
  ( ! [X17] : (s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) | sP11('const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X17)) ) | $spl0),
  inference(forward_demodulation,[],[f27115,f886])).
fof(f27115,plain,(
  ( ! [X17] : (sP11('const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X17) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) ) | $spl0),
  inference(superposition,[],[f1970,f875])).
fof(f27130,plain,(
  ( ! [X11] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | sP11('const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X11)) ) | $spl0),
  inference(forward_demodulation,[],[f27129,f892])).
fof(f27129,plain,(
  ( ! [X11] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | sP11('const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X11)) ) | $spl0),
  inference(forward_demodulation,[],[f27111,f885])).
fof(f27111,plain,(
  ( ! [X11] : (sP11('const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X11) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl0),
  inference(superposition,[],[f1970,f845])).
fof(f27127,plain,(
  ( ! [X10] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | sP11('const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X10)) ) | $spl0),
  inference(subsumption_resolution,[],[f27126,f889])).
fof(f27126,plain,(
  ( ! [X10] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | sP11('const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X10)) ) | $spl0),
  inference(forward_demodulation,[],[f27125,f883])).
fof(f27125,plain,(
  ( ! [X10] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | sP11('const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X10) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))))) ) | $spl0),
  inference(forward_demodulation,[],[f27124,f882])).
fof(f27124,plain,(
  ( ! [X10] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | sP11('const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X10) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))))) ) | $spl0),
  inference(forward_demodulation,[],[f27110,f883])).
fof(f27110,plain,(
  ( ! [X10] : (sP11('const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X10) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))))) ) | $spl0),
  inference(superposition,[],[f1970,f825])).
fof(f26912,plain,(
  ( ! [X20] : (sP11(X20,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))))) ) | $spl0),
  inference(superposition,[],[f1969,f879])).
fof(f26910,plain,(
  ( ! [X18] : (sP11(X18,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))))) ) | $spl0),
  inference(superposition,[],[f1969,f877])).
fof(f26928,plain,(
  ( ! [X17] : (s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) | sP11(X17,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl0),
  inference(forward_demodulation,[],[f26909,f886])).
fof(f26909,plain,(
  ( ! [X17] : (sP11(X17,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) ) | $spl0),
  inference(superposition,[],[f1969,f875])).
fof(f26924,plain,(
  ( ! [X11] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | sP11(X11,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl0),
  inference(forward_demodulation,[],[f26923,f892])).
fof(f26923,plain,(
  ( ! [X11] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | sP11(X11,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl0),
  inference(forward_demodulation,[],[f26905,f885])).
fof(f26905,plain,(
  ( ! [X11] : (sP11(X11,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl0),
  inference(superposition,[],[f1969,f845])).
fof(f26921,plain,(
  ( ! [X10] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | sP11(X10,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl0),
  inference(subsumption_resolution,[],[f26920,f889])).
fof(f26920,plain,(
  ( ! [X10] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | sP11(X10,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl0),
  inference(forward_demodulation,[],[f26919,f883])).
fof(f26919,plain,(
  ( ! [X10] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | sP11(X10,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))))) ) | $spl0),
  inference(forward_demodulation,[],[f26918,f882])).
fof(f26918,plain,(
  ( ! [X10] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | sP11(X10,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))))) ) | $spl0),
  inference(forward_demodulation,[],[f26904,f883])).
fof(f26904,plain,(
  ( ! [X10] : (sP11(X10,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))))) ) | $spl0),
  inference(superposition,[],[f1969,f825])).
fof(f26498,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl167),
  inference(resolution,[],[f26476,f861])).
fof(f26478,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | $spl168),
  inference(cnf_transformation,[],[f26478_D])).
fof(f26478_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f26476,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK13)))) | $spl167),
  inference(cnf_transformation,[],[f26476_D])).
fof(f26476_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK13)))) <=> ~$spl167),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl167])])).
fof(f26483,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | $spl57),
  inference(forward_demodulation,[],[f26482,f871])).
fof(f26482,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl57),
  inference(forward_demodulation,[],[f26361,f885])).
fof(f26361,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl57),
  inference(resolution,[],[f7161,f866])).
fof(f7161,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | $spl57),
  inference(resolution,[],[f2050,f707])).
fof(f26481,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | $spl57),
  inference(forward_demodulation,[],[f26480,f886])).
fof(f26480,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))))) | $spl57),
  inference(forward_demodulation,[],[f26359,f871])).
fof(f26359,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))))) | $spl57),
  inference(resolution,[],[f7161,f860])).
fof(f26347,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK13)))) | $spl57),
  inference(resolution,[],[f7161,f866])).
fof(f25874,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int',sK12)))) | $spl163),
  inference(cnf_transformation,[],[f25874_D])).
fof(f25874_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int',sK12)))) <=> ~$spl163),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl163])])).
fof(f25757,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int',sK12)))) | $spl157),
  inference(resolution,[],[f25735,f855])).
fof(f25744,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl160),
  inference(cnf_transformation,[],[f25744_D])).
fof(f25744_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f25735,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)))) | $spl157),
  inference(cnf_transformation,[],[f25735_D])).
fof(f25735_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)))) <=> ~$spl157),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl157])])).
fof(f25742,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl53),
  inference(forward_demodulation,[],[f25620,f886])).
fof(f25620,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))))),s('type.integer.int',sK12)))) | $spl53),
  inference(resolution,[],[f6930,f866])).
fof(f25740,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl53),
  inference(forward_demodulation,[],[f25739,f871])).
fof(f25739,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl53),
  inference(forward_demodulation,[],[f25618,f885])).
fof(f25618,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl53),
  inference(resolution,[],[f6930,f860])).
fof(f25607,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)))) | $spl53),
  inference(resolution,[],[f6930,f860])).
fof(f25589,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int',sK12)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl53),
  inference(forward_demodulation,[],[f25471,f892])).
fof(f25471,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))),s('type.integer.int',sK12)))) | $spl53),
  inference(resolution,[],[f6928,f833])).
fof(f25175,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl16),
  inference(resolution,[],[f5718,f861])).
fof(f25294,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl16),
  inference(forward_demodulation,[],[f25173,f892])).
fof(f25173,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))),s('type.integer.int',sK13)))) | $spl16),
  inference(resolution,[],[f5718,f855])).
fof(f24728,plain,(
  $false | ($spl53 | $spl83)),
  inference(subsumption_resolution,[],[f24604,f2046])).
fof(f24604,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | $spl83),
  inference(resolution,[],[f3199,f767])).
fof(f24727,plain,(
  $false | ($spl53 | $spl83)),
  inference(subsumption_resolution,[],[f24603,f2046])).
fof(f24603,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | $spl83),
  inference(resolution,[],[f3199,f767])).
fof(f1300,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl21),
  inference(cnf_transformation,[],[f1300_D])).
fof(f1300_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f24556,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl150),
  inference(cnf_transformation,[],[f24556_D])).
fof(f24556_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f24440,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl80),
  inference(resolution,[],[f2894,f860])).
fof(f24418,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl149),
  inference(cnf_transformation,[],[f24418_D])).
fof(f24418_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) <=> ~$spl149),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl149])])).
fof(f24296,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK12)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl77),
  inference(resolution,[],[f2488,f855])).
fof(f24280,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl43),
  inference(forward_demodulation,[],[f24160,f892])).
fof(f24160,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))),s('type.integer.int',sK13)))) | $spl43),
  inference(resolution,[],[f1723,f833])).
fof(f13330,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl143),
  inference(cnf_transformation,[],[f13330_D])).
fof(f13330_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) <=> ~$spl143),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl143])])).
fof(f1164,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK12)))) | $spl14),
  inference(cnf_transformation,[],[f1164_D])).
fof(f1164_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK12)))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f23721,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK12)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | ($spl3 | $spl33)),
  inference(resolution,[],[f5274,f833])).
fof(f22670,plain,(
  ( ! [X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) | p(s(bool,X1))) ) | $spl3),
  inference(subsumption_resolution,[],[f22657,f790])).
fof(f22657,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) | p(s(bool,X1))) ) | $spl3),
  inference(resolution,[],[f12464,f727])).
fof(f12464,plain,(
  ( ! [X128] : (sP11('const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)),X128)) ) | $spl3),
  inference(resolution,[],[f1736,f729])).
fof(f22669,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) | p(s(bool,X0))) ) | $spl3),
  inference(subsumption_resolution,[],[f22656,f790])).
fof(f22656,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) | p(s(bool,X0))) ) | $spl3),
  inference(resolution,[],[f12464,f731])).
fof(f22653,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12))))) ) | $spl3),
  inference(subsumption_resolution,[],[f22640,f790])).
fof(f22640,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12))))) ) | $spl3),
  inference(resolution,[],[f12463,f727])).
fof(f12463,plain,(
  ( ! [X127] : (sP11(X127,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) ) | $spl3),
  inference(resolution,[],[f1736,f728])).
fof(f22652,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12))))) ) | $spl3),
  inference(subsumption_resolution,[],[f22639,f790])).
fof(f22639,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12))))) ) | $spl3),
  inference(resolution,[],[f12463,f731])).
fof(f22633,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | ($spl0 | $spl57)),
  inference(subsumption_resolution,[],[f22632,f2050])).
fof(f22631,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | ($spl0 | $spl57)),
  inference(subsumption_resolution,[],[f22630,f2050])).
fof(f22608,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(forward_demodulation,[],[f22476,f886])).
fof(f22476,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))),s('type.integer.int',sK13)))) | $spl0),
  inference(resolution,[],[f905,f841])).
fof(f22473,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(resolution,[],[f905,f749])).
fof(f1302,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) | $spl22),
  inference(cnf_transformation,[],[f1302_D])).
fof(f1302_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f22465,plain,(
  ( ! [X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) | p(s(bool,X1))) ) | $spl3),
  inference(subsumption_resolution,[],[f22452,f790])).
fof(f22452,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) | p(s(bool,X1))) ) | $spl3),
  inference(resolution,[],[f12232,f727])).
fof(f12232,plain,(
  ( ! [X128] : (sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)),X128)) ) | $spl3),
  inference(resolution,[],[f1731,f729])).
fof(f22464,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) | p(s(bool,X0))) ) | $spl3),
  inference(subsumption_resolution,[],[f22451,f790])).
fof(f22451,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) | p(s(bool,X0))) ) | $spl3),
  inference(resolution,[],[f12232,f731])).
fof(f22448,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12))))) ) | $spl3),
  inference(subsumption_resolution,[],[f22435,f790])).
fof(f22435,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12))))) ) | $spl3),
  inference(resolution,[],[f12231,f727])).
fof(f12231,plain,(
  ( ! [X127] : (sP11(X127,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) ) | $spl3),
  inference(resolution,[],[f1731,f728])).
fof(f22447,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12))))) ) | $spl3),
  inference(subsumption_resolution,[],[f22434,f790])).
fof(f22434,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12))))) ) | $spl3),
  inference(resolution,[],[f12231,f731])).
fof(f22211,plain,(
  $false | ($spl3 | $spl33 | $spl50)),
  inference(subsumption_resolution,[],[f13201,f2044])).
fof(f13201,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | ($spl3 | $spl33)),
  inference(resolution,[],[f5274,f838])).
fof(f22210,plain,(
  $false | ($spl11 | $spl12)),
  inference(subsumption_resolution,[],[f11702,f1032])).
fof(f11702,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | $spl12),
  inference(resolution,[],[f1034,f840])).
fof(f22191,plain,(
  $false | ($spl38 | $spl42)),
  inference(subsumption_resolution,[],[f22071,f1596])).
fof(f22071,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl42),
  inference(resolution,[],[f1724,f747])).
fof(f22190,plain,(
  $false | ($spl38 | $spl42)),
  inference(subsumption_resolution,[],[f22070,f1596])).
fof(f22070,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl42),
  inference(resolution,[],[f1724,f710])).
fof(f22189,plain,(
  $false | ($spl38 | $spl42)),
  inference(subsumption_resolution,[],[f22069,f1596])).
fof(f22069,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl42),
  inference(resolution,[],[f1724,f708])).
fof(f22188,plain,(
  $false | ($spl41 | $spl42)),
  inference(subsumption_resolution,[],[f22067,f1722])).
fof(f22067,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) | $spl42),
  inference(resolution,[],[f1724,f840])).
fof(f22187,plain,(
  $false | ($spl6 | $spl42)),
  inference(subsumption_resolution,[],[f22066,f2775])).
fof(f2775,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | $spl6),
  inference(resolution,[],[f1028,f708])).
fof(f22066,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | $spl42),
  inference(resolution,[],[f1724,f839])).
fof(f21788,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int',sK13)))) | $spl38),
  inference(resolution,[],[f1596,f866])).
fof(f21900,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl38),
  inference(forward_demodulation,[],[f21786,f892])).
fof(f21786,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl38),
  inference(resolution,[],[f1596,f860])).
fof(f21627,plain,(
  $false | ($spl2 | $spl22)),
  inference(subsumption_resolution,[],[f1302,f1329])).
fof(f1329,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) | $spl2),
  inference(resolution,[],[f907,f708])).
fof(f20996,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK12)))) | $spl2),
  inference(resolution,[],[f1456,f841])).
fof(f8401,plain,(
  ( ! [X0] : (sP11('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)),X0)) ) | $spl126),
  inference(cnf_transformation,[],[f8401_D])).
fof(f8401_D,plain,(
  ( ! [X0] : (sP11('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)),X0)) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f20987,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | sP11('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)),X0)) ) | $spl76),
  inference(forward_demodulation,[],[f20982,f852])).
fof(f20982,plain,(
  ( ! [X0] : (sP11('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)),X0) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) ) | $spl76),
  inference(superposition,[],[f8320,f848])).
fof(f8320,plain,(
  ( ! [X128] : (sP11('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))),X128)) ) | $spl76),
  inference(resolution,[],[f2487,f729])).
fof(f8388,plain,(
  ( ! [X0] : (sP11(X0,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) | $spl124),
  inference(cnf_transformation,[],[f8388_D])).
fof(f8388_D,plain,(
  ( ! [X0] : (sP11(X0,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f20975,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | sP11(X0,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) | $spl76),
  inference(forward_demodulation,[],[f20970,f852])).
fof(f20970,plain,(
  ( ! [X0] : (sP11(X0,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) ) | $spl76),
  inference(superposition,[],[f8319,f848])).
fof(f8319,plain,(
  ( ! [X127] : (sP11(X127,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) ) | $spl76),
  inference(resolution,[],[f2487,f728])).
fof(f20710,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | sP11('const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)),X0)) ) | $spl2),
  inference(forward_demodulation,[],[f20705,f852])).
fof(f20705,plain,(
  ( ! [X0] : (sP11('const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)),X0) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) ) | $spl2),
  inference(superposition,[],[f6859,f848])).
fof(f20709,plain,(
  ( ! [X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | p(s(bool,X1))) ) | $spl2),
  inference(subsumption_resolution,[],[f20700,f790])).
fof(f20700,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | p(s(bool,X1))) ) | $spl2),
  inference(resolution,[],[f6859,f727])).
fof(f20708,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | p(s(bool,X0))) ) | $spl2),
  inference(subsumption_resolution,[],[f20699,f790])).
fof(f20699,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | p(s(bool,X0))) ) | $spl2),
  inference(resolution,[],[f6859,f731])).
fof(f20698,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | sP11(X0,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) | $spl2),
  inference(forward_demodulation,[],[f20693,f852])).
fof(f20693,plain,(
  ( ! [X0] : (sP11(X0,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) ) | $spl2),
  inference(superposition,[],[f6858,f848])).
fof(f20697,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))) ) | $spl2),
  inference(subsumption_resolution,[],[f20688,f790])).
fof(f20688,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))) ) | $spl2),
  inference(resolution,[],[f6858,f727])).
fof(f20696,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))) ) | $spl2),
  inference(subsumption_resolution,[],[f20687,f790])).
fof(f20687,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))) ) | $spl2),
  inference(resolution,[],[f6858,f731])).
fof(f20532,plain,(
  $false | ($spl4 | $spl132)),
  inference(subsumption_resolution,[],[f10394,f10705])).
fof(f2050,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl57),
  inference(cnf_transformation,[],[f2050_D])).
fof(f2050_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) <=> ~$spl57),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl57])])).
fof(f20209,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl1),
  inference(forward_demodulation,[],[f20162,f853])).
fof(f20162,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl1),
  inference(superposition,[],[f906,f879])).
fof(f20208,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl1),
  inference(forward_demodulation,[],[f20160,f853])).
fof(f20160,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl1),
  inference(superposition,[],[f906,f877])).
fof(f20207,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl1),
  inference(forward_demodulation,[],[f20206,f886])).
fof(f20206,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))) | $spl1),
  inference(forward_demodulation,[],[f20159,f853])).
fof(f20159,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))) | $spl1),
  inference(superposition,[],[f906,f875])).
fof(f20201,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl1),
  inference(forward_demodulation,[],[f20200,f892])).
fof(f20200,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl1),
  inference(forward_demodulation,[],[f20199,f885])).
fof(f20199,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl1),
  inference(forward_demodulation,[],[f20155,f853])).
fof(f20155,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl1),
  inference(superposition,[],[f906,f845])).
fof(f20198,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl1),
  inference(subsumption_resolution,[],[f20197,f889])).
fof(f20197,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl1),
  inference(forward_demodulation,[],[f20196,f883])).
fof(f20196,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl1),
  inference(forward_demodulation,[],[f20195,f882])).
fof(f20195,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl1),
  inference(forward_demodulation,[],[f20194,f883])).
fof(f20194,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl1),
  inference(forward_demodulation,[],[f20154,f853])).
fof(f20154,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl1),
  inference(superposition,[],[f906,f825])).
fof(f2662,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl79),
  inference(cnf_transformation,[],[f2662_D])).
fof(f2662_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) <=> ~$spl79),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl79])])).
fof(f20187,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | $spl1),
  inference(forward_demodulation,[],[f20186,f843])).
fof(f20186,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl1),
  inference(forward_demodulation,[],[f20049,f886])).
fof(f20049,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl1),
  inference(resolution,[],[f906,f833])).
fof(f20041,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl1),
  inference(resolution,[],[f906,f836])).
fof(f19854,plain,(
  $false | ($spl6 | $spl28 | $spl53)),
  inference(subsumption_resolution,[],[f18191,f2775])).
fof(f18191,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | ($spl28 | $spl53)),
  inference(forward_demodulation,[],[f18190,f1454])).
fof(f18190,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | ($spl28 | $spl53)),
  inference(forward_demodulation,[],[f18189,f852])).
fof(f18189,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | ($spl28 | $spl53)),
  inference(forward_demodulation,[],[f18188,f851])).
fof(f851,plain,(
  ( ! [X0] : (s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(cnf_transformation,[],[f370])).
fof(f370,plain,(
  ! [X0] : s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(rectify,[],[f137])).
fof(f137,axiom,(
  ! [X9] : s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X9))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.integer.INT_ABS_LE|split|3','thm.integer.INT_NEG_LE|48|')).
fof(f18188,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl28 | $spl53)),
  inference(forward_demodulation,[],[f18187,f850])).
fof(f18187,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))))) | ($spl28 | $spl53)),
  inference(subsumption_resolution,[],[f18186,f2046])).
fof(f18186,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))))) | $spl28),
  inference(forward_demodulation,[],[f17938,f852])).
fof(f17938,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))))) | $spl28),
  inference(superposition,[],[f824,f1454])).
fof(f19853,plain,(
  $false | ($spl4 | $spl28 | $spl95)),
  inference(subsumption_resolution,[],[f18273,f3291])).
fof(f18273,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) | ($spl4 | $spl28)),
  inference(subsumption_resolution,[],[f18272,f909])).
fof(f18272,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | s('type.integer.int',sK12) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) | $spl28),
  inference(forward_demodulation,[],[f18271,f1454])).
fof(f18271,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.integer.int',sK12) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) | $spl28),
  inference(forward_demodulation,[],[f18270,f853])).
fof(f18270,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',sK12) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) | $spl28),
  inference(forward_demodulation,[],[f17996,f850])).
fof(f17996,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl28),
  inference(superposition,[],[f848,f1454])).
fof(f19852,plain,(
  $false | ($spl6 | $spl28 | $spl53)),
  inference(subsumption_resolution,[],[f18111,f2775])).
fof(f18111,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | ($spl28 | $spl53)),
  inference(forward_demodulation,[],[f18110,f1454])).
fof(f18110,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | ($spl28 | $spl53)),
  inference(forward_demodulation,[],[f18109,f852])).
fof(f18109,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | ($spl28 | $spl53)),
  inference(subsumption_resolution,[],[f18108,f2046])).
fof(f18108,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | $spl28),
  inference(forward_demodulation,[],[f17842,f852])).
fof(f17842,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | $spl28),
  inference(superposition,[],[f824,f1454])).
fof(f19346,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl27),
  inference(resolution,[],[f1451,f749])).
fof(f19839,plain,(
  $false | ($spl3 | $spl4 | $spl13)),
  inference(subsumption_resolution,[],[f19838,f909])).
fof(f19838,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | ($spl3 | $spl13)),
  inference(forward_demodulation,[],[f19837,f843])).
fof(f19837,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | ($spl3 | $spl13)),
  inference(subsumption_resolution,[],[f19836,f1033])).
fof(f19836,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl3),
  inference(forward_demodulation,[],[f19719,f885])).
fof(f19719,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK12)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl3),
  inference(resolution,[],[f1736,f841])).
fof(f19202,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int',sK13)))) | $spl16),
  inference(resolution,[],[f1166,f841])).
fof(f1596,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl38),
  inference(cnf_transformation,[],[f1596_D])).
fof(f1596_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f19199,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl16),
  inference(resolution,[],[f1166,f749])).
fof(f19071,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl6),
  inference(resolution,[],[f2775,f836])).
fof(f19070,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int',sK13)))) | $spl6),
  inference(resolution,[],[f2775,f833])).
fof(f18911,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl13),
  inference(resolution,[],[f1033,f836])).
fof(f19030,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK12)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl13),
  inference(forward_demodulation,[],[f18910,f892])).
fof(f18910,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))),s('type.integer.int',sK12)))) | $spl13),
  inference(resolution,[],[f1033,f833])).
fof(f18838,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | $spl11),
  inference(forward_demodulation,[],[f18837,f885])).
fof(f18837,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))))))),s('type.integer.int',sK13)))) | $spl11),
  inference(forward_demodulation,[],[f18713,f843])).
fof(f18713,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))))))),s('type.integer.int',sK13)))) | $spl11),
  inference(resolution,[],[f1032,f833])).
fof(f18675,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | $spl3),
  inference(forward_demodulation,[],[f18674,f885])).
fof(f18674,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))))) | $spl3),
  inference(forward_demodulation,[],[f18550,f843])).
fof(f18550,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))))) | $spl3),
  inference(resolution,[],[f908,f836])).
fof(f1298,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl18),
  inference(cnf_transformation,[],[f1298_D])).
fof(f1298_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f18417,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl4),
  inference(resolution,[],[f10705,f861])).
fof(f18415,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int',sK13)))) | $spl4),
  inference(resolution,[],[f10705,f855])).
fof(f18100,plain,(
  $false | ($spl3 | $spl28)),
  inference(subsumption_resolution,[],[f17832,f828])).
fof(f17832,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)))) | ($spl3 | $spl28)),
  inference(backward_demodulation,[],[f1454,f908])).
fof(f1594,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | $spl37),
  inference(cnf_transformation,[],[f1594_D])).
fof(f1594_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) <=> ~$spl37),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl37])])).
fof(f2894,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) | $spl80),
  inference(cnf_transformation,[],[f2894_D])).
fof(f2894_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f17719,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | $spl6),
  inference(resolution,[],[f1028,f866])).
fof(f1451,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl27),
  inference(cnf_transformation,[],[f1451_D])).
fof(f1451_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) <=> ~$spl27),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl27])])).
fof(f17717,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl6),
  inference(resolution,[],[f1028,f860])).
fof(f17657,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(forward_demodulation,[],[f17612,f853])).
fof(f17612,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(superposition,[],[f905,f879])).
fof(f17656,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(forward_demodulation,[],[f17610,f853])).
fof(f17610,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(superposition,[],[f905,f877])).
fof(f17655,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl0),
  inference(forward_demodulation,[],[f17654,f886])).
fof(f17654,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))) | $spl0),
  inference(forward_demodulation,[],[f17609,f853])).
fof(f17609,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))) | $spl0),
  inference(superposition,[],[f905,f875])).
fof(f17649,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl0),
  inference(forward_demodulation,[],[f17648,f892])).
fof(f17648,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl0),
  inference(forward_demodulation,[],[f17647,f885])).
fof(f17647,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl0),
  inference(forward_demodulation,[],[f17605,f853])).
fof(f17605,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl0),
  inference(superposition,[],[f905,f845])).
fof(f17646,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl0),
  inference(subsumption_resolution,[],[f17645,f889])).
fof(f17645,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl0),
  inference(forward_demodulation,[],[f17644,f883])).
fof(f17644,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(forward_demodulation,[],[f17643,f882])).
fof(f17643,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(forward_demodulation,[],[f17642,f883])).
fof(f17642,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(forward_demodulation,[],[f17604,f853])).
fof(f17604,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(superposition,[],[f905,f825])).
fof(f17638,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(forward_demodulation,[],[f17506,f886])).
fof(f17506,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))),s('type.integer.int',sK13)))) | $spl0),
  inference(resolution,[],[f905,f841])).
fof(f17503,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(resolution,[],[f905,f749])).
fof(f17407,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | s('type.integer.int',sK13) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl3),
  inference(superposition,[],[f908,f876])).
fof(f17437,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | $spl3),
  inference(forward_demodulation,[],[f17436,f885])).
fof(f17436,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))))) | $spl3),
  inference(forward_demodulation,[],[f17312,f843])).
fof(f17312,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))))) | $spl3),
  inference(resolution,[],[f908,f836])).
fof(f1032,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | $spl11),
  inference(cnf_transformation,[],[f1032_D])).
fof(f1032_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f17165,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | $spl4),
  inference(resolution,[],[f909,f841])).
fof(f17143,plain,(
  $false | ($spl6 | $spl8)),
  inference(subsumption_resolution,[],[f12022,f1030])).
fof(f12022,plain,(
  s('type.integer.int',sK12) != s('type.integer.int',sK13) | $spl6),
  inference(resolution,[],[f2775,f751])).
fof(f16984,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | ($spl8 | $spl53)),
  inference(subsumption_resolution,[],[f16983,f7065])).
fof(f7065,plain,(
  s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl53),
  inference(subsumption_resolution,[],[f7029,f828])).
fof(f7029,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK12)))) | s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl53),
  inference(superposition,[],[f2046,f876])).
fof(f16983,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | $spl8),
  inference(forward_demodulation,[],[f16982,f1030])).
fof(f16982,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl8),
  inference(forward_demodulation,[],[f16981,f1030])).
fof(f16981,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl8),
  inference(forward_demodulation,[],[f16747,f883])).
fof(f16747,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl8),
  inference(superposition,[],[f825,f1030])).
fof(f16980,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl8 | $spl53)),
  inference(subsumption_resolution,[],[f16979,f7065])).
fof(f16979,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl8),
  inference(forward_demodulation,[],[f16978,f1030])).
fof(f16978,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl8),
  inference(forward_demodulation,[],[f16977,f1030])).
fof(f16977,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl8),
  inference(forward_demodulation,[],[f16746,f883])).
fof(f16746,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl8),
  inference(superposition,[],[f825,f1030])).
fof(f16966,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27(sK13)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl8),
  inference(forward_demodulation,[],[f16965,f1030])).
fof(f16965,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27(sK13)))))) | $spl8),
  inference(forward_demodulation,[],[f16742,f853])).
fof(f16742,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27(sK13)))))) | $spl8),
  inference(superposition,[],[f820,f1030])).
fof(f1029,plain,(
  s('type.integer.int',sK12) != s('type.integer.int',sK13) | $spl9),
  inference(cnf_transformation,[],[f1029_D])).
fof(f1029_D,plain,(
  s('type.integer.int',sK12) != s('type.integer.int',sK13) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f16669,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | s('type.integer.int',sK12) != s('type.integer.int',sK13) | $spl13),
  inference(forward_demodulation,[],[f15226,f852])).
fof(f15226,plain,(
  s('type.integer.int',sK12) != s('type.integer.int',sK13) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | $spl13),
  inference(superposition,[],[f14849,f848])).
fof(f14849,plain,(
  s('type.integer.int',sK12) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) | $spl13),
  inference(resolution,[],[f1033,f751])).
fof(f16668,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | $spl13),
  inference(forward_demodulation,[],[f14946,f852])).
fof(f14946,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | $spl13),
  inference(superposition,[],[f1033,f848])).
fof(f16667,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | sP11(X0,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) | $spl76),
  inference(forward_demodulation,[],[f15427,f852])).
fof(f15427,plain,(
  ( ! [X0] : (sP11(X0,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) ) | $spl76),
  inference(superposition,[],[f8319,f848])).
fof(f16666,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | sP11('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)),X0)) ) | $spl76),
  inference(forward_demodulation,[],[f15438,f852])).
fof(f15438,plain,(
  ( ! [X0] : (sP11('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)),X0) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) ) | $spl76),
  inference(superposition,[],[f8320,f848])).
fof(f14844,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | $spl13),
  inference(resolution,[],[f1033,f841])).
fof(f16649,plain,(
  $false | ($spl7 | $spl19 | $spl38)),
  inference(subsumption_resolution,[],[f13590,f1027])).
fof(f13590,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | ($spl19 | $spl38)),
  inference(subsumption_resolution,[],[f13589,f1596])).
fof(f13589,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | $spl19),
  inference(forward_demodulation,[],[f13565,f885])).
fof(f13565,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))))) | $spl19),
  inference(forward_demodulation,[],[f13437,f871])).
fof(f13437,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))))) | $spl19),
  inference(resolution,[],[f1299,f861])).
fof(f16648,plain,(
  $false | ($spl5 | $spl16 | $spl51)),
  inference(subsumption_resolution,[],[f16647,f829])).
fof(f829,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f357])).
fof(f357,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.integer.INT_ABS_LE|split|3',aHL_FALSITY)).
fof(f16647,plain,(
  p(s(bool,'F_0')) | ($spl5 | $spl16 | $spl51)),
  inference(forward_demodulation,[],[f13587,f13718])).
fof(f13718,plain,(
  s(bool,'F_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | $spl5),
  inference(resolution,[],[f910,f787])).
fof(f787,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f556])).
fof(f556,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X0)) & (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f555])).
fof(f555,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X0)) & (s(bool,'F_0') != s(bool,X0) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f432])).
fof(f432,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> s(bool,'F_0') = s(bool,X0))),
  inference(ennf_transformation,[],[f327])).
fof(f327,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> s(bool,'F_0') = s(bool,X0))),
  inference(rectify,[],[f78])).
fof(f78,axiom,(
  ! [X12] : ((p(s(bool,X12)) => p(s(bool,'F_0'))) <=> s(bool,'F_0') = s(bool,X12))),
  file('thm.integer.INT_ABS_LE|split|3','thm.bool.IMP_F_EQ_F')).
fof(f13587,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | ($spl16 | $spl51)),
  inference(subsumption_resolution,[],[f13586,f1166])).
fof(f13586,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | $spl51),
  inference(forward_demodulation,[],[f13074,f885])).
fof(f13074,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))))) | $spl51),
  inference(forward_demodulation,[],[f12947,f843])).
fof(f12947,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))))) | $spl51),
  inference(resolution,[],[f2043,f836])).
fof(f11211,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(forward_demodulation,[],[f11079,f886])).
fof(f11079,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))),s('type.integer.int',sK13)))) | $spl0),
  inference(resolution,[],[f905,f841])).
fof(f14032,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) | $spl5),
  inference(resolution,[],[f1045,f841])).
fof(f16178,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK12)))) | $spl5),
  inference(subsumption_resolution,[],[f15939,f790])).
fof(f15939,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK12)))) | $spl5),
  inference(superposition,[],[f833,f13718])).
fof(f16630,plain,(
  $false | ($spl18 | $spl36)),
  inference(subsumption_resolution,[],[f16503,f1298])).
fof(f16503,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl36),
  inference(resolution,[],[f1595,f766])).
fof(f16629,plain,(
  $false | ($spl18 | $spl36)),
  inference(subsumption_resolution,[],[f16502,f1298])).
fof(f16502,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl36),
  inference(resolution,[],[f1595,f766])).
fof(f16628,plain,(
  $false | ($spl11 | $spl36)),
  inference(subsumption_resolution,[],[f16501,f1032])).
fof(f16501,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | $spl36),
  inference(resolution,[],[f1595,f750])).
fof(f16627,plain,(
  $false | ($spl36 | $spl50)),
  inference(subsumption_resolution,[],[f16500,f2044])).
fof(f16500,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl36),
  inference(resolution,[],[f1595,f747])).
fof(f16626,plain,(
  $false | ($spl36 | $spl50)),
  inference(subsumption_resolution,[],[f16499,f2044])).
fof(f16499,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl36),
  inference(resolution,[],[f1595,f710])).
fof(f16625,plain,(
  $false | ($spl36 | $spl50)),
  inference(subsumption_resolution,[],[f16498,f2044])).
fof(f16498,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl36),
  inference(resolution,[],[f1595,f708])).
fof(f16624,plain,(
  $false | ($spl18 | $spl36)),
  inference(subsumption_resolution,[],[f16497,f1298])).
fof(f16497,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl36),
  inference(resolution,[],[f1595,f698])).
fof(f16623,plain,(
  $false | ($spl18 | $spl36)),
  inference(subsumption_resolution,[],[f16496,f1298])).
fof(f16496,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl36),
  inference(resolution,[],[f1595,f698])).
fof(f15613,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK12)))) | $spl2),
  inference(resolution,[],[f1456,f841])).
fof(f15312,plain,(
  ( ! [X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | p(s(bool,X1))) ) | $spl2),
  inference(subsumption_resolution,[],[f15303,f790])).
fof(f15303,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | p(s(bool,X1))) ) | $spl2),
  inference(resolution,[],[f6859,f727])).
fof(f15311,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | p(s(bool,X0))) ) | $spl2),
  inference(subsumption_resolution,[],[f15302,f790])).
fof(f15302,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | p(s(bool,X0))) ) | $spl2),
  inference(resolution,[],[f6859,f731])).
fof(f15301,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))) ) | $spl2),
  inference(subsumption_resolution,[],[f15292,f790])).
fof(f15292,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))) ) | $spl2),
  inference(resolution,[],[f6858,f727])).
fof(f15300,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))) ) | $spl2),
  inference(subsumption_resolution,[],[f15291,f790])).
fof(f15291,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))) ) | $spl2),
  inference(resolution,[],[f6858,f731])).
fof(f14987,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) | $spl2),
  inference(resolution,[],[f1329,f749])).
fof(f14853,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl13),
  inference(resolution,[],[f1033,f836])).
fof(f14972,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK12)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl13),
  inference(forward_demodulation,[],[f14852,f892])).
fof(f14852,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))),s('type.integer.int',sK12)))) | $spl13),
  inference(resolution,[],[f1033,f833])).
fof(f14151,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl2),
  inference(resolution,[],[f907,f749])).
fof(f13912,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) | $spl5),
  inference(resolution,[],[f1040,f866])).
fof(f13910,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl5),
  inference(resolution,[],[f1040,f860])).
fof(f13774,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK12)))) | $spl7),
  inference(resolution,[],[f1027,f855])).
fof(f13768,plain,(
  ( ! [X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | p(s(bool,X1))) ) | $spl5),
  inference(subsumption_resolution,[],[f13762,f790])).
fof(f13762,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | p(s(bool,X1))) ) | $spl5),
  inference(resolution,[],[f1681,f727])).
fof(f13767,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | p(s(bool,X0))) ) | $spl5),
  inference(subsumption_resolution,[],[f13761,f790])).
fof(f13761,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | p(s(bool,X0))) ) | $spl5),
  inference(resolution,[],[f1681,f731])).
fof(f13636,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK12)))) | $spl5),
  inference(resolution,[],[f910,f833])).
fof(f13628,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))) ) | $spl5),
  inference(subsumption_resolution,[],[f13622,f790])).
fof(f13622,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))) ) | $spl5),
  inference(resolution,[],[f1680,f727])).
fof(f13627,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))) ) | $spl5),
  inference(subsumption_resolution,[],[f13621,f790])).
fof(f13621,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))) ) | $spl5),
  inference(resolution,[],[f1680,f731])).
fof(f13620,plain,(
  ( ! [X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | p(s(bool,X1))) ) | $spl5),
  inference(subsumption_resolution,[],[f13614,f790])).
fof(f13614,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | p(s(bool,X1))) ) | $spl5),
  inference(resolution,[],[f1258,f727])).
fof(f13619,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | p(s(bool,X0))) ) | $spl5),
  inference(subsumption_resolution,[],[f13613,f790])).
fof(f13613,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | p(s(bool,X0))) ) | $spl5),
  inference(resolution,[],[f1258,f731])).
fof(f13612,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))) ) | $spl5),
  inference(subsumption_resolution,[],[f13606,f790])).
fof(f13606,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))) ) | $spl5),
  inference(resolution,[],[f1257,f727])).
fof(f13611,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))) ) | $spl5),
  inference(subsumption_resolution,[],[f13605,f790])).
fof(f13605,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))) ) | $spl5),
  inference(resolution,[],[f1257,f731])).
fof(f13604,plain,(
  $false | ($spl11 | $spl19)),
  inference(subsumption_resolution,[],[f4688,f1299])).
fof(f4688,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl11),
  inference(resolution,[],[f1032,f707])).
fof(f13603,plain,(
  $false | ($spl11 | $spl19)),
  inference(subsumption_resolution,[],[f4689,f1299])).
fof(f4689,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl11),
  inference(resolution,[],[f1032,f709])).
fof(f13602,plain,(
  $false | ($spl11 | $spl19)),
  inference(subsumption_resolution,[],[f4690,f1299])).
fof(f4690,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl11),
  inference(resolution,[],[f1032,f748])).
fof(f13601,plain,(
  $false | ($spl11 | $spl51)),
  inference(subsumption_resolution,[],[f4693,f2043])).
fof(f4693,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl11),
  inference(resolution,[],[f1032,f767])).
fof(f13600,plain,(
  $false | ($spl11 | $spl51)),
  inference(subsumption_resolution,[],[f4694,f2043])).
fof(f4694,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl11),
  inference(resolution,[],[f1032,f767])).
fof(f13597,plain,(
  $false | ($spl10 | $spl15)),
  inference(subsumption_resolution,[],[f10214,f1031])).
fof(f10214,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | $spl15),
  inference(resolution,[],[f1165,f834])).
fof(f834,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0))))) )),
  inference(cnf_transformation,[],[f577])).
fof(f13596,plain,(
  $false | ($spl12 | $spl15)),
  inference(subsumption_resolution,[],[f10369,f1034])).
fof(f10369,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) | $spl15),
  inference(forward_demodulation,[],[f10368,f892])).
fof(f10368,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK12)))) | $spl15),
  inference(forward_demodulation,[],[f10227,f885])).
fof(f10227,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))),s('type.integer.int',sK12)))) | $spl15),
  inference(resolution,[],[f1165,f839])).
fof(f13595,plain,(
  $false | ($spl12 | $spl41)),
  inference(subsumption_resolution,[],[f8061,f1034])).
fof(f8061,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) | $spl41),
  inference(forward_demodulation,[],[f7935,f885])).
fof(f7935,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK12)))) | $spl41),
  inference(resolution,[],[f1722,f839])).
fof(f13594,plain,(
  $false | ($spl12 | $spl76)),
  inference(subsumption_resolution,[],[f8242,f1034])).
fof(f8242,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) | $spl76),
  inference(resolution,[],[f2487,f708])).
fof(f13593,plain,(
  $false | ($spl12 | $spl76)),
  inference(subsumption_resolution,[],[f8243,f1034])).
fof(f8243,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) | $spl76),
  inference(resolution,[],[f2487,f710])).
fof(f13592,plain,(
  $false | ($spl12 | $spl76)),
  inference(subsumption_resolution,[],[f8244,f1034])).
fof(f8244,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) | $spl76),
  inference(resolution,[],[f2487,f747])).
fof(f12938,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl51),
  inference(resolution,[],[f2043,f836])).
fof(f13591,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | $spl51),
  inference(forward_demodulation,[],[f13071,f843])).
fof(f13071,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl51),
  inference(forward_demodulation,[],[f12946,f886])).
fof(f12946,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl51),
  inference(resolution,[],[f2043,f833])).
fof(f13588,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | $spl12),
  inference(forward_demodulation,[],[f11797,f852])).
fof(f11797,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | $spl12),
  inference(superposition,[],[f1034,f848])).
fof(f13584,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | sP11(X0,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) | $spl12),
  inference(forward_demodulation,[],[f13114,f852])).
fof(f13114,plain,(
  ( ! [X0] : (sP11(X0,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) ) | $spl12),
  inference(superposition,[],[f11778,f848])).
fof(f13582,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | sP11('const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)),X0)) ) | $spl12),
  inference(forward_demodulation,[],[f13125,f852])).
fof(f13125,plain,(
  ( ! [X0] : (sP11('const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)),X0) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) ) | $spl12),
  inference(superposition,[],[f11779,f848])).
fof(f13428,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl19),
  inference(resolution,[],[f1299,f861])).
fof(f13581,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | $spl19),
  inference(forward_demodulation,[],[f13561,f871])).
fof(f13561,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl19),
  inference(forward_demodulation,[],[f13435,f886])).
fof(f13435,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl19),
  inference(resolution,[],[f1299,f855])).
fof(f13575,plain,(
  $false | ($spl3 | $spl23)),
  inference(subsumption_resolution,[],[f1303,f1731])).
fof(f13209,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK12)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | ($spl3 | $spl33)),
  inference(resolution,[],[f5274,f833])).
fof(f12792,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl50),
  inference(resolution,[],[f2044,f749])).
fof(f12026,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl6),
  inference(resolution,[],[f2775,f836])).
fof(f12025,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int',sK13)))) | $spl6),
  inference(resolution,[],[f2775,f833])).
fof(f11707,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) | $spl12),
  inference(resolution,[],[f1034,f749])).
fof(f11517,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | $spl10),
  inference(resolution,[],[f1031,f749])).
fof(f11387,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl4),
  inference(resolution,[],[f10705,f861])).
fof(f11639,plain,(
  $false | ($spl10 | $spl18)),
  inference(subsumption_resolution,[],[f11516,f1298])).
fof(f11516,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl10),
  inference(resolution,[],[f1031,f747])).
fof(f11638,plain,(
  $false | ($spl10 | $spl18)),
  inference(subsumption_resolution,[],[f11515,f1298])).
fof(f11515,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl10),
  inference(resolution,[],[f1031,f710])).
fof(f11637,plain,(
  $false | ($spl10 | $spl18)),
  inference(subsumption_resolution,[],[f11514,f1298])).
fof(f11514,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl10),
  inference(resolution,[],[f1031,f708])).
fof(f11385,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int',sK13)))) | $spl4),
  inference(resolution,[],[f10705,f855])).
fof(f11258,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | $spl6),
  inference(resolution,[],[f1028,f866])).
fof(f11256,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl6),
  inference(resolution,[],[f1028,f860])).
fof(f11240,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | ($spl0 | $spl57)),
  inference(subsumption_resolution,[],[f11239,f2050])).
fof(f11239,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl0),
  inference(forward_demodulation,[],[f11186,f853])).
fof(f11186,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl0),
  inference(superposition,[],[f905,f880])).
fof(f11238,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | ($spl0 | $spl57)),
  inference(subsumption_resolution,[],[f11237,f2050])).
fof(f11237,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(forward_demodulation,[],[f11185,f853])).
fof(f11185,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(superposition,[],[f905,f879])).
fof(f11236,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | ($spl0 | $spl57)),
  inference(subsumption_resolution,[],[f11235,f2050])).
fof(f11235,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl0),
  inference(forward_demodulation,[],[f11184,f853])).
fof(f11184,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl0),
  inference(superposition,[],[f905,f878])).
fof(f11234,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | ($spl0 | $spl57)),
  inference(subsumption_resolution,[],[f11233,f2050])).
fof(f11233,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(forward_demodulation,[],[f11183,f853])).
fof(f11183,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(superposition,[],[f905,f877])).
fof(f11076,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(resolution,[],[f905,f749])).
fof(f11008,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | $spl3),
  inference(forward_demodulation,[],[f11007,f885])).
fof(f11007,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))))) | $spl3),
  inference(forward_demodulation,[],[f10883,f843])).
fof(f10883,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))))) | $spl3),
  inference(resolution,[],[f908,f836])).
fof(f10711,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | $spl4),
  inference(resolution,[],[f909,f841])).
fof(f8080,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | $spl43),
  inference(resolution,[],[f1723,f841])).
fof(f10679,plain,(
  $false | ($spl3 | $spl41)),
  inference(subsumption_resolution,[],[f1736,f1722])).
fof(f10678,plain,(
  $false | ($spl3 | $spl41)),
  inference(subsumption_resolution,[],[f1737,f1722])).
fof(f1737,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) | $spl3),
  inference(resolution,[],[f908,f767])).
fof(f10677,plain,(
  $false | ($spl3 | $spl41)),
  inference(subsumption_resolution,[],[f5145,f1722])).
fof(f5145,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) | $spl3),
  inference(resolution,[],[f908,f767])).
fof(f10676,plain,(
  $false | ($spl3 | $spl41)),
  inference(subsumption_resolution,[],[f5146,f1722])).
fof(f5146,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) | $spl3),
  inference(resolution,[],[f908,f767])).
fof(f7754,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | $spl38),
  inference(resolution,[],[f1596,f861])).
fof(f10395,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK12)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | $spl41),
  inference(forward_demodulation,[],[f8058,f885])).
fof(f8058,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))))))),s('type.integer.int',sK12)))) | $spl41),
  inference(forward_demodulation,[],[f7933,f843])).
fof(f7933,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))))))),s('type.integer.int',sK12)))) | $spl41),
  inference(resolution,[],[f1722,f833])).
fof(f7613,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | $spl37),
  inference(resolution,[],[f1594,f749])).
fof(f5716,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | $spl16),
  inference(resolution,[],[f1166,f836])).
fof(f10661,plain,(
  $false | $spl72),
  inference(subsumption_resolution,[],[f10551,f2097])).
fof(f10551,plain,(
  ( ! [X288,X287] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X287),s('type.integer.int',X288))))) ) | $spl72),
  inference(resolution,[],[f2097,f766])).
fof(f10660,plain,(
  $false | $spl72),
  inference(subsumption_resolution,[],[f10550,f2097])).
fof(f10550,plain,(
  ( ! [X285,X286] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X285),s('type.integer.int',X286))))) ) | $spl72),
  inference(resolution,[],[f2097,f766])).
fof(f10659,plain,(
  $false | $spl72),
  inference(subsumption_resolution,[],[f10548,f2097])).
fof(f10548,plain,(
  ( ! [X282,X281] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X281),s('type.integer.int',X282))))) ) | $spl72),
  inference(resolution,[],[f2097,f747])).
fof(f10658,plain,(
  $false | $spl72),
  inference(subsumption_resolution,[],[f10547,f2097])).
fof(f10547,plain,(
  ( ! [X280,X279] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X279),s('type.integer.int',X280))))) ) | $spl72),
  inference(resolution,[],[f2097,f710])).
fof(f10657,plain,(
  $false | $spl72),
  inference(subsumption_resolution,[],[f10546,f2097])).
fof(f10546,plain,(
  ( ! [X277,X278] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X277),s('type.integer.int',X278))))) ) | $spl72),
  inference(resolution,[],[f2097,f708])).
fof(f10656,plain,(
  $false | $spl72),
  inference(subsumption_resolution,[],[f10545,f2097])).
fof(f10545,plain,(
  ( ! [X275,X276] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X275),s('type.integer.int',X276))))) ) | $spl72),
  inference(resolution,[],[f2097,f698])).
fof(f10655,plain,(
  $false | $spl72),
  inference(subsumption_resolution,[],[f10544,f2097])).
fof(f10544,plain,(
  ( ! [X273,X274] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X273),s('type.integer.int',X274))))) ) | $spl72),
  inference(resolution,[],[f2097,f698])).
fof(f10654,plain,(
  $false | $spl72),
  inference(subsumption_resolution,[],[f10520,f2097])).
fof(f10520,plain,(
  ( ! [X243,X242] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X242),s('type.integer.int',X243))))) ) | $spl72),
  inference(resolution,[],[f2097,f747])).
fof(f10653,plain,(
  $false | $spl72),
  inference(subsumption_resolution,[],[f10519,f2097])).
fof(f10519,plain,(
  ( ! [X241,X240] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X240),s('type.integer.int',X241))))) ) | $spl72),
  inference(resolution,[],[f2097,f710])).
fof(f10652,plain,(
  $false | $spl72),
  inference(subsumption_resolution,[],[f10518,f2097])).
fof(f10518,plain,(
  ( ! [X239,X238] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X238),s('type.integer.int',X239))))) ) | $spl72),
  inference(resolution,[],[f2097,f708])).
fof(f10651,plain,(
  $false | $spl72),
  inference(subsumption_resolution,[],[f10513,f2097])).
fof(f10513,plain,(
  ( ! [X232] : (~p(s(bool,X232))) ) | $spl72),
  inference(resolution,[],[f2097,f819])).
fof(f819,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f438])).
fof(f438,plain,(
  ! [X0] : ((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f345])).
fof(f345,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) => ~p(s(bool,X0)))),
  inference(flattening,[],[f344])).
fof(f344,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) => ~p(s(bool,X0)))),
  inference(rectify,[],[f165])).
fof(f165,axiom,(
  ! [X12] : ((p(s(bool,X12)) => p(s(bool,'F_0'))) => ~p(s(bool,X12)))),
  file('thm.integer.INT_ABS_LE|split|3','thm.bool.IMP_F')).
fof(f10650,plain,(
  $false | $spl72),
  inference(subsumption_resolution,[],[f10507,f2097])).
fof(f10507,plain,(
  ( ! [X226] : (~p(s(bool,X226))) ) | $spl72),
  inference(resolution,[],[f2097,f785])).
fof(f785,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f554])).
fof(f554,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f553])).
fof(f553,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f431])).
fof(f431,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f326])).
fof(f326,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f325])).
fof(f325,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f164])).
fof(f164,axiom,(
  ! [X12] : ((p(s(bool,X12)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X12)))),
  file('thm.integer.INT_ABS_LE|split|3','thm.bool.IMP_CLAUSES|split|4')).
fof(f10515,plain,(
  $false | $spl72),
  inference(resolution,[],[f2097,f829])).
fof(f10531,plain,(
  $false | ($spl43 | $spl72)),
  inference(resolution,[],[f2097,f1723])).
fof(f10532,plain,(
  $false | ($spl11 | $spl72)),
  inference(resolution,[],[f2097,f4686])).
fof(f4686,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) | $spl11),
  inference(resolution,[],[f1032,f840])).
fof(f10536,plain,(
  $false | ($spl11 | $spl72)),
  inference(resolution,[],[f2097,f1032])).
fof(f10537,plain,(
  $false | ($spl53 | $spl72)),
  inference(resolution,[],[f2097,f2046])).
fof(f10538,plain,(
  $false | ($spl41 | $spl72)),
  inference(resolution,[],[f2097,f1722])).
fof(f10539,plain,(
  $false | ($spl15 | $spl72)),
  inference(resolution,[],[f2097,f1165])).
fof(f10542,plain,(
  $false | ($spl57 | $spl72)),
  inference(resolution,[],[f2097,f2050])).
fof(f10563,plain,(
  $false | $spl72),
  inference(resolution,[],[f2097,f889])).
fof(f10568,plain,(
  $false | ($spl2 | $spl72)),
  inference(resolution,[],[f2097,f1329])).
fof(f10569,plain,(
  $false | ($spl37 | $spl72)),
  inference(resolution,[],[f2097,f1594])).
fof(f10420,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))) ) | $spl5),
  inference(subsumption_resolution,[],[f10414,f790])).
fof(f10414,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))) ) | $spl5),
  inference(resolution,[],[f1257,f727])).
fof(f10419,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))) ) | $spl5),
  inference(subsumption_resolution,[],[f10413,f790])).
fof(f10413,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))) ) | $spl5),
  inference(resolution,[],[f1257,f731])).
fof(f6782,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | $spl2),
  inference(resolution,[],[f1456,f836])).
fof(f10412,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | $spl11),
  inference(forward_demodulation,[],[f4819,f885])).
fof(f4819,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))))))),s('type.integer.int',sK13)))) | $spl11),
  inference(forward_demodulation,[],[f4695,f843])).
fof(f4695,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))))))),s('type.integer.int',sK13)))) | $spl11),
  inference(resolution,[],[f1032,f833])).
fof(f10411,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | $spl2),
  inference(forward_demodulation,[],[f6879,f852])).
fof(f6879,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | $spl2),
  inference(superposition,[],[f1456,f848])).
fof(f7925,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int',sK13)))) | $spl41),
  inference(resolution,[],[f1722,f833])).
fof(f8209,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl119),
  inference(cnf_transformation,[],[f8209_D])).
fof(f8209_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) <=> ~$spl119),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl119])])).
fof(f8059,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl41),
  inference(forward_demodulation,[],[f7934,f886])).
fof(f7934,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl41),
  inference(resolution,[],[f1722,f836])).
fof(f10409,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl138),
  inference(cnf_transformation,[],[f10409_D])).
fof(f10409_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f9098,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl50),
  inference(forward_demodulation,[],[f8980,f886])).
fof(f8980,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))))),s('type.integer.int',sK13)))) | $spl50),
  inference(resolution,[],[f2044,f841])).
fof(f5899,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | $spl26),
  inference(resolution,[],[f1452,f860])).
fof(f7880,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl114),
  inference(cnf_transformation,[],[f7880_D])).
fof(f7880_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f6034,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl26),
  inference(forward_demodulation,[],[f5912,f886])).
fof(f5912,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK12)))) | $spl26),
  inference(resolution,[],[f1452,f866])).
fof(f10407,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK12)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | $spl2),
  inference(forward_demodulation,[],[f6559,f885])).
fof(f6559,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))))))),s('type.integer.int',sK12)))) | $spl2),
  inference(forward_demodulation,[],[f6433,f871])).
fof(f6433,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))))))),s('type.integer.int',sK12)))) | $spl2),
  inference(resolution,[],[f1329,f855])).
fof(f7536,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl18),
  inference(forward_demodulation,[],[f7535,f871])).
fof(f7535,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl18),
  inference(forward_demodulation,[],[f7416,f885])).
fof(f7416,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl18),
  inference(resolution,[],[f1298,f860])).
fof(f7608,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK12)))) | $spl37),
  inference(resolution,[],[f1594,f855])).
fof(f10406,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl137),
  inference(cnf_transformation,[],[f10406_D])).
fof(f10406_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) <=> ~$spl137),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl137])])).
fof(f7743,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl37),
  inference(forward_demodulation,[],[f7616,f886])).
fof(f7616,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl37),
  inference(resolution,[],[f1594,f861])).
fof(f6032,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | $spl26),
  inference(forward_demodulation,[],[f6031,f871])).
fof(f6031,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl26),
  inference(forward_demodulation,[],[f5910,f885])).
fof(f5910,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl26),
  inference(resolution,[],[f1452,f860])).
fof(f6182,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl2),
  inference(forward_demodulation,[],[f6064,f886])).
fof(f6064,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK12)))) | $spl2),
  inference(resolution,[],[f907,f841])).
fof(f6427,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int',sK13)))) | $spl2),
  inference(resolution,[],[f1329,f855])).
fof(f9684,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl131),
  inference(cnf_transformation,[],[f9684_D])).
fof(f9684_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) <=> ~$spl131),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl131])])).
fof(f6561,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl2),
  inference(forward_demodulation,[],[f6435,f886])).
fof(f6435,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl2),
  inference(resolution,[],[f1329,f861])).
fof(f7405,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | $spl18),
  inference(resolution,[],[f1298,f860])).
fof(f8364,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl120),
  inference(cnf_transformation,[],[f8364_D])).
fof(f8364_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f7539,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl18),
  inference(forward_demodulation,[],[f7418,f886])).
fof(f7418,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))))),s('type.integer.int',sK13)))) | $spl18),
  inference(resolution,[],[f1298,f866])).
fof(f10396,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | $spl37),
  inference(forward_demodulation,[],[f7739,f885])).
fof(f7739,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))))))),s('type.integer.int',sK13)))) | $spl37),
  inference(forward_demodulation,[],[f7614,f871])).
fof(f7614,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))))))),s('type.integer.int',sK13)))) | $spl37),
  inference(resolution,[],[f1594,f855])).
fof(f10392,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | $spl76),
  inference(forward_demodulation,[],[f8340,f852])).
fof(f8340,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | $spl76),
  inference(superposition,[],[f2487,f848])).
fof(f10391,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | $spl11),
  inference(forward_demodulation,[],[f9310,f852])).
fof(f9310,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | $spl11),
  inference(superposition,[],[f4686,f848])).
fof(f905,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(cnf_transformation,[],[f905_D])).
fof(f905_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f8089,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl43),
  inference(resolution,[],[f1723,f836])).
fof(f10362,plain,(
  $false | ($spl4 | $spl15)),
  inference(subsumption_resolution,[],[f10215,f909])).
fof(f9564,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl4),
  inference(resolution,[],[f5292,f861])).
fof(f5292,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | $spl4),
  inference(resolution,[],[f912,f864])).
fof(f912,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | $spl4),
  inference(resolution,[],[f909,f708])).
fof(f9683,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl4),
  inference(forward_demodulation,[],[f9562,f892])).
fof(f9562,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))),s('type.integer.int',sK13)))) | $spl4),
  inference(resolution,[],[f5292,f855])).
fof(f9217,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl11),
  inference(resolution,[],[f4686,f836])).
fof(f9336,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK12)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl11),
  inference(forward_demodulation,[],[f9216,f892])).
fof(f9216,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))))),s('type.integer.int',sK12)))) | $spl11),
  inference(resolution,[],[f4686,f833])).
fof(f8971,plain,(
  $false | ($spl2 | $spl99)),
  inference(subsumption_resolution,[],[f4364,f1456])).
fof(f8400,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | sP11('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)),X0)) ) | $spl76),
  inference(forward_demodulation,[],[f8395,f852])).
fof(f8395,plain,(
  ( ! [X0] : (sP11('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)),X0) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) ) | $spl76),
  inference(superposition,[],[f8320,f848])).
fof(f8387,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | sP11(X0,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) | $spl76),
  inference(forward_demodulation,[],[f8382,f852])).
fof(f8382,plain,(
  ( ! [X0] : (sP11(X0,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) ) | $spl76),
  inference(superposition,[],[f8319,f848])).
fof(f8251,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK12)))) | $spl76),
  inference(resolution,[],[f2487,f866])).
fof(f8363,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl76),
  inference(forward_demodulation,[],[f8249,f892])).
fof(f8249,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl76),
  inference(resolution,[],[f2487,f860])).
fof(f8208,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl43),
  inference(forward_demodulation,[],[f8088,f892])).
fof(f8088,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))),s('type.integer.int',sK13)))) | $spl43),
  inference(resolution,[],[f1723,f833])).
fof(f7767,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int',sK13)))) | $spl38),
  inference(resolution,[],[f1596,f866])).
fof(f7879,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl38),
  inference(forward_demodulation,[],[f7765,f892])).
fof(f7765,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl38),
  inference(resolution,[],[f1596,f860])).
fof(f7125,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | sP11('const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)),X0)) ) | $spl2),
  inference(forward_demodulation,[],[f7120,f852])).
fof(f7120,plain,(
  ( ! [X0] : (sP11('const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)),X0) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) ) | $spl2),
  inference(superposition,[],[f6859,f848])).
fof(f6924,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | sP11(X0,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) | $spl2),
  inference(forward_demodulation,[],[f6919,f852])).
fof(f6919,plain,(
  ( ! [X0] : (sP11(X0,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))) ) | $spl2),
  inference(superposition,[],[f6858,f848])).
fof(f6790,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK12)))) | $spl2),
  inference(resolution,[],[f1456,f841])).
fof(f2487,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl76),
  inference(cnf_transformation,[],[f2487_D])).
fof(f2487_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f6787,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl2),
  inference(resolution,[],[f1456,f749])).
fof(f6432,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) | $spl2),
  inference(resolution,[],[f1329,f749])).
fof(f6366,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl1),
  inference(forward_demodulation,[],[f6365,f892])).
fof(f6365,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl1),
  inference(forward_demodulation,[],[f6364,f885])).
fof(f6364,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))) | $spl1),
  inference(forward_demodulation,[],[f6314,f853])).
fof(f6314,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))) | $spl1),
  inference(superposition,[],[f906,f890])).
fof(f6363,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl1),
  inference(forward_demodulation,[],[f6313,f853])).
fof(f6313,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl1),
  inference(superposition,[],[f906,f880])).
fof(f6362,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl1),
  inference(forward_demodulation,[],[f6312,f853])).
fof(f6312,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl1),
  inference(superposition,[],[f906,f879])).
fof(f6361,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl1),
  inference(forward_demodulation,[],[f6311,f853])).
fof(f6311,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl1),
  inference(superposition,[],[f906,f878])).
fof(f6360,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl1),
  inference(forward_demodulation,[],[f6310,f853])).
fof(f6310,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl1),
  inference(superposition,[],[f906,f877])).
fof(f6359,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl1),
  inference(forward_demodulation,[],[f6358,f886])).
fof(f6358,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))) | $spl1),
  inference(forward_demodulation,[],[f6309,f853])).
fof(f6309,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))) | $spl1),
  inference(superposition,[],[f906,f875])).
fof(f6357,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl1),
  inference(forward_demodulation,[],[f6356,f886])).
fof(f6356,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))) | $spl1),
  inference(forward_demodulation,[],[f6308,f853])).
fof(f6308,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))) | $spl1),
  inference(superposition,[],[f906,f873])).
fof(f6353,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl1),
  inference(forward_demodulation,[],[f6352,f892])).
fof(f6352,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl1),
  inference(forward_demodulation,[],[f6351,f885])).
fof(f6351,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl1),
  inference(forward_demodulation,[],[f6305,f853])).
fof(f6305,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl1),
  inference(superposition,[],[f906,f845])).
fof(f6350,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl1),
  inference(subsumption_resolution,[],[f6349,f889])).
fof(f6349,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl1),
  inference(forward_demodulation,[],[f6348,f883])).
fof(f6348,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl1),
  inference(forward_demodulation,[],[f6347,f882])).
fof(f6347,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl1),
  inference(forward_demodulation,[],[f6346,f883])).
fof(f6346,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl1),
  inference(forward_demodulation,[],[f6304,f853])).
fof(f6304,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl1),
  inference(superposition,[],[f906,f825])).
fof(f6345,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl1),
  inference(forward_demodulation,[],[f6299,f853])).
fof(f6299,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl1),
  inference(superposition,[],[f906,f874])).
fof(f6343,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl1),
  inference(forward_demodulation,[],[f6295,f852])).
fof(f6295,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | $spl1),
  inference(superposition,[],[f906,f848])).
fof(f6337,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | $spl1),
  inference(forward_demodulation,[],[f6336,f843])).
fof(f6336,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl1),
  inference(forward_demodulation,[],[f6199,f886])).
fof(f6199,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl1),
  inference(resolution,[],[f906,f833])).
fof(f6191,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl1),
  inference(resolution,[],[f906,f836])).
fof(f6029,plain,(
  $false | ($spl3 | $spl26)),
  inference(subsumption_resolution,[],[f5906,f908])).
fof(f5906,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl26),
  inference(resolution,[],[f1452,f750])).
fof(f5724,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int',sK13)))) | $spl16),
  inference(resolution,[],[f1166,f841])).
fof(f5721,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl16),
  inference(resolution,[],[f1166,f749])).
fof(f5609,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl0),
  inference(forward_demodulation,[],[f5608,f892])).
fof(f5608,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl0),
  inference(forward_demodulation,[],[f5607,f885])).
fof(f5607,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))) | $spl0),
  inference(forward_demodulation,[],[f5560,f853])).
fof(f5560,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))) | $spl0),
  inference(superposition,[],[f905,f890])).
fof(f5606,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl0),
  inference(forward_demodulation,[],[f5559,f853])).
fof(f5559,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl0),
  inference(superposition,[],[f905,f880])).
fof(f5605,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(forward_demodulation,[],[f5558,f853])).
fof(f5558,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(superposition,[],[f905,f879])).
fof(f5604,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl0),
  inference(forward_demodulation,[],[f5557,f853])).
fof(f5557,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl0),
  inference(superposition,[],[f905,f878])).
fof(f5603,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(forward_demodulation,[],[f5556,f853])).
fof(f5556,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(superposition,[],[f905,f877])).
fof(f5602,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl0),
  inference(forward_demodulation,[],[f5601,f886])).
fof(f5601,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))) | $spl0),
  inference(forward_demodulation,[],[f5555,f853])).
fof(f5555,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))) | $spl0),
  inference(superposition,[],[f905,f875])).
fof(f5600,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl0),
  inference(forward_demodulation,[],[f5599,f886])).
fof(f5599,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))) | $spl0),
  inference(forward_demodulation,[],[f5554,f853])).
fof(f5554,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))) | $spl0),
  inference(superposition,[],[f905,f873])).
fof(f5596,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl0),
  inference(forward_demodulation,[],[f5595,f892])).
fof(f5595,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl0),
  inference(forward_demodulation,[],[f5594,f885])).
fof(f5594,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl0),
  inference(forward_demodulation,[],[f5551,f853])).
fof(f5551,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl0),
  inference(superposition,[],[f905,f845])).
fof(f5593,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl0),
  inference(subsumption_resolution,[],[f5592,f889])).
fof(f5592,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl0),
  inference(forward_demodulation,[],[f5591,f883])).
fof(f5591,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(forward_demodulation,[],[f5590,f882])).
fof(f5590,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(forward_demodulation,[],[f5589,f883])).
fof(f5589,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(forward_demodulation,[],[f5550,f853])).
fof(f5550,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(superposition,[],[f905,f825])).
fof(f5588,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl0),
  inference(forward_demodulation,[],[f5545,f853])).
fof(f5545,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl0),
  inference(superposition,[],[f905,f874])).
fof(f5544,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl0),
  inference(superposition,[],[f905,f876])).
fof(f5587,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl0),
  inference(forward_demodulation,[],[f5541,f852])).
fof(f5541,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | $spl0),
  inference(superposition,[],[f905,f848])).
fof(f5584,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(forward_demodulation,[],[f5452,f886])).
fof(f5452,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))),s('type.integer.int',sK13)))) | $spl0),
  inference(resolution,[],[f905,f841])).
fof(f5449,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(resolution,[],[f905,f749])).
fof(f5291,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl4),
  inference(resolution,[],[f912,f861])).
fof(f5289,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int',sK13)))) | $spl4),
  inference(resolution,[],[f912,f855])).
fof(f4999,plain,(
  $false | ($spl0 | $spl17)),
  inference(subsumption_resolution,[],[f2036,f1167])).
fof(f2036,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl0),
  inference(forward_demodulation,[],[f2035,f892])).
fof(f2035,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(forward_demodulation,[],[f1900,f885])).
fof(f1900,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | $spl0),
  inference(resolution,[],[f905,f837])).
fof(f4998,plain,(
  $false | ($spl5 | $spl6)),
  inference(subsumption_resolution,[],[f1043,f1028])).
fof(f1043,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | $spl5),
  inference(resolution,[],[f910,f750])).
fof(f4997,plain,(
  $false | ($spl5 | $spl33)),
  inference(subsumption_resolution,[],[f1045,f1462])).
fof(f4996,plain,(
  $false | ($spl5 | $spl6)),
  inference(subsumption_resolution,[],[f1178,f1028])).
fof(f1178,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | $spl5),
  inference(resolution,[],[f1040,f698])).
fof(f4995,plain,(
  $false | ($spl5 | $spl6)),
  inference(subsumption_resolution,[],[f1179,f1028])).
fof(f1179,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | $spl5),
  inference(resolution,[],[f1040,f698])).
fof(f4994,plain,(
  $false | ($spl5 | $spl33)),
  inference(subsumption_resolution,[],[f1183,f1462])).
fof(f1183,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | $spl5),
  inference(resolution,[],[f1040,f750])).
fof(f4993,plain,(
  $false | ($spl5 | $spl6)),
  inference(subsumption_resolution,[],[f1184,f1028])).
fof(f1184,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | $spl5),
  inference(resolution,[],[f1040,f766])).
fof(f4992,plain,(
  $false | ($spl5 | $spl6)),
  inference(subsumption_resolution,[],[f1185,f1028])).
fof(f1185,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | $spl5),
  inference(resolution,[],[f1040,f766])).
fof(f4991,plain,(
  $false | ($spl5 | $spl33)),
  inference(subsumption_resolution,[],[f1470,f1462])).
fof(f1470,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | $spl5),
  inference(resolution,[],[f1043,f707])).
fof(f4990,plain,(
  $false | ($spl5 | $spl33)),
  inference(subsumption_resolution,[],[f1471,f1462])).
fof(f1471,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | $spl5),
  inference(resolution,[],[f1043,f709])).
fof(f4989,plain,(
  $false | ($spl5 | $spl33)),
  inference(subsumption_resolution,[],[f1472,f1462])).
fof(f1472,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | $spl5),
  inference(resolution,[],[f1043,f748])).
fof(f4687,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK12)))) | $spl11),
  inference(resolution,[],[f1032,f833])).
fof(f4988,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl101),
  inference(cnf_transformation,[],[f4988_D])).
fof(f4988_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) <=> ~$spl101),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl101])])).
fof(f4822,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl11),
  inference(forward_demodulation,[],[f4696,f886])).
fof(f4696,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl11),
  inference(resolution,[],[f1032,f836])).
fof(f4986,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | $spl17),
  inference(forward_demodulation,[],[f4947,f852])).
fof(f4947,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | $spl17),
  inference(superposition,[],[f1167,f848])).
fof(f4972,plain,(
  $false | ($spl2 | $spl17)),
  inference(subsumption_resolution,[],[f4843,f907])).
fof(f4843,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl17),
  inference(resolution,[],[f1167,f838])).
fof(f4971,plain,(
  $false | ($spl4 | $spl17)),
  inference(subsumption_resolution,[],[f4842,f909])).
fof(f4842,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | $spl17),
  inference(resolution,[],[f1167,f837])).
fof(f4245,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl33),
  inference(resolution,[],[f1462,f836])).
fof(f4244,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int',sK13)))) | $spl33),
  inference(resolution,[],[f1462,f833])).
fof(f4091,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | $spl6),
  inference(resolution,[],[f1028,f866])).
fof(f4089,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl6),
  inference(resolution,[],[f1028,f860])).
fof(f4001,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl1),
  inference(forward_demodulation,[],[f4000,f892])).
fof(f4000,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl1),
  inference(forward_demodulation,[],[f3999,f885])).
fof(f3999,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))) | $spl1),
  inference(forward_demodulation,[],[f3948,f853])).
fof(f3948,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))) | $spl1),
  inference(superposition,[],[f906,f890])).
fof(f3998,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl1),
  inference(forward_demodulation,[],[f3947,f853])).
fof(f3947,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl1),
  inference(superposition,[],[f906,f880])).
fof(f3997,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl1),
  inference(forward_demodulation,[],[f3946,f853])).
fof(f3946,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl1),
  inference(superposition,[],[f906,f879])).
fof(f3996,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl1),
  inference(forward_demodulation,[],[f3945,f853])).
fof(f3945,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl1),
  inference(superposition,[],[f906,f878])).
fof(f3995,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl1),
  inference(forward_demodulation,[],[f3944,f853])).
fof(f3944,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl1),
  inference(superposition,[],[f906,f877])).
fof(f3994,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl1),
  inference(forward_demodulation,[],[f3993,f886])).
fof(f3993,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))) | $spl1),
  inference(forward_demodulation,[],[f3943,f853])).
fof(f3943,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))) | $spl1),
  inference(superposition,[],[f906,f875])).
fof(f3992,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl1),
  inference(forward_demodulation,[],[f3991,f886])).
fof(f3991,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))) | $spl1),
  inference(forward_demodulation,[],[f3942,f853])).
fof(f3942,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))) | $spl1),
  inference(superposition,[],[f906,f873])).
fof(f3988,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl1),
  inference(forward_demodulation,[],[f3987,f892])).
fof(f3987,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl1),
  inference(forward_demodulation,[],[f3986,f885])).
fof(f3986,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl1),
  inference(forward_demodulation,[],[f3939,f853])).
fof(f3939,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl1),
  inference(superposition,[],[f906,f845])).
fof(f3985,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl1),
  inference(subsumption_resolution,[],[f3984,f889])).
fof(f3984,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl1),
  inference(forward_demodulation,[],[f3983,f883])).
fof(f3983,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl1),
  inference(forward_demodulation,[],[f3982,f882])).
fof(f3982,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl1),
  inference(forward_demodulation,[],[f3981,f883])).
fof(f3981,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl1),
  inference(forward_demodulation,[],[f3938,f853])).
fof(f3938,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl1),
  inference(superposition,[],[f906,f825])).
fof(f3980,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl1),
  inference(forward_demodulation,[],[f3933,f853])).
fof(f3933,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl1),
  inference(superposition,[],[f906,f874])).
fof(f1167,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl17),
  inference(cnf_transformation,[],[f1167_D])).
fof(f1167_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f3932,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl1),
  inference(superposition,[],[f906,f876])).
fof(f3979,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl1),
  inference(forward_demodulation,[],[f3929,f852])).
fof(f3929,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | $spl1),
  inference(superposition,[],[f906,f848])).
fof(f3971,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | $spl1),
  inference(forward_demodulation,[],[f3970,f843])).
fof(f3970,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl1),
  inference(forward_demodulation,[],[f3833,f886])).
fof(f3833,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl1),
  inference(resolution,[],[f906,f833])).
fof(f3825,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl1),
  inference(resolution,[],[f906,f836])).
fof(f3681,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl4),
  inference(resolution,[],[f912,f861])).
fof(f3679,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int',sK13)))) | $spl4),
  inference(resolution,[],[f912,f855])).
fof(f3636,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | s('type.integer.int',sK13) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl2),
  inference(superposition,[],[f907,f876])).
fof(f3544,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl2),
  inference(resolution,[],[f907,f749])).
fof(f3401,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | $spl4),
  inference(resolution,[],[f909,f841])).
fof(f3297,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | s('type.integer.int',sK12) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) | $spl8),
  inference(forward_demodulation,[],[f3296,f1030])).
fof(f3296,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | s('type.integer.int',sK12) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) | $spl8),
  inference(forward_demodulation,[],[f3034,f852])).
fof(f3034,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | $spl8),
  inference(superposition,[],[f848,f1030])).
fof(f3295,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | $spl8),
  inference(forward_demodulation,[],[f3294,f1030])).
fof(f3294,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | s('type.integer.int',sK13) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) | $spl8),
  inference(forward_demodulation,[],[f3033,f852])).
fof(f3033,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | s('type.integer.int',sK13) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) | $spl8),
  inference(superposition,[],[f848,f1030])).
fof(f3293,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | s('type.integer.int',sK12) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) | $spl8),
  inference(forward_demodulation,[],[f3292,f1030])).
fof(f3292,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK13)))) | s('type.integer.int',sK12) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) | $spl8),
  inference(forward_demodulation,[],[f3032,f852])).
fof(f3032,plain,(
  s('type.integer.int',sK12) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | $spl8),
  inference(superposition,[],[f847,f1030])).
fof(f3289,plain,(
  s('type.integer.int',sK12) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | $spl8),
  inference(forward_demodulation,[],[f3288,f1030])).
fof(f3288,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | s('type.integer.int',sK13) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) | $spl8),
  inference(forward_demodulation,[],[f3031,f852])).
fof(f3031,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | s('type.integer.int',sK13) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))) | $spl8),
  inference(superposition,[],[f847,f1030])).
fof(f3225,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | $spl8),
  inference(forward_demodulation,[],[f3224,f1030])).
fof(f3224,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl8),
  inference(forward_demodulation,[],[f3223,f1030])).
fof(f3223,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl8),
  inference(forward_demodulation,[],[f2978,f883])).
fof(f2978,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl8),
  inference(superposition,[],[f825,f1030])).
fof(f3216,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl8),
  inference(forward_demodulation,[],[f3215,f1030])).
fof(f3215,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl8),
  inference(forward_demodulation,[],[f3214,f1030])).
fof(f3214,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl8),
  inference(forward_demodulation,[],[f2977,f883])).
fof(f2977,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl8),
  inference(superposition,[],[f825,f1030])).
fof(f3203,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK28(sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | $spl8),
  inference(forward_demodulation,[],[f3202,f1030])).
fof(f3202,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK28(sK13)))) | $spl8),
  inference(forward_demodulation,[],[f2974,f852])).
fof(f2974,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK28(sK13)))) | $spl8),
  inference(superposition,[],[f821,f1030])).
fof(f3197,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27(sK13)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl8),
  inference(forward_demodulation,[],[f3196,f1030])).
fof(f3196,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27(sK13)))))) | $spl8),
  inference(forward_demodulation,[],[f2973,f853])).
fof(f2973,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK27(sK13)))))) | $spl8),
  inference(superposition,[],[f820,f1030])).
fof(f2784,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | $spl6),
  inference(resolution,[],[f1028,f866])).
fof(f2782,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl6),
  inference(resolution,[],[f1028,f860])).
fof(f2694,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl1),
  inference(forward_demodulation,[],[f2693,f892])).
fof(f2693,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl1),
  inference(forward_demodulation,[],[f2692,f885])).
fof(f2692,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))) | $spl1),
  inference(forward_demodulation,[],[f2639,f853])).
fof(f2639,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))) | $spl1),
  inference(superposition,[],[f906,f890])).
fof(f2691,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl1),
  inference(forward_demodulation,[],[f2638,f853])).
fof(f2638,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl1),
  inference(superposition,[],[f906,f880])).
fof(f2690,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl1),
  inference(forward_demodulation,[],[f2637,f853])).
fof(f2637,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl1),
  inference(superposition,[],[f906,f879])).
fof(f2689,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl1),
  inference(forward_demodulation,[],[f2636,f853])).
fof(f2636,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl1),
  inference(superposition,[],[f906,f878])).
fof(f2688,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl1),
  inference(forward_demodulation,[],[f2635,f853])).
fof(f2635,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl1),
  inference(superposition,[],[f906,f877])).
fof(f2687,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl1),
  inference(forward_demodulation,[],[f2686,f886])).
fof(f2686,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))) | $spl1),
  inference(forward_demodulation,[],[f2634,f853])).
fof(f2634,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))) | $spl1),
  inference(superposition,[],[f906,f875])).
fof(f2685,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl1),
  inference(forward_demodulation,[],[f2684,f886])).
fof(f2684,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))) | $spl1),
  inference(forward_demodulation,[],[f2633,f853])).
fof(f2633,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))) | $spl1),
  inference(superposition,[],[f906,f873])).
fof(f2681,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl1),
  inference(forward_demodulation,[],[f2680,f892])).
fof(f2680,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl1),
  inference(forward_demodulation,[],[f2679,f885])).
fof(f2679,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl1),
  inference(forward_demodulation,[],[f2630,f853])).
fof(f2630,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl1),
  inference(superposition,[],[f906,f845])).
fof(f2678,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl1),
  inference(subsumption_resolution,[],[f2677,f889])).
fof(f2677,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl1),
  inference(forward_demodulation,[],[f2676,f883])).
fof(f2676,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl1),
  inference(forward_demodulation,[],[f2675,f882])).
fof(f2675,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl1),
  inference(forward_demodulation,[],[f2674,f883])).
fof(f2674,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl1),
  inference(forward_demodulation,[],[f2629,f853])).
fof(f2629,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl1),
  inference(superposition,[],[f906,f825])).
fof(f2673,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl1),
  inference(forward_demodulation,[],[f2624,f853])).
fof(f2624,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl1),
  inference(superposition,[],[f906,f874])).
fof(f2623,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl1),
  inference(superposition,[],[f906,f876])).
fof(f2672,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl1),
  inference(forward_demodulation,[],[f2620,f852])).
fof(f2620,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | $spl1),
  inference(superposition,[],[f906,f848])).
fof(f2664,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | $spl1),
  inference(forward_demodulation,[],[f2663,f843])).
fof(f2663,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl1),
  inference(forward_demodulation,[],[f2524,f886])).
fof(f2524,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl1),
  inference(resolution,[],[f906,f833])).
fof(f2516,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl1),
  inference(resolution,[],[f906,f836])).
fof(f2369,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))) | $spl4),
  inference(resolution,[],[f912,f861])).
fof(f2367,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int',sK13)))) | $spl4),
  inference(resolution,[],[f912,f855])).
fof(f2366,plain,(
  s('type.integer.int',sK12) = s('type.integer.int',sK13) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | $spl4),
  inference(resolution,[],[f912,f749])).
fof(f2324,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | s('type.integer.int',sK13) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl2),
  inference(superposition,[],[f907,f876])).
fof(f2232,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl2),
  inference(resolution,[],[f907,f749])).
fof(f2114,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | $spl4),
  inference(resolution,[],[f909,f841])).
fof(f2111,plain,(
  s('type.integer.int',sK12) = s('type.integer.int',sK13) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | $spl4),
  inference(resolution,[],[f909,f749])).
fof(f2086,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl0),
  inference(forward_demodulation,[],[f2085,f892])).
fof(f2085,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl0),
  inference(forward_demodulation,[],[f2084,f885])).
fof(f2084,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))) | $spl0),
  inference(forward_demodulation,[],[f2009,f853])).
fof(f2009,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))) | $spl0),
  inference(superposition,[],[f905,f890])).
fof(f2083,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl0),
  inference(forward_demodulation,[],[f2008,f853])).
fof(f2008,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl0),
  inference(superposition,[],[f905,f880])).
fof(f2080,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(forward_demodulation,[],[f2007,f853])).
fof(f2007,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(superposition,[],[f905,f879])).
fof(f2077,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl0),
  inference(forward_demodulation,[],[f2006,f853])).
fof(f2006,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl0),
  inference(superposition,[],[f905,f878])).
fof(f2072,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(forward_demodulation,[],[f2005,f853])).
fof(f2005,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK31('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK30('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(superposition,[],[f905,f877])).
fof(f2069,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl0),
  inference(forward_demodulation,[],[f2068,f886])).
fof(f2068,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))) | $spl0),
  inference(forward_demodulation,[],[f2004,f853])).
fof(f2004,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))) | $spl0),
  inference(superposition,[],[f905,f875])).
fof(f2067,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl0),
  inference(forward_demodulation,[],[f2066,f886])).
fof(f2066,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))) | $spl0),
  inference(forward_demodulation,[],[f2003,f853])).
fof(f2003,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))) | $spl0),
  inference(superposition,[],[f905,f873])).
fof(f2061,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl0),
  inference(forward_demodulation,[],[f2060,f892])).
fof(f2060,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl0),
  inference(forward_demodulation,[],[f2059,f885])).
fof(f2059,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl0),
  inference(forward_demodulation,[],[f2000,f853])).
fof(f2000,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl0),
  inference(superposition,[],[f905,f845])).
fof(f2056,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl0),
  inference(subsumption_resolution,[],[f2055,f889])).
fof(f2055,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl0),
  inference(forward_demodulation,[],[f2054,f883])).
fof(f2054,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(forward_demodulation,[],[f2053,f882])).
fof(f2053,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(forward_demodulation,[],[f2052,f883])).
fof(f2052,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(forward_demodulation,[],[f1999,f853])).
fof(f1999,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(superposition,[],[f905,f825])).
fof(f2049,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl0),
  inference(forward_demodulation,[],[f1994,f853])).
fof(f1994,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl0),
  inference(superposition,[],[f905,f874])).
fof(f1993,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl0),
  inference(superposition,[],[f905,f876])).
fof(f2042,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK12)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl0),
  inference(forward_demodulation,[],[f1990,f852])).
fof(f1990,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | $spl0),
  inference(superposition,[],[f905,f848])).
fof(f2037,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(forward_demodulation,[],[f1901,f886])).
fof(f1901,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))),s('type.integer.int',sK13)))) | $spl0),
  inference(resolution,[],[f905,f841])).
fof(f1898,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl0),
  inference(resolution,[],[f905,f749])).
fof(f1612,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) | $spl5),
  inference(resolution,[],[f1045,f841])).
fof(f1476,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl5),
  inference(resolution,[],[f1043,f861])).
fof(f1474,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK12)))) | $spl5),
  inference(resolution,[],[f1043,f855])).
fof(f1424,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | s('type.integer.int',sK13) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl2),
  inference(superposition,[],[f907,f876])).
fof(f1332,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | $spl2),
  inference(resolution,[],[f907,f749])).
fof(f1189,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) | $spl5),
  inference(resolution,[],[f1040,f866])).
fof(f1187,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | $spl5),
  inference(resolution,[],[f1040,f860])).
fof(f1048,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))) | $spl5),
  inference(resolution,[],[f910,f836])).
fof(f1047,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))))),s('type.integer.int',sK12)))) | $spl5),
  inference(resolution,[],[f910,f833])).
fof(f918,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK13))),s('type.integer.int',sK12)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))),s('type.integer.int',sK13)))) | $spl4),
  inference(resolution,[],[f909,f841])).
fof(f915,plain,(
  s('type.integer.int',sK12) = s('type.integer.int',sK13) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | $spl4),
  inference(resolution,[],[f909,f749])).
fof(f600,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))))),
  inference(cnf_transformation,[],[f464])).
fof(f464,plain,(
  (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))))) & (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13])],[f463])).
fof(f463,plain,(
  ? [X0,X1] : ((p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))))) | (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))) & (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0))))))),
  inference(flattening,[],[f462])).
fof(f462,plain,(
  ? [X0,X1] : ((p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))))) | (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))) & (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))))) | (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))))),
  inference(nnf_transformation,[],[f404])).
fof(f404,plain,(
  ? [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))))) <~> (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0))))))),
  inference(ennf_transformation,[],[f263])).
fof(f263,plain,(
  ~! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))))) <=> (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X6)))))))) <=> (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X6),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X7)))))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X7),s('type.integer.int',X6))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X6)))))))) <=> (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X6),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X7)))))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X7),s('type.integer.int',X6))))))),
  file('thm.integer.INT_ABS_LE|split|3','thm.integer.INT_ABS_LE|split|3')).
fof(f601,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))))),
  inference(cnf_transformation,[],[f464])).
fof(f906,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) | $spl1),
  inference(cnf_transformation,[],[f906_D])).
fof(f906_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12)))))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f602,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK12))))))))),
  inference(cnf_transformation,[],[f464])).
% SZS output end Proof for thm.integer
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 1180
% Passive clauses: 34294
% Generated clauses: 67998
% Final active clauses: 475
% Final passive clauses: 15485
% Input formulas: 261
% Initial clauses: 299
% 
% Selected by SInE selection: 104
% SInE iterations: 1
% 
% Duplicate literals: 1190
% Trivial inequalities: 236
% Fw subsumption resolutions: 3707
% Fw demodulations: 9744
% Bw demodulations: 225
% 
% Simple tautologies: 1799
% Equational tautologies: 4
% Forward subsumptions: 13598
% 
% Binary resolution: 20400
% Factoring: 1
% Forward superposition: 13091
% Backward superposition: 12885
% Self superposition: 52
% Equality resolution: 5
% 
% Split clauses: 775
% Split components: 1616
% SAT solver clauses: 934
% SAT solver unit clauses: 17
% SAT solver binary clauses: 187
% 
% Sat splits: 775
% Sat splitting refutations: 156
% 
% Memory used [KB]: 24178
% Time elapsed: 1.339 s
% ------------------------------
% ------------------------------
% Success in time 14.191 s
