Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.pred_set
% (102551)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 90
% Passive clauses: 3152
% Generated clauses: 16475
% Final active clauses: 78
% Final passive clauses: 2925
% Input formulas: 11
% Initial clauses: 21
% 
% Duplicate literals: 1017
% Trivial inequalities: 195
% Fw subsumption resolutions: 1904
% Fw demodulations: 1000
% Bw demodulations: 9
% 
% Simple tautologies: 323
% Equational tautologies: 98
% Forward subsumptions: 6373
% 
% Binary resolution: 91
% Forward superposition: 7158
% Backward superposition: 4289
% Self superposition: 347
% Equality factoring: 261
% Equality resolution: 14
% 
% Split clauses: 881
% Split components: 1774
% SAT solver clauses: 2050
% SAT solver unit clauses: 280
% SAT solver binary clauses: 622
% 
% Sat splits: 881
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 2042
% TWLsolver calls for satisfiability: 106
% 
% Memory used [KB]: 7803
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.pred_set
% (102565)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 40
% Passive clauses: 293
% Generated clauses: 6384
% Final active clauses: 36
% Final passive clauses: 245
% Input formulas: 11
% Initial clauses: 21
% 
% Duplicate literals: 1034
% Trivial inequalities: 56
% Fw subsumption resolutions: 295
% Fw demodulations: 50
% Bw demodulations: 6
% Condensations: 1545
% 
% Simple tautologies: 118
% Equational tautologies: 67
% Forward subsumptions: 2033
% 
% Binary resolution: 33
% Forward superposition: 1474
% Backward superposition: 1230
% Self superposition: 214
% Equality factoring: 196
% Equality resolution: 7
% 
% Split clauses: 74
% Split components: 149
% SAT solver clauses: 702
% SAT solver unit clauses: 17
% SAT solver binary clauses: 81
% 
% Sat splits: 74
% Sat splitting refutations: 14
% 
% TWLsolver clauses: 702
% TWLsolver calls for satisfiability: 81
% 
% Memory used [KB]: 4989
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.pred_set
% (102581)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 420
% Passive clauses: 5308
% Generated clauses: 141440
% Final active clauses: 78
% Final passive clauses: 588
% Input formulas: 11
% Initial clauses: 21
% 
% Duplicate literals: 7158
% Trivial inequalities: 876
% Fw subsumption resolutions: 12848
% Fw demodulations: 14207
% Bw demodulations: 4074
% Condensations: 22737
% Global subsumptions: 71
% 
% Simple tautologies: 1133
% Equational tautologies: 2585
% Forward subsumptions: 74046
% Bw demodulations to eq. taut.: 153
% 
% Unit resulting resolution: 7637
% Forward superposition: 33859
% Backward superposition: 24800
% Self superposition: 999
% Equality factoring: 594
% Equality resolution: 107
% 
% Split clauses: 435
% Split components: 875
% SAT solver clauses: 6480
% SAT solver unit clauses: 65
% SAT solver binary clauses: 253
% 
% Sat splits: 435
% Sat splitting refutations: 229
% 
% Memory used [KB]: 57696
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f19172,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f114,f117,f115,f1909,f1913,f1911,f7670,f7692,f7669,f7694,f7690,f7688,f7686,f15613,f15617,f15615,f15683,f15714,f15684,f15718,f15716,f15720,f16234,f16235,f16322,f16380,f16397,f16737,f16752,f16754,f16755,f16761,f17385,f17386,f17389,f16671,f16427,f16425,f17679,f17848,f17873,f17871,f17849,f17877,f17875,f18329,f18334,f18268,f16384,f16382,f18573,f18576,f17680,f18583,f16259,f16257,f18584,f16263,f16261,f19133,f19171])).
fof(f19171,plain,(
  $false | ($spl36 | $spl38 | $spl40 | $spl42)),
  inference(subsumption_resolution,[],[f19131,f18513])).
fof(f18513,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2)))) | $spl38),
  inference(resolution,[],[f16381,f48])).
fof(f48,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))) )),
  inference(cnf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))) & (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))))),
  inference(flattening,[],[f34])).
fof(f34,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))) & ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))))),
  inference(nnf_transformation,[],[f22])).
fof(f22,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(flattening,[],[f21])).
fof(f21,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6] : ! [X7,X8,X10] : (p(s(bool,'const.bool.IN_2'(s(X6,X10),s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X6,X10),s(fun(X6,bool),X7)))) & p(s(bool,'const.bool.IN_2'(s(X6,X10),s(fun(X6,bool),X8))))))),
  file('thm.pred_set.INTER_ASSOC','thm.pred_set.IN_INTER')).
fof(f16381,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) | $spl38),
  inference(cnf_transformation,[],[f16381_D])).
fof(f16381_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f19131,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2)))) | ($spl36 | $spl40 | $spl42)),
  inference(subsumption_resolution,[],[f19130,f61])).
fof(f61,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f44])).
fof(f44,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f33])).
fof(f33,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f20])).
fof(f20,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f19])).
fof(f19,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X6] : ! [X10] : (s(X6,X10) = s(X6,X10) <=> p(s(bool,'T_0')))),
  file('thm.pred_set.INTER_ASSOC','thm.bool.REFL_CLAUSE')).
fof(f19130,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2)))) | ($spl36 | $spl40 | $spl42)),
  inference(forward_demodulation,[],[f19103,f18631])).
fof(f18631,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3))) | $spl40),
  inference(subsumption_resolution,[],[f18630,f60])).
fof(f60,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f25])).
fof(f25,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.pred_set.INTER_ASSOC',aHL_FALSITY)).
fof(f18630,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3))) | $spl40),
  inference(superposition,[],[f16383,f46])).
fof(f46,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.pred_set.INTER_ASSOC',aHL_BOOL_CASES)).
fof(f16383,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3)))) | $spl40),
  inference(cnf_transformation,[],[f16383_D])).
fof(f16383_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3)))) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f19103,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2)))) | ($spl36 | $spl42)),
  inference(resolution,[],[f19090,f49])).
fof(f49,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f35])).
fof(f19090,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))) | ($spl36 | $spl42)),
  inference(subsumption_resolution,[],[f19089,f61])).
fof(f19089,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))) | ($spl36 | $spl42)),
  inference(forward_demodulation,[],[f19088,f17429])).
fof(f17429,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1))) | $spl42),
  inference(subsumption_resolution,[],[f17428,f60])).
fof(f17428,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1))) | $spl42),
  inference(superposition,[],[f16424,f46])).
fof(f16424,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1)))) | $spl42),
  inference(cnf_transformation,[],[f16424_D])).
fof(f16424_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1)))) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f19088,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1)))) | $spl36),
  inference(subsumption_resolution,[],[f19072,f60])).
fof(f19072,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1)))) | $spl36),
  inference(superposition,[],[f49,f16262])).
fof(f16262,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) | $spl36),
  inference(cnf_transformation,[],[f16262_D])).
fof(f16262_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f19133,plain,(
  $false | ($spl36 | $spl40 | $spl42 | $spl44)),
  inference(subsumption_resolution,[],[f19132,f59])).
fof(f59,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.INTER_ASSOC',aHL_TRUTH)).
fof(f19132,plain,(
  ~p(s(bool,'T_0')) | ($spl36 | $spl40 | $spl42 | $spl44)),
  inference(forward_demodulation,[],[f19131,f18409])).
fof(f18409,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2))) | $spl44),
  inference(subsumption_resolution,[],[f18408,f60])).
fof(f18408,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2))) | $spl44),
  inference(superposition,[],[f16426,f46])).
fof(f16426,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2)))) | $spl44),
  inference(cnf_transformation,[],[f16426_D])).
fof(f16426_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2)))) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f16261,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3))))) | $spl34),
  inference(cnf_transformation,[],[f16261_D])).
fof(f16261_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3))))) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f16263,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) | $spl37),
  inference(cnf_transformation,[],[f16263_D])).
fof(f16263_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) <=> ~$spl37),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl37])])).
fof(f18584,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)))))),
  inference(superposition,[],[f816,f46])).
fof(f816,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3))))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))))),
  inference(equality_resolution,[],[f68])).
fof(f68,plain,(
  ( ! [X2] : (s(fun(sK0,bool),X2) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),X2,sK0)),s(fun(sK0,bool),X2))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),X2,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)))))) )),
  inference(superposition,[],[f40,f43])).
fof(f43,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'const.bool.IN_2'(s(X0,sK5(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK5(X2,X1,X0)),s(fun(X0,bool),X2)))) )),
  inference(cnf_transformation,[],[f32])).
fof(f32,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (s(bool,'const.bool.IN_2'(s(X0,sK5(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK5(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5])],[f31])).
fof(f31,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(rectify,[],[f30])).
fof(f30,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(nnf_transformation,[],[f18])).
fof(f18,plain,(
  ! [X0,X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(flattening,[],[f17])).
fof(f17,plain,(
  ! [X0] : ! [X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X6] : ! [X7,X8] : (s(fun(X6,bool),X7) = s(fun(X6,bool),X8) <=> ! [X10] : s(bool,'const.bool.IN_2'(s(X6,X10),s(fun(X6,bool),X7))) = s(bool,'const.bool.IN_2'(s(X6,X10),s(fun(X6,bool),X8))))),
  file('thm.pred_set.INTER_ASSOC','thm.pred_set.EXTENSION')).
fof(f40,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))),
  inference(cnf_transformation,[],[f28])).
fof(f28,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f26])).
fof(f26,plain,(
  ? [X0,X1,X2,X3] : s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))) != s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X3)))),
  inference(ennf_transformation,[],[f14])).
fof(f14,plain,(
  ~! [X0,X1,X2,X3] : s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))) = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X3)))),
  inference(flattening,[],[f13])).
fof(f13,plain,(
  ~! [X0] : ! [X1,X2,X3] : s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))) = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))),s(fun(X0,bool),X3)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),X7),s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X9))))) = s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8))),s(fun(X6,bool),X9)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),X7),s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X9))))) = s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8))),s(fun(X6,bool),X9)))),
  file('thm.pred_set.INTER_ASSOC','thm.pred_set.INTER_ASSOC')).
fof(f16257,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3))))) | $spl30),
  inference(cnf_transformation,[],[f16257_D])).
fof(f16257_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3))))) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f16259,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) | $spl33),
  inference(cnf_transformation,[],[f16259_D])).
fof(f16259_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f18583,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)))))),
  inference(superposition,[],[f816,f46])).
fof(f17680,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)))))),
  inference(superposition,[],[f816,f46])).
fof(f18576,plain,(
  $false | ($spl34 | $spl41)),
  inference(subsumption_resolution,[],[f18575,f60])).
fof(f18575,plain,(
  p(s(bool,'F_0')) | ($spl34 | $spl41)),
  inference(forward_demodulation,[],[f17390,f18373])).
fof(f18373,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3))) | $spl41),
  inference(subsumption_resolution,[],[f18371,f61])).
fof(f18371,plain,(
  ~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3))) | $spl41),
  inference(superposition,[],[f16384,f46])).
fof(f17390,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3)))) | $spl34),
  inference(subsumption_resolution,[],[f17366,f61])).
fof(f17366,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3)))) | $spl34),
  inference(superposition,[],[f48,f16261])).
fof(f18573,plain,(
  $false | ($spl32 | $spl41)),
  inference(subsumption_resolution,[],[f18572,f60])).
fof(f18572,plain,(
  p(s(bool,'F_0')) | ($spl32 | $spl41)),
  inference(forward_demodulation,[],[f18544,f18373])).
fof(f18544,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3)))) | $spl32),
  inference(resolution,[],[f18332,f48])).
fof(f18332,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))) | $spl32),
  inference(subsumption_resolution,[],[f18312,f61])).
fof(f18312,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))) | $spl32),
  inference(superposition,[],[f48,f16258])).
fof(f16258,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) | $spl32),
  inference(cnf_transformation,[],[f16258_D])).
fof(f16258_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f16382,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) | $spl39),
  inference(cnf_transformation,[],[f16382_D])).
fof(f16382_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) <=> ~$spl39),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl39])])).
fof(f16384,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3)))) | $spl41),
  inference(cnf_transformation,[],[f16384_D])).
fof(f16384_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3)))) <=> ~$spl41),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl41])])).
fof(f18268,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) | $spl30),
  inference(subsumption_resolution,[],[f18250,f60])).
fof(f18250,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) | $spl30),
  inference(superposition,[],[f49,f16257])).
fof(f18334,plain,(
  $false | ($spl32 | $spl45)),
  inference(subsumption_resolution,[],[f18333,f60])).
fof(f18333,plain,(
  p(s(bool,'F_0')) | ($spl32 | $spl45)),
  inference(forward_demodulation,[],[f18332,f17624])).
fof(f17624,plain,(
  ( ! [X4] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X4)))))) ) | $spl45),
  inference(subsumption_resolution,[],[f17621,f61])).
fof(f17621,plain,(
  ( ! [X4] : (~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X4)))))) ) | $spl45),
  inference(superposition,[],[f17430,f46])).
fof(f17430,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X0))))))) ) | $spl45),
  inference(resolution,[],[f16427,f47])).
fof(f47,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))) )),
  inference(cnf_transformation,[],[f35])).
fof(f18329,plain,(
  $false | ($spl32 | $spl45)),
  inference(subsumption_resolution,[],[f18310,f59])).
fof(f18310,plain,(
  ~p(s(bool,'T_0')) | ($spl32 | $spl45)),
  inference(superposition,[],[f17590,f16258])).
fof(f17590,plain,(
  ( ! [X4,X3] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X3),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X4))))))))) ) | $spl45),
  inference(resolution,[],[f17430,f48])).
fof(f17875,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3))))) | $spl50),
  inference(cnf_transformation,[],[f17875_D])).
fof(f17875_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3))))) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f17877,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) | $spl53),
  inference(cnf_transformation,[],[f17877_D])).
fof(f17877_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) <=> ~$spl53),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl53])])).
fof(f17849,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)))))),
  inference(superposition,[],[f890,f46])).
fof(f890,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3))))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))))),
  inference(equality_resolution,[],[f69])).
fof(f69,plain,(
  ( ! [X3] : (s(fun(sK0,bool),X3) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | s(bool,'const.bool.IN_2'(s(sK0,sK5(X3,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),sK0)),s(fun(sK0,bool),X3))) != s(bool,'const.bool.IN_2'(s(sK0,sK5(X3,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)))))) )),
  inference(superposition,[],[f40,f43])).
fof(f17871,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3))))) | $spl46),
  inference(cnf_transformation,[],[f17871_D])).
fof(f17871_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3))))) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f17873,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) | $spl49),
  inference(cnf_transformation,[],[f17873_D])).
fof(f17873_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) <=> ~$spl49),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl49])])).
fof(f17848,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)))))),
  inference(superposition,[],[f890,f46])).
fof(f17679,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)))))),
  inference(superposition,[],[f816,f46])).
fof(f16425,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1)))) | $spl43),
  inference(cnf_transformation,[],[f16425_D])).
fof(f16425_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1)))) <=> ~$spl43),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl43])])).
fof(f16427,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2)))) | $spl45),
  inference(cnf_transformation,[],[f16427_D])).
fof(f16427_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2)))) <=> ~$spl45),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl45])])).
fof(f16671,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1)))) | $spl39),
  inference(subsumption_resolution,[],[f16655,f60])).
fof(f16655,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1)))) | $spl39),
  inference(superposition,[],[f49,f16428])).
fof(f16428,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl39),
  inference(subsumption_resolution,[],[f16422,f61])).
fof(f16422,plain,(
  ~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl39),
  inference(superposition,[],[f16382,f46])).
fof(f17389,plain,(
  $false | ($spl34 | $spl39)),
  inference(subsumption_resolution,[],[f17388,f60])).
fof(f17388,plain,(
  p(s(bool,'F_0')) | ($spl34 | $spl39)),
  inference(forward_demodulation,[],[f17387,f16428])).
fof(f17387,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) | $spl34),
  inference(subsumption_resolution,[],[f17365,f61])).
fof(f17365,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) | $spl34),
  inference(superposition,[],[f47,f16261])).
fof(f17386,plain,(
  $false | ($spl34 | $spl43)),
  inference(subsumption_resolution,[],[f17352,f59])).
fof(f17352,plain,(
  ~p(s(bool,'T_0')) | ($spl34 | $spl43)),
  inference(superposition,[],[f16527,f16261])).
fof(f16527,plain,(
  ( ! [X2,X1] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X1))),s(fun(sK0,bool),X2))))))) ) | $spl43),
  inference(resolution,[],[f16429,f47])).
fof(f16429,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X0))))))) ) | $spl43),
  inference(resolution,[],[f16425,f47])).
fof(f17385,plain,(
  $false | ($spl34 | $spl39)),
  inference(subsumption_resolution,[],[f17351,f59])).
fof(f17351,plain,(
  ~p(s(bool,'T_0')) | ($spl34 | $spl39)),
  inference(superposition,[],[f16398,f16261])).
fof(f16398,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),X0))))))) ) | $spl39),
  inference(resolution,[],[f16382,f47])).
fof(f16761,plain,(
  $false | ($spl30 | $spl43)),
  inference(subsumption_resolution,[],[f16323,f16562])).
fof(f16562,plain,(
  ( ! [X4] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X4)))))) ) | $spl43),
  inference(subsumption_resolution,[],[f16559,f61])).
fof(f16559,plain,(
  ( ! [X4] : (~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X4)))))) ) | $spl43),
  inference(superposition,[],[f16429,f46])).
fof(f16323,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) | $spl30),
  inference(backward_demodulation,[],[f16257,f816])).
fof(f16755,plain,(
  $false | ($spl30 | $spl43)),
  inference(subsumption_resolution,[],[f16379,f16562])).
fof(f16379,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) | $spl30),
  inference(trivial_inequality_removal,[],[f16358])).
fof(f16358,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) | s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3))))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3))))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | $spl30),
  inference(superposition,[],[f794,f16257])).
fof(f794,plain,(
  ( ! [X4,X5] : (s(bool,'const.bool.IN_2'(s(sK0,sK5(X4,X5,sK0)),s(fun(sK0,bool),X4))) != s(bool,'const.bool.IN_2'(s(sK0,sK5(X4,X5,sK0)),s(fun(sK0,bool),X5))) | s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),X4,sK0)),s(fun(sK0,bool),X4))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),X4,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3))))) | s(fun(sK0,bool),X5) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))) )),
  inference(superposition,[],[f68,f43])).
fof(f16754,plain,(
  $false | ($spl30 | $spl43)),
  inference(subsumption_resolution,[],[f16753,f40])).
fof(f16753,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3))) = s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | ($spl30 | $spl43)),
  inference(subsumption_resolution,[],[f16360,f16562])).
fof(f16360,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3))) = s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | $spl30),
  inference(superposition,[],[f43,f16257])).
fof(f16752,plain,(
  $false | ($spl30 | $spl43)),
  inference(subsumption_resolution,[],[f16474,f16562])).
fof(f16474,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) | $spl30),
  inference(trivial_inequality_removal,[],[f16460])).
fof(f16460,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) | s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3))))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3))))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | $spl30),
  inference(superposition,[],[f868,f16257])).
fof(f868,plain,(
  ( ! [X4,X5] : (s(bool,'const.bool.IN_2'(s(sK0,sK5(X4,X5,sK0)),s(fun(sK0,bool),X4))) != s(bool,'const.bool.IN_2'(s(sK0,sK5(X4,X5,sK0)),s(fun(sK0,bool),X5))) | s(bool,'const.bool.IN_2'(s(sK0,sK5(X4,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),sK0)),s(fun(sK0,bool),X4))) != s(bool,'const.bool.IN_2'(s(sK0,sK5(X4,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3))))) | s(fun(sK0,bool),X5) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))) )),
  inference(superposition,[],[f69,f43])).
fof(f16737,plain,(
  $false | ($spl37 | $spl43)),
  inference(trivial_inequality_removal,[],[f16717])).
fof(f16717,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl37 | $spl43)),
  inference(superposition,[],[f16263,f16562])).
fof(f16397,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1)))) | $spl39),
  inference(resolution,[],[f16382,f49])).
fof(f16380,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) | $spl30),
  inference(subsumption_resolution,[],[f16363,f60])).
fof(f16363,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) | $spl30),
  inference(superposition,[],[f49,f16257])).
fof(f16322,plain,(
  $false | ($spl33 | $spl37)),
  inference(subsumption_resolution,[],[f16321,f16259])).
fof(f16321,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) | $spl37),
  inference(trivial_inequality_removal,[],[f16320])).
fof(f16320,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) | $spl37),
  inference(superposition,[],[f16263,f46])).
fof(f16235,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)))))),
  inference(superposition,[],[f816,f46])).
fof(f16234,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)))))),
  inference(superposition,[],[f816,f46])).
fof(f15720,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK5(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)),sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK5(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | $spl29),
  inference(cnf_transformation,[],[f15720_D])).
fof(f15720_D,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK5(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)),sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK5(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f15716,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3))),s(fun(sK0,bool),sK3))) | $spl25),
  inference(cnf_transformation,[],[f15716_D])).
fof(f15716_D,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3))),s(fun(sK0,bool),sK3))) <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f15718,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),bool,sK0)))) | $spl27),
  inference(cnf_transformation,[],[f15718_D])).
fof(f15718_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),bool,sK0)))) <=> ~$spl27),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl27])])).
fof(f15684,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),bool,sK0)))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3))),s(fun(sK0,bool),sK3))) | s(bool,'const.bool.IN_2'(s(sK0,sK5(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)),sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK5(sK2,'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | $spl18),
  inference(superposition,[],[f485,f15651])).
fof(f15651,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),bool,sK0)))) | $spl18),
  inference(trivial_inequality_removal,[],[f15639])).
fof(f15639,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),bool,sK0)))) | $spl18),
  inference(superposition,[],[f606,f15615])).
fof(f606,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,f(s(sK0,sK4(X1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),bool,sK0)))) | s(fun(sK0,bool),X1) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK4(X1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),bool,sK0))))) )),
  inference(superposition,[],[f67,f46])).
fof(f67,plain,(
  ( ! [X1] : (s(bool,f(s(sK0,sK4(X1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),bool,sK0)))) != s(bool,g(s(sK0,sK4(X1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),bool,sK0)))) | s(fun(sK0,bool),X1) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))) )),
  inference(superposition,[],[f40,f41])).
fof(f41,plain,(
  ( ! [X2,X0,X3,X1] : (s(fun(X0,X1),X2) = s(fun(X0,X1),X3) | s(X1,f(s(X0,sK4(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK4(X3,X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f29])).
fof(f29,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK4(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK4(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f27])).
fof(f27,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f16])).
fof(f16,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f15])).
fof(f15,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.pred_set.INTER_ASSOC',aHL_EXT)).
fof(f485,plain,(
  ( ! [X4,X5] : (s(bool,f(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X4)),X5,bool,sK0)))) != s(bool,g(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X4)),X5,bool,sK0)))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X5),s(fun(sK0,bool),sK3))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | s(bool,'const.bool.IN_2'(s(sK0,sK5(sK2,X4,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK5(sK2,X4,sK0)),s(fun(sK0,bool),X4)))) )),
  inference(superposition,[],[f62,f43])).
fof(f62,plain,(
  ( ! [X0] : (s(bool,f(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),X0,bool,sK0)))) != s(bool,g(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),X0,bool,sK0)))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X0),s(fun(sK0,bool),sK3))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))) )),
  inference(superposition,[],[f40,f41])).
fof(f15714,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)),sK2,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | $spl23),
  inference(cnf_transformation,[],[f15714_D])).
fof(f15714_D,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)),sK2,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f15683,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),bool,sK0)))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3))),s(fun(sK0,bool),sK3))) | s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)),sK2,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | $spl18),
  inference(superposition,[],[f486,f15651])).
fof(f486,plain,(
  ( ! [X6,X7] : (s(bool,f(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X6)),X7,bool,sK0)))) != s(bool,g(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X6)),X7,bool,sK0)))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X7),s(fun(sK0,bool),sK3))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | s(bool,'const.bool.IN_2'(s(sK0,sK5(X6,sK2,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK5(X6,sK2,sK0)),s(fun(sK0,bool),X6)))) )),
  inference(superposition,[],[f62,f43])).
fof(f15615,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),bool,sK0)))) | $spl18),
  inference(cnf_transformation,[],[f15615_D])).
fof(f15615_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),bool,sK0)))) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f15617,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),bool,sK0)))) | $spl20),
  inference(cnf_transformation,[],[f15617_D])).
fof(f15617_D,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),bool,sK0)))) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f15613,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),bool,sK0))))),
  inference(equality_resolution,[],[f1465])).
fof(f1465,plain,(
  ( ! [X0] : (s(fun(sK0,bool),X0) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | s(bool,'F_0') = s(bool,g(s(sK0,sK4(X0,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK4(X0,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),bool,sK0))))) )),
  inference(trivial_inequality_removal,[],[f1463])).
fof(f1463,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(fun(sK0,bool),X0) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | s(bool,'F_0') = s(bool,g(s(sK0,sK4(X0,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK4(X0,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),bool,sK0))))) )),
  inference(superposition,[],[f605,f46])).
fof(f605,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(sK0,sK4(X0,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),bool,sK0)))) | s(fun(sK0,bool),X0) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | s(bool,'F_0') = s(bool,g(s(sK0,sK4(X0,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),bool,sK0))))) )),
  inference(superposition,[],[f67,f46])).
fof(f7686,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) | $spl8),
  inference(cnf_transformation,[],[f7686_D])).
fof(f7686_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f7688,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) | $spl10),
  inference(cnf_transformation,[],[f7688_D])).
fof(f7688_D,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f7690,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),bool,sK0)))) | $spl12),
  inference(cnf_transformation,[],[f7690_D])).
fof(f7690_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),bool,sK0)))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f7694,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))),bool,sK0)))) | $spl16),
  inference(cnf_transformation,[],[f7694_D])).
fof(f7694_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))),bool,sK0)))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f7669,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),bool,sK0)))) | s(bool,'F_0') = s(bool,g(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) | $spl4),
  inference(trivial_inequality_removal,[],[f7659])).
fof(f7659,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,g(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),bool,sK0)))) | s(bool,'F_0') = s(bool,g(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) | $spl4),
  inference(superposition,[],[f4105,f7389])).
fof(f7389,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,f(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),X0),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),bool,sK0)))) | s(bool,'F_0') = s(bool,g(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),X0,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),X0,bool,sK0))))) ) | $spl4),
  inference(trivial_inequality_removal,[],[f7387])).
fof(f7387,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'F_0') = s(bool,f(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),X0),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),bool,sK0)))) | s(bool,'F_0') = s(bool,g(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),X0,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),X0,bool,sK0))))) ) | $spl4),
  inference(superposition,[],[f4027,f46])).
fof(f4027,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),X0,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),X0),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),bool,sK0)))) | s(bool,'F_0') = s(bool,g(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),X0,bool,sK0))))) ) | $spl4),
  inference(superposition,[],[f1914,f46])).
fof(f1914,plain,(
  ( ! [X0] : (s(bool,f(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),X0,bool,sK0)))) != s(bool,g(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),X0,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),X0),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),bool,sK0))))) ) | $spl4),
  inference(superposition,[],[f1911,f41])).
fof(f4105,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,f(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)),X1,bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X1)),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)),X1,bool,sK0))))) ) | $spl4),
  inference(superposition,[],[f1974,f46])).
fof(f1974,plain,(
  ( ! [X0] : (s(bool,f(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)),X0,bool,sK0)))) != s(bool,g(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)),X0,bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X0)),bool,sK0))))) ) | $spl4),
  inference(superposition,[],[f1947,f41])).
fof(f1947,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),bool,sK0)))) | $spl4),
  inference(trivial_inequality_removal,[],[f1935])).
fof(f1935,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),bool,sK0)))) | $spl4),
  inference(superposition,[],[f543,f1911])).
fof(f543,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,f(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),X1,bool,sK0)))) | s(fun(sK0,bool),X1) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | s(bool,'T_0') = s(bool,g(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),X1,bool,sK0))))) )),
  inference(superposition,[],[f66,f46])).
fof(f66,plain,(
  ( ! [X0] : (s(bool,f(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),X0,bool,sK0)))) != s(bool,g(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),X0,bool,sK0)))) | s(fun(sK0,bool),X0) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))) )),
  inference(superposition,[],[f40,f41])).
fof(f7692,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))),bool,sK0)))) | $spl14),
  inference(cnf_transformation,[],[f7692_D])).
fof(f7692_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))),bool,sK0)))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f7670,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),bool,sK0)))) | s(bool,'F_0') = s(bool,g(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) | $spl4),
  inference(trivial_inequality_removal,[],[f7657])).
fof(f7657,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,f(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),bool,sK0)))) | s(bool,'F_0') = s(bool,g(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) | $spl4),
  inference(superposition,[],[f4048,f7389])).
fof(f4048,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,f(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)),X1,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X1)),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)),X1,bool,sK0))))) ) | $spl4),
  inference(superposition,[],[f1930,f46])).
fof(f1930,plain,(
  ( ! [X0] : (s(bool,f(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)),X0,bool,sK0)))) != s(bool,g(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)),X0,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X0)),bool,sK0))))) ) | $spl4),
  inference(superposition,[],[f1911,f41])).
fof(f1911,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),bool,sK0)))) | $spl4),
  inference(cnf_transformation,[],[f1911_D])).
fof(f1911_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),bool,sK0)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f1913,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),bool,sK0)))) | $spl6),
  inference(cnf_transformation,[],[f1913_D])).
fof(f1913_D,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),bool,sK0)))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f1909,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),bool,sK0))))),
  inference(equality_resolution,[],[f1382])).
fof(f1382,plain,(
  ( ! [X0] : (s(fun(sK0,bool),X0) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | s(bool,'F_0') = s(bool,g(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),X0,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),X0,bool,sK0))))) )),
  inference(trivial_inequality_removal,[],[f1380])).
fof(f1380,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(fun(sK0,bool),X0) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | s(bool,'F_0') = s(bool,g(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),X0,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),X0,bool,sK0))))) )),
  inference(superposition,[],[f542,f46])).
fof(f542,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),X0,bool,sK0)))) | s(fun(sK0,bool),X0) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | s(bool,'F_0') = s(bool,g(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)),X0,bool,sK0))))) )),
  inference(superposition,[],[f66,f46])).
fof(f115,plain,(
  ( ! [X1] : (s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3))) | s(bool,'F_0') = s(bool,g(s(sK0,sK4(sK1,X1,bool,sK0))))) ) | $spl0),
  inference(cnf_transformation,[],[f115_D])).
fof(f115_D,plain,(
  ( ! [X1] : (s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3))) | s(bool,'F_0') = s(bool,g(s(sK0,sK4(sK1,X1,bool,sK0))))) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f117,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl3),
  inference(cnf_transformation,[],[f117_D])).
fof(f117_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f114,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3))) | s(bool,'F_0') = s(bool,g(s(sK0,sK4(sK1,X1,bool,sK0))))) )),
  inference(subsumption_resolution,[],[f112,f82])).
fof(f82,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(sK0,sK4(sK1,X0,bool,sK0)))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X0),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X0),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3))) | s(bool,'F_0') = s(bool,g(s(sK0,sK4(sK1,X0,bool,sK0))))) )),
  inference(superposition,[],[f70,f46])).
fof(f70,plain,(
  ( ! [X0] : (s(bool,f(s(sK0,sK4(sK1,X0,bool,sK0)))) != s(bool,g(s(sK0,sK4(sK1,X0,bool,sK0)))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X0),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X0),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)))) )),
  inference(superposition,[],[f40,f41])).
fof(f112,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3))) | s(bool,'F_0') = s(bool,g(s(sK0,sK4(sK1,X1,bool,sK0)))) | s(bool,'T_0') = s(bool,f(s(sK0,sK4(sK1,X1,bool,sK0))))) )),
  inference(superposition,[],[f82,f46])).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 1072
% Passive clauses: 12963
% Generated clauses: 19220
% Final active clauses: 749
% Final passive clauses: 7506
% Input formulas: 11
% Initial clauses: 21
% 
% Selected by SInE selection: 11
% SInE iterations: 1
% 
% Duplicate literals: 1192
% Trivial inequalities: 718
% Fw subsumption resolutions: 126
% Fw demodulations: 322
% Bw demodulations: 19
% 
% Simple tautologies: 10
% Equational tautologies: 1
% Forward subsumptions: 4315
% 
% Binary resolution: 70
% Forward superposition: 13275
% Backward superposition: 3419
% Equality resolution: 233
% 
% Split clauses: 19
% Split components: 44
% SAT solver clauses: 35
% SAT solver binary clauses: 24
% 
% Sat splits: 19
% Sat splitting refutations: 15
% 
% Memory used [KB]: 9850
% Time elapsed: 1.167 s
% ------------------------------
% ------------------------------
% Success in time 12.895 s
