Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.quantHeuristics
% (60928)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 709
% Passive clauses: 10994
% Generated clauses: 20186
% Final active clauses: 704
% Final passive clauses: 9909
% Input formulas: 55
% Initial clauses: 226
% 
% Duplicate literals: 321
% Trivial inequalities: 15
% Fw subsumption resolutions: 1599
% Fw demodulations: 512
% Bw demodulations: 5
% 
% Simple tautologies: 264
% Equational tautologies: 57
% Forward subsumptions: 6122
% 
% Binary resolution: 12064
% Factoring: 991
% Forward superposition: 2232
% Backward superposition: 2075
% Self superposition: 27
% Equality factoring: 26
% Equality resolution: 4
% 
% Split clauses: 347
% Split components: 696
% SAT solver clauses: 687
% SAT solver unit clauses: 17
% SAT solver binary clauses: 203
% 
% Sat splits: 347
% 
% TWLsolver clauses: 687
% TWLsolver calls for satisfiability: 384
% 
% Memory used [KB]: 9722
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 294 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.quantHeuristics
% (60939)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 449
% Passive clauses: 3241
% Generated clauses: 13776
% Final active clauses: 394
% Final passive clauses: 2012
% Input formulas: 55
% Initial clauses: 226
% 
% Duplicate literals: 192
% Trivial inequalities: 30
% Fw subsumption resolutions: 1231
% Fw demodulations: 702
% Bw demodulations: 825
% Condensations: 294
% Global subsumptions: 13
% 
% Simple tautologies: 255
% Equational tautologies: 23
% Forward subsumptions: 4496
% Bw demodulations to eq. taut.: 6
% 
% Binary resolution: 4122
% Factoring: 191
% Forward superposition: 1168
% Backward superposition: 1409
% Self superposition: 27
% Equality factoring: 28
% Equality resolution: 8
% 
% Split clauses: 14
% Split components: 28
% SAT solver clauses: 2994
% SAT solver unit clauses: 211
% SAT solver binary clauses: 85
% SAT solver learnt clauses: 26
% SAT solver learnt literals: 10
% 
% Sat splits: 14
% Sat splitting refutations: 23
% 
% TWLsolver clauses: 2942
% TWLsolver calls for satisfiability: 2072
% 
% Memory used [KB]: 12792
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 287 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.quantHeuristics
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.quantHeuristics
% SZS output start Proof for thm.quantHeuristics
fof(f49749,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f321,f569,f320,f571,f567,f319,f572,f1520,f1559,f1557,f4569,f4573,f4576,f4579,f4583,f4586,f4589,f4592,f4595,f4598,f4601,f4604,f4607,f4610,f4613,f4616,f4619,f4622,f4625,f4677,f4570,f4679,f4574,f4681,f4577,f4683,f4580,f4686,f4584,f4688,f4587,f4690,f4590,f4692,f4593,f4694,f4596,f4696,f4599,f4698,f4602,f4700,f4605,f4702,f4608,f4704,f4611,f4706,f4614,f4708,f4617,f4710,f4620,f4712,f4623,f4714,f4626,f4942,f4943,f4944,f4945,f4946,f4947,f4948,f4949,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f4957,f4967,f4973,f4983,f4987,f4996,f5006,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f5030,f5032,f5034,f5036,f5038,f5040,f5042,f5044,f5046,f5048,f5050,f5052,f5057,f5058,f5059,f5060,f6808,f6811,f6809,f6840,f6843,f6841,f6941,f6944,f6942,f6975,f6978,f6976,f7021,f7024,f7022,f7067,f7070,f7068,f7099,f7102,f7100,f7130,f7133,f7131,f8140,f8143,f8141,f8145,f8166,f8177,f8181,f8184,f8187,f8191,f8195,f8202,f8204,f8208,f8212,f8216,f8221,f8225,f8258,f8273,f8322,f8358,f8384,f8392,f8394,f8398,f8406,f8408,f8416,f8418,f8423,f8428,f8436,f8438,f8443,f8448,f8456,f8458,f8463,f8468,f8476,f8478,f8940,f8941,f8497,f8506,f8516,f8524,f8526,f8532,f8534,f8538,f8540,f8546,f8548,f8552,f8554,f8558,f8561,f8563,f8566,f8568,f8607,f8610,f8616,f8620,f8623,f8635,f8638,f8644,f8648,f8651,f8661,f8663,f8668,f8677,f8682,f8687,f8694,f8698,f8705,f8709,f8716,f8721,f8726,f8731,f8736,f8740,f8742,f8749,f8752,f8755,f8757,f8779,f8789,f8799,f8807,f8809,f8814,f8821,f8824,f8831,f8837,f8849,f8851,f8855,f8859,f8863,f8867,f8877,f8883,f8889,f8897,f8901,f8905,f8909,f8918,f8928,f8930,f8944,f8945,f8946,f8947,f8948,f8949,f8956,f8958,f8959,f8960,f8961,f8962,f8963,f8964,f8965,f8966,f8967,f8968,f8969,f8970,f8971,f8972,f8973,f8974,f8975,f8976,f8977,f8978,f8979,f8980,f8981,f8982,f8983,f8984,f8985,f8986,f8987,f8988,f8989,f8990,f8991,f8992,f8993,f8994,f8995,f8996,f8997,f8998,f8999,f9000,f9001,f9002,f9003,f9004,f9005,f9006,f9007,f9008,f9009,f9010,f9011,f9012,f9013,f9014,f9015,f9016,f9017,f9018,f9019,f9020,f9021,f9022,f9023,f9024,f9025,f9026,f9027,f9028,f9029,f9030,f9031,f9032,f9033,f9034,f9035,f9036,f9037,f9038,f9039,f9040,f9041,f9042,f9043,f9044,f9045,f9046,f9047,f9048,f9049,f9050,f9051,f9052,f9053,f9054,f9055,f9056,f9057,f9058,f9059,f9060,f9061,f9062,f9063,f9064,f9065,f9066,f9067,f9068,f9069,f9070,f9071,f9072,f9073,f9074,f9075,f9076,f9077,f9078,f9079,f9080,f9083,f9084,f9085,f9086,f7467,f9088,f7464,f9090,f566,f9994,f12286,f12289,f12287,f12333,f12336,f12334,f12385,f12388,f12386,f12437,f12440,f12438,f12476,f12479,f12477,f12532,f12535,f12533,f12575,f12578,f12576,f12636,f12639,f12637,f12675,f12678,f12676,f15727,f15730,f15728,f15746,f15748,f15747,f15750,f15752,f15754,f15760,f15762,f15764,f15766,f15770,f15772,f15774,f15776,f15778,f15780,f15798,f15800,f15818,f15820,f15840,f15860,f15870,f15872,f15883,f15894,f15904,f15906,f15916,f15918,f15928,f15930,f15969,f16034,f16073,f16138,f16694,f16698,f16708,f16712,f16716,f16720,f16724,f16772,f16776,f16782,f16787,f16792,f16797,f16801,f16820,f16832,f16846,f16858,f16860,f16874,f16886,f16900,f16912,f16914,f16916,f16918,f16920,f19420,f19421,f19422,f19423,f19424,f19425,f16974,f16976,f16978,f16980,f16982,f16984,f16986,f16988,f16991,f16993,f16995,f16997,f16999,f17001,f17003,f17005,f17070,f17072,f17074,f17076,f17108,f17140,f17180,f17182,f17189,f17191,f17193,f17195,f17208,f17210,f17212,f17214,f17228,f17230,f17232,f17234,f17236,f17238,f17240,f17242,f17283,f17285,f17293,f17295,f17297,f17299,f17312,f17314,f17316,f17318,f17332,f17334,f17336,f17338,f17340,f17342,f17344,f17346,f17361,f17373,f17387,f17399,f17401,f17444,f17446,f17473,f17475,f17519,f17521,f17548,f17550,f17602,f17604,f17631,f17633,f17655,f17687,f17709,f17741,f17776,f17811,f17877,f17879,f17881,f17883,f17885,f17892,f17967,f17970,f17972,f17974,f17976,f17978,f17980,f17998,f18000,f18002,f18025,f18027,f18050,f18052,f18075,f18077,f18100,f18102,f18138,f18170,f18210,f18212,f18214,f18216,f18218,f18224,f18242,f18244,f18246,f18248,f18250,f18252,f18254,f18262,f18264,f18266,f18352,f18354,f18356,f18358,f18360,f18362,f18364,f18366,f18368,f18370,f18419,f18421,f18423,f18425,f18474,f18476,f18478,f18480,f18504,f18506,f18530,f18532,f18534,f18558,f18582,f18606,f18608,f18632,f18634,f18659,f18691,f18725,f18757,f18769,f18776,f18779,f18803,f18805,f18807,f18809,f18811,f18813,f18815,f18817,f18819,f18821,f18844,f18846,f18848,f18850,f18852,f18854,f18856,f18858,f18860,f18862,f18864,f18866,f18868,f18870,f18872,f18874,f18876,f18878,f18884,f18896,f18898,f18907,f18909,f18918,f18920,f18923,f18925,f18936,f18938,f18950,f18952,f18954,f18956,f18960,f18962,f18964,f18966,f18971,f18976,f18982,f18990,f18998,f19008,f19023,f19026,f19038,f19042,f19053,f19056,f19068,f19072,f19083,f19086,f19098,f19102,f19107,f19110,f19117,f19120,f19155,f19158,f19169,f19171,f19174,f19176,f19189,f19191,f19193,f19195,f19202,f19204,f19206,f19208,f19210,f19212,f19214,f19216,f19218,f19223,f19238,f19240,f19249,f19251,f19260,f19262,f19267,f19269,f19275,f19277,f19286,f19288,f19290,f19293,f19295,f19299,f19301,f19303,f19305,f19320,f19335,f19349,f19355,f19377,f19379,f19381,f19383,f19394,f19426,f19427,f19428,f19429,f19430,f19431,f19432,f19433,f19434,f19435,f19436,f19437,f19438,f19439,f19440,f19441,f19442,f19443,f19444,f19445,f19446,f19447,f19448,f19449,f19450,f19451,f19452,f19453,f19454,f19455,f19456,f19457,f19458,f19459,f19460,f19461,f19462,f19463,f19464,f19465,f19466,f19467,f19468,f19469,f19470,f19471,f19472,f19473,f19474,f19475,f19476,f19477,f19478,f19479,f19480,f19481,f19482,f19483,f19484,f19485,f19486,f19487,f19488,f19489,f19490,f19491,f19492,f19493,f19494,f19495,f19496,f19497,f19498,f19499,f19500,f19501,f19502,f19503,f19504,f19505,f19506,f19507,f19508,f19509,f19510,f19511,f19512,f19513,f19514,f19515,f19516,f19517,f19518,f19519,f19520,f19521,f19522,f19523,f19524,f19525,f19526,f19527,f19528,f19529,f19530,f19531,f19532,f19533,f19534,f19535,f19536,f19537,f19538,f19539,f19540,f19541,f19542,f19543,f19544,f19545,f19546,f19547,f19548,f19549,f19550,f19551,f19552,f19553,f19554,f19555,f19556,f19557,f19558,f19559,f19560,f19561,f19562,f19563,f19564,f19565,f19566,f19567,f19568,f19569,f19570,f19571,f19572,f19573,f19574,f19575,f19576,f19577,f19578,f19579,f19580,f19581,f19582,f19583,f19584,f19585,f19586,f19587,f19588,f19589,f19590,f19591,f19592,f19593,f19594,f19595,f19596,f19597,f19598,f19599,f19600,f19601,f19602,f19603,f19604,f19605,f19606,f19607,f19608,f19609,f19610,f19611,f19612,f19613,f19614,f19615,f19616,f19617,f19618,f19619,f19620,f19621,f19622,f19623,f19624,f19625,f19626,f19627,f19628,f19629,f19630,f19631,f19632,f19633,f19634,f19635,f19636,f19637,f19638,f19639,f19640,f19641,f19642,f19643,f19644,f19645,f19646,f19647,f19648,f19649,f19650,f19651,f19652,f19653,f19654,f19655,f19656,f19657,f19658,f19659,f19660,f19661,f19662,f19663,f19664,f19665,f19666,f19667,f19668,f19669,f19670,f19671,f19672,f19673,f19674,f19675,f19676,f19677,f19678,f19679,f19680,f19681,f19682,f19683,f19684,f19685,f19686,f19687,f19688,f19689,f19690,f19691,f19692,f19693,f19694,f19695,f19696,f19697,f19698,f19699,f19700,f19701,f19702,f19703,f19704,f19705,f19706,f19707,f19708,f19709,f19710,f19711,f19712,f19713,f19714,f19715,f19716,f19717,f19718,f19719,f19720,f19721,f19722,f19723,f19724,f19725,f19726,f19727,f19728,f19729,f19730,f19731,f19732,f19733,f19734,f19735,f19736,f19737,f19738,f19739,f19740,f19741,f19742,f19743,f19744,f19745,f19746,f19747,f19748,f19749,f19750,f19751,f19752,f19753,f19754,f19755,f19756,f19757,f19758,f19759,f19760,f19761,f19762,f19763,f19764,f19765,f19766,f19767,f19768,f19769,f19770,f19771,f19772,f19773,f19774,f19775,f19776,f19777,f19778,f19779,f19780,f19781,f19782,f19783,f19784,f19785,f19786,f19787,f19788,f19789,f19790,f19791,f19792,f19793,f19794,f19795,f19796,f19797,f19798,f19799,f19800,f19801,f19802,f19803,f19804,f19805,f19806,f19807,f19808,f19809,f19810,f19811,f19812,f19813,f19814,f19815,f19816,f19817,f19818,f19819,f19820,f19821,f19822,f19823,f19824,f19825,f19826,f19827,f19828,f19829,f19830,f19831,f19832,f19833,f19834,f19835,f19836,f19837,f19838,f19839,f19840,f19841,f19842,f19843,f19844,f19845,f19846,f19847,f19848,f19849,f19850,f19851,f19852,f19853,f19854,f19855,f19856,f19857,f19858,f19859,f19860,f19861,f19862,f19863,f19864,f19865,f19866,f19867,f19868,f19869,f19870,f19871,f19872,f19873,f19874,f19875,f19876,f19877,f19878,f19879,f19880,f19881,f19882,f19883,f19884,f19885,f19886,f19887,f19888,f19889,f19890,f19891,f19892,f19893,f19894,f19895,f19896,f19897,f19898,f19899,f19900,f19901,f19902,f19903,f19904,f19905,f19906,f19907,f19908,f19909,f19910,f19911,f19912,f19913,f19914,f19915,f19916,f19917,f19918,f19919,f19920,f19921,f19922,f19923,f19924,f19925,f19926,f19927,f19928,f19929,f19930,f19931,f19932,f19933,f19934,f19935,f19936,f19937,f19938,f19939,f19940,f19941,f19942,f19943,f19944,f19945,f19946,f19947,f19948,f19949,f19950,f19951,f19952,f19953,f19954,f19955,f19956,f19957,f19958,f19959,f19960,f19961,f19962,f19963,f19964,f19965,f19966,f19967,f19968,f19969,f19970,f19971,f19972,f19973,f19974,f19975,f19976,f19977,f19978,f19979,f19980,f19981,f19982,f19983,f19984,f19985,f19986,f19987,f19988,f19989,f19990,f19991,f19992,f19993,f19994,f19995,f19996,f19997,f19998,f19999,f20000,f20001,f20002,f20003,f20004,f20005,f20006,f20007,f20008,f20009,f20010,f20011,f20012,f20013,f20014,f20015,f20016,f20017,f20018,f20019,f20020,f20021,f20022,f20023,f20024,f20025,f20026,f20027,f20028,f20029,f20030,f20031,f20032,f20033,f20034,f20035,f20036,f20037,f20038,f20039,f20040,f20041,f20042,f20043,f20044,f20045,f20046,f20047,f20048,f20049,f20050,f20051,f20052,f20053,f20054,f20055,f20056,f20057,f20058,f20059,f20060,f20061,f20062,f20063,f20064,f20065,f20066,f20067,f20068,f20069,f20070,f20071,f20072,f20073,f20074,f20075,f20076,f20077,f20078,f20079,f20080,f20081,f20082,f20083,f20084,f20085,f20086,f20087,f20088,f20089,f20090,f20091,f20092,f20093,f20094,f20095,f20096,f20097,f20098,f20099,f20100,f20101,f20102,f20103,f20104,f20105,f20106,f20107,f20108,f20109,f20110,f20111,f20112,f20113,f20114,f20115,f20116,f20117,f20118,f20119,f20120,f20121,f20122,f20123,f20124,f20125,f20126,f20127,f20128,f20129,f20130,f20131,f20132,f20133,f20134,f20135,f20136,f20137,f20138,f20139,f20140,f20141,f20142,f20143,f20144,f20145,f20146,f20147,f20148,f20149,f20150,f20151,f20152,f20153,f20154,f20155,f20156,f20157,f20158,f20159,f20160,f20161,f20162,f20163,f20164,f20165,f20215,f20218,f20216,f20289,f20292,f20290,f20529,f20530,f20537,f20534,f20531,f20532,f20540,f20533,f20542,f20535,f20539,f20742,f20744,f20743,f20960,f20961,f20973,f20971,f21260,f21262,f21261,f21990,f21992,f21991,f22016,f22018,f15729,f22017,f39776,f39779,f39777,f39924,f39927,f39925,f40055,f40058,f40056,f40492,f40495,f40493,f40595,f40598,f40605,f40606,f40609,f40610,f40596,f40697,f40700,f40698,f40787,f40790,f4571,f40788,f48537,f48540])).
fof(f48540,plain,(
  $false | ($spl0 | $spl3 | $spl4)),
  inference(subsumption_resolution,[],[f48539,f569])).
fof(f48539,plain,(
  s(sK14,sK15) = s(sK14,sK17) | ($spl0 | $spl4)),
  inference(subsumption_resolution,[],[f48525,f322])).
fof(f322,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.quantHeuristics.GUESSES_UEXISTS_THM3',aHL_TRUTH)).
fof(f48525,plain,(
  ~p(s(bool,'T_0')) | s(sK14,sK15) = s(sK14,sK17) | ($spl0 | $spl4)),
  inference(superposition,[],[f48518,f3354])).
fof(f3354,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))),
  inference(unit_resulting_resolution,[],[f322,f2082,f387])).
fof(f387,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f159])).
fof(f159,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f158])).
fof(f158,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f99])).
fof(f99,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X12,X13] : ((p(s(bool,X12)) => p(s(bool,X13))) => ((p(s(bool,X13)) => p(s(bool,X12))) => s(bool,X12) = s(bool,X13)))),
  file('thm.quantHeuristics.GUESSES_UEXISTS_THM3','thm.bool.IMP_ANTISYM_AX')).
fof(f2082,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15))))),
  inference(forward_demodulation,[],[f2081,f421])).
fof(f421,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f116])).
fof(f116,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))),
  inference(flattening,[],[f115])).
fof(f115,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))),
  inference(rectify,[],[f56])).
fof(f56,axiom,(
  ! [X27,X28] : ! [X29,X30] : s(X28,X29) = s(X28,happ(s(fun(X27,X28),'_dst_x0_1'(s(X28,X29))),s(X27,X30)))),
  file('thm.quantHeuristics.GUESSES_UEXISTS_THM3','_dst_x0_ax')).
fof(f2081,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,happ(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(sK13,X0))))))) )),
  inference(subsumption_resolution,[],[f2077,f322])).
fof(f2077,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,happ(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(sK13,X0))))))) )),
  inference(superposition,[],[f424,f621])).
fof(f621,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))),
  inference(unit_resulting_resolution,[],[f318,f337])).
fof(f337,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f209])).
fof(f209,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f68])).
fof(f68,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X11] : (s(bool,'T_0') = s(bool,X11) <=> p(s(bool,X11)))),
  file('thm.quantHeuristics.GUESSES_UEXISTS_THM3','thm.bool.EQ_CLAUSES|split|1')).
fof(f318,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16))))),
  inference(cnf_transformation,[],[f206])).
fof(f206,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) & (p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ! [X5] : (~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,X5)))) | s(sK14,sK15) = s(sK14,X5))) & (~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | (p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) & s(sK14,sK15) != s(sK14,sK17)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13,sK14,sK15,sK16,sK17])],[f205])).
fof(f205,plain,(
  ? [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(fun(X1,bool),X3)))) & (p(s(bool,'const.bool.?|33|_1'(s(fun(X1,bool),X3)))) | ! [X5] : (~p(s(bool,happ(s(fun(X1,bool),X3),s(X1,X5)))) | s(X1,X2) = s(X1,X5))) & (~p(s(bool,'const.bool.?|33|_1'(s(fun(X1,bool),X3)))) | ? [X4] : (p(s(bool,happ(s(fun(X1,bool),X3),s(X1,X4)))) & s(X1,X2) != s(X1,X4))))),
  inference(rectify,[],[f204])).
fof(f204,plain,(
  ? [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(fun(X1,bool),X3)))) & (p(s(bool,'const.bool.?|33|_1'(s(fun(X1,bool),X3)))) | ! [X4] : (~p(s(bool,happ(s(fun(X1,bool),X3),s(X1,X4)))) | s(X1,X2) = s(X1,X4))) & (~p(s(bool,'const.bool.?|33|_1'(s(fun(X1,bool),X3)))) | ? [X4] : (p(s(bool,happ(s(fun(X1,bool),X3),s(X1,X4)))) & s(X1,X2) != s(X1,X4))))),
  inference(flattening,[],[f203])).
fof(f203,plain,(
  ? [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(fun(X1,bool),X3)))) & ((p(s(bool,'const.bool.?|33|_1'(s(fun(X1,bool),X3)))) | ! [X4] : (~p(s(bool,happ(s(fun(X1,bool),X3),s(X1,X4)))) | s(X1,X2) = s(X1,X4))) & (~p(s(bool,'const.bool.?|33|_1'(s(fun(X1,bool),X3)))) | ? [X4] : (p(s(bool,happ(s(fun(X1,bool),X3),s(X1,X4)))) & s(X1,X2) != s(X1,X4)))))),
  inference(nnf_transformation,[],[f145])).
fof(f145,plain,(
  ? [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(fun(X1,bool),X3)))) & (p(s(bool,'const.bool.?|33|_1'(s(fun(X1,bool),X3)))) <~> ! [X4] : (~p(s(bool,happ(s(fun(X1,bool),X3),s(X1,X4)))) | s(X1,X2) = s(X1,X4))))),
  inference(ennf_transformation,[],[f58])).
fof(f58,plain,(
  ~! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(fun(X1,bool),X3)))) => (p(s(bool,'const.bool.?|33|_1'(s(fun(X1,bool),X3)))) <=> ! [X4] : (p(s(bool,happ(s(fun(X1,bool),X3),s(X1,X4)))) => s(X1,X2) = s(X1,X4))))),
  inference(flattening,[],[f57])).
fof(f57,plain,(
  ~! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(fun(X1,bool),X3)))) => (p(s(bool,'const.bool.?|33|_1'(s(fun(X1,bool),X3)))) <=> ! [X4] : (p(s(bool,happ(s(fun(X1,bool),X3),s(X1,X4)))) => s(X1,X2) = s(X1,X4))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X6,X7),'_dst_x0_1'(s(X7,X8))),s(fun(X7,bool),X9)))) => (p(s(bool,'const.bool.?|33|_1'(s(fun(X7,bool),X9)))) <=> ! [X10] : (p(s(bool,happ(s(fun(X7,bool),X9),s(X7,X10)))) => s(X7,X8) = s(X7,X10))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X6,X7),'_dst_x0_1'(s(X7,X8))),s(fun(X7,bool),X9)))) => (p(s(bool,'const.bool.?|33|_1'(s(fun(X7,bool),X9)))) <=> ! [X10] : (p(s(bool,happ(s(fun(X7,bool),X9),s(X7,X10)))) => s(X7,X8) = s(X7,X10))))),
  file('thm.quantHeuristics.GUESSES_UEXISTS_THM3','thm.quantHeuristics.GUESSES_UEXISTS_THM3')).
fof(f424,plain,(
  ( ! [X2,X0,X5,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) )),
  inference(cnf_transformation,[],[f256])).
fof(f256,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) & (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK26(X3,X2,X1,X0))))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK26])],[f255])).
fof(f255,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) & (? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f254])).
fof(f254,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))))))) & (? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(nnf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))))),
  inference(flattening,[],[f121])).
fof(f121,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X6,X7] : ! [X22,X23] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X7,X6),X22),s(fun(X6,bool),X23)))) <=> ! [X21] : p(s(bool,happ(s(fun(X6,bool),X23),s(X6,happ(s(fun(X7,X6),X22),s(X7,X21)))))))),
  file('thm.quantHeuristics.GUESSES_UEXISTS_THM3','thm.quantHeuristics.GUESS_REWRITES|split|2')).
fof(f48518,plain,(
  ( ! [X1] : (~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,X1)))) | s(sK14,sK17) = s(sK14,X1)) ) | ($spl0 | $spl4)),
  inference(subsumption_resolution,[],[f48506,f322])).
fof(f48506,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,X1)))) | s(sK14,sK17) = s(sK14,X1)) ) | ($spl0 | $spl4)),
  inference(superposition,[],[f10426,f9540])).
fof(f9540,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17))) | $spl4),
  inference(unit_resulting_resolution,[],[f322,f571,f387])).
fof(f10426,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,X1)))) | ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,X0)))) | s(sK14,X0) = s(sK14,X1)) ) | $spl0),
  inference(subsumption_resolution,[],[f10343,f322])).
fof(f10343,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,X0)))) | ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,X1)))) | s(sK14,X0) = s(sK14,X1)) ) | $spl0),
  inference(superposition,[],[f368,f9096])).
fof(f9096,plain,(
  s(bool,'T_0') = s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | $spl0),
  inference(unit_resulting_resolution,[],[f322,f566,f387])).
fof(f368,plain,(
  ( ! [X6,X0,X5,X1] : (~p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | s(X0,X5) = s(X0,X6)) )),
  inference(cnf_transformation,[],[f229])).
fof(f229,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK18(X1,X0))))) & ! [X5,X6] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) | s(X0,X5) = s(X0,X6)))) & (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK19(X1,X0))))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK20(X1,X0))))) & s(X0,sK19(X1,X0)) != s(X0,sK20(X1,X0))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK18,sK19,sK20])],[f228])).
fof(f228,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) | (? [X7] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X7)))) & ! [X5,X6] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) | s(X0,X5) = s(X0,X6)))) & (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ? [X2,X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & s(X0,X2) != s(X0,X3)) | p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1))))))),
  inference(rectify,[],[f227])).
fof(f227,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) | (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ! [X2,X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | s(X0,X2) = s(X0,X3)))) & (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ? [X2,X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & s(X0,X2) != s(X0,X3)) | p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1))))))),
  inference(flattening,[],[f226])).
fof(f226,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) | (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ! [X2,X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | s(X0,X2) = s(X0,X3)))) & ((! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ? [X2,X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & s(X0,X2) != s(X0,X3))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1))))))),
  inference(nnf_transformation,[],[f157])).
fof(f157,plain,(
  ! [X0,X1] : (p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) <=> (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ! [X2,X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | s(X0,X2) = s(X0,X3))))),
  inference(flattening,[],[f156])).
fof(f156,plain,(
  ! [X0,X1] : (p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) <=> (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ! [X2,X3] : ((~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | s(X0,X2) = s(X0,X3))))),
  inference(ennf_transformation,[],[f91])).
fof(f91,plain,(
  ! [X0,X1] : (p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) <=> (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ! [X2,X3] : ((p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => s(X0,X2) = s(X0,X3))))),
  inference(flattening,[],[f90])).
fof(f90,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) <=> (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ! [X2,X3] : ((p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => s(X0,X2) = s(X0,X3))))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X7] : ! [X9] : (p(s(bool,'const.bool.?|33|_1'(s(fun(X7,bool),X9)))) <=> (? [X14] : p(s(bool,happ(s(fun(X7,bool),X9),s(X7,X14)))) & ! [X14,X16] : ((p(s(bool,happ(s(fun(X7,bool),X9),s(X7,X14)))) & p(s(bool,happ(s(fun(X7,bool),X9),s(X7,X16))))) => s(X7,X14) = s(X7,X16))))),
  file('thm.quantHeuristics.GUESSES_UEXISTS_THM3','thm.bool.EXISTS_UNIQUE_THM')).
fof(f48537,plain,(
  $false | ($spl0 | $spl3 | $spl4)),
  inference(subsumption_resolution,[],[f48536,f322])).
fof(f48536,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl3 | $spl4)),
  inference(forward_demodulation,[],[f48523,f3354])).
fof(f48523,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | ($spl0 | $spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f569,f48518])).
fof(f40788,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~sP12(happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)),X1,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) ) | $spl144),
  inference(cnf_transformation,[],[f40788_D])).
fof(f40788_D,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~sP12(happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)),X1,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f4571,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl14),
  inference(cnf_transformation,[],[f4571_D])).
fof(f4571_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f40790,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP12(happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)),X1,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) ) | $spl0),
  inference(condensation,[],[f40789])).
fof(f40789,plain,(
  ( ! [X37,X35,X33,X36,X34] : (p(s(bool,X34)) | ~sP12(happ(s(fun(X35,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X35,X36)),X34,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | s(bool,X33) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) ) | $spl0),
  inference(subsumption_resolution,[],[f40763,f902])).
fof(f902,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f894,f573])).
fof(f573,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f325,f322])).
fof(f325,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f207])).
fof(f207,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f59])).
fof(f59,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f20])).
fof(f20,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.quantHeuristics.GUESSES_UEXISTS_THM3','thm.bool.NOT_CLAUSES|split|2')).
fof(f894,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f384,f329])).
fof(f329,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.quantHeuristics.GUESSES_UEXISTS_THM3',aHL_BOOL_CASES)).
fof(f384,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f159])).
fof(f40763,plain,(
  ( ! [X37,X35,X33,X36,X34] : (~p(s(bool,X33)) | p(s(bool,X34)) | ~sP12(happ(s(fun(X35,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X35,X36)),X34,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | s(bool,X33) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) ) | $spl0),
  inference(superposition,[],[f12680,f836])).
fof(f836,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f329,f329])).
fof(f12680,plain,(
  ( ! [X52,X50,X53,X51] : (~p(s(bool,X50)) | p(s(bool,X51)) | ~sP12(happ(s(fun(X52,bool),'_dst_x0_1'(s(bool,X50))),s(X52,X53)),X51,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) ) | $spl0),
  inference(subsumption_resolution,[],[f12665,f322])).
fof(f12665,plain,(
  ( ! [X52,X50,X53,X51] : (~p(s(bool,'T_0')) | ~p(s(bool,X50)) | p(s(bool,X51)) | ~sP12(happ(s(fun(X52,bool),'_dst_x0_1'(s(bool,X50))),s(X52,X53)),X51,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) ) | $spl0),
  inference(superposition,[],[f1353,f9096])).
fof(f1353,plain,(
  ( ! [X121,X118,X122,X120,X119] : (~p(s(bool,X121)) | ~p(s(bool,X119)) | p(s(bool,X122)) | ~sP12(happ(s(fun(X118,bool),'_dst_x0_1'(s(bool,X119))),s(X118,X120)),X122,X121)) )),
  inference(superposition,[],[f533,f421])).
fof(f533,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~sP12(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f315])).
fof(f315,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP12(X2,X0,X1)) & (~sP12(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f314])).
fof(f314,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP12(X2,X0,X1)) & (~sP12(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f202])).
fof(f202,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP12(X2,X0,X1))),
  inference(definition_folding,[],[f180,f201,f200])).
fof(f200,plain,(
  ! [X2,X0,X1] : (sP11(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f201,plain,(
  ! [X2,X0,X1] : (sP12(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP11(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP12])])).
fof(f180,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f143])).
fof(f143,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f142])).
fof(f142,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f49])).
fof(f49,axiom,(
  ! [X24,X25,X26] : ((p(s(bool,X26)) <=> (p(s(bool,X25)) => p(s(bool,X24)))) <=> ((p(s(bool,X26)) | p(s(bool,X25))) & (p(s(bool,X26)) | ~p(s(bool,X24))) & (~p(s(bool,X25)) | p(s(bool,X24)) | ~p(s(bool,X26)))))),
  file('thm.quantHeuristics.GUESSES_UEXISTS_THM3','thm.sat.sat_27_prove_ltb')).
fof(f40787,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP12(happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)),X1,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) ) | $spl0),
  inference(condensation,[],[f40786])).
fof(f40786,plain,(
  ( ! [X30,X28,X31,X29,X32] : (p(s(bool,X29)) | ~sP12(happ(s(fun(X30,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X30,X31)),X29,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | s(bool,X28) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) ) | $spl0),
  inference(subsumption_resolution,[],[f40762,f902])).
fof(f40762,plain,(
  ( ! [X30,X28,X31,X29,X32] : (~p(s(bool,X28)) | p(s(bool,X29)) | ~sP12(happ(s(fun(X30,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X30,X31)),X29,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | s(bool,X28) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) ) | $spl0),
  inference(superposition,[],[f12680,f836])).
fof(f40698,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~sP11('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X1,happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)))) ) | $spl142),
  inference(cnf_transformation,[],[f40698_D])).
fof(f40698_D,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~sP11('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X1,happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)))) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f40700,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP11('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X1,happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)))) ) | $spl0),
  inference(condensation,[],[f40699])).
fof(f40699,plain,(
  ( ! [X37,X35,X33,X36,X34] : (p(s(bool,X34)) | ~sP11('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X34,happ(s(fun(X35,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X35,X36))) | s(bool,X33) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) ) | $spl0),
  inference(subsumption_resolution,[],[f40665,f902])).
fof(f40665,plain,(
  ( ! [X37,X35,X33,X36,X34] : (~p(s(bool,X33)) | p(s(bool,X34)) | ~sP11('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X34,happ(s(fun(X35,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X35,X36))) | s(bool,X33) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) ) | $spl0),
  inference(superposition,[],[f12641,f836])).
fof(f12641,plain,(
  ( ! [X52,X50,X53,X51] : (~p(s(bool,X51)) | p(s(bool,X50)) | ~sP11('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X50,happ(s(fun(X52,bool),'_dst_x0_1'(s(bool,X51))),s(X52,X53)))) ) | $spl0),
  inference(subsumption_resolution,[],[f12620,f322])).
fof(f12620,plain,(
  ( ! [X52,X50,X53,X51] : (~p(s(bool,'T_0')) | p(s(bool,X50)) | ~p(s(bool,X51)) | ~sP11('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X50,happ(s(fun(X52,bool),'_dst_x0_1'(s(bool,X51))),s(X52,X53)))) ) | $spl0),
  inference(superposition,[],[f1350,f9096])).
fof(f1350,plain,(
  ( ! [X103,X107,X105,X106,X104] : (~p(s(bool,X107)) | p(s(bool,X106)) | ~p(s(bool,X104)) | ~sP11(X107,X106,happ(s(fun(X103,bool),'_dst_x0_1'(s(bool,X104))),s(X103,X105)))) )),
  inference(superposition,[],[f526,f421])).
fof(f526,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~sP11(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f313])).
fof(f313,plain,(
  ! [X0,X1,X2] : ((~sP11(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP11(X0,X1,X2)))),
  inference(rectify,[],[f312])).
fof(f312,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP11(X2,X0,X1)))),
  inference(flattening,[],[f311])).
fof(f311,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP11(X2,X0,X1)))),
  inference(nnf_transformation,[],[f200])).
fof(f40697,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP11('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X1,happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)))) ) | $spl0),
  inference(condensation,[],[f40696])).
fof(f40696,plain,(
  ( ! [X30,X28,X31,X29,X32] : (p(s(bool,X29)) | ~sP11('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X29,happ(s(fun(X30,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X30,X31))) | s(bool,X28) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) ) | $spl0),
  inference(subsumption_resolution,[],[f40664,f902])).
fof(f40664,plain,(
  ( ! [X30,X28,X31,X29,X32] : (~p(s(bool,X28)) | p(s(bool,X29)) | ~sP11('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X29,happ(s(fun(X30,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X30,X31))) | s(bool,X28) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) ) | $spl0),
  inference(superposition,[],[f12641,f836])).
fof(f40596,plain,(
  ( ! [X2,X3,X1] : (~sP8(X1,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3))) | p(s(bool,X1))) ) | $spl136),
  inference(cnf_transformation,[],[f40596_D])).
fof(f40596_D,plain,(
  ( ! [X2,X3,X1] : (~sP8(X1,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3))) | p(s(bool,X1))) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f40610,plain,(
  ( ! [X66] : (~p(s(bool,X66))) ) | $spl140),
  inference(cnf_transformation,[],[f40610_D])).
fof(f40610_D,plain,(
  ( ! [X66] : (~p(s(bool,X66))) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f40609,plain,(
  ( ! [X70,X68,X66,X69] : (~p(s(bool,X66)) | ~sP8(X68,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X69,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X69,X70))) | p(s(bool,X68))) ) | $spl0),
  inference(subsumption_resolution,[],[f40608,f340])).
fof(f340,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f211])).
fof(f211,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X11] : (s(bool,'F_0') = s(bool,X11) <=> ~p(s(bool,X11)))),
  file('thm.quantHeuristics.GUESSES_UEXISTS_THM3','thm.bool.EQ_CLAUSES|split|3')).
fof(f40608,plain,(
  ( ! [X70,X68,X66,X69] : (~p(s(bool,X66)) | ~sP8(X68,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X69,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X69,X70))) | p(s(bool,X68)) | s(bool,'F_0') = s(bool,X66)) ) | $spl0),
  inference(forward_demodulation,[],[f40555,f421])).
fof(f40555,plain,(
  ( ! [X70,X68,X66,X69,X67,X65] : (~sP8(X68,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X69,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X69,X70))) | p(s(bool,X68)) | ~p(s(bool,happ(s(fun(X65,bool),'_dst_x0_1'(s(bool,X66))),s(X65,X67)))) | s(bool,'F_0') = s(bool,X66)) ) | $spl0),
  inference(superposition,[],[f12481,f1325])).
fof(f1325,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,X1))),s(X0,X2))) | s(bool,'F_0') = s(bool,X1)) )),
  inference(superposition,[],[f421,f329])).
fof(f12481,plain,(
  ( ! [X52,X50,X53,X51] : (~sP8(X51,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X52,bool),'_dst_x0_1'(s(bool,X50))),s(X52,X53))) | p(s(bool,X51)) | ~p(s(bool,X50))) ) | $spl0),
  inference(subsumption_resolution,[],[f12465,f322])).
fof(f12465,plain,(
  ( ! [X52,X50,X53,X51] : (~p(s(bool,'T_0')) | ~p(s(bool,X50)) | p(s(bool,X51)) | ~sP8(X51,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X52,bool),'_dst_x0_1'(s(bool,X50))),s(X52,X53)))) ) | $spl0),
  inference(superposition,[],[f1345,f9096])).
fof(f1345,plain,(
  ( ! [X80,X78,X81,X79,X82] : (~p(s(bool,X81)) | ~p(s(bool,X79)) | p(s(bool,X82)) | ~sP8(X82,X81,happ(s(fun(X78,bool),'_dst_x0_1'(s(bool,X79))),s(X78,X80)))) )),
  inference(superposition,[],[f501,f421])).
fof(f501,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP8(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f299])).
fof(f299,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP8(X2,X0,X1)) & (~sP8(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f298])).
fof(f298,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP8(X2,X0,X1)) & (~sP8(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f196])).
fof(f196,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP8(X2,X0,X1))),
  inference(definition_folding,[],[f139,f195,f194])).
fof(f194,plain,(
  ! [X0,X1,X2] : (sP7(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f195,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (sP7(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f139,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f138])).
fof(f138,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X24,X25,X26] : ((p(s(bool,X26)) <=> (p(s(bool,X25)) & p(s(bool,X24)))) <=> ((p(s(bool,X26)) | ~p(s(bool,X25)) | ~p(s(bool,X24))) & (p(s(bool,X25)) | ~p(s(bool,X26))) & (p(s(bool,X24)) | ~p(s(bool,X26)))))),
  file('thm.quantHeuristics.GUESSES_UEXISTS_THM3','thm.sat.sat_23_prove_ltb')).
fof(f40606,plain,(
  ( ! [X56,X55] : (~p(s(bool,'|dollar|exists_1'(s(fun(X55,bool),X56))))) ) | $spl138),
  inference(cnf_transformation,[],[f40606_D])).
fof(f40606_D,plain,(
  ( ! [X56,X55] : (~p(s(bool,'|dollar|exists_1'(s(fun(X55,bool),X56))))) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f40605,plain,(
  ( ! [X59,X57,X58,X56,X55] : (~p(s(bool,'|dollar|exists_1'(s(fun(X55,bool),X56)))) | ~sP8(X57,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X58,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X58,X59))) | p(s(bool,X57))) ) | $spl0),
  inference(subsumption_resolution,[],[f40604,f340])).
fof(f40604,plain,(
  ( ! [X59,X57,X58,X56,X55] : (~p(s(bool,'|dollar|exists_1'(s(fun(X55,bool),X56)))) | ~sP8(X57,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X58,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X58,X59))) | p(s(bool,X57)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X55,bool),X56)))) ) | $spl0),
  inference(forward_demodulation,[],[f40553,f364])).
fof(f364,plain,(
  ( ! [X0,X1] : (s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) )),
  inference(cnf_transformation,[],[f87])).
fof(f87,plain,(
  ! [X0,X1] : s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f86])).
fof(f86,plain,(
  ! [X0] : ! [X1] : s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X7] : ! [X5] : s(bool,'|dollar|exists_1'(s(fun(X7,bool),X5))) = s(bool,happ(s(fun(X7,bool),X5),s(X7,'const.min.@_1'(s(fun(X7,bool),X5)))))),
  file('thm.quantHeuristics.GUESSES_UEXISTS_THM3','thm.bool.EXISTS_DEF')).
fof(f40553,plain,(
  ( ! [X59,X57,X58,X56,X55] : (~sP8(X57,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X58,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X58,X59))) | p(s(bool,X57)) | ~p(s(bool,happ(s(fun(X55,bool),X56),s(X55,'const.min.@_1'(s(fun(X55,bool),X56)))))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X55,bool),X56)))) ) | $spl0),
  inference(superposition,[],[f12481,f1518])).
fof(f1518,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,happ(s(fun(X2,bool),X3),s(X2,'const.min.@_1'(s(fun(X2,bool),X3))))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X2,bool),X3)))) )),
  inference(superposition,[],[f364,f329])).
fof(f40598,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8(X1,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3))) | p(s(bool,X1))) ) | $spl0),
  inference(condensation,[],[f40597])).
fof(f40597,plain,(
  ( ! [X12,X10,X13,X11,X9] : (~sP8(X10,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X11,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X11,X12))) | p(s(bool,X10)) | s(bool,X9) = s(bool,X13) | s(bool,'T_0') = s(bool,X13)) ) | $spl0),
  inference(subsumption_resolution,[],[f40542,f902])).
fof(f40542,plain,(
  ( ! [X12,X10,X13,X11,X9] : (~sP8(X10,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X11,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X11,X12))) | p(s(bool,X10)) | ~p(s(bool,X9)) | s(bool,X9) = s(bool,X13) | s(bool,'T_0') = s(bool,X13)) ) | $spl0),
  inference(superposition,[],[f12481,f836])).
fof(f40595,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8(X1,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3))) | p(s(bool,X1))) ) | $spl0),
  inference(condensation,[],[f40594])).
fof(f40594,plain,(
  ( ! [X6,X4,X8,X7,X5] : (~sP8(X5,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X6,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X6,X7))) | p(s(bool,X5)) | s(bool,X4) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) ) | $spl0),
  inference(subsumption_resolution,[],[f40541,f902])).
fof(f40541,plain,(
  ( ! [X6,X4,X8,X7,X5] : (~sP8(X5,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X6,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X6,X7))) | p(s(bool,X5)) | ~p(s(bool,X4)) | s(bool,X4) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) ) | $spl0),
  inference(superposition,[],[f12481,f836])).
fof(f40493,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~sP7('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)),X1)) ) | $spl134),
  inference(cnf_transformation,[],[f40493_D])).
fof(f40493_D,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~sP7('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)),X1)) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f40495,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP7('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)),X1)) ) | $spl0),
  inference(condensation,[],[f40494])).
fof(f40494,plain,(
  ( ! [X37,X35,X33,X36,X34] : (p(s(bool,X34)) | ~sP7('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X35,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X35,X36)),X34) | s(bool,X33) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) ) | $spl0),
  inference(subsumption_resolution,[],[f40460,f902])).
fof(f40460,plain,(
  ( ! [X37,X35,X33,X36,X34] : (~p(s(bool,X33)) | p(s(bool,X34)) | ~sP7('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X35,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X35,X36)),X34) | s(bool,X33) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) ) | $spl0),
  inference(superposition,[],[f12442,f836])).
fof(f12442,plain,(
  ( ! [X52,X50,X53,X51] : (~p(s(bool,X50)) | p(s(bool,X51)) | ~sP7('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X52,bool),'_dst_x0_1'(s(bool,X50))),s(X52,X53)),X51)) ) | $spl0),
  inference(subsumption_resolution,[],[f12421,f322])).
fof(f12421,plain,(
  ( ! [X52,X50,X53,X51] : (~p(s(bool,'T_0')) | ~p(s(bool,X50)) | p(s(bool,X51)) | ~sP7('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X52,bool),'_dst_x0_1'(s(bool,X50))),s(X52,X53)),X51)) ) | $spl0),
  inference(superposition,[],[f1344,f9096])).
fof(f1344,plain,(
  ( ! [X76,X74,X77,X75,X73] : (~p(s(bool,X76)) | ~p(s(bool,X74)) | p(s(bool,X77)) | ~sP7(X76,happ(s(fun(X73,bool),'_dst_x0_1'(s(bool,X74))),s(X73,X75)),X77)) )),
  inference(superposition,[],[f492,f421])).
fof(f492,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP7(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f297])).
fof(f297,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(flattening,[],[f296])).
fof(f296,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(nnf_transformation,[],[f194])).
fof(f40492,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP7('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)),X1)) ) | $spl0),
  inference(condensation,[],[f40491])).
fof(f40491,plain,(
  ( ! [X30,X28,X31,X29,X32] : (p(s(bool,X29)) | ~sP7('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X30,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X30,X31)),X29) | s(bool,X28) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) ) | $spl0),
  inference(subsumption_resolution,[],[f40459,f902])).
fof(f40459,plain,(
  ( ! [X30,X28,X31,X29,X32] : (~p(s(bool,X28)) | p(s(bool,X29)) | ~sP7('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X30,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X30,X31)),X29) | s(bool,X28) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) ) | $spl0),
  inference(superposition,[],[f12442,f836])).
fof(f40056,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~sP4('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)),X1)) ) | $spl132),
  inference(cnf_transformation,[],[f40056_D])).
fof(f40056_D,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~sP4('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)),X1)) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f40058,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP4('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)),X1)) ) | $spl0),
  inference(condensation,[],[f40057])).
fof(f40057,plain,(
  ( ! [X37,X35,X33,X36,X34] : (p(s(bool,X34)) | ~sP4('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X35,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X35,X36)),X34) | s(bool,X33) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) ) | $spl0),
  inference(subsumption_resolution,[],[f40024,f902])).
fof(f40024,plain,(
  ( ! [X37,X35,X33,X36,X34] : (~p(s(bool,X33)) | p(s(bool,X34)) | ~sP4('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X35,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X35,X36)),X34) | s(bool,X33) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) ) | $spl0),
  inference(superposition,[],[f12390,f836])).
fof(f12390,plain,(
  ( ! [X52,X50,X53,X51] : (~p(s(bool,X50)) | p(s(bool,X51)) | ~sP4('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X52,bool),'_dst_x0_1'(s(bool,X50))),s(X52,X53)),X51)) ) | $spl0),
  inference(subsumption_resolution,[],[f12369,f322])).
fof(f12369,plain,(
  ( ! [X52,X50,X53,X51] : (~p(s(bool,'T_0')) | ~p(s(bool,X50)) | p(s(bool,X51)) | ~sP4('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X52,bool),'_dst_x0_1'(s(bool,X50))),s(X52,X53)),X51)) ) | $spl0),
  inference(superposition,[],[f1343,f9096])).
fof(f1343,plain,(
  ( ! [X70,X68,X72,X71,X69] : (~p(s(bool,X71)) | ~p(s(bool,X69)) | p(s(bool,X72)) | ~sP4(X71,happ(s(fun(X68,bool),'_dst_x0_1'(s(bool,X69))),s(X68,X70)),X72)) )),
  inference(superposition,[],[f473,f421])).
fof(f473,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f289])).
fof(f289,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f288])).
fof(f288,plain,(
  ! [X1,X0,X2] : ((~sP4(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP4(X1,X0,X2)))),
  inference(flattening,[],[f287])).
fof(f287,plain,(
  ! [X1,X0,X2] : ((~sP4(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP4(X1,X0,X2)))),
  inference(nnf_transformation,[],[f190])).
fof(f190,plain,(
  ! [X1,X0,X2] : (sP4(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f40055,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP4('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)),X1)) ) | $spl0),
  inference(condensation,[],[f40054])).
fof(f40054,plain,(
  ( ! [X30,X28,X31,X29,X32] : (p(s(bool,X29)) | ~sP4('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X30,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X30,X31)),X29) | s(bool,X28) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) ) | $spl0),
  inference(subsumption_resolution,[],[f40023,f902])).
fof(f40023,plain,(
  ( ! [X30,X28,X31,X29,X32] : (~p(s(bool,X28)) | p(s(bool,X29)) | ~sP4('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X30,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X30,X31)),X29) | s(bool,X28) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) ) | $spl0),
  inference(superposition,[],[f12390,f836])).
fof(f39925,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~sP5('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)),X1)) ) | $spl130),
  inference(cnf_transformation,[],[f39925_D])).
fof(f39925_D,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~sP5('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)),X1)) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f39927,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP5('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)),X1)) ) | $spl0),
  inference(condensation,[],[f39926])).
fof(f39926,plain,(
  ( ! [X37,X35,X33,X36,X34] : (p(s(bool,X34)) | ~sP5('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X35,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X35,X36)),X34) | s(bool,X33) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) ) | $spl0),
  inference(subsumption_resolution,[],[f39888,f902])).
fof(f39888,plain,(
  ( ! [X37,X35,X33,X36,X34] : (~p(s(bool,X33)) | p(s(bool,X34)) | ~sP5('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X35,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X35,X36)),X34) | s(bool,X33) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) ) | $spl0),
  inference(superposition,[],[f12338,f836])).
fof(f12338,plain,(
  ( ! [X52,X50,X53,X51] : (~p(s(bool,X50)) | p(s(bool,X51)) | ~sP5('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X52,bool),'_dst_x0_1'(s(bool,X50))),s(X52,X53)),X51)) ) | $spl0),
  inference(subsumption_resolution,[],[f12317,f322])).
fof(f12317,plain,(
  ( ! [X52,X50,X53,X51] : (~p(s(bool,'T_0')) | ~p(s(bool,X50)) | p(s(bool,X51)) | ~sP5('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X52,bool),'_dst_x0_1'(s(bool,X50))),s(X52,X53)),X51)) ) | $spl0),
  inference(superposition,[],[f1342,f9096])).
fof(f1342,plain,(
  ( ! [X66,X64,X67,X65,X63] : (~p(s(bool,X66)) | ~p(s(bool,X64)) | p(s(bool,X67)) | ~sP5(X66,happ(s(fun(X63,bool),'_dst_x0_1'(s(bool,X64))),s(X63,X65)),X67)) )),
  inference(superposition,[],[f469,f421])).
fof(f469,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f286])).
fof(f286,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f285])).
fof(f285,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP5(X2,X0,X1)))),
  inference(flattening,[],[f284])).
fof(f284,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP5(X2,X0,X1)))),
  inference(nnf_transformation,[],[f191])).
fof(f191,plain,(
  ! [X2,X0,X1] : (sP5(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f39924,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP5('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)),X1)) ) | $spl0),
  inference(condensation,[],[f39923])).
fof(f39923,plain,(
  ( ! [X30,X28,X31,X29,X32] : (p(s(bool,X29)) | ~sP5('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X30,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X30,X31)),X29) | s(bool,X28) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) ) | $spl0),
  inference(subsumption_resolution,[],[f39887,f902])).
fof(f39887,plain,(
  ( ! [X30,X28,X31,X29,X32] : (~p(s(bool,X28)) | p(s(bool,X29)) | ~sP5('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X30,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X30,X31)),X29) | s(bool,X28) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) ) | $spl0),
  inference(superposition,[],[f12338,f836])).
fof(f39777,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~sP6('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)),X1)) ) | $spl128),
  inference(cnf_transformation,[],[f39777_D])).
fof(f39777_D,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~sP6('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)),X1)) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f39779,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP6('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)),X1)) ) | $spl0),
  inference(condensation,[],[f39778])).
fof(f39778,plain,(
  ( ! [X37,X35,X33,X36,X34] : (p(s(bool,X34)) | ~sP6('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X35,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X35,X36)),X34) | s(bool,X33) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) ) | $spl0),
  inference(subsumption_resolution,[],[f39748,f902])).
fof(f39748,plain,(
  ( ! [X37,X35,X33,X36,X34] : (~p(s(bool,X33)) | p(s(bool,X34)) | ~sP6('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X35,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X35,X36)),X34) | s(bool,X33) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) ) | $spl0),
  inference(superposition,[],[f12291,f836])).
fof(f12291,plain,(
  ( ! [X52,X50,X53,X51] : (~p(s(bool,X50)) | p(s(bool,X51)) | ~sP6('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X52,bool),'_dst_x0_1'(s(bool,X50))),s(X52,X53)),X51)) ) | $spl0),
  inference(subsumption_resolution,[],[f12273,f322])).
fof(f12273,plain,(
  ( ! [X52,X50,X53,X51] : (~p(s(bool,'T_0')) | ~p(s(bool,X50)) | p(s(bool,X51)) | ~sP6('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X52,bool),'_dst_x0_1'(s(bool,X50))),s(X52,X53)),X51)) ) | $spl0),
  inference(superposition,[],[f1341,f9096])).
fof(f1341,plain,(
  ( ! [X61,X59,X62,X60,X58] : (~p(s(bool,X61)) | ~p(s(bool,X59)) | p(s(bool,X62)) | ~sP6(X61,happ(s(fun(X58,bool),'_dst_x0_1'(s(bool,X59))),s(X58,X60)),X62)) )),
  inference(superposition,[],[f465,f421])).
fof(f465,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f283])).
fof(f283,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (sP3(X2,X1,X0) & sP4(X1,X2,X0) & sP5(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP3(X2,X1,X0) | ~sP4(X1,X2,X0) | ~sP5(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(rectify,[],[f282])).
fof(f282,plain,(
  ! [X2,X1,X0] : ((~sP6(X2,X1,X0) | (sP3(X0,X1,X2) & sP4(X1,X0,X2) & sP5(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP3(X0,X1,X2) | ~sP4(X1,X0,X2) | ~sP5(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP6(X2,X1,X0)))),
  inference(flattening,[],[f281])).
fof(f281,plain,(
  ! [X2,X1,X0] : ((~sP6(X2,X1,X0) | (sP3(X0,X1,X2) & sP4(X1,X0,X2) & sP5(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP3(X0,X1,X2) | ~sP4(X1,X0,X2) | ~sP5(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP6(X2,X1,X0)))),
  inference(nnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X2,X1,X0] : (sP6(X2,X1,X0) <=> (sP3(X0,X1,X2) & sP4(X1,X0,X2) & sP5(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f39776,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP6('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)),X1)) ) | $spl0),
  inference(condensation,[],[f39775])).
fof(f39775,plain,(
  ( ! [X30,X28,X31,X29,X32] : (p(s(bool,X29)) | ~sP6('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X30,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X30,X31)),X29) | s(bool,X28) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) ) | $spl0),
  inference(subsumption_resolution,[],[f39747,f902])).
fof(f39747,plain,(
  ( ! [X30,X28,X31,X29,X32] : (~p(s(bool,X28)) | p(s(bool,X29)) | ~sP6('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X30,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X30,X31)),X29) | s(bool,X28) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) ) | $spl0),
  inference(superposition,[],[f12291,f836])).
fof(f22017,plain,(
  ( ! [X23,X21,X19,X17,X22,X18,X16] : (s(X17,happ(s(fun(X16,X17),X18),s(X16,X21))) != s(X17,sK32(X19,happ(s(fun(X22,fun(X16,X17)),'_dst_x0_1'(s(fun(X16,X17),X18))),s(X22,X23)),X16,X17)) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X16,X17),X18),s(fun(X17,bool),X19)))) ) | $spl126),
  inference(cnf_transformation,[],[f22017_D])).
fof(f22017_D,plain,(
  ( ! [X23,X21,X19,X17,X22,X18,X16] : (s(X17,happ(s(fun(X16,X17),X18),s(X16,X21))) != s(X17,sK32(X19,happ(s(fun(X22,fun(X16,X17)),'_dst_x0_1'(s(fun(X16,X17),X18))),s(X22,X23)),X16,X17)) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X16,X17),X18),s(fun(X17,bool),X19)))) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f15729,plain,(
  ( ! [X14] : (p(s(bool,X14))) ) | $spl96),
  inference(cnf_transformation,[],[f15729_D])).
fof(f15729_D,plain,(
  ( ! [X14] : (p(s(bool,X14))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f22018,plain,(
  ( ! [X30,X28,X26,X24,X31,X29,X27,X25] : (p(s(bool,X28)) | s(X25,happ(s(fun(X24,X25),X26),s(X24,X29))) != s(X25,sK32(X27,happ(s(fun(X30,fun(X24,X25)),'_dst_x0_1'(s(fun(X24,X25),X26))),s(X30,X31)),X24,X25)) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X24,X25),X26),s(fun(X25,bool),X27)))) )),
  inference(subsumption_resolution,[],[f22013,f336])).
fof(f336,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f209])).
fof(f22013,plain,(
  ( ! [X30,X28,X26,X24,X31,X29,X27,X25] : (p(s(bool,X28)) | s(X25,happ(s(fun(X24,X25),X26),s(X24,X29))) != s(X25,sK32(X27,happ(s(fun(X30,fun(X24,X25)),'_dst_x0_1'(s(fun(X24,X25),X26))),s(X30,X31)),X24,X25)) | s(bool,'T_0') = s(bool,X28) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X24,X25),X26),s(fun(X25,bool),X27)))) )),
  inference(superposition,[],[f2062,f836])).
fof(f2062,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X0,X1),X3),s(fun(X1,bool),X5)))) | s(X1,happ(s(fun(X0,X1),X3),s(X0,X6))) != s(X1,sK32(X5,happ(s(fun(X2,fun(X0,X1)),'_dst_x0_1'(s(fun(X0,X1),X3))),s(X2,X4)),X0,X1))) )),
  inference(forward_demodulation,[],[f2052,f421])).
fof(f2052,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X1,happ(s(fun(X0,X1),X3),s(X0,X6))) != s(X1,sK32(X5,happ(s(fun(X2,fun(X0,X1)),'_dst_x0_1'(s(fun(X0,X1),X3))),s(X2,X4)),X0,X1)) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X0,X1),happ(s(fun(X2,fun(X0,X1)),'_dst_x0_1'(s(fun(X0,X1),X3))),s(X2,X4))),s(fun(X1,bool),X5))))) )),
  inference(superposition,[],[f434,f421])).
fof(f434,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(X0,sK32(X3,X2,X1,X0)) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f265])).
fof(f265,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X6] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | s(X0,X6) = s(X0,happ(s(fun(X1,X0),X2),s(X1,sK31(X6,X3,X2,X1,X0)))))) & ((~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,sK32(X3,X2,X1,X0))))) & ! [X5] : s(X0,sK32(X3,X2,X1,X0)) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK31,sK32])],[f264])).
fof(f264,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X6] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | ? [X7] : s(X0,X6) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))))) & (? [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) & ! [X5] : s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f263])).
fof(f263,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) | ? [X5] : s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))) & (? [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) & ! [X5] : s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(nnf_transformation,[],[f176])).
fof(f176,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) | ? [X5] : s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))),
  inference(ennf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) => ? [X5] : s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))),
  inference(flattening,[],[f127])).
fof(f127,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) => ? [X5] : s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X6,X7] : ! [X22,X23] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X7,X6),X22),s(fun(X6,bool),X23)))) <=> ! [X10] : (~p(s(bool,happ(s(fun(X6,bool),X23),s(X6,X10)))) => ? [X21] : s(X6,X10) = s(X6,happ(s(fun(X7,X6),X22),s(X7,X21)))))),
  file('thm.quantHeuristics.GUESSES_UEXISTS_THM3','thm.quantHeuristics.GUESS_REWRITES|split|5')).
fof(f22016,plain,(
  ( ! [X23,X21,X19,X17,X22,X20,X18,X16] : (p(s(bool,X20)) | s(X17,happ(s(fun(X16,X17),X18),s(X16,X21))) != s(X17,sK32(X19,happ(s(fun(X22,fun(X16,X17)),'_dst_x0_1'(s(fun(X16,X17),X18))),s(X22,X23)),X16,X17)) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X16,X17),X18),s(fun(X17,bool),X19)))) )),
  inference(subsumption_resolution,[],[f22012,f336])).
fof(f22012,plain,(
  ( ! [X23,X21,X19,X17,X22,X20,X18,X16] : (p(s(bool,X20)) | s(X17,happ(s(fun(X16,X17),X18),s(X16,X21))) != s(X17,sK32(X19,happ(s(fun(X22,fun(X16,X17)),'_dst_x0_1'(s(fun(X16,X17),X18))),s(X22,X23)),X16,X17)) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X16,X17),X18),s(fun(X17,bool),X19))) | s(bool,'T_0') = s(bool,X20)) )),
  inference(superposition,[],[f2062,f836])).
fof(f21991,plain,(
  ( ! [X23,X21,X19,X17,X22,X18,X16] : (s(X17,happ(s(fun(X16,X17),X18),s(X16,X21))) != s(X17,sK28(X19,happ(s(fun(X22,fun(X16,X17)),'_dst_x0_1'(s(fun(X16,X17),X18))),s(X22,X23)),X16,X17)) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X16,X17),X18),s(fun(X17,bool),X19)))) ) | $spl124),
  inference(cnf_transformation,[],[f21991_D])).
fof(f21991_D,plain,(
  ( ! [X23,X21,X19,X17,X22,X18,X16] : (s(X17,happ(s(fun(X16,X17),X18),s(X16,X21))) != s(X17,sK28(X19,happ(s(fun(X22,fun(X16,X17)),'_dst_x0_1'(s(fun(X16,X17),X18))),s(X22,X23)),X16,X17)) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X16,X17),X18),s(fun(X17,bool),X19)))) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f21992,plain,(
  ( ! [X30,X28,X26,X24,X31,X29,X27,X25] : (p(s(bool,X28)) | s(X25,happ(s(fun(X24,X25),X26),s(X24,X29))) != s(X25,sK28(X27,happ(s(fun(X30,fun(X24,X25)),'_dst_x0_1'(s(fun(X24,X25),X26))),s(X30,X31)),X24,X25)) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X24,X25),X26),s(fun(X25,bool),X27)))) )),
  inference(subsumption_resolution,[],[f21987,f336])).
fof(f21987,plain,(
  ( ! [X30,X28,X26,X24,X31,X29,X27,X25] : (p(s(bool,X28)) | s(X25,happ(s(fun(X24,X25),X26),s(X24,X29))) != s(X25,sK28(X27,happ(s(fun(X30,fun(X24,X25)),'_dst_x0_1'(s(fun(X24,X25),X26))),s(X30,X31)),X24,X25)) | s(bool,'T_0') = s(bool,X28) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X24,X25),X26),s(fun(X25,bool),X27)))) )),
  inference(superposition,[],[f2038,f836])).
fof(f2038,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X0,X1),X3),s(fun(X1,bool),X5)))) | s(X1,happ(s(fun(X0,X1),X3),s(X0,X6))) != s(X1,sK28(X5,happ(s(fun(X2,fun(X0,X1)),'_dst_x0_1'(s(fun(X0,X1),X3))),s(X2,X4)),X0,X1))) )),
  inference(forward_demodulation,[],[f2028,f421])).
fof(f2028,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X1,happ(s(fun(X0,X1),X3),s(X0,X6))) != s(X1,sK28(X5,happ(s(fun(X2,fun(X0,X1)),'_dst_x0_1'(s(fun(X0,X1),X3))),s(X2,X4)),X0,X1)) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X0,X1),happ(s(fun(X2,fun(X0,X1)),'_dst_x0_1'(s(fun(X0,X1),X3))),s(X2,X4))),s(fun(X1,bool),X5))))) )),
  inference(superposition,[],[f428,f421])).
fof(f428,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(X0,sK28(X3,X2,X1,X0)) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f259])).
fof(f259,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X6] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | s(X0,X6) = s(X0,happ(s(fun(X1,X0),X2),s(X1,sK27(X6,X3,X2,X1,X0)))))) & ((p(s(bool,happ(s(fun(X0,bool),X3),s(X0,sK28(X3,X2,X1,X0))))) & ! [X5] : s(X0,sK28(X3,X2,X1,X0)) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK27,sK28])],[f258])).
fof(f258,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X6] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | ? [X7] : s(X0,X6) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))))) & (? [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) & ! [X5] : s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f257])).
fof(f257,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) | ? [X5] : s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))) & (? [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) & ! [X5] : s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(nnf_transformation,[],[f174])).
fof(f174,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) | ? [X5] : s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))),
  inference(ennf_transformation,[],[f124])).
fof(f124,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) => ? [X5] : s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))),
  inference(flattening,[],[f123])).
fof(f123,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) => ? [X5] : s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X6,X7] : ! [X22,X23] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X7,X6),X22),s(fun(X6,bool),X23)))) <=> ! [X10] : (p(s(bool,happ(s(fun(X6,bool),X23),s(X6,X10)))) => ? [X21] : s(X6,X10) = s(X6,happ(s(fun(X7,X6),X22),s(X7,X21)))))),
  file('thm.quantHeuristics.GUESSES_UEXISTS_THM3','thm.quantHeuristics.GUESS_REWRITES|split|4')).
fof(f21990,plain,(
  ( ! [X23,X21,X19,X17,X22,X20,X18,X16] : (p(s(bool,X20)) | s(X17,happ(s(fun(X16,X17),X18),s(X16,X21))) != s(X17,sK28(X19,happ(s(fun(X22,fun(X16,X17)),'_dst_x0_1'(s(fun(X16,X17),X18))),s(X22,X23)),X16,X17)) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X16,X17),X18),s(fun(X17,bool),X19)))) )),
  inference(subsumption_resolution,[],[f21986,f336])).
fof(f21986,plain,(
  ( ! [X23,X21,X19,X17,X22,X20,X18,X16] : (p(s(bool,X20)) | s(X17,happ(s(fun(X16,X17),X18),s(X16,X21))) != s(X17,sK28(X19,happ(s(fun(X22,fun(X16,X17)),'_dst_x0_1'(s(fun(X16,X17),X18))),s(X22,X23)),X16,X17)) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X16,X17),X18),s(fun(X17,bool),X19))) | s(bool,'T_0') = s(bool,X20)) )),
  inference(superposition,[],[f2038,f836])).
fof(f21261,plain,(
  ( ! [X12,X10,X13,X11] : (s(bool,'T_0') = s(bool,happ(s(fun(X11,bool),X13),s(X11,happ(s(fun(X10,X11),X12),s(X10,sK35(X13,X12,X10,X11)))))) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X10,X11),X12),s(fun(X11,bool),X13)))) ) | $spl122),
  inference(cnf_transformation,[],[f21261_D])).
fof(f21261_D,plain,(
  ( ! [X12,X10,X13,X11] : (s(bool,'T_0') = s(bool,happ(s(fun(X11,bool),X13),s(X11,happ(s(fun(X10,X11),X12),s(X10,sK35(X13,X12,X10,X11)))))) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X10,X11),X12),s(fun(X11,bool),X13)))) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f21262,plain,(
  ( ! [X19,X17,X15,X18,X16] : (p(s(bool,X19)) | s(bool,'T_0') = s(bool,happ(s(fun(X16,bool),X18),s(X16,happ(s(fun(X15,X16),X17),s(X15,sK35(X18,X17,X15,X16)))))) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X15,X16),X17),s(fun(X16,bool),X18)))) )),
  inference(subsumption_resolution,[],[f21256,f336])).
fof(f21256,plain,(
  ( ! [X19,X17,X15,X18,X16] : (p(s(bool,X19)) | s(bool,'T_0') = s(bool,happ(s(fun(X16,bool),X18),s(X16,happ(s(fun(X15,X16),X17),s(X15,sK35(X18,X17,X15,X16)))))) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X15,X16),X17),s(fun(X16,bool),X18)))) )),
  inference(superposition,[],[f2251,f836])).
fof(f2251,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X2,X0),X3),s(fun(X0,bool),X1)))) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,sK35(X1,X3,X2,X0))))))) )),
  inference(subsumption_resolution,[],[f2247,f573])).
fof(f2247,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'F_0')) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X2,X0),X3),s(fun(X0,bool),X1)))) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,sK35(X1,X3,X2,X0))))))) )),
  inference(superposition,[],[f439,f329])).
fof(f439,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK35(X3,X2,X1,X0))))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f271])).
fof(f271,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) & (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK35(X3,X2,X1,X0))))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK35])],[f270])).
fof(f270,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) & (? [X4] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f269])).
fof(f269,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))))))) & (? [X4] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(nnf_transformation,[],[f132])).
fof(f132,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))))),
  inference(flattening,[],[f131])).
fof(f131,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X6,X7] : ! [X22,X23] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X7,X6),X22),s(fun(X6,bool),X23)))) <=> ! [X21] : ~p(s(bool,happ(s(fun(X6,bool),X23),s(X6,happ(s(fun(X7,X6),X22),s(X7,X21)))))))),
  file('thm.quantHeuristics.GUESSES_UEXISTS_THM3','thm.quantHeuristics.GUESS_REWRITES|split|3')).
fof(f21260,plain,(
  ( ! [X14,X12,X10,X13,X11] : (p(s(bool,X14)) | s(bool,'T_0') = s(bool,happ(s(fun(X11,bool),X13),s(X11,happ(s(fun(X10,X11),X12),s(X10,sK35(X13,X12,X10,X11)))))) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X10,X11),X12),s(fun(X11,bool),X13)))) )),
  inference(subsumption_resolution,[],[f21255,f336])).
fof(f21255,plain,(
  ( ! [X14,X12,X10,X13,X11] : (p(s(bool,X14)) | s(bool,'T_0') = s(bool,happ(s(fun(X11,bool),X13),s(X11,happ(s(fun(X10,X11),X12),s(X10,sK35(X13,X12,X10,X11)))))) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X10,X11),X12),s(fun(X11,bool),X13))) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f2251,f836])).
fof(f20971,plain,(
  ( ! [X2,X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X2) = s(bool,sK25(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK25(X0,X1,bool,bool))))) ) | $spl118),
  inference(cnf_transformation,[],[f20971_D])).
fof(f20971_D,plain,(
  ( ! [X2,X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X2) = s(bool,sK25(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK25(X0,X1,bool,bool))))) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f20973,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | $spl121),
  inference(cnf_transformation,[],[f20973_D])).
fof(f20973_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl121),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl121])])).
fof(f20961,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,g(s(bool,sK25(X3,X4,bool,bool)))) | s(bool,X5) = s(bool,sK25(X3,X4,bool,bool)) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f2016,f836])).
fof(f2016,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK25(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK25(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f423,f329])).
fof(f423,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK25(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK25(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f253])).
fof(f253,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK25(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK25(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK25])],[f173])).
fof(f173,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f119])).
fof(f119,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.quantHeuristics.GUESSES_UEXISTS_THM3',aHL_EXT)).
fof(f20960,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,g(s(bool,sK25(X0,X1,bool,bool)))) | s(bool,X2) = s(bool,sK25(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f2016,f836])).
fof(f20743,plain,(
  ( ! [X14,X12,X17,X15,X13] : (s(bool,'T_0') = s(bool,happ(s(fun(X13,bool),X15),s(X13,happ(s(fun(X12,X13),X14),s(X12,X17))))) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X12,X13),X14),s(fun(X13,bool),X15)))) ) | $spl116),
  inference(cnf_transformation,[],[f20743_D])).
fof(f20743_D,plain,(
  ( ! [X14,X12,X17,X15,X13] : (s(bool,'T_0') = s(bool,happ(s(fun(X13,bool),X15),s(X13,happ(s(fun(X12,X13),X14),s(X12,X17))))) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X12,X13),X14),s(fun(X13,bool),X15)))) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f20744,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (p(s(bool,X22)) | s(bool,'T_0') = s(bool,happ(s(fun(X19,bool),X21),s(X19,happ(s(fun(X18,X19),X20),s(X18,X23))))) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X18,X19),X20),s(fun(X19,bool),X21)))) )),
  inference(subsumption_resolution,[],[f20738,f336])).
fof(f20738,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (p(s(bool,X22)) | s(bool,'T_0') = s(bool,happ(s(fun(X19,bool),X21),s(X19,happ(s(fun(X18,X19),X20),s(X18,X23))))) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X18,X19),X20),s(fun(X19,bool),X21)))) )),
  inference(superposition,[],[f2141,f836])).
fof(f2141,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X2,X0),X3),s(fun(X0,bool),X1)))) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,X4)))))) )),
  inference(subsumption_resolution,[],[f2137,f573])).
fof(f2137,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'F_0')) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X2,X0),X3),s(fun(X0,bool),X1)))) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,X4)))))) )),
  inference(superposition,[],[f437,f329])).
fof(f437,plain,(
  ( ! [X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f268])).
fof(f268,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X6] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK33(X6,X3,X2,X1,X0))))))))) & ((~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,sK34(X3,X2,X1,X0))))) & ! [X5] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK33,sK34])],[f267])).
fof(f267,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X6] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | ? [X7] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X7)))))))) & (? [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) & ! [X5] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f266])).
fof(f266,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) | ? [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))))) & (? [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) & ! [X5] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(nnf_transformation,[],[f177])).
fof(f177,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) | ? [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))))),
  inference(ennf_transformation,[],[f130])).
fof(f130,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) => ? [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))))),
  inference(flattening,[],[f129])).
fof(f129,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) => ? [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X6,X7] : ! [X8,X9] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X7,X6),X8),s(fun(X6,bool),X9)))) <=> ! [X10] : (~p(s(bool,happ(s(fun(X6,bool),X9),s(X6,X10)))) => ? [X21] : ~p(s(bool,happ(s(fun(X6,bool),X9),s(X6,happ(s(fun(X7,X6),X8),s(X7,X21))))))))),
  file('thm.quantHeuristics.GUESSES_UEXISTS_THM3','thm.quantHeuristics.GUESS_EXISTS_FORALL_REWRITES|split|1')).
fof(f20742,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (p(s(bool,X16)) | s(bool,'T_0') = s(bool,happ(s(fun(X13,bool),X15),s(X13,happ(s(fun(X12,X13),X14),s(X12,X17))))) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X12,X13),X14),s(fun(X13,bool),X15)))) )),
  inference(subsumption_resolution,[],[f20737,f336])).
fof(f20737,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (p(s(bool,X16)) | s(bool,'T_0') = s(bool,happ(s(fun(X13,bool),X15),s(X13,happ(s(fun(X12,X13),X14),s(X12,X17))))) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X12,X13),X14),s(fun(X13,bool),X15))) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f2141,f836])).
fof(f20539,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl110),
  inference(cnf_transformation,[],[f20539_D])).
fof(f20539_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f20535,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK25(X0,X1,bool,bool))) ) | $spl106),
  inference(cnf_transformation,[],[f20535_D])).
fof(f20535_D,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK25(X0,X1,bool,bool))) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f20542,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl115),
  inference(cnf_transformation,[],[f20542_D])).
fof(f20542_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl115),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl115])])).
fof(f20533,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X12) = s(fun(bool,bool),X13) | s(bool,'T_0') = s(bool,sK25(X12,X13,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f2015,f329])).
fof(f2015,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK25(X0,X1,X2,bool))) )),
  inference(superposition,[],[f423,f329])).
fof(f20540,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl112),
  inference(cnf_transformation,[],[f20540_D])).
fof(f20540_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f20532,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X10) = s(fun(bool,bool),X11) | s(bool,'T_0') = s(bool,sK25(X10,X11,bool,bool)) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f2015,f836])).
fof(f20531,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X7) = s(fun(bool,bool),X8) | s(bool,'T_0') = s(bool,sK25(X7,X8,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f2015,f836])).
fof(f20534,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl104),
  inference(cnf_transformation,[],[f20534_D])).
fof(f20534_D,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f20537,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | $spl109),
  inference(cnf_transformation,[],[f20537_D])).
fof(f20537_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl109),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl109])])).
fof(f20530,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,sK25(X3,X4,bool,bool)) | s(bool,X5) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f2015,f836])).
fof(f20529,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK25(X0,X1,bool,bool)) | s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f2015,f836])).
fof(f20290,plain,(
  ( ! [X4,X2,X3,X1] : (~sP9(happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X1,X2)),X3,X4) | sP10(happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X1,X2)),X3,X4)) ) | $spl102),
  inference(cnf_transformation,[],[f20290_D])).
fof(f20290_D,plain,(
  ( ! [X4,X2,X3,X1] : (~sP9(happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X1,X2)),X3,X4) | sP10(happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X1,X2)),X3,X4)) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f20292,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP9(happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X1,X2)),X3,X4) | sP10(happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X1,X2)),X3,X4)) )),
  inference(condensation,[],[f20291])).
fof(f20291,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~sP9(happ(s(fun(X36,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X36,X37)),X38,X39) | sP10(happ(s(fun(X36,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X36,X37)),X38,X39) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f20273,f902])).
fof(f20273,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X35)) | ~sP9(happ(s(fun(X36,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X36,X37)),X38,X39) | sP10(happ(s(fun(X36,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X36,X37)),X38,X39) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1356,f836])).
fof(f1356,plain,(
  ( ! [X134,X132,X136,X135,X133] : (~p(s(bool,X133)) | ~sP9(happ(s(fun(X132,bool),'_dst_x0_1'(s(bool,X133))),s(X132,X134)),X135,X136) | sP10(happ(s(fun(X132,bool),'_dst_x0_1'(s(bool,X133))),s(X132,X134)),X135,X136)) )),
  inference(superposition,[],[f557,f421])).
fof(f557,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP9(X0,X1,X2) | sP10(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f505])).
fof(f505,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~sP9(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f302])).
fof(f302,plain,(
  ! [X0,X1,X2] : ((~sP10(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP9(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP9(X0,X1,X2) | sP10(X0,X1,X2)))),
  inference(rectify,[],[f301])).
fof(f301,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP9(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP9(X2,X0,X1) | sP10(X2,X0,X1)))),
  inference(flattening,[],[f300])).
fof(f300,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP9(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP9(X2,X0,X1)) | sP10(X2,X0,X1)))),
  inference(nnf_transformation,[],[f198])).
fof(f198,plain,(
  ! [X2,X0,X1] : (sP10(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP9(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f20289,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP9(happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X1,X2)),X3,X4) | sP10(happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X1,X2)),X3,X4)) )),
  inference(condensation,[],[f20288])).
fof(f20288,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~sP9(happ(s(fun(X30,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X30,X31)),X32,X33) | sP10(happ(s(fun(X30,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X30,X31)),X32,X33) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f20272,f902])).
fof(f20272,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X29)) | ~sP9(happ(s(fun(X30,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X30,X31)),X32,X33) | sP10(happ(s(fun(X30,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X30,X31)),X32,X33) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f1356,f836])).
fof(f20216,plain,(
  ( ! [X4,X2,X3,X1] : (~sP11(happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X1,X2)),X3,X4) | sP12(happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X1,X2)),X3,X4)) ) | $spl100),
  inference(cnf_transformation,[],[f20216_D])).
fof(f20216_D,plain,(
  ( ! [X4,X2,X3,X1] : (~sP11(happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X1,X2)),X3,X4) | sP12(happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X1,X2)),X3,X4)) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f20218,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP11(happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X1,X2)),X3,X4) | sP12(happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X1,X2)),X3,X4)) )),
  inference(condensation,[],[f20217])).
fof(f20217,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~sP11(happ(s(fun(X36,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X36,X37)),X38,X39) | sP12(happ(s(fun(X36,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X36,X37)),X38,X39) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f20199,f902])).
fof(f20199,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X35)) | ~sP11(happ(s(fun(X36,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X36,X37)),X38,X39) | sP12(happ(s(fun(X36,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X36,X37)),X38,X39) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1355,f836])).
fof(f1355,plain,(
  ( ! [X127,X130,X128,X131,X129] : (~p(s(bool,X128)) | ~sP11(happ(s(fun(X127,bool),'_dst_x0_1'(s(bool,X128))),s(X127,X129)),X130,X131) | sP12(happ(s(fun(X127,bool),'_dst_x0_1'(s(bool,X128))),s(X127,X129)),X130,X131)) )),
  inference(superposition,[],[f556,f421])).
fof(f556,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP11(X0,X1,X2) | sP12(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f522])).
fof(f522,plain,(
  ( ! [X2,X0,X1] : (sP12(X0,X1,X2) | ~sP11(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f310])).
fof(f310,plain,(
  ! [X0,X1,X2] : ((~sP12(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP11(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP11(X0,X1,X2) | sP12(X0,X1,X2)))),
  inference(rectify,[],[f309])).
fof(f309,plain,(
  ! [X2,X0,X1] : ((~sP12(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP11(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP11(X2,X0,X1) | sP12(X2,X0,X1)))),
  inference(flattening,[],[f308])).
fof(f308,plain,(
  ! [X2,X0,X1] : ((~sP12(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP11(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP11(X2,X0,X1)) | sP12(X2,X0,X1)))),
  inference(nnf_transformation,[],[f201])).
fof(f20215,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP11(happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X1,X2)),X3,X4) | sP12(happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X1,X2)),X3,X4)) )),
  inference(condensation,[],[f20214])).
fof(f20214,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~sP11(happ(s(fun(X30,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X30,X31)),X32,X33) | sP12(happ(s(fun(X30,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X30,X31)),X32,X33) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f20198,f902])).
fof(f20198,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X29)) | ~sP11(happ(s(fun(X30,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X30,X31)),X32,X33) | sP12(happ(s(fun(X30,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X30,X31)),X32,X33) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f1355,f836])).
fof(f20165,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f15753,f19348])).
fof(f19348,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,X2)) ) | $spl96),
  inference(unit_resulting_resolution,[],[f15729,f15729,f15729,f368])).
fof(f15753,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) | $spl96),
  inference(unit_resulting_resolution,[],[f15729,f340])).
fof(f20164,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f15759,f322])).
fof(f15759,plain,(
  ~p(s(bool,'T_0')) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1777,f15729,f387])).
fof(f1777,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16))) | $spl11),
  inference(unit_resulting_resolution,[],[f322,f1559,f555])).
fof(f555,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s(bool,X0) != s(bool,X1) | p(s(bool,X1))) )),
  inference(condensation,[],[f554])).
fof(f554,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2)) )),
  inference(condensation,[],[f543])).
fof(f543,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f317])).
fof(f317,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f316])).
fof(f316,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f182])).
fof(f182,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f181])).
fof(f181,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f144])).
fof(f144,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X14,X19,X16,X20] : ((s(bool,X14) = s(bool,X19) & (p(s(bool,X19)) => s(bool,X16) = s(bool,X20))) => ((p(s(bool,X14)) => p(s(bool,X16))) <=> (p(s(bool,X19)) => p(s(bool,X20)))))),
  file('thm.quantHeuristics.GUESSES_UEXISTS_THM3','thm.bool.IMP_CONG')).
fof(f20163,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f15761,f322])).
fof(f15761,plain,(
  ~p(s(bool,'T_0')) | $spl96),
  inference(unit_resulting_resolution,[],[f617,f15729,f387])).
fof(f617,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f573,f336])).
fof(f20162,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f15763,f15729])).
fof(f15763,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1777,f15729,f387])).
fof(f20161,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f15769,f322])).
fof(f15769,plain,(
  ~p(s(bool,'T_0')) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1777,f15729,f387])).
fof(f20160,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f15771,f322])).
fof(f15771,plain,(
  ~p(s(bool,'T_0')) | $spl96),
  inference(unit_resulting_resolution,[],[f617,f15729,f387])).
fof(f20159,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f15773,f15729])).
fof(f15773,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1777,f15729,f387])).
fof(f20158,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f15782,f15729])).
fof(f15782,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1578,f15729,f408])).
fof(f408,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f249])).
fof(f249,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1)))),
  inference(rectify,[],[f248])).
fof(f248,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0)))),
  inference(flattening,[],[f247])).
fof(f247,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP1(X1,X0)))),
  inference(nnf_transformation,[],[f185])).
fof(f185,plain,(
  ! [X1,X0] : (sP1(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f1578,plain,(
  sP1('T_0','|dollar|exists_1'(s(fun(sK14,bool),sK16))) | $spl11),
  inference(unit_resulting_resolution,[],[f322,f1559,f409])).
fof(f409,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f249])).
fof(f20157,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f15783,f322])).
fof(f15783,plain,(
  ~p(s(bool,'T_0')) | $spl96),
  inference(unit_resulting_resolution,[],[f649,f15729,f408])).
fof(f649,plain,(
  sP1('F_0','T_0')),
  inference(unit_resulting_resolution,[],[f573,f322,f412])).
fof(f412,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f249])).
fof(f20156,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f15784,f12857])).
fof(f12857,plain,(
  ( ! [X7] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X7,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16))))) )),
  inference(subsumption_resolution,[],[f12839,f322])).
fof(f12839,plain,(
  ( ! [X7] : (~p(s(bool,'T_0')) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X7,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16))))) )),
  inference(superposition,[],[f2230,f3354])).
fof(f2230,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,happ(s(fun(X3,bool),X6),s(X3,X5)))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X4,X3),'_dst_x0_1'(s(X3,X5))),s(fun(X3,bool),X6))))) )),
  inference(superposition,[],[f425,f421])).
fof(f425,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK26(X3,X2,X1,X0))))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f256])).
fof(f15784,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) | $spl96),
  inference(unit_resulting_resolution,[],[f695,f15729,f408])).
fof(f695,plain,(
  sP1('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))),
  inference(unit_resulting_resolution,[],[f573,f682])).
fof(f682,plain,(
  ( ! [X3] : (sP1(X3,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16))) | p(s(bool,X3))) )),
  inference(subsumption_resolution,[],[f675,f322])).
fof(f675,plain,(
  ( ! [X3] : (~p(s(bool,'T_0')) | p(s(bool,X3)) | sP1(X3,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) )),
  inference(superposition,[],[f412,f621])).
fof(f20155,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f15785,f15729])).
fof(f15785,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9105,f15729,f408])).
fof(f9105,plain,(
  sP1('F_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | $spl0),
  inference(unit_resulting_resolution,[],[f573,f566,f412])).
fof(f20154,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f15786,f15729])).
fof(f15786,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | $spl96),
  inference(unit_resulting_resolution,[],[f3363,f15729,f408])).
fof(f3363,plain,(
  sP1('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))),
  inference(unit_resulting_resolution,[],[f573,f2082,f412])).
fof(f20153,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f15787,f15729])).
fof(f15787,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9549,f15729,f408])).
fof(f9549,plain,(
  sP1('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17))) | $spl4),
  inference(unit_resulting_resolution,[],[f573,f571,f412])).
fof(f20152,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f15788,f15729])).
fof(f15788,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f10792,f15729,f408])).
fof(f10792,plain,(
  sP1('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14)))) | $spl0),
  inference(unit_resulting_resolution,[],[f573,f9091,f412])).
fof(f9091,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) | $spl0),
  inference(unit_resulting_resolution,[],[f566,f367])).
fof(f367,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK18(X1,X0)))))) )),
  inference(cnf_transformation,[],[f229])).
fof(f20151,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f15789,f15729])).
fof(f15789,plain,(
  ( ! [X0] : (~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,sK14),'_dst_x0_1'(s(sK14,sK17))),s(fun(sK14,bool),sK16))))) ) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f13529,f15729,f408])).
fof(f13529,plain,(
  ( ! [X0] : (sP1('F_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,sK14),'_dst_x0_1'(s(sK14,sK17))),s(fun(sK14,bool),sK16)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f573,f12758,f412])).
fof(f12758,plain,(
  ( ! [X8] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X8,sK14),'_dst_x0_1'(s(sK14,sK17))),s(fun(sK14,bool),sK16))))) ) | $spl4),
  inference(subsumption_resolution,[],[f12738,f322])).
fof(f12738,plain,(
  ( ! [X8] : (~p(s(bool,'T_0')) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X8,sK14),'_dst_x0_1'(s(sK14,sK17))),s(fun(sK14,bool),sK16))))) ) | $spl4),
  inference(superposition,[],[f2106,f9540])).
fof(f2106,plain,(
  ( ! [X8,X7,X11,X9] : (~p(s(bool,happ(s(fun(X7,bool),X11),s(X7,X9)))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X8,X7),'_dst_x0_1'(s(X7,X9))),s(fun(X7,bool),X11))))) )),
  inference(superposition,[],[f431,f421])).
fof(f431,plain,(
  ( ! [X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f262])).
fof(f262,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X6] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK29(X6,X3,X2,X1,X0))))))))) & ((p(s(bool,happ(s(fun(X0,bool),X3),s(X0,sK30(X3,X2,X1,X0))))) & ! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK29,sK30])],[f261])).
fof(f261,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X6] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | ? [X7] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X7)))))))) & (? [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) & ! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f260])).
fof(f260,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) | ? [X5] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))))) & (? [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) & ! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(nnf_transformation,[],[f175])).
fof(f175,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) | ? [X5] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))))),
  inference(ennf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) => ? [X5] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))))),
  inference(flattening,[],[f125])).
fof(f125,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) => ? [X5] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X6,X7] : ! [X8,X9] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X7,X6),X8),s(fun(X6,bool),X9)))) <=> ! [X10] : (p(s(bool,happ(s(fun(X6,bool),X9),s(X6,X10)))) => ? [X21] : p(s(bool,happ(s(fun(X6,bool),X9),s(X6,happ(s(fun(X7,X6),X8),s(X7,X21))))))))),
  file('thm.quantHeuristics.GUESSES_UEXISTS_THM3','thm.quantHeuristics.GUESS_EXISTS_FORALL_REWRITES|split|0')).
fof(f20150,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f15795,f322])).
fof(f15795,plain,(
  ~p(s(bool,'T_0')) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1579,f15729,f408])).
fof(f1579,plain,(
  sP1('|dollar|exists_1'(s(fun(sK14,bool),sK16)),'T_0') | $spl11),
  inference(unit_resulting_resolution,[],[f322,f1559,f412])).
fof(f20149,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f15801,f322])).
fof(f15801,plain,(
  ~p(s(bool,'T_0')) | $spl96),
  inference(unit_resulting_resolution,[],[f643,f15729,f408])).
fof(f643,plain,(
  sP1('T_0','F_0')),
  inference(unit_resulting_resolution,[],[f573,f322,f409])).
fof(f20148,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f15802,f322])).
fof(f15802,plain,(
  ~p(s(bool,'T_0')) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1578,f15729,f408])).
fof(f20147,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f15810,f12857])).
fof(f15810,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) | $spl96),
  inference(unit_resulting_resolution,[],[f688,f15729,f408])).
fof(f688,plain,(
  sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'F_0')),
  inference(unit_resulting_resolution,[],[f573,f681])).
fof(f681,plain,(
  ( ! [X2] : (sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X2) | p(s(bool,X2))) )),
  inference(subsumption_resolution,[],[f674,f322])).
fof(f674,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | p(s(bool,X2)) | sP1('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X2)) )),
  inference(superposition,[],[f409,f621])).
fof(f20146,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f15811,f15729])).
fof(f15811,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9104,f15729,f408])).
fof(f9104,plain,(
  sP1('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f573,f566,f409])).
fof(f20145,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f15812,f15729])).
fof(f15812,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | $spl96),
  inference(unit_resulting_resolution,[],[f3362,f15729,f408])).
fof(f3362,plain,(
  sP1(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0')),
  inference(unit_resulting_resolution,[],[f573,f2082,f409])).
fof(f20144,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f15813,f15729])).
fof(f15813,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9548,f15729,f408])).
fof(f9548,plain,(
  sP1(happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f573,f571,f409])).
fof(f20143,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f15814,f15729])).
fof(f15814,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f10791,f15729,f408])).
fof(f10791,plain,(
  sP1(happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f573,f9091,f409])).
fof(f20142,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f15815,f15729])).
fof(f15815,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1579,f15729,f408])).
fof(f20141,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f15816,f15729])).
fof(f15816,plain,(
  ( ! [X0] : (~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,sK14),'_dst_x0_1'(s(sK14,sK17))),s(fun(sK14,bool),sK16))))) ) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f13528,f15729,f408])).
fof(f13528,plain,(
  ( ! [X0] : (sP1('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,sK14),'_dst_x0_1'(s(sK14,sK17))),s(fun(sK14,bool),sK16)),'F_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f573,f12758,f409])).
fof(f20140,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f15862,f15729])).
fof(f15862,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1584,f15729,f449])).
fof(f449,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f276])).
fof(f276,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | ((p(s(bool,X1)) | p(s(bool,X2))) & (p(s(bool,X1)) | p(s(bool,X0))))) & ((~p(s(bool,X1)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(rectify,[],[f275])).
fof(f275,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(flattening,[],[f274])).
fof(f274,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2))))) & (((~p(s(bool,X0)) & ~p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X2)))) | sP2(X2,X0,X1)))),
  inference(nnf_transformation,[],[f187])).
fof(f187,plain,(
  ! [X2,X0,X1] : (sP2(X2,X0,X1) <=> ((p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f1584,plain,(
  ( ! [X0] : (~sP2(X0,'F_0','|dollar|exists_1'(s(fun(sK14,bool),sK16)))) ) | $spl11),
  inference(unit_resulting_resolution,[],[f573,f1559,f446])).
fof(f446,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f276])).
fof(f20139,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f15863,f15729])).
fof(f15863,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f685,f15729,f449])).
fof(f685,plain,(
  ( ! [X0] : (~sP2('F_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f573,f573,f447])).
fof(f447,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f276])).
fof(f20138,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f15864,f15729])).
fof(f15864,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1588,f15729,f449])).
fof(f1588,plain,(
  ( ! [X0] : (~sP2('|dollar|exists_1'(s(fun(sK14,bool),sK16)),'F_0',X0)) ) | $spl11),
  inference(unit_resulting_resolution,[],[f573,f1559,f447])).
fof(f20137,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f15866,f15729])).
fof(f15866,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1585,f15729,f449])).
fof(f1585,plain,(
  ( ! [X0] : (~sP2(X0,'|dollar|exists_1'(s(fun(sK14,bool),sK16)),'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) ) | $spl11),
  inference(unit_resulting_resolution,[],[f1559,f1559,f446])).
fof(f20136,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f15867,f15729])).
fof(f15867,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1592,f15729,f449])).
fof(f1592,plain,(
  ( ! [X0] : (~sP2('F_0','|dollar|exists_1'(s(fun(sK14,bool),sK16)),X0)) ) | $spl11),
  inference(unit_resulting_resolution,[],[f573,f1559,f447])).
fof(f20135,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f15868,f15729])).
fof(f15868,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1593,f15729,f449])).
fof(f1593,plain,(
  ( ! [X0] : (~sP2('|dollar|exists_1'(s(fun(sK14,bool),sK16)),'|dollar|exists_1'(s(fun(sK14,bool),sK16)),X0)) ) | $spl11),
  inference(unit_resulting_resolution,[],[f1559,f1559,f447])).
fof(f20134,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f15877,f15729])).
fof(f15877,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1580,f15729,f449])).
fof(f1580,plain,(
  ( ! [X0] : (~sP2(X0,'|dollar|exists_1'(s(fun(sK14,bool),sK16)),'F_0')) ) | $spl11),
  inference(unit_resulting_resolution,[],[f573,f1559,f446])).
fof(f20133,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f15878,f15729])).
fof(f15878,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1585,f15729,f449])).
fof(f20132,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f15879,f15729])).
fof(f15879,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1592,f15729,f449])).
fof(f20131,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f15880,f15729])).
fof(f15880,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1593,f15729,f449])).
fof(f20130,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f15888,f15729])).
fof(f15888,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1580,f15729,f450])).
fof(f450,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f276])).
fof(f20129,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f15889,f15729])).
fof(f15889,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1585,f15729,f450])).
fof(f20128,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f15890,f15729])).
fof(f15890,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1592,f15729,f450])).
fof(f20127,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f15891,f15729])).
fof(f15891,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1593,f15729,f450])).
fof(f20126,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f15895,f15729])).
fof(f15895,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f664,f15729,f450])).
fof(f664,plain,(
  ( ! [X0] : (~sP2(X0,'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f573,f573,f446])).
fof(f20125,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f15896,f15729])).
fof(f15896,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1584,f15729,f450])).
fof(f20124,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f15898,f15729])).
fof(f15898,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1588,f15729,f450])).
fof(f20123,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f15899,f15729])).
fof(f15899,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1580,f15729,f450])).
fof(f20122,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f15900,f15729])).
fof(f15900,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1585,f15729,f450])).
fof(f20121,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f15902,f15729])).
fof(f15902,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1593,f15729,f450])).
fof(f20120,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f15908,f15729])).
fof(f15908,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1584,f15729,f451])).
fof(f451,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f276])).
fof(f20119,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f15909,f15729])).
fof(f15909,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f685,f15729,f451])).
fof(f20118,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f15910,f15729])).
fof(f15910,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1588,f15729,f451])).
fof(f20117,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f15912,f15729])).
fof(f15912,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1585,f15729,f451])).
fof(f20116,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f15913,f15729])).
fof(f15913,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1592,f15729,f451])).
fof(f20115,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f15914,f15729])).
fof(f15914,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1593,f15729,f451])).
fof(f20114,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f15919,f15729])).
fof(f15919,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f664,f15729,f451])).
fof(f20113,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f15920,f15729])).
fof(f15920,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1584,f15729,f451])).
fof(f20112,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f15922,f15729])).
fof(f15922,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1588,f15729,f451])).
fof(f20111,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f15923,f15729])).
fof(f15923,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1580,f15729,f451])).
fof(f20110,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f15924,f15729])).
fof(f15924,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1585,f15729,f451])).
fof(f20109,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f15926,f15729])).
fof(f15926,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1593,f15729,f451])).
fof(f20108,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f15932,f15729])).
fof(f15932,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X0,X1))))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f1311,f573,f15729,f465])).
fof(f1311,plain,(
  ( ! [X2,X0,X1] : (sP6('T_0',happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,X1))),s(X0,X2)),X1)) )),
  inference(unit_resulting_resolution,[],[f322,f421,f482])).
fof(f482,plain,(
  ( ! [X2,X0,X1] : (sP6(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f292])).
fof(f292,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP6(X2,X1,X0)) & (~sP6(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f193])).
fof(f193,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP6(X2,X1,X0))),
  inference(definition_folding,[],[f137,f192,f191,f190,f189])).
fof(f189,plain,(
  ! [X0,X1,X2] : (sP3(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f137,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f136])).
fof(f136,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X24,X25,X26] : ((p(s(bool,X26)) <=> s(bool,X24) = s(bool,X25)) <=> ((p(s(bool,X26)) | p(s(bool,X25)) | p(s(bool,X24))) & (p(s(bool,X26)) | ~p(s(bool,X24)) | ~p(s(bool,X25))) & (p(s(bool,X25)) | ~p(s(bool,X24)) | ~p(s(bool,X26))) & (p(s(bool,X24)) | ~p(s(bool,X25)) | ~p(s(bool,X26)))))),
  file('thm.quantHeuristics.GUESSES_UEXISTS_THM3','thm.sat.sat_2|49|_prove_ltb')).
fof(f20107,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f15933,f15729])).
fof(f15933,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1599,f573,f15729,f465])).
fof(f1599,plain,(
  sP6('T_0','|dollar|exists_1'(s(fun(sK14,bool),sK16)),'F_0') | $spl11),
  inference(unit_resulting_resolution,[],[f587,f582,f578,f1559,f467])).
fof(f467,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X2,X1) | p(s(bool,X1)) | sP6(X0,X1,X2) | ~sP4(X1,X2,X0) | ~sP3(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f283])).
fof(f578,plain,(
  ( ! [X0,X1] : (sP5(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f573,f471])).
fof(f471,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f286])).
fof(f582,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f322,f474])).
fof(f474,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f289])).
fof(f587,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f322,f478])).
fof(f478,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f291])).
fof(f291,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(flattening,[],[f290])).
fof(f290,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(nnf_transformation,[],[f189])).
fof(f20106,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f15934,f322])).
fof(f15934,plain,(
  ~p(s(bool,'T_0')) | $spl96),
  inference(unit_resulting_resolution,[],[f925,f573,f15729,f465])).
fof(f925,plain,(
  sP6('F_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f573,f617,f481])).
fof(f481,plain,(
  ( ! [X2,X0,X1] : (sP6(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f292])).
fof(f20105,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f15935,f12857])).
fof(f15935,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) | $spl96),
  inference(unit_resulting_resolution,[],[f1484,f573,f15729,f465])).
fof(f1484,plain,(
  sP6('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'F_0')),
  inference(unit_resulting_resolution,[],[f610,f583,f578,f573,f468])).
fof(f468,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X2,X1) | p(s(bool,X0)) | sP6(X0,X1,X2) | ~sP4(X1,X2,X0) | ~sP3(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f283])).
fof(f583,plain,(
  ( ! [X0,X1] : (sP4(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f573,f475])).
fof(f475,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f289])).
fof(f610,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X1)) )),
  inference(unit_resulting_resolution,[],[f318,f479])).
fof(f479,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f291])).
fof(f20104,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f15936,f15729])).
fof(f15936,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f10072,f573,f15729,f465])).
fof(f10072,plain,(
  sP6('F_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f573,f580,f583,f9179,f468])).
fof(f9179,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X1)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f566,f479])).
fof(f580,plain,(
  ( ! [X0,X1] : (sP5('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f573,f472])).
fof(f472,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f286])).
fof(f20103,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f15937,f15729])).
fof(f15937,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | $spl96),
  inference(unit_resulting_resolution,[],[f5953,f573,f15729,f465])).
fof(f5953,plain,(
  sP6('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0')),
  inference(unit_resulting_resolution,[],[f573,f580,f583,f3428,f468])).
fof(f3428,plain,(
  ( ! [X0,X1] : (sP3(X0,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X1)) )),
  inference(unit_resulting_resolution,[],[f2082,f479])).
fof(f20102,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f15938,f15729])).
fof(f15938,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f10564,f573,f15729,f465])).
fof(f10564,plain,(
  sP6('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f573,f580,f583,f9623,f468])).
fof(f9623,plain,(
  ( ! [X0,X1] : (sP3(X0,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f571,f479])).
fof(f20101,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f15939,f15729])).
fof(f15939,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f11836,f573,f15729,f465])).
fof(f11836,plain,(
  sP6('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f573,f580,f583,f10890,f468])).
fof(f10890,plain,(
  ( ! [X0,X1] : (sP3(X0,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),X1)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f9091,f479])).
fof(f20100,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f15942,f15729])).
fof(f15942,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X0,X1))))) ) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9186,f573,f15729,f465])).
fof(f9186,plain,(
  ( ! [X2,X0,X1] : (sP6('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,X1))),s(X0,X2)),X1)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f421,f566,f482])).
fof(f20099,plain,(
  $false | ($spl0 | $spl11 | $spl96)),
  inference(subsumption_resolution,[],[f15943,f15729])).
fof(f15943,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f9184,f573,f15729,f465])).
fof(f9184,plain,(
  sP6('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'|dollar|exists_1'(s(fun(sK14,bool),sK16)),'F_0') | ($spl0 | $spl11)),
  inference(unit_resulting_resolution,[],[f1572,f566,f482])).
fof(f1572,plain,(
  s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16))) | $spl11),
  inference(unit_resulting_resolution,[],[f573,f1559,f384])).
fof(f20098,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f15944,f15729])).
fof(f15944,plain,(
  ( ! [X0] : (~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),'_dst_x0_1'(s(bool,'F_0'))))))) ) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9181,f573,f15729,f465])).
fof(f9181,plain,(
  ( ! [X0,X1] : (sP6('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'|dollar|exists_1'(s(fun(X0,bool),'_dst_x0_1'(s(bool,X1)))),X1)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1519,f566,f482])).
fof(f1519,plain,(
  ( ! [X4,X5] : (s(bool,X5) = s(bool,'|dollar|exists_1'(s(fun(X4,bool),'_dst_x0_1'(s(bool,X5)))))) )),
  inference(superposition,[],[f364,f421])).
fof(f20097,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f15947,f15729])).
fof(f15947,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X0,X1))))) ) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9630,f573,f15729,f465])).
fof(f9630,plain,(
  ( ! [X2,X0,X1] : (sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,X1))),s(X0,X2)),X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f421,f571,f482])).
fof(f20096,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f15949,f322])).
fof(f15949,plain,(
  ~p(s(bool,'T_0')) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1654,f573,f15729,f465])).
fof(f1654,plain,(
  sP6('|dollar|exists_1'(s(fun(sK14,bool),sK16)),'T_0','F_0') | $spl11),
  inference(unit_resulting_resolution,[],[f617,f1559,f481])).
fof(f20095,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f16035,f322])).
fof(f16035,plain,(
  ~p(s(bool,'T_0')) | $spl96),
  inference(unit_resulting_resolution,[],[f1428,f573,f15729,f465])).
fof(f1428,plain,(
  sP6('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f587,f582,f578,f573,f467])).
fof(f20094,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f16036,f322])).
fof(f16036,plain,(
  ~p(s(bool,'T_0')) | $spl96),
  inference(unit_resulting_resolution,[],[f1311,f573,f15729,f465])).
fof(f20093,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f16037,f322])).
fof(f16037,plain,(
  ~p(s(bool,'T_0')) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1599,f573,f15729,f465])).
fof(f20092,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f16044,f12857])).
fof(f16044,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) | $spl96),
  inference(unit_resulting_resolution,[],[f1429,f573,f15729,f465])).
fof(f1429,plain,(
  sP6('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f609,f608,f578,f573,f467])).
fof(f608,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) )),
  inference(unit_resulting_resolution,[],[f318,f474])).
fof(f609,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) )),
  inference(unit_resulting_resolution,[],[f318,f478])).
fof(f20091,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f16045,f15729])).
fof(f16045,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f10062,f573,f15729,f465])).
fof(f10062,plain,(
  sP6('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'F_0','F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f573,f578,f585,f9178,f467])).
fof(f9178,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f566,f478])).
fof(f585,plain,(
  ( ! [X0,X1] : (sP4('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f573,f476])).
fof(f476,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f289])).
fof(f20090,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f16046,f15729])).
fof(f16046,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9186,f573,f15729,f465])).
fof(f20089,plain,(
  $false | ($spl0 | $spl11 | $spl96)),
  inference(subsumption_resolution,[],[f16047,f15729])).
fof(f16047,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f9184,f573,f15729,f465])).
fof(f20088,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f16048,f15729])).
fof(f16048,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9181,f573,f15729,f465])).
fof(f20087,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f16049,f15729])).
fof(f16049,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | $spl96),
  inference(unit_resulting_resolution,[],[f5943,f573,f15729,f465])).
fof(f5943,plain,(
  sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f573,f578,f585,f3427,f467])).
fof(f3427,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) )),
  inference(unit_resulting_resolution,[],[f2082,f478])).
fof(f20086,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f16050,f15729])).
fof(f16050,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f10553,f573,f15729,f465])).
fof(f10553,plain,(
  sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),'F_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f573,f578,f585,f9622,f467])).
fof(f9622,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f571,f478])).
fof(f20085,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f16051,f15729])).
fof(f16051,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9630,f573,f15729,f465])).
fof(f20084,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f16052,f15729])).
fof(f16052,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f11819,f573,f15729,f465])).
fof(f11819,plain,(
  sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),'F_0','F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f573,f578,f585,f10889,f467])).
fof(f10889,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f9091,f478])).
fof(f20083,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f16053,f15729])).
fof(f16053,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1654,f573,f15729,f465])).
fof(f20082,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f16187,f17071])).
fof(f17071,plain,(
  ( ! [X0] : (~sP6(X0,'T_0','F_0')) ) | $spl96),
  inference(unit_resulting_resolution,[],[f617,f15729,f483])).
fof(f483,plain,(
  ( ! [X2,X0,X1] : (~sP6(X2,X1,X0) | ~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f292])).
fof(f16187,plain,(
  ( ! [X0] : (sP6(X0,'T_0','F_0')) ) | $spl96),
  inference(unit_resulting_resolution,[],[f588,f583,f577,f15729,f466])).
fof(f466,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP6(X0,X1,X2) | ~sP5(X0,X2,X1) | ~sP4(X1,X2,X0) | ~sP3(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f283])).
fof(f577,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f322,f470])).
fof(f470,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f286])).
fof(f588,plain,(
  ( ! [X0,X1] : (sP3(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f322,f479])).
fof(f20081,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f16194,f17069])).
fof(f17069,plain,(
  ( ! [X0] : (~sP6(X0,'T_0','|dollar|exists_1'(s(fun(sK14,bool),sK16)))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1777,f15729,f483])).
fof(f16194,plain,(
  ( ! [X0] : (sP6(X0,'T_0','|dollar|exists_1'(s(fun(sK14,bool),sK16)))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f588,f1625,f577,f15729,f466])).
fof(f1625,plain,(
  ( ! [X0,X1] : (sP4(X0,'|dollar|exists_1'(s(fun(sK14,bool),sK16)),X1)) ) | $spl11),
  inference(unit_resulting_resolution,[],[f1559,f475])).
fof(f20080,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f16249,f18880])).
fof(f18880,plain,(
  ( ! [X0] : (~sP6(X0,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'F_0')) ) | $spl96),
  inference(unit_resulting_resolution,[],[f573,f15729,f992])).
fof(f992,plain,(
  ( ! [X6,X7] : (~sP6(X6,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X7) | p(s(bool,X7)) | ~p(s(bool,X6))) )),
  inference(subsumption_resolution,[],[f989,f322])).
fof(f989,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | ~p(s(bool,X6)) | p(s(bool,X7)) | ~sP6(X6,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X7)) )),
  inference(superposition,[],[f465,f621])).
fof(f16249,plain,(
  ( ! [X0] : (sP6(X0,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'F_0')) ) | $spl96),
  inference(unit_resulting_resolution,[],[f610,f583,f607,f15729,f466])).
fof(f607,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) )),
  inference(unit_resulting_resolution,[],[f318,f470])).
fof(f20079,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f16311,f19219])).
fof(f19219,plain,(
  ( ! [X0] : (~sP6(X0,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'F_0')) ) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f573,f15729,f10435])).
fof(f10435,plain,(
  ( ! [X15,X16] : (~sP6(X15,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X16) | p(s(bool,X16)) | ~p(s(bool,X15))) ) | $spl0),
  inference(subsumption_resolution,[],[f10360,f322])).
fof(f10360,plain,(
  ( ! [X15,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X15)) | p(s(bool,X16)) | ~sP6(X15,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X16)) ) | $spl0),
  inference(superposition,[],[f465,f9096])).
fof(f16311,plain,(
  ( ! [X0] : (sP6(X0,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'F_0')) ) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9179,f583,f9162,f15729,f466])).
fof(f9162,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f566,f470])).
fof(f20078,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f16517,f17071])).
fof(f16517,plain,(
  ( ! [X0] : (sP6(X0,'T_0','F_0')) ) | $spl96),
  inference(unit_resulting_resolution,[],[f588,f583,f578,f15729,f466])).
fof(f20077,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f16518,f18880])).
fof(f16518,plain,(
  ( ! [X0] : (sP6(X0,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'F_0')) ) | $spl96),
  inference(unit_resulting_resolution,[],[f610,f583,f578,f15729,f466])).
fof(f20076,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f16519,f19219])).
fof(f16519,plain,(
  ( ! [X0] : (sP6(X0,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'F_0')) ) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9179,f583,f578,f15729,f466])).
fof(f20075,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f16524,f941])).
fof(f941,plain,(
  ~sP6('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f322,f617,f483])).
fof(f16524,plain,(
  sP6('T_0','T_0','F_0') | $spl96),
  inference(unit_resulting_resolution,[],[f588,f582,f578,f15729,f466])).
fof(f20074,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f16525,f1124])).
fof(f1124,plain,(
  ~sP6('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'F_0')),
  inference(unit_resulting_resolution,[],[f322,f573,f992])).
fof(f16525,plain,(
  sP6('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'F_0') | $spl96),
  inference(unit_resulting_resolution,[],[f610,f582,f578,f15729,f466])).
fof(f20073,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f16526,f9136])).
fof(f9136,plain,(
  ~sP6('T_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f322,f573,f566,f465])).
fof(f16526,plain,(
  sP6('T_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'F_0') | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9179,f582,f578,f15729,f466])).
fof(f20072,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f16527,f3390])).
fof(f3390,plain,(
  ~sP6('T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0')),
  inference(unit_resulting_resolution,[],[f322,f573,f2082,f465])).
fof(f16527,plain,(
  sP6('T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0') | $spl96),
  inference(unit_resulting_resolution,[],[f3428,f582,f578,f15729,f466])).
fof(f20071,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f16528,f9580])).
fof(f9580,plain,(
  ~sP6('T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f322,f573,f571,f465])).
fof(f16528,plain,(
  sP6('T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),'F_0') | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9623,f582,f578,f15729,f466])).
fof(f20070,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f16529,f10823])).
fof(f10823,plain,(
  ~sP6('T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f322,f573,f9091,f465])).
fof(f16529,plain,(
  sP6('T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),'F_0') | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f10890,f582,f578,f15729,f466])).
fof(f20069,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f16531,f17071])).
fof(f16531,plain,(
  sP6('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'T_0','F_0') | $spl96),
  inference(unit_resulting_resolution,[],[f588,f608,f578,f15729,f466])).
fof(f20068,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f16532,f18880])).
fof(f16532,plain,(
  sP6('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'F_0') | $spl96),
  inference(unit_resulting_resolution,[],[f610,f608,f578,f15729,f466])).
fof(f20067,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f16533,f19219])).
fof(f16533,plain,(
  sP6('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'F_0') | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9179,f608,f578,f15729,f466])).
fof(f20066,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f16538,f9194])).
fof(f9194,plain,(
  ~sP6('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'T_0','F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f617,f566,f483])).
fof(f16538,plain,(
  sP6('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'T_0','F_0') | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f588,f9177,f578,f15729,f466])).
fof(f9177,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f566,f474])).
fof(f20065,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f16539,f9382])).
fof(f9382,plain,(
  ~sP6('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f573,f566,f992])).
fof(f16539,plain,(
  sP6('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'F_0') | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f610,f9177,f578,f15729,f466])).
fof(f20064,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f16540,f9137])).
fof(f9137,plain,(
  ~sP6('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f566,f573,f566,f465])).
fof(f16540,plain,(
  sP6('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'F_0') | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9179,f9177,f578,f15729,f466])).
fof(f20063,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f16545,f3443])).
fof(f3443,plain,(
  ~sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'T_0','F_0')),
  inference(unit_resulting_resolution,[],[f617,f2082,f483])).
fof(f16545,plain,(
  sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'T_0','F_0') | $spl96),
  inference(unit_resulting_resolution,[],[f588,f3426,f578,f15729,f466])).
fof(f3426,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) )),
  inference(unit_resulting_resolution,[],[f2082,f474])).
fof(f20062,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f16546,f3609])).
fof(f3609,plain,(
  ~sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'F_0')),
  inference(unit_resulting_resolution,[],[f573,f2082,f992])).
fof(f16546,plain,(
  sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'F_0') | $spl96),
  inference(unit_resulting_resolution,[],[f610,f3426,f578,f15729,f466])).
fof(f20061,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f16547,f19219])).
fof(f16547,plain,(
  sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'F_0') | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9179,f3426,f578,f15729,f466])).
fof(f20060,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f16548,f3391])).
fof(f3391,plain,(
  ~sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0')),
  inference(unit_resulting_resolution,[],[f2082,f573,f2082,f465])).
fof(f16548,plain,(
  sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0') | $spl96),
  inference(unit_resulting_resolution,[],[f3428,f3426,f578,f15729,f466])).
fof(f20059,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f16552,f9638])).
fof(f9638,plain,(
  ~sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f617,f571,f483])).
fof(f16552,plain,(
  sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),'T_0','F_0') | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f588,f9621,f578,f15729,f466])).
fof(f9621,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f571,f474])).
fof(f20058,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f16553,f9826])).
fof(f9826,plain,(
  ~sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f573,f571,f992])).
fof(f16553,plain,(
  sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'F_0') | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f610,f9621,f578,f15729,f466])).
fof(f20057,plain,(
  $false | ($spl0 | $spl4 | $spl96)),
  inference(subsumption_resolution,[],[f16554,f19219])).
fof(f16554,plain,(
  sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'F_0') | ($spl0 | $spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9179,f9621,f578,f15729,f466])).
fof(f20056,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f16556,f9581])).
fof(f9581,plain,(
  ~sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f571,f573,f571,f465])).
fof(f16556,plain,(
  sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),'F_0') | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9623,f9621,f578,f15729,f466])).
fof(f20055,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f16559,f10907])).
fof(f10907,plain,(
  ~sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),'T_0','F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f617,f9091,f483])).
fof(f16559,plain,(
  sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),'T_0','F_0') | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f588,f10888,f578,f15729,f466])).
fof(f10888,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f9091,f474])).
fof(f20054,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f16560,f11159])).
fof(f11159,plain,(
  ~sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f573,f9091,f992])).
fof(f16560,plain,(
  sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'F_0') | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f610,f10888,f578,f15729,f466])).
fof(f20053,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f16561,f19219])).
fof(f16561,plain,(
  sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'F_0') | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9179,f10888,f578,f15729,f466])).
fof(f20052,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f16564,f10824])).
fof(f10824,plain,(
  ~sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f9091,f573,f9091,f465])).
fof(f16564,plain,(
  sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),'F_0') | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f10890,f10888,f578,f15729,f466])).
fof(f20051,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f16583,f17069])).
fof(f16583,plain,(
  ( ! [X0] : (sP6(X0,'T_0','|dollar|exists_1'(s(fun(sK14,bool),sK16)))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f588,f1625,f1618,f15729,f466])).
fof(f1618,plain,(
  ( ! [X0,X1] : (sP5(X0,'|dollar|exists_1'(s(fun(sK14,bool),sK16)),X1)) ) | $spl11),
  inference(unit_resulting_resolution,[],[f1559,f471])).
fof(f20050,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f16590,f1596])).
fof(f1596,plain,(
  ~sP6('T_0','T_0','|dollar|exists_1'(s(fun(sK14,bool),sK16))) | $spl11),
  inference(unit_resulting_resolution,[],[f322,f322,f1559,f465])).
fof(f16590,plain,(
  sP6('T_0','T_0','|dollar|exists_1'(s(fun(sK14,bool),sK16))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f588,f582,f1618,f15729,f466])).
fof(f20049,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f16591,f1795])).
fof(f1795,plain,(
  ~sP6('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'|dollar|exists_1'(s(fun(sK14,bool),sK16))) | $spl11),
  inference(unit_resulting_resolution,[],[f322,f1559,f992])).
fof(f16591,plain,(
  sP6('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'|dollar|exists_1'(s(fun(sK14,bool),sK16))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f610,f582,f1618,f15729,f466])).
fof(f20048,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f16597,f17069])).
fof(f16597,plain,(
  sP6('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'T_0','|dollar|exists_1'(s(fun(sK14,bool),sK16))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f588,f608,f1618,f15729,f466])).
fof(f20047,plain,(
  $false | ($spl0 | $spl11 | $spl96)),
  inference(subsumption_resolution,[],[f16604,f9193])).
fof(f9193,plain,(
  ~sP6('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'T_0','|dollar|exists_1'(s(fun(sK14,bool),sK16))) | ($spl0 | $spl11)),
  inference(unit_resulting_resolution,[],[f1777,f566,f483])).
fof(f16604,plain,(
  sP6('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'T_0','|dollar|exists_1'(s(fun(sK14,bool),sK16))) | ($spl0 | $spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f588,f9177,f1618,f15729,f466])).
fof(f20046,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f16611,f3444])).
fof(f3444,plain,(
  ~sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'T_0','|dollar|exists_1'(s(fun(sK14,bool),sK16))) | $spl11),
  inference(unit_resulting_resolution,[],[f1777,f2082,f483])).
fof(f16611,plain,(
  sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'T_0','|dollar|exists_1'(s(fun(sK14,bool),sK16))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f588,f3426,f1618,f15729,f466])).
fof(f20045,plain,(
  $false | ($spl4 | $spl11 | $spl96)),
  inference(subsumption_resolution,[],[f16618,f9637])).
fof(f9637,plain,(
  ~sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),'T_0','|dollar|exists_1'(s(fun(sK14,bool),sK16))) | ($spl4 | $spl11)),
  inference(unit_resulting_resolution,[],[f1777,f571,f483])).
fof(f16618,plain,(
  sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),'T_0','|dollar|exists_1'(s(fun(sK14,bool),sK16))) | ($spl4 | $spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f588,f9621,f1618,f15729,f466])).
fof(f20044,plain,(
  $false | ($spl0 | $spl11 | $spl96)),
  inference(subsumption_resolution,[],[f16625,f10906])).
fof(f10906,plain,(
  ~sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),'T_0','|dollar|exists_1'(s(fun(sK14,bool),sK16))) | ($spl0 | $spl11)),
  inference(unit_resulting_resolution,[],[f1777,f9091,f483])).
fof(f16625,plain,(
  sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),'T_0','|dollar|exists_1'(s(fun(sK14,bool),sK16))) | ($spl0 | $spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f588,f10888,f1618,f15729,f466])).
fof(f20043,plain,(
  $false | ($spl0 | $spl11 | $spl96)),
  inference(subsumption_resolution,[],[f16669,f19219])).
fof(f16669,plain,(
  sP6('|dollar|exists_1'(s(fun(sK14,bool),sK16)),'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'F_0') | ($spl0 | $spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f9179,f583,f1619,f15729,f466])).
fof(f1619,plain,(
  ( ! [X0,X1] : (sP5('|dollar|exists_1'(s(fun(sK14,bool),sK16)),X0,X1)) ) | $spl11),
  inference(unit_resulting_resolution,[],[f1559,f472])).
fof(f20042,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f16673,f17069])).
fof(f16673,plain,(
  sP6('|dollar|exists_1'(s(fun(sK14,bool),sK16)),'T_0','|dollar|exists_1'(s(fun(sK14,bool),sK16))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f588,f1625,f1619,f15729,f466])).
fof(f20041,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f16691,f16859])).
fof(f16859,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2)) ) | $spl96),
  inference(unit_resulting_resolution,[],[f15729,f470])).
fof(f16691,plain,(
  ~sP5('T_0','F_0','T_0') | $spl96),
  inference(unit_resulting_resolution,[],[f588,f583,f941,f15729,f466])).
fof(f20040,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f16692,f16913])).
fof(f16913,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2)) ) | $spl96),
  inference(unit_resulting_resolution,[],[f15729,f474])).
fof(f16692,plain,(
  ~sP4('T_0','F_0','T_0') | $spl96),
  inference(unit_resulting_resolution,[],[f588,f578,f941,f15729,f466])).
fof(f20039,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f16693,f16919])).
fof(f16919,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2)) ) | $spl96),
  inference(unit_resulting_resolution,[],[f15729,f480])).
fof(f480,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f291])).
fof(f16693,plain,(
  ~sP3('F_0','T_0','T_0') | $spl96),
  inference(unit_resulting_resolution,[],[f583,f578,f941,f15729,f466])).
fof(f20038,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f16695,f16859])).
fof(f16695,plain,(
  ~sP5('T_0','|dollar|exists_1'(s(fun(sK14,bool),sK16)),'T_0') | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f588,f1625,f1596,f15729,f466])).
fof(f20037,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f16696,f16913])).
fof(f16696,plain,(
  ~sP4('T_0','|dollar|exists_1'(s(fun(sK14,bool),sK16)),'T_0') | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f588,f1618,f1596,f15729,f466])).
fof(f20036,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f16697,f16919])).
fof(f16697,plain,(
  ~sP3('|dollar|exists_1'(s(fun(sK14,bool),sK16)),'T_0','T_0') | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1625,f1618,f1596,f15729,f466])).
fof(f20035,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f16705,f16859])).
fof(f16705,plain,(
  ~sP5('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16))) | $spl96),
  inference(unit_resulting_resolution,[],[f610,f583,f1124,f15729,f466])).
fof(f20034,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f16706,f16913])).
fof(f16706,plain,(
  ~sP4('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'F_0','T_0') | $spl96),
  inference(unit_resulting_resolution,[],[f610,f578,f1124,f15729,f466])).
fof(f20033,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f16707,f16919])).
fof(f16707,plain,(
  ~sP3('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'T_0') | $spl96),
  inference(unit_resulting_resolution,[],[f583,f578,f1124,f15729,f466])).
fof(f20032,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f16709,f16859])).
fof(f16709,plain,(
  ~sP5('T_0','F_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9179,f583,f9136,f15729,f466])).
fof(f20031,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f16710,f16913])).
fof(f16710,plain,(
  ~sP4('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'F_0','T_0') | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9179,f578,f9136,f15729,f466])).
fof(f20030,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f16711,f16919])).
fof(f16711,plain,(
  ~sP3('F_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'T_0') | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f583,f578,f9136,f15729,f466])).
fof(f20029,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f16713,f16859])).
fof(f16713,plain,(
  ~sP5('T_0','F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15))) | $spl96),
  inference(unit_resulting_resolution,[],[f3428,f583,f3390,f15729,f466])).
fof(f20028,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f16714,f16913])).
fof(f16714,plain,(
  ~sP4(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0','T_0') | $spl96),
  inference(unit_resulting_resolution,[],[f3428,f578,f3390,f15729,f466])).
fof(f20027,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f16715,f16919])).
fof(f16715,plain,(
  ~sP3('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'T_0') | $spl96),
  inference(unit_resulting_resolution,[],[f583,f578,f3390,f15729,f466])).
fof(f20026,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f16717,f16859])).
fof(f16717,plain,(
  ~sP5('T_0','F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17))) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9623,f583,f9580,f15729,f466])).
fof(f20025,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f16718,f16913])).
fof(f16718,plain,(
  ~sP4(happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),'F_0','T_0') | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9623,f578,f9580,f15729,f466])).
fof(f20024,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f16719,f16919])).
fof(f16719,plain,(
  ~sP3('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),'T_0') | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f583,f578,f9580,f15729,f466])).
fof(f20023,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f16721,f16859])).
fof(f16721,plain,(
  ~sP5('T_0','F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14)))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f10890,f583,f10823,f15729,f466])).
fof(f20022,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f16722,f16913])).
fof(f16722,plain,(
  ~sP4(happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),'F_0','T_0') | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f10890,f578,f10823,f15729,f466])).
fof(f20021,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f16723,f16919])).
fof(f16723,plain,(
  ~sP3('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),'T_0') | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f583,f578,f10823,f15729,f466])).
fof(f20020,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f16733,f16919])).
fof(f16733,plain,(
  ( ! [X0,X1] : (~sP3(happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16))))),s(X0,X1)),'|dollar|exists_1'(s(fun(sK14,bool),sK16)),'F_0')) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1626,f580,f1313,f15729,f466])).
fof(f1313,plain,(
  ( ! [X2,X0,X1] : (~sP6('F_0',X0,happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,X0))),s(X1,X2)))) )),
  inference(unit_resulting_resolution,[],[f573,f421,f484])).
fof(f484,plain,(
  ( ! [X2,X0,X1] : (~sP6(X2,X1,X0) | s(bool,X0) != s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f292])).
fof(f1626,plain,(
  ( ! [X0,X1] : (sP4('|dollar|exists_1'(s(fun(sK14,bool),sK16)),X0,X1)) ) | $spl11),
  inference(unit_resulting_resolution,[],[f1559,f476])).
fof(f20019,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f16754,f16913])).
fof(f16754,plain,(
  ( ! [X0,X1] : (~sP4(happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15))))),s(X0,X1)),happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0')) ) | $spl96),
  inference(unit_resulting_resolution,[],[f3429,f580,f1310,f15729,f466])).
fof(f1310,plain,(
  ( ! [X2,X0,X1] : (~sP6('F_0',happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,X1))),s(X0,X2)),X1)) )),
  inference(unit_resulting_resolution,[],[f573,f421,f484])).
fof(f3429,plain,(
  ( ! [X0,X1] : (sP3(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f2082,f480])).
fof(f20018,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f16755,f16913])).
fof(f16755,plain,(
  ( ! [X0,X1] : (~sP4(happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17))))),s(X0,X1)),happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),'F_0')) ) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9624,f580,f1310,f15729,f466])).
fof(f9624,plain,(
  ( ! [X0,X1] : (sP3(happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),X0,X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f571,f480])).
fof(f20017,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f16756,f16913])).
fof(f16756,plain,(
  ( ! [X0,X1] : (~sP4(happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14)))))),s(X0,X1)),happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),'F_0')) ) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f10891,f580,f1310,f15729,f466])).
fof(f10891,plain,(
  ( ! [X0,X1] : (sP3(happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),X0,X1)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f9091,f480])).
fof(f20016,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f16758,f16919])).
fof(f16758,plain,(
  ( ! [X0,X1] : (~sP3('|dollar|exists_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16))))),s(X0,X1)),'F_0')) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1625,f580,f1310,f15729,f466])).
fof(f20015,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f16760,f16919])).
fof(f16760,plain,(
  ( ! [X0,X1] : (~sP3('|dollar|exists_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16))))),s(X0,X1)),'F_0')) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1625,f1618,f1310,f15729,f466])).
fof(f20014,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f16769,f16859])).
fof(f16769,plain,(
  ~sP5('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'F_0','T_0') | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f588,f583,f9194,f15729,f466])).
fof(f20013,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f16770,f16913])).
fof(f16770,plain,(
  ~sP4('T_0','F_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f588,f578,f9194,f15729,f466])).
fof(f20012,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f16771,f16919])).
fof(f16771,plain,(
  ~sP3('F_0','T_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f583,f578,f9194,f15729,f466])).
fof(f20011,plain,(
  $false | ($spl0 | $spl11 | $spl96)),
  inference(subsumption_resolution,[],[f16773,f16859])).
fof(f16773,plain,(
  ~sP5('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'|dollar|exists_1'(s(fun(sK14,bool),sK16)),'T_0') | ($spl0 | $spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f588,f1625,f9193,f15729,f466])).
fof(f20010,plain,(
  $false | ($spl0 | $spl11 | $spl96)),
  inference(subsumption_resolution,[],[f16774,f16913])).
fof(f16774,plain,(
  ~sP4('T_0','|dollar|exists_1'(s(fun(sK14,bool),sK16)),'const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | ($spl0 | $spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f588,f1618,f9193,f15729,f466])).
fof(f20009,plain,(
  $false | ($spl0 | $spl11 | $spl96)),
  inference(subsumption_resolution,[],[f16775,f16919])).
fof(f16775,plain,(
  ~sP3('|dollar|exists_1'(s(fun(sK14,bool),sK16)),'T_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | ($spl0 | $spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1625,f1618,f9193,f15729,f466])).
fof(f20008,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f16779,f16859])).
fof(f16779,plain,(
  ~sP5('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'F_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9179,f583,f9137,f15729,f466])).
fof(f20007,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f16780,f16913])).
fof(f16780,plain,(
  ~sP4('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'F_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9179,f578,f9137,f15729,f466])).
fof(f20006,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f16781,f16919])).
fof(f16781,plain,(
  ~sP3('F_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f583,f578,f9137,f15729,f466])).
fof(f20005,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f16784,f16859])).
fof(f16784,plain,(
  ~sP5(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0','T_0') | $spl96),
  inference(unit_resulting_resolution,[],[f588,f583,f3443,f15729,f466])).
fof(f20004,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f16785,f16913])).
fof(f16785,plain,(
  ~sP4('T_0','F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15))) | $spl96),
  inference(unit_resulting_resolution,[],[f588,f578,f3443,f15729,f466])).
fof(f20003,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f16786,f16919])).
fof(f16786,plain,(
  ~sP3('F_0','T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15))) | $spl96),
  inference(unit_resulting_resolution,[],[f583,f578,f3443,f15729,f466])).
fof(f20002,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f16789,f16859])).
fof(f16789,plain,(
  ~sP5(happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),'F_0','T_0') | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f588,f583,f9638,f15729,f466])).
fof(f20001,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f16790,f16913])).
fof(f16790,plain,(
  ~sP4('T_0','F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17))) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f588,f578,f9638,f15729,f466])).
fof(f20000,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f16791,f16919])).
fof(f16791,plain,(
  ~sP3('F_0','T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17))) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f583,f578,f9638,f15729,f466])).
fof(f19999,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f16794,f16859])).
fof(f16794,plain,(
  ~sP5(happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),'F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17))) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9623,f583,f9581,f15729,f466])).
fof(f19998,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f16795,f16913])).
fof(f16795,plain,(
  ~sP4(happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),'F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17))) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9623,f578,f9581,f15729,f466])).
fof(f19997,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f16796,f16919])).
fof(f16796,plain,(
  ~sP3('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),happ(s(fun(sK14,bool),sK16),s(sK14,sK17))) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f583,f578,f9581,f15729,f466])).
fof(f19996,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f16798,f16859])).
fof(f16798,plain,(
  ~sP5(happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),'F_0','T_0') | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f588,f583,f10907,f15729,f466])).
fof(f19995,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f16799,f16913])).
fof(f16799,plain,(
  ~sP4('T_0','F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14)))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f588,f578,f10907,f15729,f466])).
fof(f19994,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f16800,f16919])).
fof(f16800,plain,(
  ~sP3('F_0','T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14)))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f583,f578,f10907,f15729,f466])).
fof(f19993,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f16808,f15729])).
fof(f16808,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1618,f573,f15729,f469])).
fof(f19992,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f16809,f15729])).
fof(f16809,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f580,f573,f15729,f469])).
fof(f19991,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f16810,f15729])).
fof(f16810,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1619,f573,f15729,f469])).
fof(f19990,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f16819,f16859])).
fof(f16819,plain,(
  ( ! [X0] : (~sP5(X0,'T_0','F_0')) ) | $spl96),
  inference(unit_resulting_resolution,[],[f573,f322,f15729,f469])).
fof(f19989,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f16831,f16859])).
fof(f16831,plain,(
  ( ! [X0,X1] : (~sP5(X0,X1,'F_0')) ) | $spl96),
  inference(unit_resulting_resolution,[],[f573,f15729,f15729,f469])).
fof(f19988,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f16833,f15729])).
fof(f16833,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f578,f573,f15729,f469])).
fof(f19987,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f16834,f15729])).
fof(f16834,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1618,f573,f15729,f469])).
fof(f19986,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f16836,f15729])).
fof(f16836,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1619,f573,f15729,f469])).
fof(f19985,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f16845,f16859])).
fof(f16845,plain,(
  ( ! [X0] : (~sP5('T_0',X0,'F_0')) ) | $spl96),
  inference(unit_resulting_resolution,[],[f573,f322,f15729,f469])).
fof(f19984,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f16857,f16859])).
fof(f16857,plain,(
  ( ! [X0,X1] : (~sP5(X0,X1,'F_0')) ) | $spl96),
  inference(unit_resulting_resolution,[],[f573,f15729,f15729,f469])).
fof(f19983,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f16862,f15729])).
fof(f16862,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1625,f573,f15729,f473])).
fof(f19982,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f16863,f15729])).
fof(f16863,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f585,f573,f15729,f473])).
fof(f19981,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f16864,f15729])).
fof(f16864,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1626,f573,f15729,f473])).
fof(f19980,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f16873,f16913])).
fof(f16873,plain,(
  ( ! [X0] : (~sP4(X0,'T_0','F_0')) ) | $spl96),
  inference(unit_resulting_resolution,[],[f573,f322,f15729,f473])).
fof(f19979,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f16885,f16913])).
fof(f16885,plain,(
  ( ! [X0,X1] : (~sP4(X0,X1,'F_0')) ) | $spl96),
  inference(unit_resulting_resolution,[],[f573,f15729,f15729,f473])).
fof(f19978,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f16887,f15729])).
fof(f16887,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f583,f573,f15729,f473])).
fof(f19977,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f16888,f15729])).
fof(f16888,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1625,f573,f15729,f473])).
fof(f19976,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f16890,f15729])).
fof(f16890,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1626,f573,f15729,f473])).
fof(f19975,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f16899,f16913])).
fof(f16899,plain,(
  ( ! [X0] : (~sP4('T_0',X0,'F_0')) ) | $spl96),
  inference(unit_resulting_resolution,[],[f573,f322,f15729,f473])).
fof(f19974,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f16911,f16913])).
fof(f16911,plain,(
  ( ! [X0,X1] : (~sP4(X0,X1,'F_0')) ) | $spl96),
  inference(unit_resulting_resolution,[],[f573,f15729,f15729,f473])).
fof(f19973,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f16924,f19348])).
fof(f16924,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16))) | $spl96),
  inference(unit_resulting_resolution,[],[f1178,f15729,f482])).
fof(f1178,plain,(
  ~sP6('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))),
  inference(unit_resulting_resolution,[],[f1131,f464])).
fof(f464,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | sP5(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f283])).
fof(f1131,plain,(
  ~sP5('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'F_0')),
  inference(unit_resulting_resolution,[],[f322,f573,f1011])).
fof(f1011,plain,(
  ( ! [X6,X7] : (~sP5(X6,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X7) | p(s(bool,X7)) | ~p(s(bool,X6))) )),
  inference(subsumption_resolution,[],[f1006,f322])).
fof(f1006,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | ~p(s(bool,X6)) | p(s(bool,X7)) | ~sP5(X6,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X7)) )),
  inference(superposition,[],[f469,f621])).
fof(f19972,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f16925,f19348])).
fof(f16925,plain,(
  s(bool,'F_0') != s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f10021,f15729,f482])).
fof(f10021,plain,(
  ~sP6('T_0','F_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | $spl0),
  inference(unit_resulting_resolution,[],[f9155,f464])).
fof(f9155,plain,(
  ~sP5('T_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f322,f573,f566,f469])).
fof(f19971,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f16926,f19348])).
fof(f16926,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15))) | $spl96),
  inference(unit_resulting_resolution,[],[f5799,f15729,f482])).
fof(f5799,plain,(
  ~sP6('T_0','F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))),
  inference(unit_resulting_resolution,[],[f3404,f464])).
fof(f3404,plain,(
  ~sP5('T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0')),
  inference(unit_resulting_resolution,[],[f322,f573,f2082,f469])).
fof(f19970,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f16927,f19348])).
fof(f16927,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17))) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f10508,f15729,f482])).
fof(f10508,plain,(
  ~sP6('T_0','F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17))) | $spl4),
  inference(unit_resulting_resolution,[],[f9599,f464])).
fof(f9599,plain,(
  ~sP5('T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f322,f573,f571,f469])).
fof(f19969,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f16928,f19348])).
fof(f16928,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14)))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f11723,f15729,f482])).
fof(f11723,plain,(
  ~sP6('T_0','F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14)))) | $spl0),
  inference(unit_resulting_resolution,[],[f10860,f464])).
fof(f10860,plain,(
  ~sP5('T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f322,f573,f9091,f469])).
fof(f19968,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f16929,f19348])).
fof(f16929,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16))) | $spl96),
  inference(unit_resulting_resolution,[],[f1124,f15729,f482])).
fof(f19967,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f16930,f19348])).
fof(f16930,plain,(
  s(bool,'F_0') != s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9136,f15729,f482])).
fof(f19966,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f16931,f19348])).
fof(f16931,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15))) | $spl96),
  inference(unit_resulting_resolution,[],[f3390,f15729,f482])).
fof(f19965,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f16932,f19348])).
fof(f16932,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17))) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9580,f15729,f482])).
fof(f19964,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f16933,f19348])).
fof(f16933,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14)))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f10823,f15729,f482])).
fof(f19963,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f16935,f421])).
fof(f16935,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,X0))),s(X1,X2)))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f1313,f15729,f482])).
fof(f19962,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f16937,f621])).
fof(f16937,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16))) | $spl96),
  inference(unit_resulting_resolution,[],[f951,f15729,f482])).
fof(f951,plain,(
  ~sP6('F_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))),
  inference(unit_resulting_resolution,[],[f573,f621,f484])).
fof(f19961,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f16938,f9096])).
fof(f16938,plain,(
  s(bool,'T_0') != s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f10045,f15729,f482])).
fof(f10045,plain,(
  ~sP6('F_0','T_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | $spl0),
  inference(unit_resulting_resolution,[],[f9170,f463])).
fof(f463,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | sP4(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f283])).
fof(f9170,plain,(
  ~sP4('T_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f322,f573,f566,f473])).
fof(f19960,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f16939,f3354])).
fof(f16939,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15))) | $spl96),
  inference(unit_resulting_resolution,[],[f5909,f15729,f482])).
fof(f5909,plain,(
  ~sP6('F_0','T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))),
  inference(unit_resulting_resolution,[],[f3419,f463])).
fof(f3419,plain,(
  ~sP4('T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0')),
  inference(unit_resulting_resolution,[],[f322,f573,f2082,f473])).
fof(f19959,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f16940,f9540])).
fof(f16940,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17))) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f10534,f15729,f482])).
fof(f10534,plain,(
  ~sP6('F_0','T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17))) | $spl4),
  inference(unit_resulting_resolution,[],[f9614,f463])).
fof(f9614,plain,(
  ~sP4('T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f322,f573,f571,f473])).
fof(f19958,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f16941,f10783])).
fof(f10783,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14)))) | $spl0),
  inference(unit_resulting_resolution,[],[f322,f9091,f387])).
fof(f16941,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14)))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f11794,f15729,f482])).
fof(f11794,plain,(
  ~sP6('F_0','T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14)))) | $spl0),
  inference(unit_resulting_resolution,[],[f10879,f463])).
fof(f10879,plain,(
  ~sP4('T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f322,f573,f9091,f473])).
fof(f19957,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f16943,f1572])).
fof(f16943,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f2013,f15729,f482])).
fof(f2013,plain,(
  ~sP6('F_0','F_0','|dollar|exists_1'(s(fun(sK14,bool),sK16))) | $spl11),
  inference(unit_resulting_resolution,[],[f1627,f462])).
fof(f462,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | sP3(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f283])).
fof(f1627,plain,(
  ~sP3('|dollar|exists_1'(s(fun(sK14,bool),sK16)),'F_0','F_0') | $spl11),
  inference(unit_resulting_resolution,[],[f573,f573,f1559,f477])).
fof(f477,plain,(
  ( ! [X2,X0,X1] : (~sP3(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f291])).
fof(f19956,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f16944,f621])).
fof(f16944,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16))) | $spl96),
  inference(unit_resulting_resolution,[],[f949,f15729,f482])).
fof(f949,plain,(
  ~sP6('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'T_0')),
  inference(unit_resulting_resolution,[],[f573,f621,f484])).
fof(f19955,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f16945,f9096])).
fof(f16945,plain,(
  s(bool,'T_0') != s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f10040,f15729,f482])).
fof(f10040,plain,(
  ~sP6('F_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f9163,f463])).
fof(f9163,plain,(
  ~sP4('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'T_0','F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f322,f573,f566,f473])).
fof(f19954,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f16947,f421])).
fof(f16947,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,X0))),s(X1,X2)))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f1310,f15729,f482])).
fof(f19953,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f16948,f3354])).
fof(f16948,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15))) | $spl96),
  inference(unit_resulting_resolution,[],[f5824,f15729,f482])).
fof(f5824,plain,(
  ~sP6('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'T_0')),
  inference(unit_resulting_resolution,[],[f3412,f463])).
fof(f3412,plain,(
  ~sP4(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'T_0','F_0')),
  inference(unit_resulting_resolution,[],[f322,f573,f2082,f473])).
fof(f19952,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f16949,f9540])).
fof(f16949,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17))) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f10529,f15729,f482])).
fof(f10529,plain,(
  ~sP6('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f9607,f463])).
fof(f9607,plain,(
  ~sP4(happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f322,f573,f571,f473])).
fof(f19951,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f16950,f10783])).
fof(f16950,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14)))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f11789,f15729,f482])).
fof(f11789,plain,(
  ~sP6('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),'T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f10870,f463])).
fof(f10870,plain,(
  ~sP4(happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),'T_0','F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f322,f573,f9091,f473])).
fof(f19950,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f16951,f364])).
fof(f16951,plain,(
  ( ! [X0,X1] : (s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) != s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f1512,f15729,f482])).
fof(f1512,plain,(
  ( ! [X0,X1] : (~sP6('F_0',happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))),'|dollar|exists_1'(s(fun(X0,bool),X1)))) )),
  inference(unit_resulting_resolution,[],[f573,f364,f484])).
fof(f19949,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f16952,f1572])).
fof(f16952,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f2020,f15729,f482])).
fof(f2020,plain,(
  ~sP6('F_0','|dollar|exists_1'(s(fun(sK14,bool),sK16)),'F_0') | $spl11),
  inference(unit_resulting_resolution,[],[f1635,f462])).
fof(f1635,plain,(
  ~sP3('F_0','|dollar|exists_1'(s(fun(sK14,bool),sK16)),'F_0') | $spl11),
  inference(unit_resulting_resolution,[],[f573,f573,f1559,f477])).
fof(f19948,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f16953,f364])).
fof(f16953,plain,(
  ( ! [X0,X1] : (s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) != s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f1514,f15729,f482])).
fof(f1514,plain,(
  ( ! [X0,X1] : (~sP6('F_0','|dollar|exists_1'(s(fun(X0,bool),X1)),happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f573,f364,f484])).
fof(f19947,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f16958,f19348])).
fof(f16958,plain,(
  s(bool,'F_0') != s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f11395,f15729,f482])).
fof(f11395,plain,(
  ~sP6('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'F_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | $spl0),
  inference(unit_resulting_resolution,[],[f9156,f464])).
fof(f9156,plain,(
  ~sP5('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f566,f573,f566,f469])).
fof(f19946,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f16959,f19348])).
fof(f16959,plain,(
  s(bool,'F_0') != s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9137,f15729,f482])).
fof(f19945,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f16965,f19348])).
fof(f16965,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17))) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9581,f15729,f482])).
fof(f19944,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f16970,f1572])).
fof(f16970,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f6713,f15729,f482])).
fof(f6713,plain,(
  ~sP6('|dollar|exists_1'(s(fun(sK14,bool),sK16)),'F_0','|dollar|exists_1'(s(fun(sK14,bool),sK16))) | $spl11),
  inference(unit_resulting_resolution,[],[f1644,f462])).
fof(f1644,plain,(
  ~sP3('|dollar|exists_1'(s(fun(sK14,bool),sK16)),'F_0','|dollar|exists_1'(s(fun(sK14,bool),sK16))) | $spl11),
  inference(unit_resulting_resolution,[],[f1559,f573,f1559,f477])).
fof(f19943,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f16971,f1572])).
fof(f16971,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f6719,f15729,f482])).
fof(f6719,plain,(
  ~sP6('|dollar|exists_1'(s(fun(sK14,bool),sK16)),'|dollar|exists_1'(s(fun(sK14,bool),sK16)),'F_0') | $spl11),
  inference(unit_resulting_resolution,[],[f1645,f462])).
fof(f1645,plain,(
  ~sP3('F_0','|dollar|exists_1'(s(fun(sK14,bool),sK16)),'|dollar|exists_1'(s(fun(sK14,bool),sK16))) | $spl11),
  inference(unit_resulting_resolution,[],[f573,f1559,f1559,f477])).
fof(f19942,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17107,f17632])).
fof(f17632,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2)) ) | $spl96),
  inference(unit_resulting_resolution,[],[f15729,f15729,f15729,f498])).
fof(f498,plain,(
  ( ! [X2,X0,X1] : (sP8(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f299])).
fof(f17107,plain,(
  ( ! [X0,X1] : (~sP8(X0,X1,'F_0')) ) | $spl96),
  inference(unit_resulting_resolution,[],[f573,f15729,f486])).
fof(f486,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f295])).
fof(f295,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | (sP7(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP7(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f294])).
fof(f294,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (sP7(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP7(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f293])).
fof(f293,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (sP7(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP7(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f195])).
fof(f19941,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17139,f17632])).
fof(f17139,plain,(
  ( ! [X0,X1] : (~sP8(X0,'F_0',X1)) ) | $spl96),
  inference(unit_resulting_resolution,[],[f573,f15729,f487])).
fof(f487,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f295])).
fof(f19940,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f17142,f15729])).
fof(f17142,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1674,f1663,f15729,f488])).
fof(f488,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP7(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f295])).
fof(f1663,plain,(
  ( ! [X0] : (~sP8('T_0',X0,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) ) | $spl11),
  inference(unit_resulting_resolution,[],[f322,f1559,f486])).
fof(f1674,plain,(
  ( ! [X0,X1] : (sP7(X0,'|dollar|exists_1'(s(fun(sK14,bool),sK16)),X1)) ) | $spl11),
  inference(unit_resulting_resolution,[],[f1559,f494])).
fof(f494,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f297])).
fof(f19939,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17143,f15729])).
fof(f17143,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f593,f711,f15729,f488])).
fof(f711,plain,(
  ( ! [X0] : (~sP8('T_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f322,f573,f487])).
fof(f593,plain,(
  ( ! [X0,X1] : (sP7('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f573,f495])).
fof(f495,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f297])).
fof(f19938,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17144,f15729])).
fof(f17144,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f590,f711,f15729,f488])).
fof(f590,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f322,f493])).
fof(f493,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f297])).
fof(f19937,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f17146,f15729])).
fof(f17146,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1674,f711,f15729,f488])).
fof(f19936,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f17147,f15729])).
fof(f17147,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1675,f1666,f15729,f488])).
fof(f1666,plain,(
  ( ! [X0] : (~sP8('T_0','|dollar|exists_1'(s(fun(sK14,bool),sK16)),X0)) ) | $spl11),
  inference(unit_resulting_resolution,[],[f322,f1559,f487])).
fof(f1675,plain,(
  ( ! [X0,X1] : (sP7('|dollar|exists_1'(s(fun(sK14,bool),sK16)),X0,X1)) ) | $spl11),
  inference(unit_resulting_resolution,[],[f1559,f495])).
fof(f19935,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f17148,f15729])).
fof(f17148,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f590,f1666,f15729,f488])).
fof(f19934,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f17150,f15729])).
fof(f17150,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1674,f1666,f15729,f488])).
fof(f19933,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17152,f15729])).
fof(f17152,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f593,f9198,f15729,f488])).
fof(f9198,plain,(
  ( ! [X0] : (~sP8('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'F_0',X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f573,f566,f487])).
fof(f19932,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17153,f15729])).
fof(f17153,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9231,f9198,f15729,f488])).
fof(f9231,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f566,f493])).
fof(f19931,plain,(
  $false | ($spl0 | $spl11 | $spl96)),
  inference(subsumption_resolution,[],[f17155,f15729])).
fof(f17155,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1674,f9198,f15729,f488])).
fof(f19930,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17157,f15729])).
fof(f17157,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f593,f3449,f15729,f488])).
fof(f3449,plain,(
  ( ! [X0] : (~sP8(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f573,f2082,f487])).
fof(f19929,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17158,f15729])).
fof(f17158,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f3476,f3449,f15729,f488])).
fof(f3476,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) )),
  inference(unit_resulting_resolution,[],[f2082,f493])).
fof(f19928,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f17160,f15729])).
fof(f17160,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1674,f3449,f15729,f488])).
fof(f19927,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f17162,f15729])).
fof(f17162,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f593,f9642,f15729,f488])).
fof(f9642,plain,(
  ( ! [X0] : (~sP8(happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),'F_0',X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f573,f571,f487])).
fof(f19926,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f17163,f15729])).
fof(f17163,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9675,f9642,f15729,f488])).
fof(f9675,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f571,f493])).
fof(f19925,plain,(
  $false | ($spl4 | $spl11 | $spl96)),
  inference(subsumption_resolution,[],[f17165,f15729])).
fof(f17165,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl4 | $spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1674,f9642,f15729,f488])).
fof(f19924,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17167,f15729])).
fof(f17167,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f593,f10911,f15729,f488])).
fof(f10911,plain,(
  ( ! [X0] : (~sP8(happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),'F_0',X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f573,f9091,f487])).
fof(f19923,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17168,f15729])).
fof(f17168,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f10960,f10911,f15729,f488])).
fof(f10960,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f9091,f493])).
fof(f19922,plain,(
  $false | ($spl0 | $spl11 | $spl96)),
  inference(subsumption_resolution,[],[f17170,f15729])).
fof(f17170,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1674,f10911,f15729,f488])).
fof(f19921,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17179,f17400])).
fof(f17400,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2)) ) | $spl96),
  inference(unit_resulting_resolution,[],[f15729,f493])).
fof(f17179,plain,(
  ~sP7('F_0','T_0','T_0') | $spl96),
  inference(unit_resulting_resolution,[],[f711,f322,f15729,f488])).
fof(f19920,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f17181,f17400])).
fof(f17181,plain,(
  ~sP7('|dollar|exists_1'(s(fun(sK14,bool),sK16)),'T_0','T_0') | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1666,f322,f15729,f488])).
fof(f19919,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17188,f17400])).
fof(f17188,plain,(
  ~sP7('F_0','T_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9198,f322,f15729,f488])).
fof(f19918,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17190,f17400])).
fof(f17190,plain,(
  ~sP7('F_0','T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15))) | $spl96),
  inference(unit_resulting_resolution,[],[f3449,f322,f15729,f488])).
fof(f19917,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f17192,f17400])).
fof(f17192,plain,(
  ~sP7('F_0','T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17))) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9642,f322,f15729,f488])).
fof(f19916,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17194,f17400])).
fof(f17194,plain,(
  ~sP7('F_0','T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14)))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f10911,f322,f15729,f488])).
fof(f19915,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17207,f17400])).
fof(f17207,plain,(
  ( ! [X0] : (~sP7(X0,'F_0','T_0')) ) | $spl96),
  inference(unit_resulting_resolution,[],[f708,f15729,f15729,f488])).
fof(f708,plain,(
  ( ! [X0] : (~sP8('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f322,f573,f486])).
fof(f19914,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f17209,f17400])).
fof(f17209,plain,(
  ( ! [X0] : (~sP7(X0,'|dollar|exists_1'(s(fun(sK14,bool),sK16)),'T_0')) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1663,f15729,f15729,f488])).
fof(f19913,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17211,f17400])).
fof(f17211,plain,(
  ( ! [X0] : (~sP7('F_0',X0,'T_0')) ) | $spl96),
  inference(unit_resulting_resolution,[],[f711,f15729,f15729,f488])).
fof(f19912,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f17213,f17400])).
fof(f17213,plain,(
  ( ! [X0] : (~sP7('|dollar|exists_1'(s(fun(sK14,bool),sK16)),X0,'T_0')) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1666,f15729,f15729,f488])).
fof(f19911,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17227,f17400])).
fof(f17227,plain,(
  ( ! [X0] : (~sP7(X0,'F_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) ) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9197,f15729,f15729,f488])).
fof(f9197,plain,(
  ( ! [X0] : (~sP8('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X0,'F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f573,f566,f486])).
fof(f19910,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17229,f17400])).
fof(f17229,plain,(
  ( ! [X0] : (~sP7('F_0',X0,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) ) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9198,f15729,f15729,f488])).
fof(f19909,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17231,f17400])).
fof(f17231,plain,(
  ( ! [X0] : (~sP7(X0,'F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f3448,f15729,f15729,f488])).
fof(f3448,plain,(
  ( ! [X0] : (~sP8(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f573,f2082,f486])).
fof(f19908,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17233,f17400])).
fof(f17233,plain,(
  ( ! [X0] : (~sP7('F_0',X0,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f3449,f15729,f15729,f488])).
fof(f19907,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f17235,f17400])).
fof(f17235,plain,(
  ( ! [X0] : (~sP7(X0,'F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) ) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9641,f15729,f15729,f488])).
fof(f9641,plain,(
  ( ! [X0] : (~sP8(happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),X0,'F_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f573,f571,f486])).
fof(f19906,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f17237,f17400])).
fof(f17237,plain,(
  ( ! [X0] : (~sP7('F_0',X0,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) ) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9642,f15729,f15729,f488])).
fof(f19905,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17239,f17400])).
fof(f17239,plain,(
  ( ! [X0] : (~sP7(X0,'F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) ) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f10910,f15729,f15729,f488])).
fof(f10910,plain,(
  ( ! [X0] : (~sP8(happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),X0,'F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f573,f9091,f486])).
fof(f19904,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17241,f17400])).
fof(f17241,plain,(
  ( ! [X0] : (~sP7('F_0',X0,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) ) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f10911,f15729,f15729,f488])).
fof(f19903,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17244,f15729])).
fof(f17244,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f591,f708,f15729,f488])).
fof(f591,plain,(
  ( ! [X0,X1] : (sP7(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f573,f494])).
fof(f19902,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17245,f15729])).
fof(f17245,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f590,f708,f15729,f488])).
fof(f19901,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f17247,f15729])).
fof(f17247,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1675,f708,f15729,f488])).
fof(f19900,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f17248,f15729])).
fof(f17248,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1674,f1663,f15729,f488])).
fof(f19899,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f17249,f15729])).
fof(f17249,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f590,f1663,f15729,f488])).
fof(f19898,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f17251,f15729])).
fof(f17251,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1675,f1663,f15729,f488])).
fof(f19897,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f17253,f15729])).
fof(f17253,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1675,f1666,f15729,f488])).
fof(f19896,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17254,f15729])).
fof(f17254,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f591,f9197,f15729,f488])).
fof(f19895,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17255,f15729])).
fof(f17255,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9231,f9197,f15729,f488])).
fof(f19894,plain,(
  $false | ($spl0 | $spl11 | $spl96)),
  inference(subsumption_resolution,[],[f17257,f15729])).
fof(f17257,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1675,f9197,f15729,f488])).
fof(f19893,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17259,f15729])).
fof(f17259,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f591,f3448,f15729,f488])).
fof(f19892,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17260,f15729])).
fof(f17260,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f3476,f3448,f15729,f488])).
fof(f19891,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f17262,f15729])).
fof(f17262,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1675,f3448,f15729,f488])).
fof(f19890,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f17264,f15729])).
fof(f17264,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f591,f9641,f15729,f488])).
fof(f19889,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f17265,f15729])).
fof(f17265,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9675,f9641,f15729,f488])).
fof(f19888,plain,(
  $false | ($spl4 | $spl11 | $spl96)),
  inference(subsumption_resolution,[],[f17267,f15729])).
fof(f17267,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl4 | $spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1675,f9641,f15729,f488])).
fof(f19887,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17269,f15729])).
fof(f17269,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f591,f10910,f15729,f488])).
fof(f19886,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17270,f15729])).
fof(f17270,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f10960,f10910,f15729,f488])).
fof(f19885,plain,(
  $false | ($spl0 | $spl11 | $spl96)),
  inference(subsumption_resolution,[],[f17272,f15729])).
fof(f17272,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1675,f10910,f15729,f488])).
fof(f19884,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17282,f17400])).
fof(f17282,plain,(
  ~sP7('T_0','F_0','T_0') | $spl96),
  inference(unit_resulting_resolution,[],[f708,f322,f15729,f488])).
fof(f19883,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f17284,f17400])).
fof(f17284,plain,(
  ~sP7('T_0','|dollar|exists_1'(s(fun(sK14,bool),sK16)),'T_0') | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1663,f322,f15729,f488])).
fof(f19882,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17292,f17400])).
fof(f17292,plain,(
  ~sP7('T_0','F_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9197,f322,f15729,f488])).
fof(f19881,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17294,f17400])).
fof(f17294,plain,(
  ~sP7('T_0','F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15))) | $spl96),
  inference(unit_resulting_resolution,[],[f3448,f322,f15729,f488])).
fof(f19880,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f17296,f17400])).
fof(f17296,plain,(
  ~sP7('T_0','F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17))) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9641,f322,f15729,f488])).
fof(f19879,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17298,f17400])).
fof(f17298,plain,(
  ~sP7('T_0','F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14)))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f10910,f322,f15729,f488])).
fof(f19878,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17311,f17400])).
fof(f17311,plain,(
  ( ! [X0] : (~sP7(X0,'F_0','T_0')) ) | $spl96),
  inference(unit_resulting_resolution,[],[f708,f15729,f15729,f488])).
fof(f19877,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f17313,f17400])).
fof(f17313,plain,(
  ( ! [X0] : (~sP7(X0,'|dollar|exists_1'(s(fun(sK14,bool),sK16)),'T_0')) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1663,f15729,f15729,f488])).
fof(f19876,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17315,f17400])).
fof(f17315,plain,(
  ( ! [X0] : (~sP7('F_0',X0,'T_0')) ) | $spl96),
  inference(unit_resulting_resolution,[],[f711,f15729,f15729,f488])).
fof(f19875,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f17317,f17400])).
fof(f17317,plain,(
  ( ! [X0] : (~sP7('|dollar|exists_1'(s(fun(sK14,bool),sK16)),X0,'T_0')) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1666,f15729,f15729,f488])).
fof(f19874,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17331,f17400])).
fof(f17331,plain,(
  ( ! [X0] : (~sP7(X0,'F_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) ) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9197,f15729,f15729,f488])).
fof(f19873,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17333,f17400])).
fof(f17333,plain,(
  ( ! [X0] : (~sP7('F_0',X0,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) ) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9198,f15729,f15729,f488])).
fof(f19872,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17335,f17400])).
fof(f17335,plain,(
  ( ! [X0] : (~sP7(X0,'F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f3448,f15729,f15729,f488])).
fof(f19871,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17337,f17400])).
fof(f17337,plain,(
  ( ! [X0] : (~sP7('F_0',X0,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f3449,f15729,f15729,f488])).
fof(f19870,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f17339,f17400])).
fof(f17339,plain,(
  ( ! [X0] : (~sP7(X0,'F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) ) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9641,f15729,f15729,f488])).
fof(f19869,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f17341,f17400])).
fof(f17341,plain,(
  ( ! [X0] : (~sP7('F_0',X0,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) ) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9642,f15729,f15729,f488])).
fof(f19868,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17343,f17400])).
fof(f17343,plain,(
  ( ! [X0] : (~sP7(X0,'F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) ) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f10910,f15729,f15729,f488])).
fof(f19867,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17345,f17400])).
fof(f17345,plain,(
  ( ! [X0] : (~sP7('F_0',X0,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) ) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f10911,f15729,f15729,f488])).
fof(f19866,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f17349,f15729])).
fof(f17349,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1674,f573,f15729,f492])).
fof(f19865,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17350,f15729])).
fof(f17350,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f593,f573,f15729,f492])).
fof(f19864,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f17351,f15729])).
fof(f17351,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1675,f573,f15729,f492])).
fof(f19863,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17360,f17400])).
fof(f17360,plain,(
  ( ! [X0] : (~sP7(X0,'T_0','F_0')) ) | $spl96),
  inference(unit_resulting_resolution,[],[f573,f322,f15729,f492])).
fof(f19862,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17372,f17400])).
fof(f17372,plain,(
  ( ! [X0,X1] : (~sP7(X0,X1,'F_0')) ) | $spl96),
  inference(unit_resulting_resolution,[],[f573,f15729,f15729,f492])).
fof(f19861,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17374,f15729])).
fof(f17374,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f591,f573,f15729,f492])).
fof(f19860,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f17375,f15729])).
fof(f17375,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1674,f573,f15729,f492])).
fof(f19859,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f17377,f15729])).
fof(f17377,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1675,f573,f15729,f492])).
fof(f19858,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17386,f17400])).
fof(f17386,plain,(
  ( ! [X0] : (~sP7('T_0',X0,'F_0')) ) | $spl96),
  inference(unit_resulting_resolution,[],[f573,f322,f15729,f492])).
fof(f19857,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17398,f17400])).
fof(f17398,plain,(
  ( ! [X0,X1] : (~sP7(X0,X1,'F_0')) ) | $spl96),
  inference(unit_resulting_resolution,[],[f573,f15729,f15729,f492])).
fof(f19856,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f17403,f15729])).
fof(f17403,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1663,f322,f15729,f498])).
fof(f19855,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17404,f15729])).
fof(f17404,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f711,f322,f15729,f498])).
fof(f19854,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f17405,f15729])).
fof(f17405,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1666,f322,f15729,f498])).
fof(f19853,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f17407,f15729])).
fof(f17407,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1663,f15729,f15729,f498])).
fof(f19852,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17408,f15729])).
fof(f17408,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f711,f15729,f15729,f498])).
fof(f19851,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f17409,f15729])).
fof(f17409,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1666,f15729,f15729,f498])).
fof(f19850,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17410,f322])).
fof(f17410,plain,(
  ~p(s(bool,'T_0')) | $spl96),
  inference(unit_resulting_resolution,[],[f1080,f15729,f15729,f498])).
fof(f1080,plain,(
  ~sP8('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f1068,f485])).
fof(f485,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X1,X2) | sP7(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f295])).
fof(f1068,plain,(
  ~sP7('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f573,f322,f322,f492])).
fof(f19849,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17411,f12857])).
fof(f17411,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) | $spl96),
  inference(unit_resulting_resolution,[],[f1222,f15729,f15729,f498])).
fof(f1222,plain,(
  ~sP8('F_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))),
  inference(unit_resulting_resolution,[],[f1212,f485])).
fof(f1212,plain,(
  ~sP7('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'F_0')),
  inference(unit_resulting_resolution,[],[f322,f573,f1076])).
fof(f1076,plain,(
  ( ! [X6,X7] : (~sP7(X6,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X7) | p(s(bool,X7)) | ~p(s(bool,X6))) )),
  inference(subsumption_resolution,[],[f1071,f322])).
fof(f1071,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | ~p(s(bool,X6)) | p(s(bool,X7)) | ~sP7(X6,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X7)) )),
  inference(superposition,[],[f492,f621])).
fof(f19848,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17412,f15729])).
fof(f17412,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9252,f15729,f15729,f498])).
fof(f9252,plain,(
  ~sP8('F_0','T_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | $spl0),
  inference(unit_resulting_resolution,[],[f322,f573,f566,f501])).
fof(f19847,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17413,f15729])).
fof(f17413,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | $spl96),
  inference(unit_resulting_resolution,[],[f3498,f15729,f15729,f498])).
fof(f3498,plain,(
  ~sP8('F_0','T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))),
  inference(unit_resulting_resolution,[],[f322,f573,f2082,f501])).
fof(f19846,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f17414,f15729])).
fof(f17414,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9696,f15729,f15729,f498])).
fof(f9696,plain,(
  ~sP8('F_0','T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17))) | $spl4),
  inference(unit_resulting_resolution,[],[f322,f573,f571,f501])).
fof(f19845,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17415,f15729])).
fof(f17415,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f10981,f15729,f15729,f498])).
fof(f10981,plain,(
  ~sP8('F_0','T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14)))) | $spl0),
  inference(unit_resulting_resolution,[],[f322,f573,f9091,f501])).
fof(f19844,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17416,f322])).
fof(f17416,plain,(
  ~p(s(bool,'T_0')) | $spl96),
  inference(unit_resulting_resolution,[],[f3493,f15729,f15729,f498])).
fof(f3493,plain,(
  ~sP8('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'T_0')),
  inference(unit_resulting_resolution,[],[f322,f573,f2082,f501])).
fof(f19843,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f17417,f322])).
fof(f17417,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9692,f15729,f15729,f498])).
fof(f9692,plain,(
  ~sP8('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f322,f573,f571,f501])).
fof(f19842,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17418,f322])).
fof(f17418,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f10977,f15729,f15729,f498])).
fof(f10977,plain,(
  ~sP8('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),'T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f322,f573,f9091,f501])).
fof(f19841,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f17419,f15729])).
fof(f17419,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9697,f15729,f15729,f498])).
fof(f9697,plain,(
  ~sP8('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),happ(s(fun(sK14,bool),sK16),s(sK14,sK17))) | $spl4),
  inference(unit_resulting_resolution,[],[f571,f573,f571,f501])).
fof(f19840,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17420,f322])).
fof(f17420,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9248,f15729,f15729,f498])).
fof(f9248,plain,(
  ~sP8('F_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f322,f573,f566,f501])).
fof(f19839,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17421,f15729])).
fof(f17421,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9253,f15729,f15729,f498])).
fof(f9253,plain,(
  ~sP8('F_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | $spl0),
  inference(unit_resulting_resolution,[],[f566,f573,f566,f501])).
fof(f19838,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17423,f15729])).
fof(f17423,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9198,f15729,f15729,f498])).
fof(f19837,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17425,f15729])).
fof(f17425,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f3449,f15729,f15729,f498])).
fof(f19836,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f17427,f15729])).
fof(f17427,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9642,f15729,f15729,f498])).
fof(f19835,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17429,f15729])).
fof(f17429,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f10911,f15729,f15729,f498])).
fof(f19834,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f17430,f322])).
fof(f17430,plain,(
  ~p(s(bool,'T_0')) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1686,f15729,f15729,f498])).
fof(f1686,plain,(
  ~sP8('|dollar|exists_1'(s(fun(sK14,bool),sK16)),'T_0','T_0') | $spl11),
  inference(unit_resulting_resolution,[],[f322,f322,f1559,f501])).
fof(f19833,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17431,f322])).
fof(f17431,plain,(
  ~p(s(bool,'T_0')) | $spl96),
  inference(unit_resulting_resolution,[],[f711,f322,f15729,f498])).
fof(f19832,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f17432,f322])).
fof(f17432,plain,(
  ~p(s(bool,'T_0')) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1666,f322,f15729,f498])).
fof(f19831,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17438,f15729])).
fof(f17438,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9198,f322,f15729,f498])).
fof(f19830,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17439,f15729])).
fof(f17439,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | $spl96),
  inference(unit_resulting_resolution,[],[f3449,f322,f15729,f498])).
fof(f19829,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f17440,f15729])).
fof(f17440,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9642,f322,f15729,f498])).
fof(f19828,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17441,f15729])).
fof(f17441,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f10911,f322,f15729,f498])).
fof(f19827,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f17442,f15729])).
fof(f17442,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1686,f322,f15729,f498])).
fof(f19826,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17447,f322])).
fof(f17447,plain,(
  ~p(s(bool,'T_0')) | $spl96),
  inference(unit_resulting_resolution,[],[f708,f15729,f15729,f498])).
fof(f19825,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f17448,f322])).
fof(f17448,plain,(
  ~p(s(bool,'T_0')) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1663,f15729,f15729,f498])).
fof(f19824,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17449,f322])).
fof(f17449,plain,(
  ~p(s(bool,'T_0')) | $spl96),
  inference(unit_resulting_resolution,[],[f711,f15729,f15729,f498])).
fof(f19823,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f17450,f322])).
fof(f17450,plain,(
  ~p(s(bool,'T_0')) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1666,f15729,f15729,f498])).
fof(f19822,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17463,f15729])).
fof(f17463,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9197,f15729,f15729,f498])).
fof(f19821,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17464,f15729])).
fof(f17464,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9198,f15729,f15729,f498])).
fof(f19820,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17465,f15729])).
fof(f17465,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | $spl96),
  inference(unit_resulting_resolution,[],[f3448,f15729,f15729,f498])).
fof(f19819,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17466,f15729])).
fof(f17466,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | $spl96),
  inference(unit_resulting_resolution,[],[f3449,f15729,f15729,f498])).
fof(f19818,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f17467,f15729])).
fof(f17467,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9641,f15729,f15729,f498])).
fof(f19817,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f17468,f15729])).
fof(f17468,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9642,f15729,f15729,f498])).
fof(f19816,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17469,f15729])).
fof(f17469,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f10910,f15729,f15729,f498])).
fof(f19815,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17470,f15729])).
fof(f17470,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f10911,f15729,f15729,f498])).
fof(f19814,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f17471,f15729])).
fof(f17471,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1686,f15729,f15729,f498])).
fof(f19813,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17476,f15729])).
fof(f17476,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f708,f322,f15729,f498])).
fof(f19812,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f17477,f15729])).
fof(f17477,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1663,f322,f15729,f498])).
fof(f19811,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f17479,f15729])).
fof(f17479,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1666,f322,f15729,f498])).
fof(f19810,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17480,f15729])).
fof(f17480,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f708,f15729,f15729,f498])).
fof(f19809,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f17481,f15729])).
fof(f17481,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1663,f15729,f15729,f498])).
fof(f19808,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f17483,f15729])).
fof(f17483,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1666,f15729,f15729,f498])).
fof(f19807,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17484,f322])).
fof(f17484,plain,(
  ~p(s(bool,'T_0')) | $spl96),
  inference(unit_resulting_resolution,[],[f1080,f15729,f15729,f498])).
fof(f19806,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17485,f322])).
fof(f17485,plain,(
  ~p(s(bool,'T_0')) | $spl96),
  inference(unit_resulting_resolution,[],[f1222,f15729,f15729,f498])).
fof(f19805,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17486,f322])).
fof(f17486,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9252,f15729,f15729,f498])).
fof(f19804,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17487,f322])).
fof(f17487,plain,(
  ~p(s(bool,'T_0')) | $spl96),
  inference(unit_resulting_resolution,[],[f3498,f15729,f15729,f498])).
fof(f19803,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f17488,f322])).
fof(f17488,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9696,f15729,f15729,f498])).
fof(f19802,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17489,f322])).
fof(f17489,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f10981,f15729,f15729,f498])).
fof(f19801,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17490,f15729])).
fof(f17490,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | $spl96),
  inference(unit_resulting_resolution,[],[f3493,f15729,f15729,f498])).
fof(f19800,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f17491,f15729])).
fof(f17491,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9692,f15729,f15729,f498])).
fof(f19799,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17492,f15729])).
fof(f17492,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f10977,f15729,f15729,f498])).
fof(f19798,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f17493,f15729])).
fof(f17493,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9697,f15729,f15729,f498])).
fof(f19797,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17494,f15729])).
fof(f17494,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9248,f15729,f15729,f498])).
fof(f19796,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17495,f15729])).
fof(f17495,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9253,f15729,f15729,f498])).
fof(f19795,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17496,f15729])).
fof(f17496,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9197,f15729,f15729,f498])).
fof(f19794,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17498,f15729])).
fof(f17498,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f3448,f15729,f15729,f498])).
fof(f19793,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f17500,f15729])).
fof(f17500,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9641,f15729,f15729,f498])).
fof(f19792,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17502,f15729])).
fof(f17502,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f10910,f15729,f15729,f498])).
fof(f19791,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f17504,f322])).
fof(f17504,plain,(
  ~p(s(bool,'T_0')) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1686,f15729,f15729,f498])).
fof(f19790,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17505,f322])).
fof(f17505,plain,(
  ~p(s(bool,'T_0')) | $spl96),
  inference(unit_resulting_resolution,[],[f708,f322,f15729,f498])).
fof(f19789,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f17506,f322])).
fof(f17506,plain,(
  ~p(s(bool,'T_0')) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1663,f322,f15729,f498])).
fof(f19788,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17513,f15729])).
fof(f17513,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9197,f322,f15729,f498])).
fof(f19787,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17514,f15729])).
fof(f17514,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | $spl96),
  inference(unit_resulting_resolution,[],[f3448,f322,f15729,f498])).
fof(f19786,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f17515,f15729])).
fof(f17515,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9641,f322,f15729,f498])).
fof(f19785,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17516,f15729])).
fof(f17516,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f10910,f322,f15729,f498])).
fof(f19784,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f17517,f15729])).
fof(f17517,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1686,f322,f15729,f498])).
fof(f19783,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17522,f322])).
fof(f17522,plain,(
  ~p(s(bool,'T_0')) | $spl96),
  inference(unit_resulting_resolution,[],[f708,f15729,f15729,f498])).
fof(f19782,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f17523,f322])).
fof(f17523,plain,(
  ~p(s(bool,'T_0')) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1663,f15729,f15729,f498])).
fof(f19781,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17524,f322])).
fof(f17524,plain,(
  ~p(s(bool,'T_0')) | $spl96),
  inference(unit_resulting_resolution,[],[f711,f15729,f15729,f498])).
fof(f19780,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f17525,f322])).
fof(f17525,plain,(
  ~p(s(bool,'T_0')) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1666,f15729,f15729,f498])).
fof(f19779,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17538,f15729])).
fof(f17538,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9197,f15729,f15729,f498])).
fof(f19778,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17539,f15729])).
fof(f17539,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9198,f15729,f15729,f498])).
fof(f19777,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17540,f15729])).
fof(f17540,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | $spl96),
  inference(unit_resulting_resolution,[],[f3448,f15729,f15729,f498])).
fof(f19776,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17541,f15729])).
fof(f17541,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | $spl96),
  inference(unit_resulting_resolution,[],[f3449,f15729,f15729,f498])).
fof(f19775,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f17542,f15729])).
fof(f17542,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9641,f15729,f15729,f498])).
fof(f19774,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f17543,f15729])).
fof(f17543,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9642,f15729,f15729,f498])).
fof(f19773,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17544,f15729])).
fof(f17544,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f10910,f15729,f15729,f498])).
fof(f19772,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17545,f15729])).
fof(f17545,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f10911,f15729,f15729,f498])).
fof(f19771,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f17546,f15729])).
fof(f17546,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1686,f15729,f15729,f498])).
fof(f19770,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f17552,f15729])).
fof(f17552,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1666,f322,f15729,f498])).
fof(f19769,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17553,f322])).
fof(f17553,plain,(
  ~p(s(bool,'T_0')) | $spl96),
  inference(unit_resulting_resolution,[],[f1080,f322,f15729,f498])).
fof(f19768,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17554,f15729])).
fof(f17554,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | $spl96),
  inference(unit_resulting_resolution,[],[f3493,f322,f15729,f498])).
fof(f19767,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f17555,f15729])).
fof(f17555,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9692,f322,f15729,f498])).
fof(f19766,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17556,f15729])).
fof(f17556,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f10977,f322,f15729,f498])).
fof(f19765,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17557,f15729])).
fof(f17557,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9248,f322,f15729,f498])).
fof(f19764,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f17562,f322])).
fof(f17562,plain,(
  ~p(s(bool,'T_0')) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1686,f322,f15729,f498])).
fof(f19763,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17563,f15729])).
fof(f17563,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f708,f15729,f15729,f498])).
fof(f19762,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f17564,f15729])).
fof(f17564,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1663,f15729,f15729,f498])).
fof(f19761,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f17566,f15729])).
fof(f17566,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1666,f15729,f15729,f498])).
fof(f19760,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17567,f322])).
fof(f17567,plain,(
  ~p(s(bool,'T_0')) | $spl96),
  inference(unit_resulting_resolution,[],[f1080,f15729,f15729,f498])).
fof(f19759,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17568,f322])).
fof(f17568,plain,(
  ~p(s(bool,'T_0')) | $spl96),
  inference(unit_resulting_resolution,[],[f1222,f15729,f15729,f498])).
fof(f19758,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17569,f322])).
fof(f17569,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9252,f15729,f15729,f498])).
fof(f19757,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17570,f322])).
fof(f17570,plain,(
  ~p(s(bool,'T_0')) | $spl96),
  inference(unit_resulting_resolution,[],[f3498,f15729,f15729,f498])).
fof(f19756,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f17571,f322])).
fof(f17571,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9696,f15729,f15729,f498])).
fof(f19755,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17572,f322])).
fof(f17572,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f10981,f15729,f15729,f498])).
fof(f19754,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17573,f15729])).
fof(f17573,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | $spl96),
  inference(unit_resulting_resolution,[],[f3493,f15729,f15729,f498])).
fof(f19753,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f17574,f15729])).
fof(f17574,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9692,f15729,f15729,f498])).
fof(f19752,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17575,f15729])).
fof(f17575,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f10977,f15729,f15729,f498])).
fof(f19751,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f17576,f15729])).
fof(f17576,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9697,f15729,f15729,f498])).
fof(f19750,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17577,f15729])).
fof(f17577,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9248,f15729,f15729,f498])).
fof(f19749,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17578,f15729])).
fof(f17578,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9253,f15729,f15729,f498])).
fof(f19748,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17579,f15729])).
fof(f17579,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9197,f15729,f15729,f498])).
fof(f19747,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17581,f15729])).
fof(f17581,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f3448,f15729,f15729,f498])).
fof(f19746,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f17583,f15729])).
fof(f17583,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9641,f15729,f15729,f498])).
fof(f19745,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17585,f15729])).
fof(f17585,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f10910,f15729,f15729,f498])).
fof(f19744,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f17587,f322])).
fof(f17587,plain,(
  ~p(s(bool,'T_0')) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1686,f15729,f15729,f498])).
fof(f19743,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f17589,f15729])).
fof(f17589,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1663,f322,f15729,f498])).
fof(f19742,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17590,f322])).
fof(f17590,plain,(
  ~p(s(bool,'T_0')) | $spl96),
  inference(unit_resulting_resolution,[],[f1080,f322,f15729,f498])).
fof(f19741,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17591,f12857])).
fof(f17591,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) | $spl96),
  inference(unit_resulting_resolution,[],[f1222,f322,f15729,f498])).
fof(f19740,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17592,f15729])).
fof(f17592,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9252,f322,f15729,f498])).
fof(f19739,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17593,f15729])).
fof(f17593,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | $spl96),
  inference(unit_resulting_resolution,[],[f3498,f322,f15729,f498])).
fof(f19738,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f17594,f15729])).
fof(f17594,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9696,f322,f15729,f498])).
fof(f19737,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17595,f15729])).
fof(f17595,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f10981,f322,f15729,f498])).
fof(f19736,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f17600,f322])).
fof(f17600,plain,(
  ~p(s(bool,'T_0')) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1686,f322,f15729,f498])).
fof(f19735,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f17606,f15729])).
fof(f17606,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1663,f15729,f15729,f498])).
fof(f19734,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17607,f15729])).
fof(f17607,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f711,f15729,f15729,f498])).
fof(f19733,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f17608,f15729])).
fof(f17608,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1666,f15729,f15729,f498])).
fof(f19732,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17609,f322])).
fof(f17609,plain,(
  ~p(s(bool,'T_0')) | $spl96),
  inference(unit_resulting_resolution,[],[f1080,f15729,f15729,f498])).
fof(f19731,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17610,f12857])).
fof(f17610,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) | $spl96),
  inference(unit_resulting_resolution,[],[f1222,f15729,f15729,f498])).
fof(f19730,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17611,f15729])).
fof(f17611,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9252,f15729,f15729,f498])).
fof(f19729,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17612,f15729])).
fof(f17612,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | $spl96),
  inference(unit_resulting_resolution,[],[f3498,f15729,f15729,f498])).
fof(f19728,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f17613,f15729])).
fof(f17613,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9696,f15729,f15729,f498])).
fof(f19727,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17614,f15729])).
fof(f17614,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f10981,f15729,f15729,f498])).
fof(f19726,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17615,f322])).
fof(f17615,plain,(
  ~p(s(bool,'T_0')) | $spl96),
  inference(unit_resulting_resolution,[],[f3493,f15729,f15729,f498])).
fof(f19725,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f17616,f322])).
fof(f17616,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9692,f15729,f15729,f498])).
fof(f19724,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17617,f322])).
fof(f17617,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f10977,f15729,f15729,f498])).
fof(f19723,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f17618,f15729])).
fof(f17618,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9697,f15729,f15729,f498])).
fof(f19722,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17619,f322])).
fof(f17619,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9248,f15729,f15729,f498])).
fof(f19721,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17620,f15729])).
fof(f17620,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9253,f15729,f15729,f498])).
fof(f19720,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17622,f15729])).
fof(f17622,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9198,f15729,f15729,f498])).
fof(f19719,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17624,f15729])).
fof(f17624,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f3449,f15729,f15729,f498])).
fof(f19718,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f17626,f15729])).
fof(f17626,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9642,f15729,f15729,f498])).
fof(f19717,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17628,f15729])).
fof(f17628,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f10911,f15729,f15729,f498])).
fof(f19716,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f17629,f322])).
fof(f17629,plain,(
  ~p(s(bool,'T_0')) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1686,f15729,f15729,f498])).
fof(f19715,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f17635,f15729])).
fof(f17635,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1676,f573,f15729,f501])).
fof(f1676,plain,(
  ( ! [X0] : (sP8('F_0',X0,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) ) | $spl11),
  inference(unit_resulting_resolution,[],[f573,f1559,f496])).
fof(f496,plain,(
  ( ! [X2,X0,X1] : (sP8(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f299])).
fof(f19714,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17636,f15729])).
fof(f17636,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f725,f573,f15729,f501])).
fof(f725,plain,(
  ( ! [X0] : (sP8('F_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f573,f573,f497])).
fof(f497,plain,(
  ( ! [X2,X0,X1] : (sP8(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f299])).
fof(f19713,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f17637,f15729])).
fof(f17637,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1681,f573,f15729,f501])).
fof(f1681,plain,(
  ( ! [X0] : (sP8('F_0','|dollar|exists_1'(s(fun(sK14,bool),sK16)),X0)) ) | $spl11),
  inference(unit_resulting_resolution,[],[f573,f1559,f497])).
fof(f19712,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17688,f15729])).
fof(f17688,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f716,f573,f15729,f501])).
fof(f716,plain,(
  ( ! [X0] : (sP8('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f573,f573,f496])).
fof(f19711,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f17689,f15729])).
fof(f17689,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1676,f573,f15729,f501])).
fof(f19710,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f17691,f15729])).
fof(f17691,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1681,f573,f15729,f501])).
fof(f19709,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17775,f18867])).
fof(f18867,plain,(
  ( ! [X0,X1] : (sP10('F_0',X0,X1)) ) | $spl96),
  inference(unit_resulting_resolution,[],[f597,f15729,f557])).
fof(f597,plain,(
  ( ! [X0,X1] : (sP9('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f573,f512])).
fof(f512,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f305])).
fof(f305,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f304])).
fof(f304,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP9(X2,X0,X1)))),
  inference(flattening,[],[f303])).
fof(f303,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP9(X2,X0,X1)))),
  inference(nnf_transformation,[],[f197])).
fof(f197,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f17775,plain,(
  ( ! [X0,X1] : (~sP10('F_0',X0,X1)) ) | $spl96),
  inference(unit_resulting_resolution,[],[f573,f15729,f502])).
fof(f502,plain,(
  ( ! [X2,X0,X1] : (~sP10(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f302])).
fof(f19708,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17810,f18867])).
fof(f17810,plain,(
  ( ! [X0,X1] : (~sP10('F_0',X0,X1)) ) | $spl96),
  inference(unit_resulting_resolution,[],[f573,f15729,f503])).
fof(f503,plain,(
  ( ! [X2,X0,X1] : (~sP10(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f302])).
fof(f19707,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17876,f18001])).
fof(f18001,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2)) ) | $spl96),
  inference(unit_resulting_resolution,[],[f15729,f511])).
fof(f511,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f305])).
fof(f17876,plain,(
  ~sP9('F_0','T_0','F_0') | $spl96),
  inference(unit_resulting_resolution,[],[f730,f573,f15729,f506])).
fof(f506,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~sP9(X0,X1,X2) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f302])).
fof(f730,plain,(
  ( ! [X0] : (~sP10('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f322,f573,f503])).
fof(f19706,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17878,f18001])).
fof(f17878,plain,(
  ~sP9('F_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'F_0') | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9262,f573,f15729,f506])).
fof(f9262,plain,(
  ( ! [X0] : (~sP10('F_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f573,f566,f503])).
fof(f19705,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17880,f18001])).
fof(f17880,plain,(
  ~sP9('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0') | $spl96),
  inference(unit_resulting_resolution,[],[f3507,f573,f15729,f506])).
fof(f3507,plain,(
  ( ! [X0] : (~sP10('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X0)) )),
  inference(unit_resulting_resolution,[],[f573,f2082,f503])).
fof(f19704,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f17882,f18001])).
fof(f17882,plain,(
  ~sP9('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),'F_0') | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9706,f573,f15729,f506])).
fof(f9706,plain,(
  ( ! [X0] : (~sP10('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f573,f571,f503])).
fof(f19703,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17884,f18001])).
fof(f17884,plain,(
  ~sP9('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),'F_0') | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f10994,f573,f15729,f506])).
fof(f10994,plain,(
  ( ! [X0] : (~sP10('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f573,f9091,f503])).
fof(f19702,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f17891,f18001])).
fof(f17891,plain,(
  ~sP9('|dollar|exists_1'(s(fun(sK14,bool),sK16)),'T_0','F_0') | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1688,f573,f15729,f506])).
fof(f1688,plain,(
  ( ! [X0] : (~sP10('|dollar|exists_1'(s(fun(sK14,bool),sK16)),'T_0',X0)) ) | $spl11),
  inference(unit_resulting_resolution,[],[f322,f1559,f503])).
fof(f19701,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17955,f18001])).
fof(f17955,plain,(
  ~sP9('F_0','F_0','T_0') | $spl96),
  inference(unit_resulting_resolution,[],[f729,f573,f15729,f507])).
fof(f507,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~sP9(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f302])).
fof(f729,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f322,f573,f502])).
fof(f19700,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17956,f18001])).
fof(f17956,plain,(
  ~sP9('F_0','F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15))) | $spl96),
  inference(unit_resulting_resolution,[],[f3503,f573,f15729,f507])).
fof(f3503,plain,(
  ( ! [X0] : (~sP10('F_0',X0,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) )),
  inference(unit_resulting_resolution,[],[f573,f2082,f502])).
fof(f19699,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f17957,f18001])).
fof(f17957,plain,(
  ~sP9('F_0','F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17))) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9700,f573,f15729,f507])).
fof(f9700,plain,(
  ( ! [X0] : (~sP10('F_0',X0,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f573,f571,f502])).
fof(f19698,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17958,f18001])).
fof(f17958,plain,(
  ~sP9('F_0','F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14)))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f10985,f573,f15729,f507])).
fof(f10985,plain,(
  ( ! [X0] : (~sP10('F_0',X0,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f573,f9091,f502])).
fof(f19697,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f17959,f18001])).
fof(f17959,plain,(
  ~sP9('F_0','F_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9256,f573,f15729,f507])).
fof(f9256,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f573,f566,f502])).
fof(f19696,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f17965,f18001])).
fof(f17965,plain,(
  ~sP9('|dollar|exists_1'(s(fun(sK14,bool),sK16)),'F_0','T_0') | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1687,f573,f15729,f507])).
fof(f1687,plain,(
  ( ! [X0] : (~sP10('|dollar|exists_1'(s(fun(sK14,bool),sK16)),X0,'T_0')) ) | $spl11),
  inference(unit_resulting_resolution,[],[f322,f1559,f502])).
fof(f19695,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f17997,f18001])).
fof(f17997,plain,(
  ( ! [X0] : (~sP9(X0,'F_0','F_0')) ) | $spl96),
  inference(unit_resulting_resolution,[],[f573,f573,f15729,f509])).
fof(f509,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f305])).
fof(f19694,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f18003,f322])).
fof(f18003,plain,(
  ~p(s(bool,'T_0')) | $spl96),
  inference(unit_resulting_resolution,[],[f1164,f15729,f514])).
fof(f514,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | sP10(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f307])).
fof(f307,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP10(X2,X0,X1)) & (~sP10(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f306])).
fof(f306,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP10(X2,X0,X1)) & (~sP10(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f199])).
fof(f199,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP10(X2,X0,X1))),
  inference(definition_folding,[],[f141,f198,f197])).
fof(f141,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f140])).
fof(f140,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f48])).
fof(f48,axiom,(
  ! [X24,X25,X26] : ((p(s(bool,X26)) <=> (p(s(bool,X25)) | p(s(bool,X24)))) <=> ((p(s(bool,X26)) | ~p(s(bool,X25))) & (p(s(bool,X26)) | ~p(s(bool,X24))) & (p(s(bool,X25)) | p(s(bool,X24)) | ~p(s(bool,X26)))))),
  file('thm.quantHeuristics.GUESSES_UEXISTS_THM3','thm.sat.sat_25_prove_ltb')).
fof(f1164,plain,(
  ~sP10('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1152,f504])).
fof(f504,plain,(
  ( ! [X2,X0,X1] : (~sP10(X0,X1,X2) | sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f302])).
fof(f1152,plain,(
  ~sP9('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f573,f573,f322,f509])).
fof(f19693,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18004,f322])).
fof(f18004,plain,(
  ~p(s(bool,'T_0')) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1736,f15729,f514])).
fof(f1736,plain,(
  ~sP10('T_0','|dollar|exists_1'(s(fun(sK14,bool),sK16)),'F_0') | $spl11),
  inference(unit_resulting_resolution,[],[f573,f322,f1559,f516])).
fof(f516,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)) | ~sP10(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f307])).
fof(f19692,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18005,f322])).
fof(f18005,plain,(
  ~p(s(bool,'T_0')) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1740,f15729,f514])).
fof(f1740,plain,(
  ~sP10('T_0','F_0','|dollar|exists_1'(s(fun(sK14,bool),sK16))) | $spl11),
  inference(unit_resulting_resolution,[],[f322,f573,f1559,f516])).
fof(f19691,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18006,f322])).
fof(f18006,plain,(
  ~p(s(bool,'T_0')) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1741,f15729,f514])).
fof(f1741,plain,(
  ~sP10('T_0','|dollar|exists_1'(s(fun(sK14,bool),sK16)),'|dollar|exists_1'(s(fun(sK14,bool),sK16))) | $spl11),
  inference(unit_resulting_resolution,[],[f322,f1559,f1559,f516])).
fof(f19690,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f18017,f12857])).
fof(f18017,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) | $spl96),
  inference(unit_resulting_resolution,[],[f1282,f15729,f514])).
fof(f1282,plain,(
  ~sP10('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1272,f504])).
fof(f1272,plain,(
  ~sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f573,f573,f1160])).
fof(f1160,plain,(
  ( ! [X6,X7] : (~sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X6,X7) | p(s(bool,X7)) | p(s(bool,X6))) )),
  inference(subsumption_resolution,[],[f1155,f322])).
fof(f1155,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | p(s(bool,X6)) | p(s(bool,X7)) | ~sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X6,X7)) )),
  inference(superposition,[],[f509,f621])).
fof(f19689,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f18018,f15729])).
fof(f18018,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9297,f15729,f514])).
fof(f9297,plain,(
  ~sP10('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'F_0','F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f573,f573,f566,f516])).
fof(f19688,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f18019,f15729])).
fof(f18019,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | $spl96),
  inference(unit_resulting_resolution,[],[f3534,f15729,f514])).
fof(f3534,plain,(
  ~sP10(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f573,f573,f2082,f516])).
fof(f19687,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f18020,f15729])).
fof(f18020,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9741,f15729,f514])).
fof(f9741,plain,(
  ~sP10(happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),'F_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f573,f573,f571,f516])).
fof(f19686,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f18021,f15729])).
fof(f18021,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f11044,f15729,f514])).
fof(f11044,plain,(
  ~sP10(happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),'F_0','F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f573,f573,f9091,f516])).
fof(f19685,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18022,f15729])).
fof(f18022,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1687,f15729,f514])).
fof(f19684,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18023,f15729])).
fof(f18023,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1688,f15729,f514])).
fof(f19683,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18030,f15729])).
fof(f18030,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1740,f15729,f514])).
fof(f19682,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18031,f15729])).
fof(f18031,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1741,f15729,f514])).
fof(f19681,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f18032,f322])).
fof(f18032,plain,(
  ~p(s(bool,'T_0')) | $spl96),
  inference(unit_resulting_resolution,[],[f729,f15729,f514])).
fof(f19680,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f18033,f15729])).
fof(f18033,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | $spl96),
  inference(unit_resulting_resolution,[],[f3503,f15729,f514])).
fof(f19679,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f18034,f15729])).
fof(f18034,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9700,f15729,f514])).
fof(f19678,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f18035,f15729])).
fof(f18035,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f10985,f15729,f514])).
fof(f19677,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f18036,f15729])).
fof(f18036,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9256,f15729,f514])).
fof(f19676,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f18037,f15729])).
fof(f18037,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f730,f15729,f514])).
fof(f19675,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f18038,f15729])).
fof(f18038,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9262,f15729,f514])).
fof(f19674,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f18039,f15729])).
fof(f18039,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f3507,f15729,f514])).
fof(f19673,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f18040,f15729])).
fof(f18040,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9706,f15729,f514])).
fof(f19672,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f18041,f15729])).
fof(f18041,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f10994,f15729,f514])).
fof(f19671,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18047,f322])).
fof(f18047,plain,(
  ~p(s(bool,'T_0')) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1687,f15729,f514])).
fof(f19670,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18048,f15729])).
fof(f18048,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1688,f15729,f514])).
fof(f19669,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f18053,f322])).
fof(f18053,plain,(
  ~p(s(bool,'T_0')) | $spl96),
  inference(unit_resulting_resolution,[],[f1164,f15729,f515])).
fof(f515,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP10(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f307])).
fof(f19668,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18054,f322])).
fof(f18054,plain,(
  ~p(s(bool,'T_0')) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1736,f15729,f515])).
fof(f19667,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18055,f322])).
fof(f18055,plain,(
  ~p(s(bool,'T_0')) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1740,f15729,f515])).
fof(f19666,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18056,f322])).
fof(f18056,plain,(
  ~p(s(bool,'T_0')) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1741,f15729,f515])).
fof(f19665,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f18067,f12857])).
fof(f18067,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) | $spl96),
  inference(unit_resulting_resolution,[],[f1282,f15729,f515])).
fof(f19664,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f18068,f15729])).
fof(f18068,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9297,f15729,f515])).
fof(f19663,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f18069,f15729])).
fof(f18069,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | $spl96),
  inference(unit_resulting_resolution,[],[f3534,f15729,f515])).
fof(f19662,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f18070,f15729])).
fof(f18070,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9741,f15729,f515])).
fof(f19661,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f18071,f15729])).
fof(f18071,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f11044,f15729,f515])).
fof(f19660,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18072,f15729])).
fof(f18072,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1687,f15729,f515])).
fof(f19659,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18073,f15729])).
fof(f18073,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1688,f15729,f515])).
fof(f19658,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18079,f15729])).
fof(f18079,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1736,f15729,f515])).
fof(f19657,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18081,f15729])).
fof(f18081,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1741,f15729,f515])).
fof(f19656,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f18082,f15729])).
fof(f18082,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f729,f15729,f515])).
fof(f19655,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f18083,f15729])).
fof(f18083,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f3503,f15729,f515])).
fof(f19654,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f18084,f15729])).
fof(f18084,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9700,f15729,f515])).
fof(f19653,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f18085,f15729])).
fof(f18085,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f10985,f15729,f515])).
fof(f19652,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f18086,f15729])).
fof(f18086,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9256,f15729,f515])).
fof(f19651,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f18087,f322])).
fof(f18087,plain,(
  ~p(s(bool,'T_0')) | $spl96),
  inference(unit_resulting_resolution,[],[f730,f15729,f515])).
fof(f19650,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f18088,f15729])).
fof(f18088,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9262,f15729,f515])).
fof(f19649,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f18089,f15729])).
fof(f18089,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | $spl96),
  inference(unit_resulting_resolution,[],[f3507,f15729,f515])).
fof(f19648,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f18090,f15729])).
fof(f18090,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9706,f15729,f515])).
fof(f19647,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f18091,f15729])).
fof(f18091,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f10994,f15729,f515])).
fof(f19646,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18097,f15729])).
fof(f18097,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1687,f15729,f515])).
fof(f19645,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18098,f322])).
fof(f18098,plain,(
  ~p(s(bool,'T_0')) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1688,f15729,f515])).
fof(f19644,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f18169,f18818])).
fof(f18818,plain,(
  ( ! [X0,X1] : (sP12('F_0',X0,X1)) ) | $spl96),
  inference(unit_resulting_resolution,[],[f602,f15729,f556])).
fof(f602,plain,(
  ( ! [X0,X1] : (sP11('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f573,f529])).
fof(f529,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f313])).
fof(f18169,plain,(
  ( ! [X0,X1] : (~sP12('F_0',X0,X1)) ) | $spl96),
  inference(unit_resulting_resolution,[],[f573,f15729,f520])).
fof(f520,plain,(
  ( ! [X2,X0,X1] : (~sP12(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f310])).
fof(f19643,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18172,f15729])).
fof(f18172,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f602,f1746,f15729,f523])).
fof(f523,plain,(
  ( ! [X2,X0,X1] : (sP12(X0,X1,X2) | ~sP11(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f310])).
fof(f1746,plain,(
  ( ! [X0] : (~sP12('F_0',X0,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) ) | $spl11),
  inference(unit_resulting_resolution,[],[f573,f1559,f519])).
fof(f519,plain,(
  ( ! [X2,X0,X1] : (~sP12(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f310])).
fof(f19642,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f18173,f15729])).
fof(f18173,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f602,f757,f15729,f523])).
fof(f757,plain,(
  ( ! [X0] : (~sP12('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f322,f573,f520])).
fof(f19641,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f18174,f15729])).
fof(f18174,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f601,f757,f15729,f523])).
fof(f601,plain,(
  ( ! [X0,X1] : (sP11(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f322,f528])).
fof(f528,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f313])).
fof(f19640,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18176,f15729])).
fof(f18176,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1765,f757,f15729,f523])).
fof(f1765,plain,(
  ( ! [X0,X1] : (sP11(X0,X1,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) ) | $spl11),
  inference(unit_resulting_resolution,[],[f1559,f527])).
fof(f527,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f313])).
fof(f19639,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f18177,f15729])).
fof(f18177,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f602,f9298,f15729,f523])).
fof(f9298,plain,(
  ( ! [X0] : (~sP12('F_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f573,f566,f520])).
fof(f19638,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f18178,f15729])).
fof(f18178,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9343,f9298,f15729,f523])).
fof(f9343,plain,(
  ( ! [X0,X1] : (sP11(X0,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X1)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f566,f528])).
fof(f19637,plain,(
  $false | ($spl0 | $spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18180,f15729])).
fof(f18180,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1765,f9298,f15729,f523])).
fof(f19636,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f18181,f15729])).
fof(f18181,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f602,f3535,f15729,f523])).
fof(f3535,plain,(
  ( ! [X0] : (~sP12('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X0)) )),
  inference(unit_resulting_resolution,[],[f573,f2082,f520])).
fof(f19635,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f18182,f15729])).
fof(f18182,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f3575,f3535,f15729,f523])).
fof(f3575,plain,(
  ( ! [X0,X1] : (sP11(X0,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X1)) )),
  inference(unit_resulting_resolution,[],[f2082,f528])).
fof(f19634,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18184,f15729])).
fof(f18184,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1765,f3535,f15729,f523])).
fof(f19633,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f18185,f15729])).
fof(f18185,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f602,f9742,f15729,f523])).
fof(f9742,plain,(
  ( ! [X0] : (~sP12('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f573,f571,f520])).
fof(f19632,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f18186,f15729])).
fof(f18186,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9787,f9742,f15729,f523])).
fof(f9787,plain,(
  ( ! [X0,X1] : (sP11(X0,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f571,f528])).
fof(f19631,plain,(
  $false | ($spl4 | $spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18188,f15729])).
fof(f18188,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl4 | $spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1765,f9742,f15729,f523])).
fof(f19630,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f18189,f15729])).
fof(f18189,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f602,f11045,f15729,f523])).
fof(f11045,plain,(
  ( ! [X0] : (~sP12('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f573,f9091,f520])).
fof(f19629,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f18190,f15729])).
fof(f18190,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f11110,f11045,f15729,f523])).
fof(f11110,plain,(
  ( ! [X0,X1] : (sP11(X0,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),X1)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f9091,f528])).
fof(f19628,plain,(
  $false | ($spl0 | $spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18192,f15729])).
fof(f18192,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1765,f11045,f15729,f523])).
fof(f19627,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18194,f15729])).
fof(f18194,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1766,f1747,f15729,f523])).
fof(f1747,plain,(
  ( ! [X0] : (~sP12('|dollar|exists_1'(s(fun(sK14,bool),sK16)),X0,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) ) | $spl11),
  inference(unit_resulting_resolution,[],[f1559,f1559,f519])).
fof(f1766,plain,(
  ( ! [X0,X1] : (sP11('|dollar|exists_1'(s(fun(sK14,bool),sK16)),X0,X1)) ) | $spl11),
  inference(unit_resulting_resolution,[],[f1559,f529])).
fof(f19626,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18195,f15729])).
fof(f18195,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1766,f1751,f15729,f523])).
fof(f1751,plain,(
  ( ! [X0] : (~sP12('|dollar|exists_1'(s(fun(sK14,bool),sK16)),'T_0',X0)) ) | $spl11),
  inference(unit_resulting_resolution,[],[f322,f1559,f520])).
fof(f19625,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18196,f15729])).
fof(f18196,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f601,f1751,f15729,f523])).
fof(f19624,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18198,f15729])).
fof(f18198,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1765,f1751,f15729,f523])).
fof(f19623,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f18209,f18533])).
fof(f18533,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2)) ) | $spl96),
  inference(unit_resulting_resolution,[],[f15729,f528])).
fof(f18209,plain,(
  ~sP11('F_0','T_0','T_0') | $spl96),
  inference(unit_resulting_resolution,[],[f757,f322,f15729,f523])).
fof(f19622,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f18211,f18533])).
fof(f18211,plain,(
  ~sP11('F_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'T_0') | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9298,f322,f15729,f523])).
fof(f19621,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f18213,f18533])).
fof(f18213,plain,(
  ~sP11('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'T_0') | $spl96),
  inference(unit_resulting_resolution,[],[f3535,f322,f15729,f523])).
fof(f19620,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f18215,f18533])).
fof(f18215,plain,(
  ~sP11('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),'T_0') | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9742,f322,f15729,f523])).
fof(f19619,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f18217,f18533])).
fof(f18217,plain,(
  ~sP11('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),'T_0') | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f11045,f322,f15729,f523])).
fof(f19618,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18223,f18533])).
fof(f18223,plain,(
  ~sP11('|dollar|exists_1'(s(fun(sK14,bool),sK16)),'T_0','T_0') | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1751,f322,f15729,f523])).
fof(f19617,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f18241,f18533])).
fof(f18241,plain,(
  ( ! [X0] : (~sP11('F_0',X0,'F_0')) ) | $spl96),
  inference(unit_resulting_resolution,[],[f755,f15729,f15729,f523])).
fof(f755,plain,(
  ( ! [X0] : (~sP12('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f573,f573,f519])).
fof(f19616,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18243,f18533])).
fof(f18243,plain,(
  ( ! [X0] : (~sP11('F_0',X0,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1746,f15729,f15729,f523])).
fof(f19615,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f18245,f18533])).
fof(f18245,plain,(
  ( ! [X0] : (~sP11('F_0','T_0',X0)) ) | $spl96),
  inference(unit_resulting_resolution,[],[f757,f15729,f15729,f523])).
fof(f19614,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f18247,f18533])).
fof(f18247,plain,(
  ( ! [X0] : (~sP11('F_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X0)) ) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9298,f15729,f15729,f523])).
fof(f19613,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f18249,f18533])).
fof(f18249,plain,(
  ( ! [X0] : (~sP11('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X0)) ) | $spl96),
  inference(unit_resulting_resolution,[],[f3535,f15729,f15729,f523])).
fof(f19612,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f18251,f18533])).
fof(f18251,plain,(
  ( ! [X0] : (~sP11('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),X0)) ) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9742,f15729,f15729,f523])).
fof(f19611,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f18253,f18533])).
fof(f18253,plain,(
  ( ! [X0] : (~sP11('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),X0)) ) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f11045,f15729,f15729,f523])).
fof(f19610,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18261,f18533])).
fof(f18261,plain,(
  ( ! [X0] : (~sP11('|dollar|exists_1'(s(fun(sK14,bool),sK16)),X0,'F_0')) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1744,f15729,f15729,f523])).
fof(f1744,plain,(
  ( ! [X0] : (~sP12('|dollar|exists_1'(s(fun(sK14,bool),sK16)),X0,'F_0')) ) | $spl11),
  inference(unit_resulting_resolution,[],[f573,f1559,f519])).
fof(f19609,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18263,f18533])).
fof(f18263,plain,(
  ( ! [X0] : (~sP11('|dollar|exists_1'(s(fun(sK14,bool),sK16)),X0,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1747,f15729,f15729,f523])).
fof(f19608,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18265,f18533])).
fof(f18265,plain,(
  ( ! [X0] : (~sP11('|dollar|exists_1'(s(fun(sK14,bool),sK16)),'T_0',X0)) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1751,f15729,f15729,f523])).
fof(f19607,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18273,f15729])).
fof(f18273,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1744,f599,f15729,f523])).
fof(f599,plain,(
  ( ! [X0,X1] : (sP11(X0,X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f573,f527])).
fof(f19606,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18274,f15729])).
fof(f18274,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1751,f599,f15729,f523])).
fof(f19605,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18281,f15729])).
fof(f18281,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1747,f1765,f15729,f523])).
fof(f19604,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18282,f15729])).
fof(f18282,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1751,f1765,f15729,f523])).
fof(f19603,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18286,f15729])).
fof(f18286,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1751,f601,f15729,f523])).
fof(f19602,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18287,f15729])).
fof(f18287,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1744,f601,f15729,f523])).
fof(f19601,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18288,f15729])).
fof(f18288,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1747,f601,f15729,f523])).
fof(f19600,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18291,f15729])).
fof(f18291,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1744,f615,f15729,f523])).
fof(f615,plain,(
  ( ! [X0,X1] : (sP11(X0,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X1)) )),
  inference(unit_resulting_resolution,[],[f318,f528])).
fof(f19599,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18292,f15729])).
fof(f18292,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1747,f615,f15729,f523])).
fof(f19598,plain,(
  $false | ($spl0 | $spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18296,f15729])).
fof(f18296,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1744,f9343,f15729,f523])).
fof(f19597,plain,(
  $false | ($spl0 | $spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18297,f15729])).
fof(f18297,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1747,f9343,f15729,f523])).
fof(f19596,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18301,f15729])).
fof(f18301,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1744,f3575,f15729,f523])).
fof(f19595,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18302,f15729])).
fof(f18302,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1747,f3575,f15729,f523])).
fof(f19594,plain,(
  $false | ($spl4 | $spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18306,f15729])).
fof(f18306,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl4 | $spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1744,f9787,f15729,f523])).
fof(f19593,plain,(
  $false | ($spl4 | $spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18307,f15729])).
fof(f18307,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl4 | $spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1747,f9787,f15729,f523])).
fof(f19592,plain,(
  $false | ($spl0 | $spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18311,f15729])).
fof(f18311,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1744,f11110,f15729,f523])).
fof(f19591,plain,(
  $false | ($spl0 | $spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18312,f15729])).
fof(f18312,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1747,f11110,f15729,f523])).
fof(f19590,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18314,f15729])).
fof(f18314,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1744,f1766,f15729,f523])).
fof(f19589,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f18335,f18533])).
fof(f18335,plain,(
  ( ! [X0] : (~sP11('F_0',X0,'F_0')) ) | $spl96),
  inference(unit_resulting_resolution,[],[f755,f15729,f15729,f523])).
fof(f19588,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18336,f18533])).
fof(f18336,plain,(
  ( ! [X0] : (~sP11('F_0',X0,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1746,f15729,f15729,f523])).
fof(f19587,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f18337,f18533])).
fof(f18337,plain,(
  ( ! [X0] : (~sP11('F_0','T_0',X0)) ) | $spl96),
  inference(unit_resulting_resolution,[],[f757,f15729,f15729,f523])).
fof(f19586,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f18338,f18533])).
fof(f18338,plain,(
  ( ! [X0] : (~sP11('F_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X0)) ) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9298,f15729,f15729,f523])).
fof(f19585,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f18339,f18533])).
fof(f18339,plain,(
  ( ! [X0] : (~sP11('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X0)) ) | $spl96),
  inference(unit_resulting_resolution,[],[f3535,f15729,f15729,f523])).
fof(f19584,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f18340,f18533])).
fof(f18340,plain,(
  ( ! [X0] : (~sP11('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),X0)) ) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9742,f15729,f15729,f523])).
fof(f19583,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f18341,f18533])).
fof(f18341,plain,(
  ( ! [X0] : (~sP11('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),X0)) ) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f11045,f15729,f15729,f523])).
fof(f19582,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18348,f18533])).
fof(f18348,plain,(
  ( ! [X0] : (~sP11('|dollar|exists_1'(s(fun(sK14,bool),sK16)),X0,'F_0')) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1744,f15729,f15729,f523])).
fof(f19581,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18349,f18533])).
fof(f18349,plain,(
  ( ! [X0] : (~sP11('|dollar|exists_1'(s(fun(sK14,bool),sK16)),X0,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1747,f15729,f15729,f523])).
fof(f19580,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18350,f18533])).
fof(f18350,plain,(
  ( ! [X0] : (~sP11('|dollar|exists_1'(s(fun(sK14,bool),sK16)),'T_0',X0)) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1751,f15729,f15729,f523])).
fof(f19579,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f18408,f18533])).
fof(f18408,plain,(
  ~sP11('F_0','F_0','F_0') | $spl96),
  inference(unit_resulting_resolution,[],[f755,f573,f15729,f524])).
fof(f524,plain,(
  ( ! [X2,X0,X1] : (sP12(X0,X1,X2) | ~sP11(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f310])).
fof(f19578,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18409,f18533])).
fof(f18409,plain,(
  ~sP11('F_0','F_0','|dollar|exists_1'(s(fun(sK14,bool),sK16))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1746,f573,f15729,f524])).
fof(f19577,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18416,f18533])).
fof(f18416,plain,(
  ~sP11('|dollar|exists_1'(s(fun(sK14,bool),sK16)),'F_0','F_0') | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1744,f573,f15729,f524])).
fof(f19576,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18417,f18533])).
fof(f18417,plain,(
  ~sP11('|dollar|exists_1'(s(fun(sK14,bool),sK16)),'F_0','|dollar|exists_1'(s(fun(sK14,bool),sK16))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1747,f573,f15729,f524])).
fof(f19575,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f18463,f18533])).
fof(f18463,plain,(
  ~sP11('F_0','F_0','F_0') | $spl96),
  inference(unit_resulting_resolution,[],[f755,f573,f15729,f525])).
fof(f525,plain,(
  ( ! [X2,X0,X1] : (sP12(X0,X1,X2) | ~sP11(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f310])).
fof(f19574,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18464,f18533])).
fof(f18464,plain,(
  ~sP11('F_0','F_0','|dollar|exists_1'(s(fun(sK14,bool),sK16))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1746,f573,f15729,f525])).
fof(f19573,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18471,f18533])).
fof(f18471,plain,(
  ~sP11('|dollar|exists_1'(s(fun(sK14,bool),sK16)),'F_0','F_0') | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1744,f573,f15729,f525])).
fof(f19572,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18472,f18533])).
fof(f18472,plain,(
  ~sP11('|dollar|exists_1'(s(fun(sK14,bool),sK16)),'F_0','|dollar|exists_1'(s(fun(sK14,bool),sK16))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1747,f573,f15729,f525])).
fof(f19571,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18500,f15729])).
fof(f18500,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1765,f573,f15729,f526])).
fof(f19570,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f18501,f15729])).
fof(f18501,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f602,f573,f15729,f526])).
fof(f19569,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18502,f15729])).
fof(f18502,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1766,f573,f15729,f526])).
fof(f19568,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f18503,f18533])).
fof(f18503,plain,(
  ( ! [X0] : (~sP11(X0,'F_0','T_0')) ) | $spl96),
  inference(unit_resulting_resolution,[],[f322,f573,f15729,f526])).
fof(f19567,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f18505,f18533])).
fof(f18505,plain,(
  ( ! [X0,X1] : (~sP11(X0,'F_0',X1)) ) | $spl96),
  inference(unit_resulting_resolution,[],[f15729,f573,f15729,f526])).
fof(f19566,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f18525,f15729])).
fof(f18525,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f599,f573,f15729,f526])).
fof(f19565,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18526,f15729])).
fof(f18526,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1765,f573,f15729,f526])).
fof(f19564,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18528,f15729])).
fof(f18528,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1766,f573,f15729,f526])).
fof(f19563,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f18529,f18533])).
fof(f18529,plain,(
  ( ! [X0] : (~sP11('T_0','F_0',X0)) ) | $spl96),
  inference(unit_resulting_resolution,[],[f322,f573,f15729,f526])).
fof(f19562,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f18531,f18533])).
fof(f18531,plain,(
  ( ! [X0,X1] : (~sP11(X0,'F_0',X1)) ) | $spl96),
  inference(unit_resulting_resolution,[],[f15729,f573,f15729,f526])).
fof(f19561,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f18583,f322])).
fof(f18583,plain,(
  ~p(s(bool,'T_0')) | $spl96),
  inference(unit_resulting_resolution,[],[f1261,f15729,f532])).
fof(f532,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP12(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f315])).
fof(f1261,plain,(
  ~sP12('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f1249,f521])).
fof(f521,plain,(
  ( ! [X2,X0,X1] : (~sP12(X0,X1,X2) | sP11(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f310])).
fof(f1249,plain,(
  ~sP11('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f322,f573,f322,f526])).
fof(f19560,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18584,f322])).
fof(f18584,plain,(
  ~p(s(bool,'T_0')) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1774,f15729,f532])).
fof(f1774,plain,(
  ~sP12('T_0','|dollar|exists_1'(s(fun(sK14,bool),sK16)),'T_0') | $spl11),
  inference(unit_resulting_resolution,[],[f322,f322,f1559,f533])).
fof(f19559,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f18585,f322])).
fof(f18585,plain,(
  ~p(s(bool,'T_0')) | $spl96),
  inference(unit_resulting_resolution,[],[f3586,f15729,f532])).
fof(f3586,plain,(
  ~sP12('T_0','F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))),
  inference(unit_resulting_resolution,[],[f322,f573,f2082,f533])).
fof(f19558,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f18586,f322])).
fof(f18586,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9800,f15729,f532])).
fof(f9800,plain,(
  ~sP12('T_0','F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17))) | $spl4),
  inference(unit_resulting_resolution,[],[f322,f573,f571,f533])).
fof(f19557,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f18587,f322])).
fof(f18587,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f11125,f15729,f532])).
fof(f11125,plain,(
  ~sP12('T_0','F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14)))) | $spl0),
  inference(unit_resulting_resolution,[],[f322,f573,f9091,f533])).
fof(f19556,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f18588,f322])).
fof(f18588,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9356,f15729,f532])).
fof(f9356,plain,(
  ~sP12('T_0','F_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | $spl0),
  inference(unit_resulting_resolution,[],[f322,f573,f566,f533])).
fof(f19555,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f18596,f15729])).
fof(f18596,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9361,f15729,f532])).
fof(f9361,plain,(
  ~sP12('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'F_0','T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f322,f573,f566,f533])).
fof(f19554,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f18597,f15729])).
fof(f18597,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9362,f15729,f532])).
fof(f9362,plain,(
  ~sP12('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'F_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | $spl0),
  inference(unit_resulting_resolution,[],[f566,f573,f566,f533])).
fof(f19553,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f18598,f15729])).
fof(f18598,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | $spl96),
  inference(unit_resulting_resolution,[],[f3590,f15729,f532])).
fof(f3590,plain,(
  ~sP12(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f322,f573,f2082,f533])).
fof(f19552,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f18599,f15729])).
fof(f18599,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9805,f15729,f532])).
fof(f9805,plain,(
  ~sP12(happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),'F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f322,f573,f571,f533])).
fof(f19551,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f18600,f15729])).
fof(f18600,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f11132,f15729,f532])).
fof(f11132,plain,(
  ~sP12(happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),'F_0','T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f322,f573,f9091,f533])).
fof(f19550,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f18601,f15729])).
fof(f18601,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9806,f15729,f532])).
fof(f9806,plain,(
  ~sP12(happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),'F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17))) | $spl4),
  inference(unit_resulting_resolution,[],[f571,f573,f571,f533])).
fof(f19549,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18602,f15729])).
fof(f18602,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1744,f15729,f532])).
fof(f19548,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18603,f15729])).
fof(f18603,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1747,f15729,f532])).
fof(f19547,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18604,f15729])).
fof(f18604,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1751,f15729,f532])).
fof(f19546,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18610,f15729])).
fof(f18610,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1774,f15729,f532])).
fof(f19545,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f18615,f15729])).
fof(f18615,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f755,f15729,f532])).
fof(f19544,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18616,f15729])).
fof(f18616,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1746,f15729,f532])).
fof(f19543,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f18617,f322])).
fof(f18617,plain,(
  ~p(s(bool,'T_0')) | $spl96),
  inference(unit_resulting_resolution,[],[f757,f15729,f532])).
fof(f19542,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f18618,f15729])).
fof(f18618,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9298,f15729,f532])).
fof(f19541,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f18619,f15729])).
fof(f18619,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | $spl96),
  inference(unit_resulting_resolution,[],[f3535,f15729,f532])).
fof(f19540,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f18620,f15729])).
fof(f18620,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9742,f15729,f532])).
fof(f19539,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f18621,f15729])).
fof(f18621,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f11045,f15729,f532])).
fof(f19538,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18628,f15729])).
fof(f18628,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1744,f15729,f532])).
fof(f19537,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18629,f15729])).
fof(f18629,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1747,f15729,f532])).
fof(f19536,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18630,f322])).
fof(f18630,plain,(
  ~p(s(bool,'T_0')) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1751,f15729,f532])).
fof(f19535,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f18635,f322])).
fof(f18635,plain,(
  ~p(s(bool,'T_0')) | $spl96),
  inference(unit_resulting_resolution,[],[f765,f573,f15729,f533])).
fof(f765,plain,(
  ( ! [X0] : (sP12('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f573,f322,f531])).
fof(f531,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | sP12(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f315])).
fof(f19534,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18636,f322])).
fof(f18636,plain,(
  ~p(s(bool,'T_0')) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1772,f573,f15729,f533])).
fof(f1772,plain,(
  ( ! [X0] : (sP12('T_0',X0,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) ) | $spl11),
  inference(unit_resulting_resolution,[],[f322,f1559,f531])).
fof(f19533,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f18642,f12857])).
fof(f18642,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) | $spl96),
  inference(unit_resulting_resolution,[],[f772,f573,f15729,f533])).
fof(f772,plain,(
  ( ! [X0] : (sP12('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f573,f768])).
fof(f768,plain,(
  ( ! [X2,X3] : (sP12('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X3,X2) | p(s(bool,X2))) )),
  inference(subsumption_resolution,[],[f767,f322])).
fof(f767,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,X2)) | sP12('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X3,X2)) )),
  inference(superposition,[],[f531,f621])).
fof(f19532,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f18643,f15729])).
fof(f18643,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9366,f573,f15729,f533])).
fof(f9366,plain,(
  ( ! [X0] : (sP12('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X0,'F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f599,f566,f556])).
fof(f19531,plain,(
  $false | ($spl0 | $spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18644,f15729])).
fof(f18644,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f9367,f573,f15729,f533])).
fof(f9367,plain,(
  ( ! [X0] : (sP12('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X0,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) ) | ($spl0 | $spl11)),
  inference(unit_resulting_resolution,[],[f1765,f566,f556])).
fof(f19530,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f18645,f15729])).
fof(f18645,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | $spl96),
  inference(unit_resulting_resolution,[],[f3595,f573,f15729,f533])).
fof(f3595,plain,(
  ( ! [X0] : (sP12(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f599,f2082,f556])).
fof(f19529,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f18646,f15729])).
fof(f18646,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9810,f573,f15729,f533])).
fof(f9810,plain,(
  ( ! [X0] : (sP12(happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),X0,'F_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f599,f571,f556])).
fof(f19528,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f18647,f15729])).
fof(f18647,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f11137,f573,f15729,f533])).
fof(f11137,plain,(
  ( ! [X0] : (sP12(happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),X0,'F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f599,f9091,f556])).
fof(f19527,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18648,f15729])).
fof(f18648,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1770,f573,f15729,f533])).
fof(f1770,plain,(
  sP12('|dollar|exists_1'(s(fun(sK14,bool),sK16)),'F_0','T_0') | $spl11),
  inference(unit_resulting_resolution,[],[f573,f322,f1559,f530])).
fof(f530,plain,(
  ( ! [X2,X0,X1] : (sP12(X2,X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f315])).
fof(f19526,plain,(
  $false | ($spl0 | $spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18649,f15729])).
fof(f18649,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f9327,f573,f15729,f533])).
fof(f9327,plain,(
  sP12('|dollar|exists_1'(s(fun(sK14,bool),sK16)),'F_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | ($spl0 | $spl11)),
  inference(unit_resulting_resolution,[],[f573,f1766,f566,f525])).
fof(f19525,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18693,f15729])).
fof(f18693,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1772,f573,f15729,f533])).
fof(f19524,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f18694,f322])).
fof(f18694,plain,(
  ~p(s(bool,'T_0')) | $spl96),
  inference(unit_resulting_resolution,[],[f1237,f573,f15729,f533])).
fof(f1237,plain,(
  sP12('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f602,f322,f573,f525])).
fof(f19523,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f18695,f15729])).
fof(f18695,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | $spl96),
  inference(unit_resulting_resolution,[],[f3576,f573,f15729,f533])).
fof(f3576,plain,(
  sP12('F_0','F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))),
  inference(unit_resulting_resolution,[],[f573,f573,f2082,f530])).
fof(f19522,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f18696,f15729])).
fof(f18696,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9788,f573,f15729,f533])).
fof(f9788,plain,(
  sP12('F_0','F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17))) | $spl4),
  inference(unit_resulting_resolution,[],[f573,f573,f571,f530])).
fof(f19521,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f18697,f15729])).
fof(f18697,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f11111,f573,f15729,f533])).
fof(f11111,plain,(
  sP12('F_0','F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14)))) | $spl0),
  inference(unit_resulting_resolution,[],[f573,f573,f9091,f530])).
fof(f19520,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f18698,f15729])).
fof(f18698,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9344,f573,f15729,f533])).
fof(f9344,plain,(
  sP12('F_0','F_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | $spl0),
  inference(unit_resulting_resolution,[],[f573,f573,f566,f530])).
fof(f19519,plain,(
  $false | ($spl0 | $spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18701,f15729])).
fof(f18701,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f9367,f573,f15729,f533])).
fof(f19518,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18705,f322])).
fof(f18705,plain,(
  ~p(s(bool,'T_0')) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1770,f573,f15729,f533])).
fof(f19517,plain,(
  $false | ($spl0 | $spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18706,f15729])).
fof(f18706,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f9327,f573,f15729,f533])).
fof(f19516,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f18758,f19348])).
fof(f18758,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) | $spl96),
  inference(unit_resulting_resolution,[],[f573,f15729,f555])).
fof(f19515,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f18786,f599])).
fof(f18786,plain,(
  ( ! [X0] : (~sP11('F_0',X0,'F_0')) ) | $spl96),
  inference(unit_resulting_resolution,[],[f755,f15729,f556])).
fof(f19514,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18787,f1765])).
fof(f18787,plain,(
  ( ! [X0] : (~sP11('F_0',X0,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1746,f15729,f556])).
fof(f19513,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f18788,f601])).
fof(f18788,plain,(
  ( ! [X0] : (~sP11('F_0','T_0',X0)) ) | $spl96),
  inference(unit_resulting_resolution,[],[f757,f15729,f556])).
fof(f19512,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f18789,f9343])).
fof(f18789,plain,(
  ( ! [X0] : (~sP11('F_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X0)) ) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9298,f15729,f556])).
fof(f19511,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f18790,f3575])).
fof(f18790,plain,(
  ( ! [X0] : (~sP11('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X0)) ) | $spl96),
  inference(unit_resulting_resolution,[],[f3535,f15729,f556])).
fof(f19510,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f18791,f9787])).
fof(f18791,plain,(
  ( ! [X0] : (~sP11('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),X0)) ) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9742,f15729,f556])).
fof(f19509,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f18792,f11110])).
fof(f18792,plain,(
  ( ! [X0] : (~sP11('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),X0)) ) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f11045,f15729,f556])).
fof(f19508,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18799,f599])).
fof(f18799,plain,(
  ( ! [X0] : (~sP11('|dollar|exists_1'(s(fun(sK14,bool),sK16)),X0,'F_0')) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1744,f15729,f556])).
fof(f19507,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18800,f1765])).
fof(f18800,plain,(
  ( ! [X0] : (~sP11('|dollar|exists_1'(s(fun(sK14,bool),sK16)),X0,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1747,f15729,f556])).
fof(f19506,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18801,f601])).
fof(f18801,plain,(
  ( ! [X0] : (~sP11('|dollar|exists_1'(s(fun(sK14,bool),sK16)),'T_0',X0)) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1751,f15729,f556])).
fof(f19505,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f18826,f595])).
fof(f595,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f322,f510])).
fof(f510,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f305])).
fof(f18826,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'T_0')) ) | $spl96),
  inference(unit_resulting_resolution,[],[f729,f15729,f557])).
fof(f19504,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f18827,f3520])).
fof(f3520,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) )),
  inference(unit_resulting_resolution,[],[f2082,f510])).
fof(f18827,plain,(
  ( ! [X0] : (~sP9('F_0',X0,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f3503,f15729,f557])).
fof(f19503,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f18828,f9725])).
fof(f9725,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f571,f510])).
fof(f18828,plain,(
  ( ! [X0] : (~sP9('F_0',X0,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) ) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9700,f15729,f557])).
fof(f19502,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f18829,f11024])).
fof(f11024,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f9091,f510])).
fof(f18829,plain,(
  ( ! [X0] : (~sP9('F_0',X0,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) ) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f10985,f15729,f557])).
fof(f19501,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f18830,f9281])).
fof(f9281,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f566,f510])).
fof(f18830,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) ) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9256,f15729,f557])).
fof(f19500,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f18831,f596])).
fof(f596,plain,(
  ( ! [X0,X1] : (sP9(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f322,f511])).
fof(f18831,plain,(
  ( ! [X0] : (~sP9('F_0','T_0',X0)) ) | $spl96),
  inference(unit_resulting_resolution,[],[f730,f15729,f557])).
fof(f19499,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f18832,f9282])).
fof(f9282,plain,(
  ( ! [X0,X1] : (sP9(X0,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X1)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f566,f511])).
fof(f18832,plain,(
  ( ! [X0] : (~sP9('F_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X0)) ) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9262,f15729,f557])).
fof(f19498,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f18833,f3521])).
fof(f3521,plain,(
  ( ! [X0,X1] : (sP9(X0,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X1)) )),
  inference(unit_resulting_resolution,[],[f2082,f511])).
fof(f18833,plain,(
  ( ! [X0] : (~sP9('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X0)) ) | $spl96),
  inference(unit_resulting_resolution,[],[f3507,f15729,f557])).
fof(f19497,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f18834,f9726])).
fof(f9726,plain,(
  ( ! [X0,X1] : (sP9(X0,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f571,f511])).
fof(f18834,plain,(
  ( ! [X0] : (~sP9('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),X0)) ) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9706,f15729,f557])).
fof(f19496,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f18835,f11025])).
fof(f11025,plain,(
  ( ! [X0,X1] : (sP9(X0,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),X1)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f9091,f511])).
fof(f18835,plain,(
  ( ! [X0] : (~sP9('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),X0)) ) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f10994,f15729,f557])).
fof(f19495,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18841,f595])).
fof(f18841,plain,(
  ( ! [X0] : (~sP9('|dollar|exists_1'(s(fun(sK14,bool),sK16)),X0,'T_0')) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1687,f15729,f557])).
fof(f19494,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18842,f596])).
fof(f18842,plain,(
  ( ! [X0] : (~sP9('|dollar|exists_1'(s(fun(sK14,bool),sK16)),'T_0',X0)) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1688,f15729,f557])).
fof(f19493,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18951,f15729])).
fof(f18951,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1588,f15729,f1340])).
fof(f1340,plain,(
  ( ! [X57,X54,X56,X55,X53] : (~p(s(bool,X56)) | ~p(s(bool,X54)) | sP2(X56,X57,happ(s(fun(X53,bool),'_dst_x0_1'(s(bool,X54))),s(X53,X55)))) )),
  inference(superposition,[],[f451,f421])).
fof(f19492,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18955,f15729])).
fof(f18955,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1593,f15729,f1340])).
fof(f19491,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f18959,f15729])).
fof(f18959,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f685,f15729,f1340])).
fof(f19490,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18961,f15729])).
fof(f18961,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1588,f15729,f1340])).
fof(f19489,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18963,f15729])).
fof(f18963,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1592,f15729,f1340])).
fof(f19488,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f18965,f15729])).
fof(f18965,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1593,f15729,f1340])).
fof(f19487,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f18969,f322])).
fof(f18969,plain,(
  ~p(s(bool,'T_0')) | $spl96),
  inference(unit_resulting_resolution,[],[f573,f1311,f15729,f1341])).
fof(f19486,plain,(
  $false | ($spl0 | $spl96)),
  inference(subsumption_resolution,[],[f18974,f15729])).
fof(f18974,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f573,f9186,f15729,f1341])).
fof(f19485,plain,(
  $false | ($spl4 | $spl96)),
  inference(subsumption_resolution,[],[f18980,f15729])).
fof(f18980,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f573,f9630,f15729,f1341])).
fof(f19484,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f19024,f15729])).
fof(f19024,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f573,f1619,f15729,f1342])).
fof(f19483,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f19037,f15729])).
fof(f19037,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f573,f580,f15729,f1342])).
fof(f19482,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f19041,f15729])).
fof(f19041,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f573,f1619,f15729,f1342])).
fof(f19481,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f19054,f15729])).
fof(f19054,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f573,f1626,f15729,f1343])).
fof(f19480,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f19067,f15729])).
fof(f19067,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f573,f585,f15729,f1343])).
fof(f19479,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f19071,f15729])).
fof(f19071,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f573,f1626,f15729,f1343])).
fof(f19478,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f19084,f15729])).
fof(f19084,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f573,f1675,f15729,f1344])).
fof(f19477,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f19097,f15729])).
fof(f19097,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f573,f593,f15729,f1344])).
fof(f19476,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f19101,f15729])).
fof(f19101,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f573,f1675,f15729,f1344])).
fof(f19475,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f19108,f15729])).
fof(f19108,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f573,f1681,f15729,f1345])).
fof(f19474,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f19116,f15729])).
fof(f19116,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f573,f725,f15729,f1345])).
fof(f19473,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f19119,f15729])).
fof(f19119,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f573,f1681,f15729,f1345])).
fof(f19472,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f19156,f15729])).
fof(f19156,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f573,f1766,f15729,f1350])).
fof(f19471,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f19167,f15729])).
fof(f19167,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f573,f602,f15729,f1350])).
fof(f19470,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f19172,f15729])).
fof(f19172,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f573,f1766,f15729,f1350])).
fof(f19469,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f19188,f15729])).
fof(f19188,plain,(
  ( ! [X0,X1] : (~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f685,f15729,f1535])).
fof(f1535,plain,(
  ( ! [X39,X41,X42,X40] : (sP2(X41,X42,happ(s(fun(X39,bool),X40),s(X39,'const.min.@_1'(s(fun(X39,bool),X40))))) | ~p(s(bool,X41)) | ~p(s(bool,'|dollar|exists_1'(s(fun(X39,bool),X40))))) )),
  inference(superposition,[],[f451,f364])).
fof(f19468,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f19190,f15729])).
fof(f19190,plain,(
  ( ! [X0,X1] : (~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1588,f15729,f1535])).
fof(f19467,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f19192,f15729])).
fof(f19192,plain,(
  ( ! [X0,X1] : (~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1592,f15729,f1535])).
fof(f19466,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f19194,f15729])).
fof(f19194,plain,(
  ( ! [X0,X1] : (~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1593,f15729,f1535])).
fof(f19465,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f19307,f15729])).
fof(f19307,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'T_0'))))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1777,f15729,f15729,f368])).
fof(f19464,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f19308,f15729])).
fof(f19308,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'T_0'))))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f617,f15729,f15729,f368])).
fof(f19463,plain,(
  $false | ($spl3 | $spl96)),
  inference(subsumption_resolution,[],[f19309,f15729])).
fof(f19309,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK14,bool),X0),s(sK14,sK15))))) ) | ($spl3 | $spl96)),
  inference(unit_resulting_resolution,[],[f569,f15729,f15729,f368])).
fof(f19462,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f19310,f15729])).
fof(f19310,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16))))))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1777,f15729,f15729,f368])).
fof(f19461,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f19311,f15729])).
fof(f19311,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0'))))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f617,f15729,f15729,f368])).
fof(f19460,plain,(
  $false | ($spl3 | $spl96)),
  inference(subsumption_resolution,[],[f19312,f15729])).
fof(f19312,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK14,bool),X0),s(sK14,sK17))))) ) | ($spl3 | $spl96)),
  inference(unit_resulting_resolution,[],[f569,f15729,f15729,f368])).
fof(f19459,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f19313,f15729])).
fof(f19313,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'T_0'))))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1777,f15729,f15729,f368])).
fof(f19458,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f19314,f15729])).
fof(f19314,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'T_0'))))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f617,f15729,f15729,f368])).
fof(f19457,plain,(
  $false | ($spl3 | $spl96)),
  inference(subsumption_resolution,[],[f19315,f15729])).
fof(f19315,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK14,bool),X0),s(sK14,sK15))))) ) | ($spl3 | $spl96)),
  inference(unit_resulting_resolution,[],[f569,f15729,f15729,f368])).
fof(f19456,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f19316,f15729])).
fof(f19316,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16))))))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1777,f15729,f15729,f368])).
fof(f19455,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f19317,f15729])).
fof(f19317,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0'))))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f617,f15729,f15729,f368])).
fof(f19454,plain,(
  $false | ($spl3 | $spl96)),
  inference(subsumption_resolution,[],[f19318,f15729])).
fof(f19318,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK14,bool),X0),s(sK14,sK17))))) ) | ($spl3 | $spl96)),
  inference(unit_resulting_resolution,[],[f569,f15729,f15729,f368])).
fof(f19453,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f19322,f15729])).
fof(f19322,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'T_0'))))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1777,f15729,f15729,f368])).
fof(f19452,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f19323,f15729])).
fof(f19323,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'T_0'))))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f617,f15729,f15729,f368])).
fof(f19451,plain,(
  $false | ($spl3 | $spl96)),
  inference(subsumption_resolution,[],[f19324,f15729])).
fof(f19324,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK14,bool),X0),s(sK14,sK15))))) ) | ($spl3 | $spl96)),
  inference(unit_resulting_resolution,[],[f569,f15729,f15729,f368])).
fof(f19450,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f19325,f15729])).
fof(f19325,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16))))))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1777,f15729,f15729,f368])).
fof(f19449,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f19326,f15729])).
fof(f19326,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0'))))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f617,f15729,f15729,f368])).
fof(f19448,plain,(
  $false | ($spl3 | $spl96)),
  inference(subsumption_resolution,[],[f19327,f15729])).
fof(f19327,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK14,bool),X0),s(sK14,sK17))))) ) | ($spl3 | $spl96)),
  inference(unit_resulting_resolution,[],[f569,f15729,f15729,f368])).
fof(f19447,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f19328,f15729])).
fof(f19328,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.?|33|_1'(s(fun(bool,bool),X0))))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1777,f15729,f15729,f368])).
fof(f19446,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f19329,f15729])).
fof(f19329,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.?|33|_1'(s(fun(bool,bool),X0))))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f617,f15729,f15729,f368])).
fof(f19445,plain,(
  $false | ($spl3 | $spl96)),
  inference(subsumption_resolution,[],[f19330,f15729])).
fof(f19330,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),X0))))) ) | ($spl3 | $spl96)),
  inference(unit_resulting_resolution,[],[f569,f15729,f15729,f368])).
fof(f19444,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f19331,f15729])).
fof(f19331,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.?|33|_1'(s(fun(bool,bool),X0))))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1777,f15729,f15729,f368])).
fof(f19443,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f19332,f15729])).
fof(f19332,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.?|33|_1'(s(fun(bool,bool),X0))))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f617,f15729,f15729,f368])).
fof(f19442,plain,(
  $false | ($spl3 | $spl96)),
  inference(subsumption_resolution,[],[f19333,f15729])).
fof(f19333,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),X0))))) ) | ($spl3 | $spl96)),
  inference(unit_resulting_resolution,[],[f569,f15729,f15729,f368])).
fof(f19441,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f19336,f15729])).
fof(f19336,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'T_0'))))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1777,f15729,f15729,f368])).
fof(f19440,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f19337,f15729])).
fof(f19337,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'T_0'))))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f617,f15729,f15729,f368])).
fof(f19439,plain,(
  $false | ($spl3 | $spl96)),
  inference(subsumption_resolution,[],[f19338,f15729])).
fof(f19338,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK14,bool),X0),s(sK14,sK15))))) ) | ($spl3 | $spl96)),
  inference(unit_resulting_resolution,[],[f569,f15729,f15729,f368])).
fof(f19438,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f19339,f15729])).
fof(f19339,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16))))))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1777,f15729,f15729,f368])).
fof(f19437,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f19340,f15729])).
fof(f19340,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0'))))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f617,f15729,f15729,f368])).
fof(f19436,plain,(
  $false | ($spl3 | $spl96)),
  inference(subsumption_resolution,[],[f19341,f15729])).
fof(f19341,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK14,bool),X0),s(sK14,sK17))))) ) | ($spl3 | $spl96)),
  inference(unit_resulting_resolution,[],[f569,f15729,f15729,f368])).
fof(f19435,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f19342,f15729])).
fof(f19342,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.?|33|_1'(s(fun(bool,bool),X0))))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1777,f15729,f15729,f368])).
fof(f19434,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f19343,f15729])).
fof(f19343,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.?|33|_1'(s(fun(bool,bool),X0))))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f617,f15729,f15729,f368])).
fof(f19433,plain,(
  $false | ($spl3 | $spl96)),
  inference(subsumption_resolution,[],[f19344,f15729])).
fof(f19344,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),X0))))) ) | ($spl3 | $spl96)),
  inference(unit_resulting_resolution,[],[f569,f15729,f15729,f368])).
fof(f19432,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f19345,f15729])).
fof(f19345,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.?|33|_1'(s(fun(bool,bool),X0))))) ) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1777,f15729,f15729,f368])).
fof(f19431,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f19346,f15729])).
fof(f19346,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.?|33|_1'(s(fun(bool,bool),X0))))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f617,f15729,f15729,f368])).
fof(f19430,plain,(
  $false | ($spl3 | $spl96)),
  inference(subsumption_resolution,[],[f19347,f15729])).
fof(f19347,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),X0))))) ) | ($spl3 | $spl96)),
  inference(unit_resulting_resolution,[],[f569,f15729,f15729,f368])).
fof(f19429,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f19354,f15729])).
fof(f19354,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X3))))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f15729,f438])).
fof(f438,plain,(
  ( ! [X2,X0,X5,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) )),
  inference(cnf_transformation,[],[f271])).
fof(f19428,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f19378,f15729])).
fof(f19378,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1588,f15729,f1535])).
fof(f19427,plain,(
  $false | ($spl11 | $spl96)),
  inference(subsumption_resolution,[],[f19382,f15729])).
fof(f19382,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1593,f15729,f1535])).
fof(f19426,plain,(
  $false | $spl96),
  inference(subsumption_resolution,[],[f19393,f15729])).
fof(f19393,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,X4))))))) ) | $spl96),
  inference(unit_resulting_resolution,[],[f15729,f438])).
fof(f19394,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f15729,f15729,f438])).
fof(f19383,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f15729,f1593,f15729,f1535])).
fof(f19381,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f15729,f1592,f15729,f1535])).
fof(f19379,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f15729,f1588,f15729,f1535])).
fof(f19377,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f15729,f685,f15729,f1535])).
fof(f19355,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f15729,f15729,f438])).
fof(f19349,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1777,f15729,f15729,f15729,f368])).
fof(f19335,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1777,f15729,f15729,f15729,f368])).
fof(f19320,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1777,f15729,f15729,f15729,f368])).
fof(f19305,plain,(
  $false | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9805,f15729,f11621])).
fof(f11621,plain,(
  ( ! [X37,X38] : (sP12(happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),X37,X38) | ~p(s(bool,X37))) ) | $spl4),
  inference(subsumption_resolution,[],[f11538,f322])).
fof(f11538,plain,(
  ( ! [X37,X38] : (~p(s(bool,'T_0')) | ~p(s(bool,X37)) | sP12(happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),X37,X38)) ) | $spl4),
  inference(superposition,[],[f532,f9540])).
fof(f19303,plain,(
  $false | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9741,f15729,f11617])).
fof(f11617,plain,(
  ( ! [X30,X29] : (sP10(happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),X29,X30) | ~p(s(bool,X29))) ) | $spl4),
  inference(subsumption_resolution,[],[f11534,f322])).
fof(f11534,plain,(
  ( ! [X30,X29] : (~p(s(bool,'T_0')) | ~p(s(bool,X29)) | sP10(happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),X29,X30)) ) | $spl4),
  inference(superposition,[],[f515,f9540])).
fof(f19301,plain,(
  $false | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9741,f15729,f11616])).
fof(f11616,plain,(
  ( ! [X28,X27] : (sP10(happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),X28,X27) | ~p(s(bool,X27))) ) | $spl4),
  inference(subsumption_resolution,[],[f11533,f322])).
fof(f11533,plain,(
  ( ! [X28,X27] : (~p(s(bool,'T_0')) | ~p(s(bool,X27)) | sP10(happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),X28,X27)) ) | $spl4),
  inference(superposition,[],[f514,f9540])).
fof(f19299,plain,(
  $false | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f685,f15729,f11607])).
fof(f11607,plain,(
  ( ! [X12,X11] : (sP2(X11,X12,happ(s(fun(sK14,bool),sK16),s(sK14,sK17))) | ~p(s(bool,X11))) ) | $spl4),
  inference(subsumption_resolution,[],[f11525,f322])).
fof(f11525,plain,(
  ( ! [X12,X11] : (~p(s(bool,'T_0')) | ~p(s(bool,X11)) | sP2(X11,X12,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) ) | $spl4),
  inference(superposition,[],[f451,f9540])).
fof(f19295,plain,(
  $false | ($spl0 | $spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f573,f9367,f15729,f10449])).
fof(f10449,plain,(
  ( ! [X41,X42] : (~sP12('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X42,X41) | p(s(bool,X42)) | ~p(s(bool,X41))) ) | $spl0),
  inference(subsumption_resolution,[],[f10373,f322])).
fof(f10373,plain,(
  ( ! [X41,X42] : (~p(s(bool,'T_0')) | ~p(s(bool,X41)) | p(s(bool,X42)) | ~sP12('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X42,X41)) ) | $spl0),
  inference(superposition,[],[f533,f9096])).
fof(f19293,plain,(
  $false | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f573,f9366,f15729,f10449])).
fof(f19290,plain,(
  $false | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9361,f15729,f10448])).
fof(f10448,plain,(
  ( ! [X39,X40] : (sP12('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X39,X40) | ~p(s(bool,X39))) ) | $spl0),
  inference(subsumption_resolution,[],[f10372,f322])).
fof(f10372,plain,(
  ( ! [X39,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,X39)) | sP12('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X39,X40)) ) | $spl0),
  inference(superposition,[],[f532,f9096])).
fof(f19288,plain,(
  $false | ($spl0 | $spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f573,f1766,f15729,f10446])).
fof(f10446,plain,(
  ( ! [X35,X36] : (~sP11(X36,X35,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | ~p(s(bool,X36)) | p(s(bool,X35))) ) | $spl0),
  inference(subsumption_resolution,[],[f10370,f322])).
fof(f10370,plain,(
  ( ! [X35,X36] : (~p(s(bool,'T_0')) | p(s(bool,X35)) | ~p(s(bool,X36)) | ~sP11(X36,X35,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) ) | $spl0),
  inference(superposition,[],[f526,f9096])).
fof(f19286,plain,(
  $false | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f573,f602,f15729,f10446])).
fof(f19277,plain,(
  $false | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9297,f15729,f10444])).
fof(f10444,plain,(
  ( ! [X31,X32] : (sP10('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X31,X32) | ~p(s(bool,X31))) ) | $spl0),
  inference(subsumption_resolution,[],[f10368,f322])).
fof(f10368,plain,(
  ( ! [X31,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X31)) | sP10('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X31,X32)) ) | $spl0),
  inference(superposition,[],[f515,f9096])).
fof(f19275,plain,(
  $false | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9297,f15729,f10443])).
fof(f10443,plain,(
  ( ! [X30,X29] : (sP10('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X30,X29) | ~p(s(bool,X29))) ) | $spl0),
  inference(subsumption_resolution,[],[f10367,f322])).
fof(f10367,plain,(
  ( ! [X30,X29] : (~p(s(bool,'T_0')) | ~p(s(bool,X29)) | sP10('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X30,X29)) ) | $spl0),
  inference(superposition,[],[f514,f9096])).
fof(f19269,plain,(
  $false | ($spl0 | $spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f573,f1681,f15729,f10441])).
fof(f10441,plain,(
  ( ! [X26,X25] : (~sP8(X26,X25,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | p(s(bool,X26)) | ~p(s(bool,X25))) ) | $spl0),
  inference(subsumption_resolution,[],[f10365,f322])).
fof(f10365,plain,(
  ( ! [X26,X25] : (~p(s(bool,'T_0')) | ~p(s(bool,X25)) | p(s(bool,X26)) | ~sP8(X26,X25,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) ) | $spl0),
  inference(superposition,[],[f501,f9096])).
fof(f19267,plain,(
  $false | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f573,f725,f15729,f10441])).
fof(f19262,plain,(
  $false | ($spl0 | $spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f573,f1675,f15729,f10440])).
fof(f10440,plain,(
  ( ! [X24,X23] : (~sP7(X23,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X24) | p(s(bool,X24)) | ~p(s(bool,X23))) ) | $spl0),
  inference(subsumption_resolution,[],[f10364,f322])).
fof(f10364,plain,(
  ( ! [X24,X23] : (~p(s(bool,'T_0')) | ~p(s(bool,X23)) | p(s(bool,X24)) | ~sP7(X23,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X24)) ) | $spl0),
  inference(superposition,[],[f492,f9096])).
fof(f19260,plain,(
  $false | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f573,f593,f15729,f10440])).
fof(f19251,plain,(
  $false | ($spl0 | $spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f573,f1626,f15729,f10439])).
fof(f10439,plain,(
  ( ! [X21,X22] : (~sP4(X21,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X22) | p(s(bool,X22)) | ~p(s(bool,X21))) ) | $spl0),
  inference(subsumption_resolution,[],[f10363,f322])).
fof(f10363,plain,(
  ( ! [X21,X22] : (~p(s(bool,'T_0')) | ~p(s(bool,X21)) | p(s(bool,X22)) | ~sP4(X21,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X22)) ) | $spl0),
  inference(superposition,[],[f473,f9096])).
fof(f19249,plain,(
  $false | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f573,f585,f15729,f10439])).
fof(f19240,plain,(
  $false | ($spl0 | $spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f573,f1619,f15729,f10438])).
fof(f10438,plain,(
  ( ! [X19,X20] : (~sP5(X19,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X20) | p(s(bool,X20)) | ~p(s(bool,X19))) ) | $spl0),
  inference(subsumption_resolution,[],[f10362,f322])).
fof(f10362,plain,(
  ( ! [X19,X20] : (~p(s(bool,'T_0')) | ~p(s(bool,X19)) | p(s(bool,X20)) | ~sP5(X19,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X20)) ) | $spl0),
  inference(superposition,[],[f469,f9096])).
fof(f19238,plain,(
  $false | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f573,f580,f15729,f10438])).
fof(f19223,plain,(
  $false | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f573,f10072,f15729,f10435])).
fof(f19218,plain,(
  $false | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f685,f15729,f10434])).
fof(f10434,plain,(
  ( ! [X14,X13] : (sP2(X13,X14,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | ~p(s(bool,X13))) ) | $spl0),
  inference(subsumption_resolution,[],[f10359,f322])).
fof(f10359,plain,(
  ( ! [X14,X13] : (~p(s(bool,'T_0')) | ~p(s(bool,X13)) | sP2(X13,X14,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) ) | $spl0),
  inference(superposition,[],[f451,f9096])).
fof(f19216,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f3590,f15729,f7491])).
fof(f7491,plain,(
  ( ! [X37,X38] : (sP12(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X37,X38) | ~p(s(bool,X37))) )),
  inference(subsumption_resolution,[],[f7406,f322])).
fof(f7406,plain,(
  ( ! [X37,X38] : (~p(s(bool,'T_0')) | ~p(s(bool,X37)) | sP12(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X37,X38)) )),
  inference(superposition,[],[f532,f3354])).
fof(f19214,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f3534,f15729,f7487])).
fof(f7487,plain,(
  ( ! [X30,X29] : (sP10(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X29,X30) | ~p(s(bool,X29))) )),
  inference(subsumption_resolution,[],[f7402,f322])).
fof(f7402,plain,(
  ( ! [X30,X29] : (~p(s(bool,'T_0')) | ~p(s(bool,X29)) | sP10(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X29,X30)) )),
  inference(superposition,[],[f515,f3354])).
fof(f19212,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f3534,f15729,f7486])).
fof(f7486,plain,(
  ( ! [X28,X27] : (sP10(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X28,X27) | ~p(s(bool,X27))) )),
  inference(subsumption_resolution,[],[f7401,f322])).
fof(f7401,plain,(
  ( ! [X28,X27] : (~p(s(bool,'T_0')) | ~p(s(bool,X27)) | sP10(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X28,X27)) )),
  inference(superposition,[],[f514,f3354])).
fof(f19210,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f685,f15729,f7477])).
fof(f7477,plain,(
  ( ! [X12,X11] : (sP2(X11,X12,happ(s(fun(sK14,bool),sK16),s(sK14,sK15))) | ~p(s(bool,X11))) )),
  inference(subsumption_resolution,[],[f7393,f322])).
fof(f7393,plain,(
  ( ! [X12,X11] : (~p(s(bool,'T_0')) | ~p(s(bool,X11)) | sP2(X11,X12,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) )),
  inference(superposition,[],[f451,f3354])).
fof(f19208,plain,(
  $false | ($spl44 | $spl96)),
  inference(unit_resulting_resolution,[],[f1261,f15729,f4617])).
fof(f19206,plain,(
  $false | ($spl36 | $spl96)),
  inference(unit_resulting_resolution,[],[f1164,f15729,f4605])).
fof(f19204,plain,(
  $false | ($spl34 | $spl96)),
  inference(unit_resulting_resolution,[],[f1164,f15729,f4602])).
fof(f19202,plain,(
  $false | ($spl18 | $spl96)),
  inference(unit_resulting_resolution,[],[f685,f15729,f4577])).
fof(f19195,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f15729,f1593,f15729,f1535])).
fof(f19193,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f15729,f1592,f15729,f1535])).
fof(f19191,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f15729,f1588,f15729,f1535])).
fof(f19189,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f15729,f685,f15729,f1535])).
fof(f19176,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f573,f15729,f1766,f15729,f1350])).
fof(f19174,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f573,f322,f1766,f15729,f1350])).
fof(f19171,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f573,f15729,f602,f15729,f1350])).
fof(f19169,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f573,f322,f602,f15729,f1350])).
fof(f19158,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f573,f15729,f1766,f15729,f1350])).
fof(f19155,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f573,f15729,f602,f15729,f1350])).
fof(f19120,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f322,f573,f1681,f15729,f1345])).
fof(f19117,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f322,f573,f725,f15729,f1345])).
fof(f19110,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f573,f15729,f1681,f15729,f1345])).
fof(f19107,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f15729,f573,f725,f15729,f1345])).
fof(f19102,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f322,f573,f1675,f15729,f1344])).
fof(f19098,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f322,f573,f593,f15729,f1344])).
fof(f19086,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f573,f15729,f1675,f15729,f1344])).
fof(f19083,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f573,f15729,f593,f15729,f1344])).
fof(f19072,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f322,f573,f1626,f15729,f1343])).
fof(f19068,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f322,f573,f585,f15729,f1343])).
fof(f19056,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f573,f15729,f1626,f15729,f1343])).
fof(f19053,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f573,f15729,f585,f15729,f1343])).
fof(f19042,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f322,f573,f1619,f15729,f1342])).
fof(f19038,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f322,f573,f580,f15729,f1342])).
fof(f19026,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f573,f15729,f1619,f15729,f1342])).
fof(f19023,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f573,f15729,f580,f15729,f1342])).
fof(f19008,plain,(
  $false | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f15729,f573,f9630,f15729,f1341])).
fof(f18998,plain,(
  $false | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f15729,f573,f9186,f15729,f1341])).
fof(f18990,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f15729,f573,f1311,f15729,f1341])).
fof(f18982,plain,(
  $false | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f573,f15729,f9630,f15729,f1341])).
fof(f18976,plain,(
  $false | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f573,f15729,f9186,f15729,f1341])).
fof(f18971,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f573,f322,f1311,f15729,f1341])).
fof(f18966,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f322,f1593,f15729,f1340])).
fof(f18964,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f322,f1592,f15729,f1340])).
fof(f18962,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f322,f1588,f15729,f1340])).
fof(f18960,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f322,f685,f15729,f1340])).
fof(f18956,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f15729,f1593,f15729,f1340])).
fof(f18954,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f15729,f1592,f15729,f1340])).
fof(f18952,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f15729,f1588,f15729,f1340])).
fof(f18950,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f15729,f685,f15729,f1340])).
fof(f18938,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f573,f1766,f15729,f1257])).
fof(f1257,plain,(
  ( ! [X6,X7] : (~sP11(X7,X6,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16))) | ~p(s(bool,X7)) | p(s(bool,X6))) )),
  inference(subsumption_resolution,[],[f1252,f322])).
fof(f1252,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | p(s(bool,X6)) | ~p(s(bool,X7)) | ~sP11(X7,X6,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) )),
  inference(superposition,[],[f526,f621])).
fof(f18936,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f573,f602,f15729,f1257])).
fof(f18925,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f573,f1681,f15729,f1113])).
fof(f1113,plain,(
  ( ! [X6,X7] : (~sP8(X7,X6,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16))) | p(s(bool,X7)) | ~p(s(bool,X6))) )),
  inference(subsumption_resolution,[],[f1109,f322])).
fof(f1109,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | ~p(s(bool,X6)) | p(s(bool,X7)) | ~sP8(X7,X6,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) )),
  inference(superposition,[],[f501,f621])).
fof(f18923,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f573,f725,f15729,f1113])).
fof(f18920,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f573,f1675,f15729,f1076])).
fof(f18918,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f573,f593,f15729,f1076])).
fof(f18909,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f573,f1626,f15729,f1032])).
fof(f1032,plain,(
  ( ! [X6,X7] : (~sP4(X6,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X7) | p(s(bool,X7)) | ~p(s(bool,X6))) )),
  inference(subsumption_resolution,[],[f1027,f322])).
fof(f1027,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | ~p(s(bool,X6)) | p(s(bool,X7)) | ~sP4(X6,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X7)) )),
  inference(superposition,[],[f473,f621])).
fof(f18907,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f573,f585,f15729,f1032])).
fof(f18898,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f573,f1619,f15729,f1011])).
fof(f18896,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f573,f580,f15729,f1011])).
fof(f18884,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f573,f1484,f15729,f992])).
fof(f18878,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f1282,f15729,f746])).
fof(f746,plain,(
  ( ! [X2,X3] : (sP10('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X2,X3) | ~p(s(bool,X2))) )),
  inference(subsumption_resolution,[],[f745,f322])).
fof(f745,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | ~p(s(bool,X2)) | sP10('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X2,X3)) )),
  inference(superposition,[],[f515,f621])).
fof(f18876,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f1282,f15729,f736])).
fof(f736,plain,(
  ( ! [X2,X3] : (sP10('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X3,X2) | ~p(s(bool,X2))) )),
  inference(subsumption_resolution,[],[f735,f322])).
fof(f735,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | ~p(s(bool,X2)) | sP10('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X3,X2)) )),
  inference(superposition,[],[f514,f621])).
fof(f18874,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f685,f15729,f707])).
fof(f707,plain,(
  ( ! [X2,X3] : (sP2(X2,X3,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16))) | ~p(s(bool,X2))) )),
  inference(subsumption_resolution,[],[f706,f322])).
fof(f706,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | ~p(s(bool,X2)) | sP2(X2,X3,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) )),
  inference(superposition,[],[f451,f621])).
fof(f18872,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f664,f15729,f559])).
fof(f559,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(duplicate_literal_removal,[],[f448])).
fof(f448,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f276])).
fof(f18870,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1687,f1721,f15729,f557])).
fof(f1721,plain,(
  ( ! [X0,X1] : (sP9('|dollar|exists_1'(s(fun(sK14,bool),sK16)),X0,X1)) ) | $spl11),
  inference(unit_resulting_resolution,[],[f1559,f512])).
fof(f18868,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f729,f597,f15729,f557])).
fof(f18866,plain,(
  $false | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f10994,f11025,f15729,f557])).
fof(f18864,plain,(
  $false | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9706,f9726,f15729,f557])).
fof(f18862,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f3507,f3521,f15729,f557])).
fof(f18860,plain,(
  $false | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9262,f9282,f15729,f557])).
fof(f18858,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f729,f614,f15729,f557])).
fof(f614,plain,(
  ( ! [X0,X1] : (sP9(X0,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X1)) )),
  inference(unit_resulting_resolution,[],[f318,f511])).
fof(f18856,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f730,f596,f15729,f557])).
fof(f18854,plain,(
  $false | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f10985,f11024,f15729,f557])).
fof(f18852,plain,(
  $false | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9700,f9725,f15729,f557])).
fof(f18850,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f3503,f3520,f15729,f557])).
fof(f18848,plain,(
  $false | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9256,f9281,f15729,f557])).
fof(f18846,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f730,f613,f15729,f557])).
fof(f613,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) )),
  inference(unit_resulting_resolution,[],[f318,f510])).
fof(f18844,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f729,f595,f15729,f557])).
fof(f18821,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1744,f1766,f15729,f556])).
fof(f18819,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f755,f602,f15729,f556])).
fof(f18817,plain,(
  $false | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f11045,f11110,f15729,f556])).
fof(f18815,plain,(
  $false | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9742,f9787,f15729,f556])).
fof(f18813,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f3535,f3575,f15729,f556])).
fof(f18811,plain,(
  $false | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9298,f9343,f15729,f556])).
fof(f18809,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f755,f615,f15729,f556])).
fof(f18807,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f757,f601,f15729,f556])).
fof(f18805,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1746,f1765,f15729,f556])).
fof(f18803,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f755,f599,f15729,f556])).
fof(f18779,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f573,f421,f15729,f555])).
fof(f18776,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f573,f1572,f15729,f555])).
fof(f18769,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f573,f1519,f15729,f555])).
fof(f18757,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f765,f573,f15729,f15729,f533])).
fof(f18725,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f1237,f573,f322,f15729,f533])).
fof(f18691,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f765,f573,f15729,f15729,f533])).
fof(f18659,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f765,f573,f322,f15729,f533])).
fof(f18634,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f1261,f15729,f15729,f532])).
fof(f18632,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f757,f322,f15729,f532])).
fof(f18608,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f1261,f15729,f15729,f532])).
fof(f18606,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f1261,f322,f15729,f532])).
fof(f18582,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f755,f573,f15729,f531])).
fof(f18558,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f755,f573,f573,f15729,f530])).
fof(f18534,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f1249,f15729,f528])).
fof(f18532,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f599,f15729,f573,f15729,f526])).
fof(f18530,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f599,f322,f573,f15729,f526])).
fof(f18506,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f599,f15729,f573,f15729,f526])).
fof(f18504,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f602,f322,f573,f15729,f526])).
fof(f18480,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1744,f1766,f573,f15729,f525])).
fof(f18478,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f755,f602,f573,f15729,f525])).
fof(f18476,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1746,f1765,f573,f15729,f525])).
fof(f18474,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f755,f599,f573,f15729,f525])).
fof(f18425,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1744,f1766,f573,f15729,f524])).
fof(f18423,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f755,f602,f573,f15729,f524])).
fof(f18421,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1746,f1765,f573,f15729,f524])).
fof(f18419,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f755,f599,f573,f15729,f524])).
fof(f18370,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1744,f1766,f15729,f15729,f523])).
fof(f18368,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f755,f602,f15729,f15729,f523])).
fof(f18366,plain,(
  $false | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f11045,f11110,f15729,f15729,f523])).
fof(f18364,plain,(
  $false | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9742,f9787,f15729,f15729,f523])).
fof(f18362,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f3535,f3575,f15729,f15729,f523])).
fof(f18360,plain,(
  $false | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9298,f9343,f15729,f15729,f523])).
fof(f18358,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f755,f615,f15729,f15729,f523])).
fof(f18356,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f757,f601,f15729,f15729,f523])).
fof(f18354,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1746,f1765,f15729,f15729,f523])).
fof(f18352,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f755,f599,f15729,f15729,f523])).
fof(f18266,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1766,f1751,f15729,f15729,f523])).
fof(f18264,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1766,f1747,f15729,f15729,f523])).
fof(f18262,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1766,f1744,f15729,f15729,f523])).
fof(f18254,plain,(
  $false | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f602,f11045,f15729,f15729,f523])).
fof(f18252,plain,(
  $false | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f602,f9742,f15729,f15729,f523])).
fof(f18250,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f602,f3535,f15729,f15729,f523])).
fof(f18248,plain,(
  $false | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f602,f9298,f15729,f15729,f523])).
fof(f18246,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f602,f757,f15729,f15729,f523])).
fof(f18244,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f602,f1746,f15729,f15729,f523])).
fof(f18242,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f602,f755,f15729,f15729,f523])).
fof(f18224,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1766,f1751,f322,f15729,f523])).
fof(f18218,plain,(
  $false | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f602,f11045,f322,f15729,f523])).
fof(f18216,plain,(
  $false | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f602,f9742,f322,f15729,f523])).
fof(f18214,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f602,f3535,f322,f15729,f523])).
fof(f18212,plain,(
  $false | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f602,f9298,f322,f15729,f523])).
fof(f18210,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f602,f757,f322,f15729,f523])).
fof(f18170,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f1237,f573,f15729,f520])).
fof(f18138,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f1140,f573,f573,f15729,f516])).
fof(f1140,plain,(
  sP10('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f597,f573,f573,f508])).
fof(f508,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~sP9(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f302])).
fof(f18102,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f1164,f15729,f15729,f515])).
fof(f18100,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f730,f322,f15729,f515])).
fof(f18077,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f1164,f15729,f15729,f515])).
fof(f18075,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f1164,f322,f15729,f515])).
fof(f18052,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f1164,f15729,f15729,f514])).
fof(f18050,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f729,f322,f15729,f514])).
fof(f18027,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f1164,f15729,f15729,f514])).
fof(f18025,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f1164,f322,f15729,f514])).
fof(f18002,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f1152,f15729,f511])).
fof(f18000,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f1152,f15729,f510])).
fof(f17998,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f597,f573,f573,f15729,f509])).
fof(f17980,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1687,f1721,f573,f15729,f507])).
fof(f17978,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f729,f597,f573,f15729,f507])).
fof(f17976,plain,(
  $false | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f10985,f11024,f573,f15729,f507])).
fof(f17974,plain,(
  $false | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9700,f9725,f573,f15729,f507])).
fof(f17972,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f3503,f3520,f573,f15729,f507])).
fof(f17970,plain,(
  $false | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9256,f9281,f573,f15729,f507])).
fof(f17967,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f729,f595,f573,f15729,f507])).
fof(f17892,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1721,f1688,f573,f15729,f506])).
fof(f17885,plain,(
  $false | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f597,f10994,f573,f15729,f506])).
fof(f17883,plain,(
  $false | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f597,f9706,f573,f15729,f506])).
fof(f17881,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f597,f3507,f573,f15729,f506])).
fof(f17879,plain,(
  $false | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f597,f9262,f573,f15729,f506])).
fof(f17877,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f597,f730,f573,f15729,f506])).
fof(f17811,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f1140,f573,f15729,f503])).
fof(f17776,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f1140,f573,f15729,f502])).
fof(f17741,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f716,f573,f15729,f15729,f501])).
fof(f17709,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f716,f573,f322,f15729,f501])).
fof(f17687,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f716,f573,f15729,f15729,f501])).
fof(f17655,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f725,f573,f322,f15729,f501])).
fof(f17633,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f708,f15729,f15729,f15729,f498])).
fof(f17631,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f711,f322,f15729,f15729,f498])).
fof(f17604,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f708,f15729,f322,f15729,f498])).
fof(f17602,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f1080,f322,f322,f15729,f498])).
fof(f17550,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f708,f15729,f15729,f15729,f498])).
fof(f17548,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f708,f322,f15729,f15729,f498])).
fof(f17521,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f708,f15729,f322,f15729,f498])).
fof(f17519,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f708,f322,f322,f15729,f498])).
fof(f17475,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f708,f15729,f15729,f15729,f498])).
fof(f17473,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f708,f322,f15729,f15729,f498])).
fof(f17446,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f711,f15729,f322,f15729,f498])).
fof(f17444,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f711,f322,f322,f15729,f498])).
fof(f17401,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f1068,f15729,f493])).
fof(f17399,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f591,f573,f15729,f15729,f492])).
fof(f17387,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f591,f573,f322,f15729,f492])).
fof(f17373,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f591,f573,f15729,f15729,f492])).
fof(f17361,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f593,f573,f322,f15729,f492])).
fof(f17346,plain,(
  $false | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f593,f10911,f15729,f15729,f488])).
fof(f17344,plain,(
  $false | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f591,f10910,f15729,f15729,f488])).
fof(f17342,plain,(
  $false | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f593,f9642,f15729,f15729,f488])).
fof(f17340,plain,(
  $false | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f591,f9641,f15729,f15729,f488])).
fof(f17338,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f593,f3449,f15729,f15729,f488])).
fof(f17336,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f591,f3448,f15729,f15729,f488])).
fof(f17334,plain,(
  $false | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f593,f9198,f15729,f15729,f488])).
fof(f17332,plain,(
  $false | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f591,f9197,f15729,f15729,f488])).
fof(f17318,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1675,f1666,f15729,f15729,f488])).
fof(f17316,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f593,f711,f15729,f15729,f488])).
fof(f17314,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1674,f1663,f15729,f15729,f488])).
fof(f17312,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f591,f708,f15729,f15729,f488])).
fof(f17299,plain,(
  $false | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f591,f10910,f322,f15729,f488])).
fof(f17297,plain,(
  $false | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f591,f9641,f322,f15729,f488])).
fof(f17295,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f591,f3448,f322,f15729,f488])).
fof(f17293,plain,(
  $false | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f591,f9197,f322,f15729,f488])).
fof(f17285,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1674,f1663,f322,f15729,f488])).
fof(f17283,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f591,f708,f322,f15729,f488])).
fof(f17242,plain,(
  $false | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f593,f10911,f15729,f15729,f488])).
fof(f17240,plain,(
  $false | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f591,f10910,f15729,f15729,f488])).
fof(f17238,plain,(
  $false | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f593,f9642,f15729,f15729,f488])).
fof(f17236,plain,(
  $false | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f591,f9641,f15729,f15729,f488])).
fof(f17234,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f593,f3449,f15729,f15729,f488])).
fof(f17232,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f591,f3448,f15729,f15729,f488])).
fof(f17230,plain,(
  $false | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f593,f9198,f15729,f15729,f488])).
fof(f17228,plain,(
  $false | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f591,f9197,f15729,f15729,f488])).
fof(f17214,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1675,f1666,f15729,f15729,f488])).
fof(f17212,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f593,f711,f15729,f15729,f488])).
fof(f17210,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1674,f1663,f15729,f15729,f488])).
fof(f17208,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f591,f708,f15729,f15729,f488])).
fof(f17195,plain,(
  $false | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f593,f10911,f322,f15729,f488])).
fof(f17193,plain,(
  $false | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f593,f9642,f322,f15729,f488])).
fof(f17191,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f593,f3449,f322,f15729,f488])).
fof(f17189,plain,(
  $false | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f593,f9198,f322,f15729,f488])).
fof(f17182,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1675,f1666,f322,f15729,f488])).
fof(f17180,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f593,f711,f322,f15729,f488])).
fof(f17140,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f725,f573,f15729,f487])).
fof(f17108,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f716,f573,f15729,f486])).
fof(f17076,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f923,f617,f15729,f483])).
fof(f923,plain,(
  sP6('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f573,f617,f481])).
fof(f17074,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f2003,f1777,f15729,f483])).
fof(f2003,plain,(
  sP6('F_0','|dollar|exists_1'(s(fun(sK14,bool),sK16)),'T_0') | $spl11),
  inference(unit_resulting_resolution,[],[f322,f589,f580,f1626,f466])).
fof(f589,plain,(
  ( ! [X0,X1] : (sP3('T_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f322,f480])).
fof(f17072,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f925,f617,f15729,f483])).
fof(f17070,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1988,f1777,f15729,f483])).
fof(f1988,plain,(
  sP6('F_0','T_0','|dollar|exists_1'(s(fun(sK14,bool),sK16))) | $spl11),
  inference(unit_resulting_resolution,[],[f573,f588,f580,f1625,f468])).
fof(f17005,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f1313,f421,f15729,f482])).
fof(f17003,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f1514,f364,f15729,f482])).
fof(f17001,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f2013,f1572,f15729,f482])).
fof(f16999,plain,(
  $false | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f10045,f9096,f15729,f482])).
fof(f16997,plain,(
  $false | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f11794,f10783,f15729,f482])).
fof(f16995,plain,(
  $false | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f10534,f9540,f15729,f482])).
fof(f16993,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f5909,f3354,f15729,f482])).
fof(f16991,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f951,f621,f15729,f482])).
fof(f16988,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f1310,f421,f15729,f482])).
fof(f16986,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f1512,f364,f15729,f482])).
fof(f16984,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f2020,f1572,f15729,f482])).
fof(f16982,plain,(
  $false | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f10040,f9096,f15729,f482])).
fof(f16980,plain,(
  $false | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f11789,f10783,f15729,f482])).
fof(f16978,plain,(
  $false | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f10529,f9540,f15729,f482])).
fof(f16976,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f5824,f3354,f15729,f482])).
fof(f16974,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f949,f621,f15729,f482])).
fof(f19425,plain,(
  $false | ($spl11 | $spl96)),
  inference(trivial_inequality_removal,[],[f16969])).
fof(f16969,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f2025,f15729,f482])).
fof(f2025,plain,(
  ~sP6('|dollar|exists_1'(s(fun(sK14,bool),sK16)),'F_0','F_0') | $spl11),
  inference(unit_resulting_resolution,[],[f1643,f462])).
fof(f1643,plain,(
  ~sP3('F_0','F_0','|dollar|exists_1'(s(fun(sK14,bool),sK16))) | $spl11),
  inference(unit_resulting_resolution,[],[f573,f573,f1559,f477])).
fof(f19424,plain,(
  $false | ($spl11 | $spl96)),
  inference(trivial_inequality_removal,[],[f16968])).
fof(f16968,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1982,f15729,f482])).
fof(f1982,plain,(
  ~sP6('|dollar|exists_1'(s(fun(sK14,bool),sK16)),'T_0','T_0') | $spl11),
  inference(unit_resulting_resolution,[],[f1620,f463])).
fof(f1620,plain,(
  ~sP4('T_0','T_0','|dollar|exists_1'(s(fun(sK14,bool),sK16))) | $spl11),
  inference(unit_resulting_resolution,[],[f322,f322,f1559,f473])).
fof(f19423,plain,(
  $false | ($spl11 | $spl96)),
  inference(trivial_inequality_removal,[],[f16954])).
fof(f16954,plain,(
  s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16))) != s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16))) | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f6686,f15729,f482])).
fof(f6686,plain,(
  ~sP6('F_0','|dollar|exists_1'(s(fun(sK14,bool),sK16)),'|dollar|exists_1'(s(fun(sK14,bool),sK16))) | $spl11),
  inference(unit_resulting_resolution,[],[f1636,f462])).
fof(f1636,plain,(
  ~sP3('|dollar|exists_1'(s(fun(sK14,bool),sK16)),'|dollar|exists_1'(s(fun(sK14,bool),sK16)),'F_0') | $spl11),
  inference(unit_resulting_resolution,[],[f1559,f573,f1559,f477])).
fof(f19422,plain,(
  $false | ($spl0 | $spl96)),
  inference(trivial_inequality_removal,[],[f16946])).
fof(f16946,plain,(
  s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) != s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f11408,f15729,f482])).
fof(f11408,plain,(
  ~sP6('F_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | $spl0),
  inference(unit_resulting_resolution,[],[f9171,f463])).
fof(f9171,plain,(
  ~sP4('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f566,f573,f566,f473])).
fof(f19421,plain,(
  $false | $spl96),
  inference(trivial_inequality_removal,[],[f16942])).
fof(f16942,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl96),
  inference(unit_resulting_resolution,[],[f1051,f15729,f482])).
fof(f1051,plain,(
  ~sP6('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1042,f462])).
fof(f1042,plain,(
  ~sP3('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f573,f573,f573,f477])).
fof(f19420,plain,(
  $false | $spl96),
  inference(trivial_inequality_removal,[],[f16936])).
fof(f16936,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl96),
  inference(unit_resulting_resolution,[],[f1036,f15729,f482])).
fof(f1036,plain,(
  ~sP6('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f1024,f463])).
fof(f1024,plain,(
  ~sP4('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f573,f322,f322,f473])).
fof(f16920,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f1042,f15729,f480])).
fof(f16918,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f1042,f15729,f479])).
fof(f16916,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f1042,f15729,f478])).
fof(f16914,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f1024,f15729,f474])).
fof(f16912,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f583,f573,f15729,f15729,f473])).
fof(f16900,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f583,f573,f322,f15729,f473])).
fof(f16886,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f583,f573,f15729,f15729,f473])).
fof(f16874,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f585,f573,f322,f15729,f473])).
fof(f16860,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f1003,f15729,f470])).
fof(f1003,plain,(
  ~sP5('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f573,f322,f322,f469])).
fof(f16858,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f578,f573,f15729,f15729,f469])).
fof(f16846,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f578,f573,f322,f15729,f469])).
fof(f16832,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f578,f573,f15729,f15729,f469])).
fof(f16820,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f580,f573,f322,f15729,f469])).
fof(f16801,plain,(
  $false | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f588,f583,f578,f10907,f15729,f466])).
fof(f16797,plain,(
  $false | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9623,f583,f578,f9581,f15729,f466])).
fof(f16792,plain,(
  $false | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f588,f583,f578,f9638,f15729,f466])).
fof(f16787,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f588,f583,f578,f3443,f15729,f466])).
fof(f16782,plain,(
  $false | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9179,f583,f578,f9137,f15729,f466])).
fof(f16776,plain,(
  $false | ($spl0 | $spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f588,f1625,f1618,f9193,f15729,f466])).
fof(f16772,plain,(
  $false | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f588,f583,f578,f9194,f15729,f466])).
fof(f16724,plain,(
  $false | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f10890,f583,f578,f10823,f15729,f466])).
fof(f16720,plain,(
  $false | ($spl4 | $spl96)),
  inference(unit_resulting_resolution,[],[f9623,f583,f578,f9580,f15729,f466])).
fof(f16716,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f3428,f583,f578,f3390,f15729,f466])).
fof(f16712,plain,(
  $false | ($spl0 | $spl96)),
  inference(unit_resulting_resolution,[],[f9179,f583,f578,f9136,f15729,f466])).
fof(f16708,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f610,f583,f578,f1124,f15729,f466])).
fof(f16698,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f588,f1625,f1618,f1596,f15729,f466])).
fof(f16694,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f588,f583,f578,f941,f15729,f466])).
fof(f16138,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f1428,f573,f15729,f15729,f465])).
fof(f16073,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f1428,f573,f322,f15729,f465])).
fof(f16034,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f1428,f573,f15729,f15729,f465])).
fof(f15969,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f925,f573,f322,f15729,f465])).
fof(f15930,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f664,f15729,f15729,f451])).
fof(f15928,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f664,f322,f15729,f451])).
fof(f15918,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f664,f15729,f15729,f451])).
fof(f15916,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f685,f322,f15729,f451])).
fof(f15906,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f664,f15729,f15729,f450])).
fof(f15904,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f664,f322,f15729,f450])).
fof(f15894,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f664,f15729,f15729,f450])).
fof(f15883,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f664,f15729,f15729,f449])).
fof(f15872,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f664,f15729,f15729,f449])).
fof(f15870,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f685,f322,f15729,f449])).
fof(f15860,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f639,f573,f15729,f412])).
fof(f639,plain,(
  ~sP1('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f573,f573,f407])).
fof(f407,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f249])).
fof(f15840,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f639,f573,f15729,f409])).
fof(f15820,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f643,f15729,f15729,f408])).
fof(f15818,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f643,f322,f15729,f408])).
fof(f15800,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f643,f15729,f15729,f408])).
fof(f15798,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f649,f322,f15729,f408])).
fof(f15780,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f657,f15729,f394])).
fof(f394,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f242])).
fof(f242,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f241])).
fof(f241,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f183])).
fof(f183,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f657,plain,(
  ~sP0('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f573,f573,f574])).
fof(f574,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f392,f573])).
fof(f392,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f242])).
fof(f15778,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f657,f15729,f393])).
fof(f393,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f242])).
fof(f15776,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f15729,f617,f15729,f387])).
fof(f15774,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f15729,f1777,f15729,f387])).
fof(f15772,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f322,f617,f15729,f387])).
fof(f15770,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f322,f1777,f15729,f387])).
fof(f15766,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f15729,f617,f15729,f387])).
fof(f15764,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f15729,f1777,f15729,f387])).
fof(f15762,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f322,f617,f15729,f387])).
fof(f15760,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f322,f1777,f15729,f387])).
fof(f15754,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1572,f15729,f340])).
fof(f15752,plain,(
  $false | ($spl11 | $spl96)),
  inference(unit_resulting_resolution,[],[f1777,f15729,f337])).
fof(f15750,plain,(
  $false | $spl96),
  inference(unit_resulting_resolution,[],[f573,f15729])).
fof(f15747,plain,(
  ( ! [X12,X10,X13,X11] : (s(bool,'T_0') = s(bool,happ(s(fun(X11,bool),X13),s(X11,sK30(X13,X12,X10,X11)))) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X10,X11),X12),s(fun(X11,bool),X13)))) ) | $spl98),
  inference(cnf_transformation,[],[f15747_D])).
fof(f15747_D,plain,(
  ( ! [X12,X10,X13,X11] : (s(bool,'T_0') = s(bool,happ(s(fun(X11,bool),X13),s(X11,sK30(X13,X12,X10,X11)))) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X10,X11),X12),s(fun(X11,bool),X13)))) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f15748,plain,(
  ( ! [X19,X17,X15,X18,X16] : (p(s(bool,X19)) | s(bool,'T_0') = s(bool,happ(s(fun(X16,bool),X18),s(X16,sK30(X18,X17,X15,X16)))) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X15,X16),X17),s(fun(X16,bool),X18)))) )),
  inference(subsumption_resolution,[],[f15743,f336])).
fof(f15743,plain,(
  ( ! [X19,X17,X15,X18,X16] : (p(s(bool,X19)) | s(bool,'T_0') = s(bool,happ(s(fun(X16,bool),X18),s(X16,sK30(X18,X17,X15,X16)))) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X15,X16),X17),s(fun(X16,bool),X18)))) )),
  inference(superposition,[],[f1877,f836])).
fof(f1877,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X3,X0),X2),s(fun(X0,bool),X1)))) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK30(X1,X2,X3,X0))))) )),
  inference(subsumption_resolution,[],[f1874,f573])).
fof(f1874,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'F_0')) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X3,X0),X2),s(fun(X0,bool),X1)))) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK30(X1,X2,X3,X0))))) )),
  inference(superposition,[],[f430,f329])).
fof(f430,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,sK30(X3,X2,X1,X0))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f262])).
fof(f15746,plain,(
  ( ! [X14,X12,X10,X13,X11] : (p(s(bool,X14)) | s(bool,'T_0') = s(bool,happ(s(fun(X11,bool),X13),s(X11,sK30(X13,X12,X10,X11)))) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X10,X11),X12),s(fun(X11,bool),X13)))) )),
  inference(subsumption_resolution,[],[f15742,f336])).
fof(f15742,plain,(
  ( ! [X14,X12,X10,X13,X11] : (p(s(bool,X14)) | s(bool,'T_0') = s(bool,happ(s(fun(X11,bool),X13),s(X11,sK30(X13,X12,X10,X11)))) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X10,X11),X12),s(fun(X11,bool),X13))) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f1877,f836])).
fof(f15728,plain,(
  ( ! [X12,X10,X13,X11] : (s(bool,'T_0') = s(bool,happ(s(fun(X11,bool),X13),s(X11,sK28(X13,X12,X10,X11)))) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X10,X11),X12),s(fun(X11,bool),X13)))) ) | $spl94),
  inference(cnf_transformation,[],[f15728_D])).
fof(f15728_D,plain,(
  ( ! [X12,X10,X13,X11] : (s(bool,'T_0') = s(bool,happ(s(fun(X11,bool),X13),s(X11,sK28(X13,X12,X10,X11)))) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X10,X11),X12),s(fun(X11,bool),X13)))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f15730,plain,(
  ( ! [X19,X17,X15,X18,X16] : (p(s(bool,X19)) | s(bool,'T_0') = s(bool,happ(s(fun(X16,bool),X18),s(X16,sK28(X18,X17,X15,X16)))) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X15,X16),X17),s(fun(X16,bool),X18)))) )),
  inference(subsumption_resolution,[],[f15724,f336])).
fof(f15724,plain,(
  ( ! [X19,X17,X15,X18,X16] : (p(s(bool,X19)) | s(bool,'T_0') = s(bool,happ(s(fun(X16,bool),X18),s(X16,sK28(X18,X17,X15,X16)))) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X15,X16),X17),s(fun(X16,bool),X18)))) )),
  inference(superposition,[],[f1867,f836])).
fof(f1867,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X3,X0),X2),s(fun(X0,bool),X1)))) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK28(X1,X2,X3,X0))))) )),
  inference(subsumption_resolution,[],[f1864,f573])).
fof(f1864,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'F_0')) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X3,X0),X2),s(fun(X0,bool),X1)))) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK28(X1,X2,X3,X0))))) )),
  inference(superposition,[],[f427,f329])).
fof(f427,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,sK28(X3,X2,X1,X0))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f259])).
fof(f15727,plain,(
  ( ! [X14,X12,X10,X13,X11] : (p(s(bool,X14)) | s(bool,'T_0') = s(bool,happ(s(fun(X11,bool),X13),s(X11,sK28(X13,X12,X10,X11)))) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X10,X11),X12),s(fun(X11,bool),X13)))) )),
  inference(subsumption_resolution,[],[f15723,f336])).
fof(f15723,plain,(
  ( ! [X14,X12,X10,X13,X11] : (p(s(bool,X14)) | s(bool,'T_0') = s(bool,happ(s(fun(X11,bool),X13),s(X11,sK28(X13,X12,X10,X11)))) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X10,X11),X12),s(fun(X11,bool),X13))) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f1867,f836])).
fof(f12676,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,X1)) | ~sP12(happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4)),X2,'T_0') | p(s(bool,X2))) ) | $spl92),
  inference(cnf_transformation,[],[f12676_D])).
fof(f12676_D,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,X1)) | ~sP12(happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4)),X2,'T_0') | p(s(bool,X2))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f12678,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP12(happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4)),X2,'T_0')) )),
  inference(condensation,[],[f12677])).
fof(f12677,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X36)) | p(s(bool,X37)) | ~sP12(happ(s(fun(X38,bool),'_dst_x0_1'(s(bool,X36))),s(X38,X39)),X37,'T_0') | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f12662,f902])).
fof(f12662,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X35)) | ~p(s(bool,X36)) | p(s(bool,X37)) | ~sP12(happ(s(fun(X38,bool),'_dst_x0_1'(s(bool,X36))),s(X38,X39)),X37,'T_0') | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1353,f836])).
fof(f12675,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP12(happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4)),X2,'T_0')) )),
  inference(condensation,[],[f12674])).
fof(f12674,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X30)) | p(s(bool,X31)) | ~sP12(happ(s(fun(X32,bool),'_dst_x0_1'(s(bool,X30))),s(X32,X33)),X31,'T_0') | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f12661,f902])).
fof(f12661,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X29)) | ~p(s(bool,X30)) | p(s(bool,X31)) | ~sP12(happ(s(fun(X32,bool),'_dst_x0_1'(s(bool,X30))),s(X32,X33)),X31,'T_0') | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f1353,f836])).
fof(f12637,plain,(
  ( ! [X4,X2,X3,X1] : (p(s(bool,X1)) | ~sP11('T_0',X1,happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X2))),s(X3,X4))) | ~p(s(bool,X2))) ) | $spl90),
  inference(cnf_transformation,[],[f12637_D])).
fof(f12637_D,plain,(
  ( ! [X4,X2,X3,X1] : (p(s(bool,X1)) | ~sP11('T_0',X1,happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X2))),s(X3,X4))) | ~p(s(bool,X2))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f12639,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP11('T_0',X1,happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X2))),s(X3,X4)))) )),
  inference(condensation,[],[f12638])).
fof(f12638,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (p(s(bool,X36)) | ~p(s(bool,X37)) | ~sP11('T_0',X36,happ(s(fun(X38,bool),'_dst_x0_1'(s(bool,X37))),s(X38,X39))) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f12617,f902])).
fof(f12617,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X35)) | p(s(bool,X36)) | ~p(s(bool,X37)) | ~sP11('T_0',X36,happ(s(fun(X38,bool),'_dst_x0_1'(s(bool,X37))),s(X38,X39))) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1350,f836])).
fof(f12636,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP11('T_0',X1,happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X2))),s(X3,X4)))) )),
  inference(condensation,[],[f12635])).
fof(f12635,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (p(s(bool,X30)) | ~p(s(bool,X31)) | ~sP11('T_0',X30,happ(s(fun(X32,bool),'_dst_x0_1'(s(bool,X31))),s(X32,X33))) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f12616,f902])).
fof(f12616,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X29)) | p(s(bool,X30)) | ~p(s(bool,X31)) | ~sP11('T_0',X30,happ(s(fun(X32,bool),'_dst_x0_1'(s(bool,X31))),s(X32,X33))) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f1350,f836])).
fof(f12576,plain,(
  ( ! [X4,X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | ~sP10(happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X3,X4)),X2,X1)) ) | $spl88),
  inference(cnf_transformation,[],[f12576_D])).
fof(f12576_D,plain,(
  ( ! [X4,X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | ~sP10(happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X3,X4)),X2,X1)) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f12578,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP10(happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X3,X4)),X2,X1)) )),
  inference(condensation,[],[f12577])).
fof(f12577,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (p(s(bool,X36)) | p(s(bool,X37)) | ~sP10(happ(s(fun(X38,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X38,X39)),X37,X36) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f12562,f902])).
fof(f12562,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X35)) | p(s(bool,X36)) | p(s(bool,X37)) | ~sP10(happ(s(fun(X38,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X38,X39)),X37,X36) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1349,f836])).
fof(f1349,plain,(
  ( ! [X101,X99,X102,X100,X98] : (~p(s(bool,X99)) | p(s(bool,X101)) | p(s(bool,X102)) | ~sP10(happ(s(fun(X98,bool),'_dst_x0_1'(s(bool,X99))),s(X98,X100)),X102,X101)) )),
  inference(superposition,[],[f516,f421])).
fof(f12575,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP10(happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X3,X4)),X2,X1)) )),
  inference(condensation,[],[f12574])).
fof(f12574,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (p(s(bool,X30)) | p(s(bool,X31)) | ~sP10(happ(s(fun(X32,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X32,X33)),X31,X30) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f12561,f902])).
fof(f12561,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X29)) | p(s(bool,X30)) | p(s(bool,X31)) | ~sP10(happ(s(fun(X32,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X32,X33)),X31,X30) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f1349,f836])).
fof(f12533,plain,(
  ( ! [X4,X2,X3,X1] : (p(s(bool,X1)) | ~sP9(happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X3,X4)),X1,X2) | p(s(bool,X2))) ) | $spl86),
  inference(cnf_transformation,[],[f12533_D])).
fof(f12533_D,plain,(
  ( ! [X4,X2,X3,X1] : (p(s(bool,X1)) | ~sP9(happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X3,X4)),X1,X2) | p(s(bool,X2))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f12535,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9(happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X3,X4)),X1,X2)) )),
  inference(condensation,[],[f12534])).
fof(f12534,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (p(s(bool,X36)) | p(s(bool,X37)) | ~sP9(happ(s(fun(X38,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X38,X39)),X36,X37) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f12509,f902])).
fof(f12509,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X35)) | p(s(bool,X36)) | p(s(bool,X37)) | ~sP9(happ(s(fun(X38,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X38,X39)),X36,X37) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1346,f836])).
fof(f1346,plain,(
  ( ! [X87,X85,X83,X86,X84] : (~p(s(bool,X84)) | p(s(bool,X86)) | p(s(bool,X87)) | ~sP9(happ(s(fun(X83,bool),'_dst_x0_1'(s(bool,X84))),s(X83,X85)),X86,X87)) )),
  inference(superposition,[],[f509,f421])).
fof(f12532,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9(happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X3,X4)),X1,X2)) )),
  inference(condensation,[],[f12531])).
fof(f12531,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (p(s(bool,X30)) | p(s(bool,X31)) | ~sP9(happ(s(fun(X32,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X32,X33)),X30,X31) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f12508,f902])).
fof(f12508,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X29)) | p(s(bool,X30)) | p(s(bool,X31)) | ~sP9(happ(s(fun(X32,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X32,X33)),X30,X31) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f1346,f836])).
fof(f12477,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,X1)) | p(s(bool,X2)) | ~sP8(X2,'T_0',happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4)))) ) | $spl84),
  inference(cnf_transformation,[],[f12477_D])).
fof(f12477_D,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,X1)) | p(s(bool,X2)) | ~sP8(X2,'T_0',happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4)))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f12479,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP8(X2,'T_0',happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4)))) )),
  inference(condensation,[],[f12478])).
fof(f12478,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X36)) | p(s(bool,X37)) | ~sP8(X37,'T_0',happ(s(fun(X38,bool),'_dst_x0_1'(s(bool,X36))),s(X38,X39))) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f12462,f902])).
fof(f12462,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X35)) | ~p(s(bool,X36)) | p(s(bool,X37)) | ~sP8(X37,'T_0',happ(s(fun(X38,bool),'_dst_x0_1'(s(bool,X36))),s(X38,X39))) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1345,f836])).
fof(f12476,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP8(X2,'T_0',happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4)))) )),
  inference(condensation,[],[f12475])).
fof(f12475,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X30)) | p(s(bool,X31)) | ~sP8(X31,'T_0',happ(s(fun(X32,bool),'_dst_x0_1'(s(bool,X30))),s(X32,X33))) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f12461,f902])).
fof(f12461,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X29)) | ~p(s(bool,X30)) | p(s(bool,X31)) | ~sP8(X31,'T_0',happ(s(fun(X32,bool),'_dst_x0_1'(s(bool,X30))),s(X32,X33))) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f1345,f836])).
fof(f12438,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,X1)) | ~sP7('T_0',happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4)),X2) | p(s(bool,X2))) ) | $spl82),
  inference(cnf_transformation,[],[f12438_D])).
fof(f12438_D,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,X1)) | ~sP7('T_0',happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4)),X2) | p(s(bool,X2))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f12440,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7('T_0',happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4)),X2)) )),
  inference(condensation,[],[f12439])).
fof(f12439,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X36)) | p(s(bool,X37)) | ~sP7('T_0',happ(s(fun(X38,bool),'_dst_x0_1'(s(bool,X36))),s(X38,X39)),X37) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f12418,f902])).
fof(f12418,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X35)) | ~p(s(bool,X36)) | p(s(bool,X37)) | ~sP7('T_0',happ(s(fun(X38,bool),'_dst_x0_1'(s(bool,X36))),s(X38,X39)),X37) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1344,f836])).
fof(f12437,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7('T_0',happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4)),X2)) )),
  inference(condensation,[],[f12436])).
fof(f12436,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X30)) | p(s(bool,X31)) | ~sP7('T_0',happ(s(fun(X32,bool),'_dst_x0_1'(s(bool,X30))),s(X32,X33)),X31) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f12417,f902])).
fof(f12417,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X29)) | ~p(s(bool,X30)) | p(s(bool,X31)) | ~sP7('T_0',happ(s(fun(X32,bool),'_dst_x0_1'(s(bool,X30))),s(X32,X33)),X31) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f1344,f836])).
fof(f12386,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,X1)) | ~sP4('T_0',happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4)),X2) | p(s(bool,X2))) ) | $spl80),
  inference(cnf_transformation,[],[f12386_D])).
fof(f12386_D,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,X1)) | ~sP4('T_0',happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4)),X2) | p(s(bool,X2))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f12388,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4('T_0',happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4)),X2)) )),
  inference(condensation,[],[f12387])).
fof(f12387,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X36)) | p(s(bool,X37)) | ~sP4('T_0',happ(s(fun(X38,bool),'_dst_x0_1'(s(bool,X36))),s(X38,X39)),X37) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f12366,f902])).
fof(f12366,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X35)) | ~p(s(bool,X36)) | p(s(bool,X37)) | ~sP4('T_0',happ(s(fun(X38,bool),'_dst_x0_1'(s(bool,X36))),s(X38,X39)),X37) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1343,f836])).
fof(f12385,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4('T_0',happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4)),X2)) )),
  inference(condensation,[],[f12384])).
fof(f12384,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X30)) | p(s(bool,X31)) | ~sP4('T_0',happ(s(fun(X32,bool),'_dst_x0_1'(s(bool,X30))),s(X32,X33)),X31) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f12365,f902])).
fof(f12365,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X29)) | ~p(s(bool,X30)) | p(s(bool,X31)) | ~sP4('T_0',happ(s(fun(X32,bool),'_dst_x0_1'(s(bool,X30))),s(X32,X33)),X31) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f1343,f836])).
fof(f12334,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,X1)) | ~sP5('T_0',happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4)),X2) | p(s(bool,X2))) ) | $spl78),
  inference(cnf_transformation,[],[f12334_D])).
fof(f12334_D,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,X1)) | ~sP5('T_0',happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4)),X2) | p(s(bool,X2))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f12336,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP5('T_0',happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4)),X2)) )),
  inference(condensation,[],[f12335])).
fof(f12335,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X36)) | p(s(bool,X37)) | ~sP5('T_0',happ(s(fun(X38,bool),'_dst_x0_1'(s(bool,X36))),s(X38,X39)),X37) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f12314,f902])).
fof(f12314,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X35)) | ~p(s(bool,X36)) | p(s(bool,X37)) | ~sP5('T_0',happ(s(fun(X38,bool),'_dst_x0_1'(s(bool,X36))),s(X38,X39)),X37) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1342,f836])).
fof(f12333,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP5('T_0',happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4)),X2)) )),
  inference(condensation,[],[f12332])).
fof(f12332,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X30)) | p(s(bool,X31)) | ~sP5('T_0',happ(s(fun(X32,bool),'_dst_x0_1'(s(bool,X30))),s(X32,X33)),X31) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f12313,f902])).
fof(f12313,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X29)) | ~p(s(bool,X30)) | p(s(bool,X31)) | ~sP5('T_0',happ(s(fun(X32,bool),'_dst_x0_1'(s(bool,X30))),s(X32,X33)),X31) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f1342,f836])).
fof(f12287,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,X1)) | ~sP6('T_0',happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4)),X2) | p(s(bool,X2))) ) | $spl76),
  inference(cnf_transformation,[],[f12287_D])).
fof(f12287_D,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,X1)) | ~sP6('T_0',happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4)),X2) | p(s(bool,X2))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f12289,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP6('T_0',happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4)),X2)) )),
  inference(condensation,[],[f12288])).
fof(f12288,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X36)) | p(s(bool,X37)) | ~sP6('T_0',happ(s(fun(X38,bool),'_dst_x0_1'(s(bool,X36))),s(X38,X39)),X37) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f12270,f902])).
fof(f12270,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X35)) | ~p(s(bool,X36)) | p(s(bool,X37)) | ~sP6('T_0',happ(s(fun(X38,bool),'_dst_x0_1'(s(bool,X36))),s(X38,X39)),X37) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1341,f836])).
fof(f12286,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP6('T_0',happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4)),X2)) )),
  inference(condensation,[],[f12285])).
fof(f12285,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X30)) | p(s(bool,X31)) | ~sP6('T_0',happ(s(fun(X32,bool),'_dst_x0_1'(s(bool,X30))),s(X32,X33)),X31) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f12269,f902])).
fof(f12269,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X29)) | ~p(s(bool,X30)) | p(s(bool,X31)) | ~sP6('T_0',happ(s(fun(X32,bool),'_dst_x0_1'(s(bool,X30))),s(X32,X33)),X31) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f1341,f836])).
fof(f9994,plain,(
  $false | ($spl3 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f9993,f322])).
fof(f9993,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f9984,f9540])).
fof(f9984,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ($spl3 | $spl6)),
  inference(unit_resulting_resolution,[],[f569,f572])).
fof(f566,plain,(
  p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | $spl0),
  inference(cnf_transformation,[],[f566_D])).
fof(f566_D,plain,(
  p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f9090,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK20(sK16,sK14))))) | $spl74),
  inference(cnf_transformation,[],[f9090_D])).
fof(f9090_D,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK20(sK16,sK14))))) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f7464,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK20(sK16,sK14))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16))))),
  inference(subsumption_resolution,[],[f7377,f322])).
fof(f7377,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK20(sK16,sK14))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16))))),
  inference(superposition,[],[f370,f3354])).
fof(f370,plain,(
  ( ! [X4,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK20(X1,X0))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f229])).
fof(f9088,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK19(sK16,sK14))))) | $spl72),
  inference(cnf_transformation,[],[f9088_D])).
fof(f9088_D,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK19(sK16,sK14))))) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f7467,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK19(sK16,sK14))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16))))),
  inference(subsumption_resolution,[],[f7378,f322])).
fof(f7378,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK19(sK16,sK14))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16))))),
  inference(superposition,[],[f369,f3354])).
fof(f369,plain,(
  ( ! [X4,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK19(X1,X0))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f229])).
fof(f9086,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8159,f8917])).
fof(f8917,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl68),
  inference(unit_resulting_resolution,[],[f3575,f8141,f1257])).
fof(f8159,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,sK14),X1),s(fun(sK14,bool),sK16))))) ) | $spl68),
  inference(unit_resulting_resolution,[],[f8141,f437])).
fof(f9085,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8161,f8917])).
fof(f8161,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X0,sK14),X1),s(fun(sK14,bool),sK16))))) ) | $spl68),
  inference(unit_resulting_resolution,[],[f8141,f439])).
fof(f9084,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8163,f8917])).
fof(f8163,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X0,sK14),X1),s(fun(sK14,bool),sK16))))) ) | $spl68),
  inference(unit_resulting_resolution,[],[f8141,f427])).
fof(f9083,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8164,f8917])).
fof(f8164,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,sK14),X1),s(fun(sK14,bool),sK16))))) ) | $spl68),
  inference(unit_resulting_resolution,[],[f8141,f430])).
fof(f9080,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8176,f8917])).
fof(f8176,plain,(
  p(s(bool,'F_0')) | $spl68),
  inference(unit_resulting_resolution,[],[f3362,f8141,f407])).
fof(f9079,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8180,f8917])).
fof(f8180,plain,(
  p(s(bool,'F_0')) | $spl68),
  inference(unit_resulting_resolution,[],[f3363,f8141,f407])).
fof(f9078,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8190,f8917])).
fof(f8190,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl68),
  inference(unit_resulting_resolution,[],[f3604,f8141,f446])).
fof(f3604,plain,(
  ( ! [X0,X1] : (sP2(X0,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X1)) )),
  inference(unit_resulting_resolution,[],[f2082,f559])).
fof(f9077,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8193,f8917])).
fof(f8193,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) | $spl68),
  inference(unit_resulting_resolution,[],[f712,f8141,f446])).
fof(f712,plain,(
  ( ! [X0] : (sP2('T_0',X0,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) )),
  inference(unit_resulting_resolution,[],[f322,f707])).
fof(f9076,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8194,f8917])).
fof(f8194,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | $spl68),
  inference(unit_resulting_resolution,[],[f3383,f8141,f446])).
fof(f3383,plain,(
  ( ! [X0] : (sP2('T_0',X0,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) )),
  inference(unit_resulting_resolution,[],[f322,f2082,f451])).
fof(f9075,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8200,f8917])).
fof(f8200,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) | $spl68),
  inference(unit_resulting_resolution,[],[f616,f8141,f446])).
fof(f616,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X1)) )),
  inference(unit_resulting_resolution,[],[f318,f559])).
fof(f9074,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8201,f8917])).
fof(f8201,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | $spl68),
  inference(unit_resulting_resolution,[],[f3604,f8141,f446])).
fof(f9073,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8203,f8917])).
fof(f8203,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl68),
  inference(unit_resulting_resolution,[],[f3383,f8141,f446])).
fof(f9072,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8207,f8917])).
fof(f8207,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl68),
  inference(unit_resulting_resolution,[],[f3378,f8141,f447])).
fof(f3378,plain,(
  ( ! [X0] : (sP2(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f322,f2082,f451])).
fof(f9071,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8210,f8917])).
fof(f8210,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) | $spl68),
  inference(unit_resulting_resolution,[],[f616,f8141,f447])).
fof(f9070,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8211,f8917])).
fof(f8211,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | $spl68),
  inference(unit_resulting_resolution,[],[f3604,f8141,f447])).
fof(f9069,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8215,f8917])).
fof(f8215,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl68),
  inference(unit_resulting_resolution,[],[f3604,f8141,f447])).
fof(f9068,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8220,f8917])).
fof(f8220,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | $spl68),
  inference(unit_resulting_resolution,[],[f3378,f8141,f447])).
fof(f9067,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8228,f3411])).
fof(f3411,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) )),
  inference(unit_resulting_resolution,[],[f2082,f470])).
fof(f8228,plain,(
  ~sP5('T_0','F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15))) | $spl68),
  inference(unit_resulting_resolution,[],[f3428,f583,f3390,f8141,f467])).
fof(f9066,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8256,f582])).
fof(f8256,plain,(
  ~sP4(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0','T_0') | $spl68),
  inference(unit_resulting_resolution,[],[f3428,f3390,f3411,f8141,f467])).
fof(f9065,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8257,f587])).
fof(f8257,plain,(
  ~sP3('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'T_0') | $spl68),
  inference(unit_resulting_resolution,[],[f583,f3390,f3411,f8141,f467])).
fof(f9064,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8260,f8407])).
fof(f8407,plain,(
  ( ! [X2,X0,X1] : (sP4(happ(s(fun(sK14,bool),sK16),s(sK14,X0)),X1,X2)) ) | $spl68),
  inference(unit_resulting_resolution,[],[f8141,f476])).
fof(f8260,plain,(
  ( ! [X0,X1] : (~sP4(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15))))),s(X0,X1)),'F_0')) ) | $spl68),
  inference(unit_resulting_resolution,[],[f3428,f1313,f3411,f8141,f467])).
fof(f9063,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8271,f582])).
fof(f8271,plain,(
  ~sP4(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0','T_0') | $spl68),
  inference(unit_resulting_resolution,[],[f3428,f3390,f578,f8141,f467])).
fof(f9062,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8272,f587])).
fof(f8272,plain,(
  ~sP3('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'T_0') | $spl68),
  inference(unit_resulting_resolution,[],[f583,f3390,f578,f8141,f467])).
fof(f9061,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8300,f8407])).
fof(f8300,plain,(
  ( ! [X0,X1] : (~sP4(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15))))),s(X0,X1)),'F_0')) ) | $spl68),
  inference(unit_resulting_resolution,[],[f3428,f1313,f580,f8141,f467])).
fof(f9060,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8307,f577])).
fof(f8307,plain,(
  ~sP5(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0','T_0') | $spl68),
  inference(unit_resulting_resolution,[],[f588,f583,f3443,f8141,f468])).
fof(f9059,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8320,f3426])).
fof(f8320,plain,(
  ~sP4('T_0','F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15))) | $spl68),
  inference(unit_resulting_resolution,[],[f588,f3443,f577,f8141,f468])).
fof(f9058,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8321,f3427])).
fof(f8321,plain,(
  ~sP3('F_0','T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15))) | $spl68),
  inference(unit_resulting_resolution,[],[f583,f3443,f577,f8141,f468])).
fof(f9057,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8351,f3609])).
fof(f8351,plain,(
  sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'F_0') | $spl68),
  inference(unit_resulting_resolution,[],[f610,f3426,f578,f8141,f468])).
fof(f9056,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8352,f3391])).
fof(f8352,plain,(
  sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0') | $spl68),
  inference(unit_resulting_resolution,[],[f3428,f3426,f578,f8141,f468])).
fof(f9055,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8356,f3426])).
fof(f8356,plain,(
  ~sP4('T_0','F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15))) | $spl68),
  inference(unit_resulting_resolution,[],[f588,f3443,f578,f8141,f468])).
fof(f9054,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8357,f3427])).
fof(f8357,plain,(
  ~sP3('F_0','T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15))) | $spl68),
  inference(unit_resulting_resolution,[],[f583,f3443,f578,f8141,f468])).
fof(f9053,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8413,f8917])).
fof(f8413,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl68),
  inference(unit_resulting_resolution,[],[f573,f3429,f8141,f477])).
fof(f9052,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8414,f8917])).
fof(f8414,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl68),
  inference(unit_resulting_resolution,[],[f8141,f3429,f8141,f477])).
fof(f9051,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8415,f8917])).
fof(f8415,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl68),
  inference(unit_resulting_resolution,[],[f573,f3429,f8141,f477])).
fof(f9050,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8417,f8917])).
fof(f8417,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl68),
  inference(unit_resulting_resolution,[],[f8141,f3429,f8141,f477])).
fof(f9049,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8420,f8917])).
fof(f8420,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) | $spl68),
  inference(unit_resulting_resolution,[],[f573,f609,f8141,f477])).
fof(f9048,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8421,f8917])).
fof(f8421,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | $spl68),
  inference(unit_resulting_resolution,[],[f573,f3427,f8141,f477])).
fof(f9047,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8422,f8917])).
fof(f8422,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl68),
  inference(unit_resulting_resolution,[],[f8141,f3427,f8141,f477])).
fof(f9046,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8425,f8917])).
fof(f8425,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) | $spl68),
  inference(unit_resulting_resolution,[],[f573,f610,f8141,f477])).
fof(f9045,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8426,f8917])).
fof(f8426,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | $spl68),
  inference(unit_resulting_resolution,[],[f573,f3428,f8141,f477])).
fof(f9044,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8427,f8917])).
fof(f8427,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl68),
  inference(unit_resulting_resolution,[],[f8141,f3428,f8141,f477])).
fof(f9043,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8433,f8917])).
fof(f8433,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl68),
  inference(unit_resulting_resolution,[],[f573,f3428,f8141,f477])).
fof(f9042,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8434,f8917])).
fof(f8434,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl68),
  inference(unit_resulting_resolution,[],[f8141,f3428,f8141,f477])).
fof(f9041,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8435,f8917])).
fof(f8435,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl68),
  inference(unit_resulting_resolution,[],[f573,f3428,f8141,f477])).
fof(f9040,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8437,f8917])).
fof(f8437,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl68),
  inference(unit_resulting_resolution,[],[f8141,f3428,f8141,f477])).
fof(f9039,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8440,f8917])).
fof(f8440,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) | $spl68),
  inference(unit_resulting_resolution,[],[f573,f609,f8141,f477])).
fof(f9038,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8441,f8917])).
fof(f8441,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | $spl68),
  inference(unit_resulting_resolution,[],[f573,f3427,f8141,f477])).
fof(f9037,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8442,f8917])).
fof(f8442,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl68),
  inference(unit_resulting_resolution,[],[f8141,f3427,f8141,f477])).
fof(f9036,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8445,f8917])).
fof(f8445,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) | $spl68),
  inference(unit_resulting_resolution,[],[f573,f611,f8141,f477])).
fof(f611,plain,(
  ( ! [X0,X1] : (sP3('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f318,f480])).
fof(f9035,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8446,f8917])).
fof(f8446,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | $spl68),
  inference(unit_resulting_resolution,[],[f573,f3429,f8141,f477])).
fof(f9034,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8447,f8917])).
fof(f8447,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl68),
  inference(unit_resulting_resolution,[],[f8141,f3429,f8141,f477])).
fof(f9033,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8453,f8917])).
fof(f8453,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl68),
  inference(unit_resulting_resolution,[],[f573,f3427,f8141,f477])).
fof(f9032,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8454,f8917])).
fof(f8454,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl68),
  inference(unit_resulting_resolution,[],[f8141,f3427,f8141,f477])).
fof(f9031,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8455,f8917])).
fof(f8455,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl68),
  inference(unit_resulting_resolution,[],[f573,f3427,f8141,f477])).
fof(f9030,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8457,f8917])).
fof(f8457,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl68),
  inference(unit_resulting_resolution,[],[f8141,f3427,f8141,f477])).
fof(f9029,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8460,f8917])).
fof(f8460,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) | $spl68),
  inference(unit_resulting_resolution,[],[f573,f610,f8141,f477])).
fof(f9028,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8461,f8917])).
fof(f8461,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | $spl68),
  inference(unit_resulting_resolution,[],[f573,f3428,f8141,f477])).
fof(f9027,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8462,f8917])).
fof(f8462,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl68),
  inference(unit_resulting_resolution,[],[f8141,f3428,f8141,f477])).
fof(f9026,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8465,f8917])).
fof(f8465,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) | $spl68),
  inference(unit_resulting_resolution,[],[f573,f611,f8141,f477])).
fof(f9025,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8466,f8917])).
fof(f8466,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | $spl68),
  inference(unit_resulting_resolution,[],[f573,f3429,f8141,f477])).
fof(f9024,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8467,f8917])).
fof(f8467,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl68),
  inference(unit_resulting_resolution,[],[f8141,f3429,f8141,f477])).
fof(f9023,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8475,f617])).
fof(f8475,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl68),
  inference(unit_resulting_resolution,[],[f3443,f8141,f481])).
fof(f9022,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8477,f617])).
fof(f8477,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl68),
  inference(unit_resulting_resolution,[],[f3446,f8141,f481])).
fof(f3446,plain,(
  ~sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f617,f2082,f483])).
fof(f9021,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8531,f8917])).
fof(f8531,plain,(
  p(s(bool,'F_0')) | $spl68),
  inference(unit_resulting_resolution,[],[f3498,f8141,f496])).
fof(f9020,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8533,f8917])).
fof(f8533,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | $spl68),
  inference(unit_resulting_resolution,[],[f3449,f8141,f496])).
fof(f9019,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8537,f8917])).
fof(f8537,plain,(
  p(s(bool,'F_0')) | $spl68),
  inference(unit_resulting_resolution,[],[f3448,f8141,f496])).
fof(f9018,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8539,f8917])).
fof(f8539,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl68),
  inference(unit_resulting_resolution,[],[f3449,f8141,f496])).
fof(f9017,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8545,f8917])).
fof(f8545,plain,(
  p(s(bool,'F_0')) | $spl68),
  inference(unit_resulting_resolution,[],[f3493,f8141,f497])).
fof(f9016,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8547,f8917])).
fof(f8547,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | $spl68),
  inference(unit_resulting_resolution,[],[f3448,f8141,f497])).
fof(f9015,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8551,f8917])).
fof(f8551,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl68),
  inference(unit_resulting_resolution,[],[f3448,f8141,f497])).
fof(f9014,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8553,f8917])).
fof(f8553,plain,(
  p(s(bool,'F_0')) | $spl68),
  inference(unit_resulting_resolution,[],[f3449,f8141,f497])).
fof(f9013,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8601,f3521])).
fof(f8601,plain,(
  ~sP9('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0') | $spl68),
  inference(unit_resulting_resolution,[],[f573,f3507,f8141,f508])).
fof(f9012,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8602,f3521])).
fof(f8602,plain,(
  ( ! [X0] : (~sP9('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),happ(s(fun(sK14,bool),sK16),s(sK14,X0)))) ) | $spl68),
  inference(unit_resulting_resolution,[],[f8141,f3507,f8141,f508])).
fof(f9011,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8603,f3520])).
fof(f8603,plain,(
  ( ! [X0] : (~sP9('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,X0)),happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) ) | $spl68),
  inference(unit_resulting_resolution,[],[f8141,f3503,f8141,f508])).
fof(f9010,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8606,f8917])).
fof(f8606,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl68),
  inference(unit_resulting_resolution,[],[f3507,f3521,f8141,f508])).
fof(f9009,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8609,f8917])).
fof(f8609,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | $spl68),
  inference(unit_resulting_resolution,[],[f3503,f3521,f8141,f508])).
fof(f9008,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8613,f8917])).
fof(f8613,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) | $spl68),
  inference(unit_resulting_resolution,[],[f3507,f613,f8141,f508])).
fof(f9007,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8615,f8917])).
fof(f8615,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | $spl68),
  inference(unit_resulting_resolution,[],[f3507,f3520,f8141,f508])).
fof(f9006,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8619,f8917])).
fof(f8619,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl68),
  inference(unit_resulting_resolution,[],[f3507,f597,f8141,f508])).
fof(f9005,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8622,f8917])).
fof(f8622,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | $spl68),
  inference(unit_resulting_resolution,[],[f3503,f597,f8141,f508])).
fof(f9004,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8629,f3520])).
fof(f8629,plain,(
  ~sP9('F_0','F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15))) | $spl68),
  inference(unit_resulting_resolution,[],[f573,f3503,f8141,f508])).
fof(f9003,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8630,f3520])).
fof(f8630,plain,(
  ( ! [X0] : (~sP9('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,X0)),happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) ) | $spl68),
  inference(unit_resulting_resolution,[],[f8141,f3503,f8141,f508])).
fof(f9002,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8631,f3521])).
fof(f8631,plain,(
  ( ! [X0] : (~sP9('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),happ(s(fun(sK14,bool),sK16),s(sK14,X0)))) ) | $spl68),
  inference(unit_resulting_resolution,[],[f8141,f3507,f8141,f508])).
fof(f9001,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8634,f8917])).
fof(f8634,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl68),
  inference(unit_resulting_resolution,[],[f3503,f3520,f8141,f508])).
fof(f9000,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8637,f8917])).
fof(f8637,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | $spl68),
  inference(unit_resulting_resolution,[],[f3507,f3520,f8141,f508])).
fof(f8999,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8641,f8917])).
fof(f8641,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) | $spl68),
  inference(unit_resulting_resolution,[],[f3503,f614,f8141,f508])).
fof(f8998,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8643,f8917])).
fof(f8643,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | $spl68),
  inference(unit_resulting_resolution,[],[f3507,f3521,f8141,f508])).
fof(f8997,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8647,f8917])).
fof(f8647,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl68),
  inference(unit_resulting_resolution,[],[f3503,f597,f8141,f508])).
fof(f8996,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8650,f8917])).
fof(f8650,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | $spl68),
  inference(unit_resulting_resolution,[],[f3507,f597,f8141,f508])).
fof(f8995,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8659,f8917])).
fof(f8659,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl68),
  inference(unit_resulting_resolution,[],[f322,f3521,f8141,f509])).
fof(f8994,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8665,f8917])).
fof(f8665,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) | $spl68),
  inference(unit_resulting_resolution,[],[f322,f613,f8141,f509])).
fof(f8993,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8666,f8917])).
fof(f8666,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | $spl68),
  inference(unit_resulting_resolution,[],[f322,f3520,f8141,f509])).
fof(f8992,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8676,f8917])).
fof(f8676,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl68),
  inference(unit_resulting_resolution,[],[f322,f3520,f8141,f509])).
fof(f8991,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8679,f8917])).
fof(f8679,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) | $spl68),
  inference(unit_resulting_resolution,[],[f322,f614,f8141,f509])).
fof(f8990,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8680,f8917])).
fof(f8680,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | $spl68),
  inference(unit_resulting_resolution,[],[f322,f3521,f8141,f509])).
fof(f8989,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8692,f8917])).
fof(f8692,plain,(
  p(s(bool,'F_0')) | $spl68),
  inference(unit_resulting_resolution,[],[f573,f3507,f8141,f513])).
fof(f513,plain,(
  ( ! [X2,X0,X1] : (sP10(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f307])).
fof(f8988,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8693,f8917])).
fof(f8693,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl68),
  inference(unit_resulting_resolution,[],[f573,f3507,f8141,f513])).
fof(f8987,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8696,f8917])).
fof(f8696,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | $spl68),
  inference(unit_resulting_resolution,[],[f573,f3503,f8141,f513])).
fof(f8986,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8697,f8917])).
fof(f8697,plain,(
  p(s(bool,'F_0')) | $spl68),
  inference(unit_resulting_resolution,[],[f8141,f3503,f8141,f513])).
fof(f8985,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8703,f8917])).
fof(f8703,plain,(
  p(s(bool,'F_0')) | $spl68),
  inference(unit_resulting_resolution,[],[f573,f3503,f8141,f513])).
fof(f8984,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8704,f8917])).
fof(f8704,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl68),
  inference(unit_resulting_resolution,[],[f573,f3503,f8141,f513])).
fof(f8983,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8707,f8917])).
fof(f8707,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | $spl68),
  inference(unit_resulting_resolution,[],[f573,f3507,f8141,f513])).
fof(f8982,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8708,f8917])).
fof(f8708,plain,(
  p(s(bool,'F_0')) | $spl68),
  inference(unit_resulting_resolution,[],[f8141,f3507,f8141,f513])).
fof(f8981,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8714,f8917])).
fof(f8714,plain,(
  p(s(bool,'F_0')) | $spl68),
  inference(unit_resulting_resolution,[],[f573,f3534,f8141,f513])).
fof(f8980,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8715,f8917])).
fof(f8715,plain,(
  p(s(bool,'F_0')) | $spl68),
  inference(unit_resulting_resolution,[],[f573,f3534,f8141,f513])).
fof(f8979,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8720,f8917])).
fof(f8720,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl68),
  inference(unit_resulting_resolution,[],[f322,f3528,f8141,f516])).
fof(f3528,plain,(
  ( ! [X0] : (sP10('T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X0)) )),
  inference(unit_resulting_resolution,[],[f322,f2082,f515])).
fof(f8978,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8723,f8917])).
fof(f8723,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) | $spl68),
  inference(unit_resulting_resolution,[],[f322,f813,f8141,f516])).
fof(f813,plain,(
  ( ! [X0] : (sP10('T_0',X0,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) )),
  inference(unit_resulting_resolution,[],[f613,f322,f557])).
fof(f8977,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8724,f8917])).
fof(f8724,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | $spl68),
  inference(unit_resulting_resolution,[],[f322,f3522,f8141,f516])).
fof(f3522,plain,(
  ( ! [X0] : (sP10('T_0',X0,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) )),
  inference(unit_resulting_resolution,[],[f322,f2082,f514])).
fof(f8976,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8730,f8917])).
fof(f8730,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl68),
  inference(unit_resulting_resolution,[],[f322,f3522,f8141,f516])).
fof(f8975,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8733,f8917])).
fof(f8733,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) | $spl68),
  inference(unit_resulting_resolution,[],[f322,f815,f8141,f516])).
fof(f815,plain,(
  ( ! [X0] : (sP10('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X0)) )),
  inference(unit_resulting_resolution,[],[f614,f322,f557])).
fof(f8974,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8734,f8917])).
fof(f8734,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | $spl68),
  inference(unit_resulting_resolution,[],[f322,f3528,f8141,f516])).
fof(f8973,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8739,f8917])).
fof(f8739,plain,(
  p(s(bool,'F_0')) | $spl68),
  inference(unit_resulting_resolution,[],[f3595,f8141,f519])).
fof(f8972,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8741,f8917])).
fof(f8741,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl68),
  inference(unit_resulting_resolution,[],[f3598,f8141,f519])).
fof(f3598,plain,(
  ( ! [X0] : (sP12(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f601,f2082,f556])).
fof(f8971,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8748,f8917])).
fof(f8748,plain,(
  p(s(bool,'F_0')) | $spl68),
  inference(unit_resulting_resolution,[],[f3576,f8141,f519])).
fof(f8970,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8750,f8917])).
fof(f8750,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) | $spl68),
  inference(unit_resulting_resolution,[],[f884,f8141,f519])).
fof(f884,plain,(
  ( ! [X0] : (sP12('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f322,f778])).
fof(f778,plain,(
  ( ! [X2,X3] : (sP12('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X2,X3) | ~p(s(bool,X2))) )),
  inference(subsumption_resolution,[],[f777,f322])).
fof(f777,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | ~p(s(bool,X2)) | sP12('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X2,X3)) )),
  inference(superposition,[],[f532,f621])).
fof(f8969,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8751,f8917])).
fof(f8751,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | $spl68),
  inference(unit_resulting_resolution,[],[f3598,f8141,f519])).
fof(f8968,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8776,f3575])).
fof(f8776,plain,(
  ~sP11('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'T_0') | $spl68),
  inference(unit_resulting_resolution,[],[f322,f3535,f8141,f525])).
fof(f8967,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8812,f8917])).
fof(f8812,plain,(
  p(s(bool,'F_0')) | $spl68),
  inference(unit_resulting_resolution,[],[f322,f3535,f8141,f530])).
fof(f8966,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8820,f8917])).
fof(f8820,plain,(
  p(s(bool,'F_0')) | $spl68),
  inference(unit_resulting_resolution,[],[f322,f3590,f8141,f530])).
fof(f8965,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8848,f3476])).
fof(f8848,plain,(
  ( ! [X0] : (~sP7(X0,'F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) ) | $spl68),
  inference(unit_resulting_resolution,[],[f3448,f8141,f558])).
fof(f558,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0)) | ~sP7(X1,X2,X0)) )),
  inference(duplicate_literal_removal,[],[f491])).
fof(f491,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X0)) | ~sP7(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f295])).
fof(f8964,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8850,f3476])).
fof(f8850,plain,(
  ( ! [X0] : (~sP7('F_0',X0,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) ) | $spl68),
  inference(unit_resulting_resolution,[],[f3449,f8141,f558])).
fof(f8963,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8854,f8917])).
fof(f8854,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl68),
  inference(unit_resulting_resolution,[],[f3357,f8141,f574])).
fof(f3357,plain,(
  ( ! [X0] : (sP0(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X0)) )),
  inference(unit_resulting_resolution,[],[f2082,f394])).
fof(f8962,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8857,f8917])).
fof(f8857,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) | $spl68),
  inference(unit_resulting_resolution,[],[f605,f8141,f574])).
fof(f605,plain,(
  ( ! [X0] : (sP0(X0,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) )),
  inference(unit_resulting_resolution,[],[f318,f393])).
fof(f8961,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8858,f8917])).
fof(f8858,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | $spl68),
  inference(unit_resulting_resolution,[],[f3356,f8141,f574])).
fof(f3356,plain,(
  ( ! [X0] : (sP0(X0,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) )),
  inference(unit_resulting_resolution,[],[f2082,f393])).
fof(f8960,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8862,f8917])).
fof(f8862,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl68),
  inference(unit_resulting_resolution,[],[f3356,f8141,f574])).
fof(f8959,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8865,f8917])).
fof(f8865,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) | $spl68),
  inference(unit_resulting_resolution,[],[f606,f8141,f574])).
fof(f606,plain,(
  ( ! [X0] : (sP0('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X0)) )),
  inference(unit_resulting_resolution,[],[f318,f394])).
fof(f8958,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8866,f8917])).
fof(f8866,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | $spl68),
  inference(unit_resulting_resolution,[],[f3357,f8141,f574])).
fof(f8956,plain,(
  $false | ($spl1 | $spl11 | $spl68)),
  inference(subsumption_resolution,[],[f8955,f8917])).
fof(f8955,plain,(
  p(s(bool,'F_0')) | ($spl1 | $spl11 | $spl68)),
  inference(backward_demodulation,[],[f8950,f7469])).
fof(f7469,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK19(sK16,sK14))))) | $spl1),
  inference(subsumption_resolution,[],[f7468,f573])).
fof(f7468,plain,(
  p(s(bool,'F_0')) | p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK19(sK16,sK14))))) | $spl1),
  inference(forward_demodulation,[],[f7467,f627])).
fof(f627,plain,(
  s(bool,'F_0') = s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | $spl1),
  inference(unit_resulting_resolution,[],[f567,f341])).
fof(f341,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f211])).
fof(f8950,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,X0)))) ) | ($spl11 | $spl68)),
  inference(forward_demodulation,[],[f8873,f1572])).
fof(f8873,plain,(
  ( ! [X0] : (s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16))) = s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,X0)))) ) | ($spl11 | $spl68)),
  inference(unit_resulting_resolution,[],[f1777,f8141,f902])).
fof(f8949,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8896,f8917])).
fof(f8896,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl68),
  inference(unit_resulting_resolution,[],[f3520,f8141,f1160])).
fof(f8948,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8899,f8917])).
fof(f8899,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) | $spl68),
  inference(unit_resulting_resolution,[],[f614,f8141,f1160])).
fof(f8947,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8900,f8917])).
fof(f8900,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | $spl68),
  inference(unit_resulting_resolution,[],[f3521,f8141,f1160])).
fof(f8946,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8904,f8917])).
fof(f8904,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl68),
  inference(unit_resulting_resolution,[],[f3521,f8141,f1160])).
fof(f8945,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8907,f8917])).
fof(f8907,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) | $spl68),
  inference(unit_resulting_resolution,[],[f613,f8141,f1160])).
fof(f8944,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f8908,f8917])).
fof(f8908,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | $spl68),
  inference(unit_resulting_resolution,[],[f3520,f8141,f1160])).
fof(f8930,plain,(
  $false | ($spl16 | $spl68)),
  inference(unit_resulting_resolution,[],[f3358,f8141,f4574])).
fof(f3358,plain,(
  ~sP1(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'T_0')),
  inference(unit_resulting_resolution,[],[f322,f2082,f408])).
fof(f8928,plain,(
  $false | ($spl12 | $spl68)),
  inference(unit_resulting_resolution,[],[f3360,f8141,f4570])).
fof(f3360,plain,(
  ~sP1('T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))),
  inference(unit_resulting_resolution,[],[f322,f2082,f408])).
fof(f8918,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f322,f3575,f8141,f1257])).
fof(f8909,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f8141,f3520,f8141,f1160])).
fof(f8905,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f573,f3521,f8141,f1160])).
fof(f8901,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f8141,f3521,f8141,f1160])).
fof(f8897,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f573,f3520,f8141,f1160])).
fof(f8889,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f322,f3476,f8141,f1076])).
fof(f8883,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f322,f3426,f8141,f1032])).
fof(f8877,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f322,f3411,f8141,f1011])).
fof(f8867,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f8141,f3357,f8141,f574])).
fof(f8863,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f573,f3356,f8141,f574])).
fof(f8859,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f8141,f3356,f8141,f574])).
fof(f8855,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f573,f3357,f8141,f574])).
fof(f8851,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f593,f3449,f8141,f558])).
fof(f8849,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f591,f3448,f8141,f558])).
fof(f8837,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f322,f3354,f8141,f555])).
fof(f8831,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f322,f322,f3579,f8141,f533])).
fof(f3579,plain,(
  ( ! [X0] : (sP12('T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X0)) )),
  inference(unit_resulting_resolution,[],[f322,f2082,f532])).
fof(f8824,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f322,f3586,f8141,f531])).
fof(f8821,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f573,f322,f3590,f8141,f530])).
fof(f8814,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f322,f573,f3535,f8141,f530])).
fof(f8809,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f3561,f8141,f529])).
fof(f3561,plain,(
  ~sP11(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f573,f322,f2082,f526])).
fof(f8807,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f3568,f8141,f527])).
fof(f3568,plain,(
  ~sP11('T_0','F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))),
  inference(unit_resulting_resolution,[],[f573,f322,f2082,f526])).
fof(f8799,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f322,f322,f3575,f8141,f526])).
fof(f8789,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f322,f3535,f602,f8141,f525])).
fof(f8779,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f322,f3535,f3575,f8141,f525])).
fof(f8757,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f322,f3598,f8141,f520])).
fof(f8755,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f322,f3595,f8141,f520])).
fof(f8752,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f8141,f3598,f8141,f519])).
fof(f8749,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f573,f3576,f8141,f519])).
fof(f8742,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f573,f3598,f8141,f519])).
fof(f8740,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f573,f3595,f8141,f519])).
fof(f8736,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f322,f8141,f3528,f8141,f516])).
fof(f8731,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f573,f322,f3522,f8141,f516])).
fof(f8726,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f322,f8141,f3522,f8141,f516])).
fof(f8721,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f573,f322,f3528,f8141,f516])).
fof(f8716,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f573,f573,f3534,f8141,f513])).
fof(f8709,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f573,f8141,f3507,f8141,f513])).
fof(f8705,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f573,f573,f3503,f8141,f513])).
fof(f8698,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f573,f8141,f3503,f8141,f513])).
fof(f8694,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f573,f573,f3507,f8141,f513])).
fof(f8687,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f3515,f8141,f512])).
fof(f3515,plain,(
  ~sP9(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f573,f573,f2082,f509])).
fof(f8682,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f322,f8141,f3521,f8141,f509])).
fof(f8677,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f573,f322,f3520,f8141,f509])).
fof(f8668,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f322,f8141,f3520,f8141,f509])).
fof(f8663,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f322,f8141,f3521,f8141,f509])).
fof(f8661,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f322,f573,f3521,f8141,f509])).
fof(f8651,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f8141,f3507,f597,f8141,f508])).
fof(f8648,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f573,f3503,f597,f8141,f508])).
fof(f8644,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f8141,f3507,f3521,f8141,f508])).
fof(f8638,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f8141,f3507,f3520,f8141,f508])).
fof(f8635,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f573,f3503,f3520,f8141,f508])).
fof(f8623,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f8141,f3503,f597,f8141,f508])).
fof(f8620,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f573,f3507,f597,f8141,f508])).
fof(f8616,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f8141,f3507,f3520,f8141,f508])).
fof(f8610,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f8141,f3503,f3521,f8141,f508])).
fof(f8607,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f573,f3507,f3521,f8141,f508])).
fof(f8568,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f322,f3602,f8141,f503])).
fof(f3602,plain,(
  ( ! [X0] : (sP10(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f596,f2082,f557])).
fof(f8566,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f322,f3600,f8141,f503])).
fof(f3600,plain,(
  ( ! [X0] : (sP10(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f595,f2082,f557])).
fof(f8563,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f322,f3602,f8141,f502])).
fof(f8561,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f322,f3600,f8141,f502])).
fof(f8558,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f322,f322,f3489,f8141,f501])).
fof(f3489,plain,(
  sP8(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'T_0','T_0')),
  inference(unit_resulting_resolution,[],[f322,f322,f2082,f498])).
fof(f8554,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f573,f3449,f8141,f497])).
fof(f8552,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f573,f3448,f8141,f497])).
fof(f8548,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f8141,f3448,f8141,f497])).
fof(f8546,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f573,f3493,f8141,f497])).
fof(f8540,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f573,f3449,f8141,f496])).
fof(f8538,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f573,f3448,f8141,f496])).
fof(f8534,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f8141,f3449,f8141,f496])).
fof(f8532,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f573,f3498,f8141,f496])).
fof(f8526,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f3462,f8141,f495])).
fof(f3462,plain,(
  ~sP7(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'T_0','F_0')),
  inference(unit_resulting_resolution,[],[f322,f573,f2082,f492])).
fof(f8524,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f3469,f8141,f494])).
fof(f3469,plain,(
  ~sP7('T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0')),
  inference(unit_resulting_resolution,[],[f322,f573,f2082,f492])).
fof(f8516,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f322,f322,f3476,f8141,f492])).
fof(f8506,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f322,f3477,f8141,f487])).
fof(f3477,plain,(
  sP8('T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'T_0')),
  inference(unit_resulting_resolution,[],[f322,f322,f2082,f498])).
fof(f8497,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f322,f3483,f8141,f486])).
fof(f3483,plain,(
  sP8('T_0','T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))),
  inference(unit_resulting_resolution,[],[f322,f322,f2082,f498])).
fof(f8941,plain,(
  $false | $spl68),
  inference(trivial_inequality_removal,[],[f8494])).
fof(f8494,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl68),
  inference(unit_resulting_resolution,[],[f5943,f8141,f484])).
fof(f8940,plain,(
  $false | $spl68),
  inference(trivial_inequality_removal,[],[f8493])).
fof(f8493,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl68),
  inference(unit_resulting_resolution,[],[f5912,f8141,f484])).
fof(f5912,plain,(
  sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'T_0','T_0')),
  inference(unit_resulting_resolution,[],[f322,f589,f577,f3426,f466])).
fof(f8478,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f617,f3446,f8141,f481])).
fof(f8476,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f617,f3443,f8141,f481])).
fof(f8468,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f573,f8141,f3429,f8141,f477])).
fof(f8463,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f573,f8141,f3428,f8141,f477])).
fof(f8458,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f573,f8141,f3427,f8141,f477])).
fof(f8456,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f573,f573,f3427,f8141,f477])).
fof(f8448,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f573,f8141,f3429,f8141,f477])).
fof(f8443,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f573,f8141,f3427,f8141,f477])).
fof(f8438,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f573,f8141,f3428,f8141,f477])).
fof(f8436,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f573,f573,f3428,f8141,f477])).
fof(f8428,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f573,f8141,f3428,f8141,f477])).
fof(f8423,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f573,f8141,f3427,f8141,f477])).
fof(f8418,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f573,f8141,f3429,f8141,f477])).
fof(f8416,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f573,f573,f3429,f8141,f477])).
fof(f8408,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f3412,f8141,f476])).
fof(f8406,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f3419,f8141,f475])).
fof(f8398,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f322,f322,f3426,f8141,f473])).
fof(f8394,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f3397,f8141,f472])).
fof(f3397,plain,(
  ~sP5(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'T_0','F_0')),
  inference(unit_resulting_resolution,[],[f322,f573,f2082,f469])).
fof(f8392,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f3404,f8141,f471])).
fof(f8384,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f322,f322,f3411,f8141,f469])).
fof(f8358,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f588,f583,f3443,f578,f8141,f468])).
fof(f8322,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f588,f583,f3443,f577,f8141,f468])).
fof(f8273,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f3428,f583,f3390,f578,f8141,f467])).
fof(f8258,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f3428,f583,f3390,f3411,f8141,f467])).
fof(f8225,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f322,f322,f3393,f8141,f465])).
fof(f3393,plain,(
  sP6('T_0','T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))),
  inference(unit_resulting_resolution,[],[f588,f582,f577,f2082,f466])).
fof(f8221,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f8141,f3378,f8141,f447])).
fof(f8216,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f573,f3604,f8141,f447])).
fof(f8212,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f8141,f3604,f8141,f447])).
fof(f8208,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f573,f3378,f8141,f447])).
fof(f8204,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f573,f3383,f8141,f446])).
fof(f8202,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f8141,f3604,f8141,f446])).
fof(f8195,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f8141,f3383,f8141,f446])).
fof(f8191,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f573,f3604,f8141,f446])).
fof(f8187,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f322,f3358,f8141,f412])).
fof(f8184,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f322,f3360,f8141,f409])).
fof(f8181,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f573,f3363,f8141,f407])).
fof(f8177,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f573,f3362,f8141,f407])).
fof(f8166,plain,(
  $false | $spl68),
  inference(unit_resulting_resolution,[],[f3354,f8141,f336])).
fof(f8145,plain,(
  $false | ($spl1 | $spl6 | $spl71)),
  inference(subsumption_resolution,[],[f8144,f7469])).
fof(f8144,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK19(sK16,sK14))))) | ($spl6 | $spl71)),
  inference(unit_resulting_resolution,[],[f8143,f572])).
fof(f8141,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,X0))))) ) | $spl68),
  inference(cnf_transformation,[],[f8141_D])).
fof(f8141_D,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,X0))))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f8143,plain,(
  s(sK14,sK15) != s(sK14,sK19(sK16,sK14)) | $spl71),
  inference(cnf_transformation,[],[f8143_D])).
fof(f8143_D,plain,(
  s(sK14,sK15) != s(sK14,sK19(sK16,sK14)) <=> ~$spl71),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl71])])).
fof(f8140,plain,(
  ( ! [X0] : (s(sK14,sK15) != s(sK14,sK19(sK16,sK14)) | ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,X0))))) ) | ($spl1 | $spl6)),
  inference(subsumption_resolution,[],[f8139,f573])).
fof(f8139,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | s(sK14,sK15) != s(sK14,sK19(sK16,sK14)) | ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,X0))))) ) | ($spl1 | $spl6)),
  inference(forward_demodulation,[],[f8125,f627])).
fof(f8125,plain,(
  ( ! [X0] : (s(sK14,sK15) != s(sK14,sK19(sK16,sK14)) | p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,X0))))) ) | ($spl1 | $spl6)),
  inference(superposition,[],[f371,f7518])).
fof(f7518,plain,(
  s(sK14,sK15) = s(sK14,sK20(sK16,sK14)) | ($spl1 | $spl6)),
  inference(unit_resulting_resolution,[],[f7466,f572])).
fof(f7466,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK20(sK16,sK14))))) | $spl1),
  inference(subsumption_resolution,[],[f7465,f573])).
fof(f7465,plain,(
  p(s(bool,'F_0')) | p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK20(sK16,sK14))))) | $spl1),
  inference(forward_demodulation,[],[f7464,f627])).
fof(f371,plain,(
  ( ! [X4,X0,X1] : (s(X0,sK19(X1,X0)) != s(X0,sK20(X1,X0)) | p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) )),
  inference(cnf_transformation,[],[f229])).
fof(f7131,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,X1)) | sP12(happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,X1))),s(X2,X3)),'T_0',X4)) ) | $spl66),
  inference(cnf_transformation,[],[f7131_D])).
fof(f7131_D,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,X1)) | sP12(happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,X1))),s(X2,X3)),'T_0',X4)) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f7133,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP12(happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,X1))),s(X2,X3)),'T_0',X4)) )),
  inference(condensation,[],[f7132])).
fof(f7132,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X36)) | sP12(happ(s(fun(X37,bool),'_dst_x0_1'(s(bool,X36))),s(X37,X38)),'T_0',X39) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f7120,f902])).
fof(f7120,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X35)) | ~p(s(bool,X36)) | sP12(happ(s(fun(X37,bool),'_dst_x0_1'(s(bool,X36))),s(X37,X38)),'T_0',X39) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1352,f836])).
fof(f1352,plain,(
  ( ! [X116,X114,X117,X115,X113] : (~p(s(bool,X116)) | ~p(s(bool,X114)) | sP12(happ(s(fun(X113,bool),'_dst_x0_1'(s(bool,X114))),s(X113,X115)),X116,X117)) )),
  inference(superposition,[],[f532,f421])).
fof(f7130,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP12(happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,X1))),s(X2,X3)),'T_0',X4)) )),
  inference(condensation,[],[f7129])).
fof(f7129,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X30)) | sP12(happ(s(fun(X31,bool),'_dst_x0_1'(s(bool,X30))),s(X31,X32)),'T_0',X33) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f7119,f902])).
fof(f7119,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X29)) | ~p(s(bool,X30)) | sP12(happ(s(fun(X31,bool),'_dst_x0_1'(s(bool,X30))),s(X31,X32)),'T_0',X33) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f1352,f836])).
fof(f7100,plain,(
  ( ! [X4,X2,X3,X1] : (p(s(bool,X1)) | sP12(happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)),X4,X1)) ) | $spl64),
  inference(cnf_transformation,[],[f7100_D])).
fof(f7100_D,plain,(
  ( ! [X4,X2,X3,X1] : (p(s(bool,X1)) | sP12(happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)),X4,X1)) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f7102,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP12(happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)),X4,X1)) )),
  inference(condensation,[],[f7101])).
fof(f7101,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (p(s(bool,X36)) | sP12(happ(s(fun(X37,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X37,X38)),X39,X36) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f7089,f902])).
fof(f7089,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X35)) | p(s(bool,X36)) | sP12(happ(s(fun(X37,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X37,X38)),X39,X36) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1351,f836])).
fof(f1351,plain,(
  ( ! [X111,X109,X112,X110,X108] : (~p(s(bool,X109)) | p(s(bool,X111)) | sP12(happ(s(fun(X108,bool),'_dst_x0_1'(s(bool,X109))),s(X108,X110)),X112,X111)) )),
  inference(superposition,[],[f531,f421])).
fof(f7099,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP12(happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)),X4,X1)) )),
  inference(condensation,[],[f7098])).
fof(f7098,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (p(s(bool,X30)) | sP12(happ(s(fun(X31,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X31,X32)),X33,X30) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f7088,f902])).
fof(f7088,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X29)) | p(s(bool,X30)) | sP12(happ(s(fun(X31,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X31,X32)),X33,X30) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f1351,f836])).
fof(f7068,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,X1)) | sP10(happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,X1))),s(X2,X3)),'T_0',X4)) ) | $spl62),
  inference(cnf_transformation,[],[f7068_D])).
fof(f7068_D,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,X1)) | sP10(happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,X1))),s(X2,X3)),'T_0',X4)) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f7070,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP10(happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,X1))),s(X2,X3)),'T_0',X4)) )),
  inference(condensation,[],[f7069])).
fof(f7069,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X36)) | sP10(happ(s(fun(X37,bool),'_dst_x0_1'(s(bool,X36))),s(X37,X38)),'T_0',X39) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f7057,f902])).
fof(f7057,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X35)) | ~p(s(bool,X36)) | sP10(happ(s(fun(X37,bool),'_dst_x0_1'(s(bool,X36))),s(X37,X38)),'T_0',X39) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1348,f836])).
fof(f1348,plain,(
  ( ! [X94,X97,X95,X93,X96] : (~p(s(bool,X96)) | ~p(s(bool,X94)) | sP10(happ(s(fun(X93,bool),'_dst_x0_1'(s(bool,X94))),s(X93,X95)),X96,X97)) )),
  inference(superposition,[],[f515,f421])).
fof(f7067,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP10(happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,X1))),s(X2,X3)),'T_0',X4)) )),
  inference(condensation,[],[f7066])).
fof(f7066,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X30)) | sP10(happ(s(fun(X31,bool),'_dst_x0_1'(s(bool,X30))),s(X31,X32)),'T_0',X33) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f7056,f902])).
fof(f7056,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X29)) | ~p(s(bool,X30)) | sP10(happ(s(fun(X31,bool),'_dst_x0_1'(s(bool,X30))),s(X31,X32)),'T_0',X33) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f1348,f836])).
fof(f7022,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,X1)) | sP10(happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,X1))),s(X2,X3)),X4,'T_0')) ) | $spl60),
  inference(cnf_transformation,[],[f7022_D])).
fof(f7022_D,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,X1)) | sP10(happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,X1))),s(X2,X3)),X4,'T_0')) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f7024,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP10(happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,X1))),s(X2,X3)),X4,'T_0')) )),
  inference(condensation,[],[f7023])).
fof(f7023,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X36)) | sP10(happ(s(fun(X37,bool),'_dst_x0_1'(s(bool,X36))),s(X37,X38)),X39,'T_0') | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f7011,f902])).
fof(f7011,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X35)) | ~p(s(bool,X36)) | sP10(happ(s(fun(X37,bool),'_dst_x0_1'(s(bool,X36))),s(X37,X38)),X39,'T_0') | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1347,f836])).
fof(f1347,plain,(
  ( ! [X92,X90,X88,X91,X89] : (~p(s(bool,X91)) | ~p(s(bool,X89)) | sP10(happ(s(fun(X88,bool),'_dst_x0_1'(s(bool,X89))),s(X88,X90)),X92,X91)) )),
  inference(superposition,[],[f514,f421])).
fof(f7021,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP10(happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,X1))),s(X2,X3)),X4,'T_0')) )),
  inference(condensation,[],[f7020])).
fof(f7020,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X30)) | sP10(happ(s(fun(X31,bool),'_dst_x0_1'(s(bool,X30))),s(X31,X32)),X33,'T_0') | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f7010,f902])).
fof(f7010,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X29)) | ~p(s(bool,X30)) | sP10(happ(s(fun(X31,bool),'_dst_x0_1'(s(bool,X30))),s(X31,X32)),X33,'T_0') | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f1347,f836])).
fof(f6976,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,X1)) | sP2('T_0',X2,happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4)))) ) | $spl58),
  inference(cnf_transformation,[],[f6976_D])).
fof(f6976_D,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,X1)) | sP2('T_0',X2,happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4)))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f6978,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP2('T_0',X2,happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4)))) )),
  inference(condensation,[],[f6977])).
fof(f6977,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X36)) | sP2('T_0',X37,happ(s(fun(X38,bool),'_dst_x0_1'(s(bool,X36))),s(X38,X39))) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f6963,f902])).
fof(f6963,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X35)) | ~p(s(bool,X36)) | sP2('T_0',X37,happ(s(fun(X38,bool),'_dst_x0_1'(s(bool,X36))),s(X38,X39))) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1340,f836])).
fof(f6975,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP2('T_0',X2,happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4)))) )),
  inference(condensation,[],[f6974])).
fof(f6974,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X30)) | sP2('T_0',X31,happ(s(fun(X32,bool),'_dst_x0_1'(s(bool,X30))),s(X32,X33))) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f6962,f902])).
fof(f6962,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X29)) | ~p(s(bool,X30)) | sP2('T_0',X31,happ(s(fun(X32,bool),'_dst_x0_1'(s(bool,X30))),s(X32,X33))) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f1340,f836])).
fof(f6942,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,X1)) | sP2('T_0',happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,X1))),s(X2,X3)),X4)) ) | $spl56),
  inference(cnf_transformation,[],[f6942_D])).
fof(f6942_D,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,X1)) | sP2('T_0',happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,X1))),s(X2,X3)),X4)) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f6944,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP2('T_0',happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,X1))),s(X2,X3)),X4)) )),
  inference(condensation,[],[f6943])).
fof(f6943,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X36)) | sP2('T_0',happ(s(fun(X37,bool),'_dst_x0_1'(s(bool,X36))),s(X37,X38)),X39) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f6931,f902])).
fof(f6931,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X35)) | ~p(s(bool,X36)) | sP2('T_0',happ(s(fun(X37,bool),'_dst_x0_1'(s(bool,X36))),s(X37,X38)),X39) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1339,f836])).
fof(f1339,plain,(
  ( ! [X52,X50,X48,X51,X49] : (~p(s(bool,X51)) | ~p(s(bool,X49)) | sP2(X51,happ(s(fun(X48,bool),'_dst_x0_1'(s(bool,X49))),s(X48,X50)),X52)) )),
  inference(superposition,[],[f450,f421])).
fof(f6941,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP2('T_0',happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,X1))),s(X2,X3)),X4)) )),
  inference(condensation,[],[f6940])).
fof(f6940,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X30)) | sP2('T_0',happ(s(fun(X31,bool),'_dst_x0_1'(s(bool,X30))),s(X31,X32)),X33) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f6930,f902])).
fof(f6930,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X29)) | ~p(s(bool,X30)) | sP2('T_0',happ(s(fun(X31,bool),'_dst_x0_1'(s(bool,X30))),s(X31,X32)),X33) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f1339,f836])).
fof(f6841,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | sP1(X1,happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)))) ) | $spl54),
  inference(cnf_transformation,[],[f6841_D])).
fof(f6841_D,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | sP1(X1,happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f6843,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)))) )),
  inference(condensation,[],[f6842])).
fof(f6842,plain,(
  ( ! [X30,X33,X31,X29,X32] : (p(s(bool,X30)) | sP1(X30,happ(s(fun(X31,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X31,X32))) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f6830,f902])).
fof(f6830,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | p(s(bool,X30)) | sP1(X30,happ(s(fun(X31,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X31,X32))) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f1337,f836])).
fof(f1337,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X40)) | p(s(bool,X42)) | sP1(X42,happ(s(fun(X39,bool),'_dst_x0_1'(s(bool,X40))),s(X39,X41)))) )),
  inference(superposition,[],[f412,f421])).
fof(f6840,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)))) )),
  inference(condensation,[],[f6839])).
fof(f6839,plain,(
  ( ! [X28,X26,X24,X27,X25] : (p(s(bool,X25)) | sP1(X25,happ(s(fun(X26,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X26,X27))) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f6829,f902])).
fof(f6829,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | p(s(bool,X25)) | sP1(X25,happ(s(fun(X26,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X26,X27))) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f1337,f836])).
fof(f6809,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | sP1(happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)),X1)) ) | $spl52),
  inference(cnf_transformation,[],[f6809_D])).
fof(f6809_D,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | sP1(happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)),X1)) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f6811,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)),X1)) )),
  inference(condensation,[],[f6810])).
fof(f6810,plain,(
  ( ! [X30,X33,X31,X29,X32] : (p(s(bool,X30)) | sP1(happ(s(fun(X31,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X31,X32)),X30) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f6798,f902])).
fof(f6798,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | p(s(bool,X30)) | sP1(happ(s(fun(X31,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X31,X32)),X30) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f1336,f836])).
fof(f1336,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X36)) | p(s(bool,X38)) | sP1(happ(s(fun(X35,bool),'_dst_x0_1'(s(bool,X36))),s(X35,X37)),X38)) )),
  inference(superposition,[],[f409,f421])).
fof(f6808,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)),X1)) )),
  inference(condensation,[],[f6807])).
fof(f6807,plain,(
  ( ! [X28,X26,X24,X27,X25] : (p(s(bool,X25)) | sP1(happ(s(fun(X26,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X26,X27)),X25) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f6797,f902])).
fof(f6797,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | p(s(bool,X25)) | sP1(happ(s(fun(X26,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X26,X27)),X25) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f1336,f836])).
fof(f5060,plain,(
  $false | $spl14),
  inference(global_subsumption,[],[f321,f320,f319,f318,f322,f573,f329,f337,f336,f341,f340,f364,f371,f370,f369,f368,f367,f379,f377,f382,f380,f387,f384,f394,f393,f574,f412,f409,f408,f407,f421,f423,f425,f424,f428,f427,f426,f431,f430,f429,f434,f433,f432,f437,f436,f435,f439,f438,f451,f450,f449,f559,f447,f446,f468,f467,f466,f465,f464,f463,f462,f472,f471,f470,f469,f476,f475,f474,f473,f480,f479,f478,f477,f484,f483,f482,f481,f558,f488,f487,f486,f485,f495,f494,f493,f492,f501,f498,f497,f496,f508,f507,f506,f557,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f525,f524,f523,f556,f521,f520,f519,f529,f528,f527,f526,f533,f532,f531,f530,f555,f575,f576,f577,f578,f580,f582,f583,f585,f587,f588,f589,f590,f591,f593,f595,f596,f597,f599,f601,f602,f604,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f617,f621,f639,f640,f643,f649,f657,f664,f681,f682,f685,f688,f695,f704,f707,f708,f711,f712,f716,f725,f729,f730,f733,f736,f740,f743,f746,f755,f756,f757,f765,f768,f772,f775,f778,f797,f800,f815,f813,f818,f836,f884,f902,f925,f923,f929,f927,f941,f938,f951,f949,f992,f1003,f1011,f1024,f1032,f1036,f1042,f1051,f1056,f1055,f1068,f1076,f1080,f1113,f1124,f1131,f1140,f1152,f1160,f1164,f1178,f1196,f1197,f1212,f1222,f1237,f1249,f1257,f1261,f1272,f1282,f1298,f1314,f1313,f1311,f1310,f1325,f1334,f1335,f1336,f1337,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1347,f1348,f1349,f1350,f1351,f1352,f1353,f1355,f1356,f1371,f1381,f1397,f1389,f1387,f1386,f1405,f1402,f1439,f1429,f1428,f1484,f1503,f1515,f1514,f1513,f1512,f1517,f1518,f1519,f1529,f1530,f1531,f1532,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1543,f1544,f1545,f1546,f1547,f1548,f1549,f1551,f1552,f1866,f1863,f1867,f1865,f1876,f1873,f1877,f1875,f1887,f1884,f1886,f1900,f1897,f1899,f1916,f1912,f1934,f1930,f1952,f1948,f1977,f1973,f2000,f2015,f2016,f2027,f2038,f2035,f2036,f2037,f2051,f2062,f2059,f2060,f2061,f2082,f2104,f2105,f2106,f2111,f2134,f2135,f2136,f2141,f2142,f2186,f2182,f2209,f2205,f2233,f2234,f2228,f2229,f2230,f2232,f2249,f2250,f2244,f2245,f2246,f2251,f2248,f2286,f2287,f2305,f2306,f2334,f2336,f2347,f2349,f2367,f2364,f2386,f2370,f2371,f2372,f2374,f2375,f2400,f2403,f2404,f2405,f2407,f2408,f2429,f2444,f2427,f2432,f2433,f2434,f2436,f2456,f2460,f2461,f2462,f2464,f2489,f2488,f2486,f2485,f2497,f2506,f2507,f2508,f2509,f2511,f2512,f2513,f2514,f2515,f2516,f2517,f2518,f2519,f2520,f2521,f2522,f2523,f2524,f2525,f2526,f2528,f2529,f2530,f2531,f2532,f2533,f2831,f2830,f2827,f2824,f2821,f2820,f2819,f2818,f2817,f2816,f2815,f2814,f2813,f2812,f2811,f2810,f2809,f2806,f2803,f2802,f2801,f2797,f2796,f2790,f2787,f2786,f2780,f2779,f2772,f2746,f2745,f2744,f2739,f2738,f2733,f2732,f2731,f2726,f2718,f2714,f2710,f2709,f2704,f2703,f2702,f2701,f2700,f2696,f2694,f2688,f2687,f2681,f2680,f2673,f2670,f2669,f2663,f2660,f2659,f2658,f2656,f2654,f2653,f2651,f2650,f2649,f2648,f2646,f2645,f2644,f2643,f2642,f2641,f2635,f2634,f2627,f2626,f2620,f2619,f2612,f2611,f2609,f2607,f2605,f2602,f2601,f2595,f2594,f2589,f2574,f2573,f2572,f2571,f2569,f2568,f2567,f2566,f2565,f3267,f3266,f3263,f3260,f3257,f3256,f3255,f3254,f3253,f3252,f3251,f3250,f3249,f3248,f3247,f3246,f3245,f3242,f3239,f3238,f3237,f3233,f3232,f3226,f3223,f3222,f3216,f3215,f3208,f3182,f3181,f3180,f3175,f3174,f3169,f3168,f3167,f3162,f3154,f3150,f3146,f3145,f3140,f3139,f3138,f3137,f3136,f3132,f3130,f3124,f3123,f3117,f3116,f3109,f3106,f3105,f3099,f3096,f3095,f3093,f3090,f3088,f3087,f3084,f3083,f3082,f3081,f3078,f3077,f3076,f3075,f3074,f3073,f3067,f3066,f3059,f3058,f3052,f3051,f3044,f3043,f3041,f3039,f3037,f3034,f3033,f3027,f3026,f3021,f3006,f3005,f3004,f3003,f3001,f3000,f2999,f2998,f2997,f3620,f3619,f3616,f3613,f3610,f3609,f3608,f3607,f3606,f3605,f3604,f3603,f3602,f3601,f3600,f3599,f3598,f3595,f3592,f3591,f3590,f3586,f3585,f3579,f3576,f3575,f3569,f3568,f3561,f3535,f3534,f3533,f3528,f3527,f3522,f3521,f3520,f3515,f3507,f3503,f3499,f3498,f3493,f3492,f3491,f3490,f3489,f3485,f3483,f3477,f3476,f3470,f3469,f3462,f3459,f3458,f3452,f3449,f3448,f3446,f3443,f3441,f3440,f3437,f3436,f3435,f3434,f3431,f3430,f3429,f3428,f3427,f3426,f3420,f3419,f3412,f3411,f3405,f3404,f3397,f3396,f3395,f3394,f3393,f3391,f3390,f3384,f3383,f3378,f3363,f3362,f3361,f3360,f3358,f3357,f3356,f3354,f3708,f3726,f3770,f3771,f3773,f3788,f3799,f3800,f3802,f3851,f3852,f3853,f3854,f3905,f3907,f3908,f3909,f3910,f3911,f3912,f3913,f3914,f3915,f3916,f3917,f3918,f3919,f3920,f3921,f3922,f3923,f3924,f3925,f3926,f3927,f4039,f4569,f4573,f4576,f4579,f4583,f4586,f4589,f4592,f4595,f4598,f4601,f4604,f4607,f4610,f4613,f4616,f4619,f4622,f4625,f4072,f4073,f4074,f4075,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4646,f4647,f4648,f4649,f4650,f4124,f4125,f4126,f4127,f4651,f4652,f4653,f4654,f4655,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4667,f4668,f4669,f4670,f4671,f4177,f4179,f4180,f4181,f4182,f4183,f4184,f4185,f4186,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4677,f4679,f4681,f4683,f4686,f4688,f4690,f4692,f4694,f4696,f4698,f4700,f4702,f4704,f4706,f4708,f4710,f4712,f4714,f4571,f4941,f4744,f4751,f4966,f4972,f4982,f4986,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f5037,f5039,f5041,f5043,f5045,f5047,f5049,f5051,f4899,f4900,f4901,f4902])).
fof(f4902,plain,(
  ( ! [X12,X10,X13,X11] : (sP6('T_0',happ(s(fun(X10,bool),X11),s(X10,'const.min.@_1'(s(fun(X10,bool),X11)))),happ(s(fun(X12,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X12,X13)))) ) | $spl14),
  inference(backward_demodulation,[],[f4571,f2464])).
fof(f4901,plain,(
  ( ! [X12,X10,X13,X11] : (~sP6('F_0',happ(s(fun(X10,bool),X11),s(X10,'const.min.@_1'(s(fun(X10,bool),X11)))),happ(s(fun(X12,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X12,X13)))) ) | $spl14),
  inference(backward_demodulation,[],[f4571,f2436])).
fof(f4900,plain,(
  ( ! [X12,X10,X13,X11] : (sP6('T_0',happ(s(fun(X12,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X12,X13)),happ(s(fun(X10,bool),X11),s(X10,'const.min.@_1'(s(fun(X10,bool),X11)))))) ) | $spl14),
  inference(backward_demodulation,[],[f4571,f2407])).
fof(f4899,plain,(
  ( ! [X12,X10,X13,X11] : (~sP6('F_0',happ(s(fun(X12,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X12,X13)),happ(s(fun(X10,bool),X11),s(X10,'const.min.@_1'(s(fun(X10,bool),X11)))))) ) | $spl14),
  inference(backward_demodulation,[],[f4571,f2374])).
fof(f5051,plain,(
  ( ! [X109,X107,X108,X106] : (~sP9(happ(s(fun(X106,bool),X107),s(X106,'const.min.@_1'(s(fun(X106,bool),X107)))),X108,X109) | sP10(happ(s(fun(X106,bool),X107),s(X106,'const.min.@_1'(s(fun(X106,bool),X107)))),X108,X109)) ) | $spl14),
  inference(subsumption_resolution,[],[f4896,f322])).
fof(f4896,plain,(
  ( ! [X109,X107,X108,X106] : (~p(s(bool,'T_0')) | ~sP9(happ(s(fun(X106,bool),X107),s(X106,'const.min.@_1'(s(fun(X106,bool),X107)))),X108,X109) | sP10(happ(s(fun(X106,bool),X107),s(X106,'const.min.@_1'(s(fun(X106,bool),X107)))),X108,X109)) ) | $spl14),
  inference(backward_demodulation,[],[f4571,f1552])).
fof(f5049,plain,(
  ( ! [X103,X105,X102,X104] : (~sP11(happ(s(fun(X102,bool),X103),s(X102,'const.min.@_1'(s(fun(X102,bool),X103)))),X104,X105) | sP12(happ(s(fun(X102,bool),X103),s(X102,'const.min.@_1'(s(fun(X102,bool),X103)))),X104,X105)) ) | $spl14),
  inference(subsumption_resolution,[],[f4895,f322])).
fof(f4895,plain,(
  ( ! [X103,X105,X102,X104] : (~p(s(bool,'T_0')) | ~sP11(happ(s(fun(X102,bool),X103),s(X102,'const.min.@_1'(s(fun(X102,bool),X103)))),X104,X105) | sP12(happ(s(fun(X102,bool),X103),s(X102,'const.min.@_1'(s(fun(X102,bool),X103)))),X104,X105)) ) | $spl14),
  inference(backward_demodulation,[],[f4571,f1551])).
fof(f5047,plain,(
  ( ! [X97,X95,X98,X96] : (~p(s(bool,X97)) | p(s(bool,X98)) | ~sP12(happ(s(fun(X95,bool),X96),s(X95,'const.min.@_1'(s(fun(X95,bool),X96)))),X98,X97)) ) | $spl14),
  inference(subsumption_resolution,[],[f4894,f322])).
fof(f4894,plain,(
  ( ! [X97,X95,X98,X96] : (~p(s(bool,'T_0')) | ~p(s(bool,X97)) | p(s(bool,X98)) | ~sP12(happ(s(fun(X95,bool),X96),s(X95,'const.min.@_1'(s(fun(X95,bool),X96)))),X98,X97)) ) | $spl14),
  inference(backward_demodulation,[],[f4571,f1549])).
fof(f5045,plain,(
  ( ! [X94,X92,X93,X91] : (~p(s(bool,X93)) | sP12(happ(s(fun(X91,bool),X92),s(X91,'const.min.@_1'(s(fun(X91,bool),X92)))),X93,X94)) ) | $spl14),
  inference(subsumption_resolution,[],[f4893,f322])).
fof(f4893,plain,(
  ( ! [X94,X92,X93,X91] : (~p(s(bool,'T_0')) | ~p(s(bool,X93)) | sP12(happ(s(fun(X91,bool),X92),s(X91,'const.min.@_1'(s(fun(X91,bool),X92)))),X93,X94)) ) | $spl14),
  inference(backward_demodulation,[],[f4571,f1548])).
fof(f5043,plain,(
  ( ! [X90,X88,X87,X89] : (p(s(bool,X89)) | sP12(happ(s(fun(X87,bool),X88),s(X87,'const.min.@_1'(s(fun(X87,bool),X88)))),X90,X89)) ) | $spl14),
  inference(subsumption_resolution,[],[f4892,f322])).
fof(f4892,plain,(
  ( ! [X90,X88,X87,X89] : (~p(s(bool,'T_0')) | p(s(bool,X89)) | sP12(happ(s(fun(X87,bool),X88),s(X87,'const.min.@_1'(s(fun(X87,bool),X88)))),X90,X89)) ) | $spl14),
  inference(backward_demodulation,[],[f4571,f1547])).
fof(f5041,plain,(
  ( ! [X85,X83,X86,X84] : (p(s(bool,X85)) | ~p(s(bool,X86)) | ~sP11(X86,X85,happ(s(fun(X83,bool),X84),s(X83,'const.min.@_1'(s(fun(X83,bool),X84)))))) ) | $spl14),
  inference(subsumption_resolution,[],[f4891,f322])).
fof(f4891,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,'T_0')) | p(s(bool,X85)) | ~p(s(bool,X86)) | ~sP11(X86,X85,happ(s(fun(X83,bool),X84),s(X83,'const.min.@_1'(s(fun(X83,bool),X84)))))) ) | $spl14),
  inference(backward_demodulation,[],[f4571,f1546])).
fof(f5039,plain,(
  ( ! [X80,X81,X79,X82] : (p(s(bool,X81)) | p(s(bool,X82)) | ~sP10(happ(s(fun(X79,bool),X80),s(X79,'const.min.@_1'(s(fun(X79,bool),X80)))),X82,X81)) ) | $spl14),
  inference(subsumption_resolution,[],[f4890,f322])).
fof(f4890,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,'T_0')) | p(s(bool,X81)) | p(s(bool,X82)) | ~sP10(happ(s(fun(X79,bool),X80),s(X79,'const.min.@_1'(s(fun(X79,bool),X80)))),X82,X81)) ) | $spl14),
  inference(backward_demodulation,[],[f4571,f1545])).
fof(f5037,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,X77)) | sP10(happ(s(fun(X75,bool),X76),s(X75,'const.min.@_1'(s(fun(X75,bool),X76)))),X77,X78)) ) | $spl14),
  inference(subsumption_resolution,[],[f4889,f322])).
fof(f4889,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,'T_0')) | ~p(s(bool,X77)) | sP10(happ(s(fun(X75,bool),X76),s(X75,'const.min.@_1'(s(fun(X75,bool),X76)))),X77,X78)) ) | $spl14),
  inference(backward_demodulation,[],[f4571,f1544])).
fof(f5035,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,X73)) | sP10(happ(s(fun(X71,bool),X72),s(X71,'const.min.@_1'(s(fun(X71,bool),X72)))),X74,X73)) ) | $spl14),
  inference(subsumption_resolution,[],[f4888,f322])).
fof(f4888,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,'T_0')) | ~p(s(bool,X73)) | sP10(happ(s(fun(X71,bool),X72),s(X71,'const.min.@_1'(s(fun(X71,bool),X72)))),X74,X73)) ) | $spl14),
  inference(backward_demodulation,[],[f4571,f1543])).
fof(f5033,plain,(
  ( ! [X70,X68,X69,X67] : (p(s(bool,X69)) | p(s(bool,X70)) | ~sP9(happ(s(fun(X67,bool),X68),s(X67,'const.min.@_1'(s(fun(X67,bool),X68)))),X69,X70)) ) | $spl14),
  inference(subsumption_resolution,[],[f4887,f322])).
fof(f4887,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,'T_0')) | p(s(bool,X69)) | p(s(bool,X70)) | ~sP9(happ(s(fun(X67,bool),X68),s(X67,'const.min.@_1'(s(fun(X67,bool),X68)))),X69,X70)) ) | $spl14),
  inference(backward_demodulation,[],[f4571,f1542])).
fof(f5031,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,X65)) | p(s(bool,X66)) | ~sP8(X66,X65,happ(s(fun(X63,bool),X64),s(X63,'const.min.@_1'(s(fun(X63,bool),X64)))))) ) | $spl14),
  inference(subsumption_resolution,[],[f4886,f322])).
fof(f4886,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,'T_0')) | ~p(s(bool,X65)) | p(s(bool,X66)) | ~sP8(X66,X65,happ(s(fun(X63,bool),X64),s(X63,'const.min.@_1'(s(fun(X63,bool),X64)))))) ) | $spl14),
  inference(backward_demodulation,[],[f4571,f1541])).
fof(f5029,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X61)) | p(s(bool,X62)) | ~sP7(X61,happ(s(fun(X59,bool),X60),s(X59,'const.min.@_1'(s(fun(X59,bool),X60)))),X62)) ) | $spl14),
  inference(subsumption_resolution,[],[f4885,f322])).
fof(f4885,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,'T_0')) | ~p(s(bool,X61)) | p(s(bool,X62)) | ~sP7(X61,happ(s(fun(X59,bool),X60),s(X59,'const.min.@_1'(s(fun(X59,bool),X60)))),X62)) ) | $spl14),
  inference(backward_demodulation,[],[f4571,f1540])).
fof(f5027,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,X57)) | p(s(bool,X58)) | ~sP4(X57,happ(s(fun(X55,bool),X56),s(X55,'const.min.@_1'(s(fun(X55,bool),X56)))),X58)) ) | $spl14),
  inference(subsumption_resolution,[],[f4884,f322])).
fof(f4884,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,'T_0')) | ~p(s(bool,X57)) | p(s(bool,X58)) | ~sP4(X57,happ(s(fun(X55,bool),X56),s(X55,'const.min.@_1'(s(fun(X55,bool),X56)))),X58)) ) | $spl14),
  inference(backward_demodulation,[],[f4571,f1539])).
fof(f5025,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X53)) | p(s(bool,X54)) | ~sP5(X53,happ(s(fun(X51,bool),X52),s(X51,'const.min.@_1'(s(fun(X51,bool),X52)))),X54)) ) | $spl14),
  inference(subsumption_resolution,[],[f4883,f322])).
fof(f4883,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,'T_0')) | ~p(s(bool,X53)) | p(s(bool,X54)) | ~sP5(X53,happ(s(fun(X51,bool),X52),s(X51,'const.min.@_1'(s(fun(X51,bool),X52)))),X54)) ) | $spl14),
  inference(backward_demodulation,[],[f4571,f1538])).
fof(f5023,plain,(
  ( ! [X47,X50,X48,X49] : (sP6(X49,X50,happ(s(fun(X47,bool),X48),s(X47,'const.min.@_1'(s(fun(X47,bool),X48))))) | ~sP5(X49,happ(s(fun(X47,bool),X48),s(X47,'const.min.@_1'(s(fun(X47,bool),X48)))),X50) | ~sP4(X50,happ(s(fun(X47,bool),X48),s(X47,'const.min.@_1'(s(fun(X47,bool),X48)))),X49) | ~sP3(happ(s(fun(X47,bool),X48),s(X47,'const.min.@_1'(s(fun(X47,bool),X48)))),X50,X49)) ) | $spl14),
  inference(subsumption_resolution,[],[f4882,f322])).
fof(f4882,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,'T_0')) | sP6(X49,X50,happ(s(fun(X47,bool),X48),s(X47,'const.min.@_1'(s(fun(X47,bool),X48))))) | ~sP5(X49,happ(s(fun(X47,bool),X48),s(X47,'const.min.@_1'(s(fun(X47,bool),X48)))),X50) | ~sP4(X50,happ(s(fun(X47,bool),X48),s(X47,'const.min.@_1'(s(fun(X47,bool),X48)))),X49) | ~sP3(happ(s(fun(X47,bool),X48),s(X47,'const.min.@_1'(s(fun(X47,bool),X48)))),X50,X49)) ) | $spl14),
  inference(backward_demodulation,[],[f4571,f1537])).
fof(f5021,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X45)) | p(s(bool,X46)) | ~sP6(X45,happ(s(fun(X43,bool),X44),s(X43,'const.min.@_1'(s(fun(X43,bool),X44)))),X46)) ) | $spl14),
  inference(subsumption_resolution,[],[f4881,f322])).
fof(f4881,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,'T_0')) | ~p(s(bool,X45)) | p(s(bool,X46)) | ~sP6(X45,happ(s(fun(X43,bool),X44),s(X43,'const.min.@_1'(s(fun(X43,bool),X44)))),X46)) ) | $spl14),
  inference(backward_demodulation,[],[f4571,f1536])).
fof(f5019,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X41)) | sP2(X41,X42,happ(s(fun(X39,bool),X40),s(X39,'const.min.@_1'(s(fun(X39,bool),X40)))))) ) | $spl14),
  inference(subsumption_resolution,[],[f4880,f322])).
fof(f4880,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,X41)) | sP2(X41,X42,happ(s(fun(X39,bool),X40),s(X39,'const.min.@_1'(s(fun(X39,bool),X40)))))) ) | $spl14),
  inference(backward_demodulation,[],[f4571,f1535])).
fof(f5017,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X37)) | sP2(X37,happ(s(fun(X35,bool),X36),s(X35,'const.min.@_1'(s(fun(X35,bool),X36)))),X38)) ) | $spl14),
  inference(subsumption_resolution,[],[f4879,f322])).
fof(f4879,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'T_0')) | ~p(s(bool,X37)) | sP2(X37,happ(s(fun(X35,bool),X36),s(X35,'const.min.@_1'(s(fun(X35,bool),X36)))),X38)) ) | $spl14),
  inference(backward_demodulation,[],[f4571,f1534])).
fof(f5015,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X30)) | sP1(X30,happ(s(fun(X28,bool),X29),s(X28,'const.min.@_1'(s(fun(X28,bool),X29)))))) ) | $spl14),
  inference(subsumption_resolution,[],[f4878,f322])).
fof(f4878,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'T_0')) | p(s(bool,X30)) | sP1(X30,happ(s(fun(X28,bool),X29),s(X28,'const.min.@_1'(s(fun(X28,bool),X29)))))) ) | $spl14),
  inference(backward_demodulation,[],[f4571,f1532])).
fof(f5013,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X27)) | sP1(happ(s(fun(X25,bool),X26),s(X25,'const.min.@_1'(s(fun(X25,bool),X26)))),X27)) ) | $spl14),
  inference(subsumption_resolution,[],[f4877,f322])).
fof(f4877,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'T_0')) | p(s(bool,X27)) | sP1(happ(s(fun(X25,bool),X26),s(X25,'const.min.@_1'(s(fun(X25,bool),X26)))),X27)) ) | $spl14),
  inference(backward_demodulation,[],[f4571,f1531])).
fof(f5011,plain,(
  ( ! [X24,X23,X22] : (sP0(happ(s(fun(X22,bool),X23),s(X22,'const.min.@_1'(s(fun(X22,bool),X23)))),X24)) ) | $spl14),
  inference(subsumption_resolution,[],[f4876,f322])).
fof(f4876,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | sP0(happ(s(fun(X22,bool),X23),s(X22,'const.min.@_1'(s(fun(X22,bool),X23)))),X24)) ) | $spl14),
  inference(backward_demodulation,[],[f4571,f1530])).
fof(f5009,plain,(
  ( ! [X21,X19,X20] : (sP0(X21,happ(s(fun(X19,bool),X20),s(X19,'const.min.@_1'(s(fun(X19,bool),X20)))))) ) | $spl14),
  inference(subsumption_resolution,[],[f4875,f322])).
fof(f4875,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | sP0(X21,happ(s(fun(X19,bool),X20),s(X19,'const.min.@_1'(s(fun(X19,bool),X20)))))) ) | $spl14),
  inference(backward_demodulation,[],[f4571,f1529])).
fof(f4986,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,X6) = s(X0,happ(s(fun(X4,X0),X5),s(X4,sK27(X6,happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X5,X4,X0))))) ) | $spl14),
  inference(subsumption_resolution,[],[f4985,f322])).
fof(f4985,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'T_0')) | s(X0,X6) = s(X0,happ(s(fun(X4,X0),X5),s(X4,sK27(X6,happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X5,X4,X0))))) ) | $spl14),
  inference(forward_demodulation,[],[f4984,f4571])).
fof(f4984,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X4,X0),X5),s(fun(X0,bool),X2)))) | s(X0,X6) = s(X0,happ(s(fun(X4,X0),X5),s(X4,sK27(X6,happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X5,X4,X0))))) ) | $spl14),
  inference(subsumption_resolution,[],[f4809,f322])).
fof(f4809,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X4,X0),X5),s(fun(X0,bool),X2)))) | s(X0,X6) = s(X0,happ(s(fun(X4,X0),X5),s(X4,sK27(X6,happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X5,X4,X0))))) ) | $spl14),
  inference(backward_demodulation,[],[f4571,f2287])).
fof(f4982,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X1,X6) = s(X1,happ(s(fun(X0,X1),X3),s(X0,sK27(X6,X5,happ(s(fun(X2,fun(X0,X1)),'_dst_x0_1'(s(fun(X0,X1),X3))),s(X2,X4)),X0,X1))))) ) | $spl14),
  inference(subsumption_resolution,[],[f4981,f322])).
fof(f4981,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'T_0')) | s(X1,X6) = s(X1,happ(s(fun(X0,X1),X3),s(X0,sK27(X6,X5,happ(s(fun(X2,fun(X0,X1)),'_dst_x0_1'(s(fun(X0,X1),X3))),s(X2,X4)),X0,X1))))) ) | $spl14),
  inference(forward_demodulation,[],[f4980,f4571])).
fof(f4980,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X1,X6) = s(X1,happ(s(fun(X0,X1),X3),s(X0,sK27(X6,X5,happ(s(fun(X2,fun(X0,X1)),'_dst_x0_1'(s(fun(X0,X1),X3))),s(X2,X4)),X0,X1)))) | ~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X0,X1),X3),s(fun(X1,bool),X5))))) ) | $spl14),
  inference(subsumption_resolution,[],[f4808,f322])).
fof(f4808,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'T_0')) | s(X1,X6) = s(X1,happ(s(fun(X0,X1),X3),s(X0,sK27(X6,X5,happ(s(fun(X2,fun(X0,X1)),'_dst_x0_1'(s(fun(X0,X1),X3))),s(X2,X4)),X0,X1)))) | ~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X0,X1),X3),s(fun(X1,bool),X5))))) ) | $spl14),
  inference(backward_demodulation,[],[f4571,f2286])).
fof(f4972,plain,(
  ( ! [X6,X2,X0,X3,X1] : (s(X0,X6) = s(X0,happ(s(fun(X1,X0),X2),s(X1,sK27(X6,X3,X2,X1,X0))))) ) | $spl14),
  inference(subsumption_resolution,[],[f4971,f322])).
fof(f4971,plain,(
  ( ! [X6,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | s(X0,X6) = s(X0,happ(s(fun(X1,X0),X2),s(X1,sK27(X6,X3,X2,X1,X0))))) ) | $spl14),
  inference(forward_demodulation,[],[f4970,f4571])).
fof(f4970,plain,(
  ( ! [X6,X2,X0,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | s(X0,X6) = s(X0,happ(s(fun(X1,X0),X2),s(X1,sK27(X6,X3,X2,X1,X0))))) ) | $spl14),
  inference(subsumption_resolution,[],[f4797,f322])).
fof(f4797,plain,(
  ( ! [X6,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | s(X0,X6) = s(X0,happ(s(fun(X1,X0),X2),s(X1,sK27(X6,X3,X2,X1,X0))))) ) | $spl14),
  inference(backward_demodulation,[],[f4571,f426])).
fof(f4966,plain,(
  ( ! [X6,X0,X5] : (s(X0,X5) = s(X0,X6)) ) | $spl14),
  inference(subsumption_resolution,[],[f4965,f322])).
fof(f4965,plain,(
  ( ! [X6,X0,X5] : (~p(s(bool,'T_0')) | s(X0,X5) = s(X0,X6)) ) | $spl14),
  inference(forward_demodulation,[],[f4964,f4571])).
fof(f4964,plain,(
  ( ! [X6,X0,X5,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | s(X0,X5) = s(X0,X6)) ) | $spl14),
  inference(subsumption_resolution,[],[f4963,f322])).
fof(f4963,plain,(
  ( ! [X6,X0,X5,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | s(X0,X5) = s(X0,X6)) ) | $spl14),
  inference(forward_demodulation,[],[f4962,f4571])).
fof(f4962,plain,(
  ( ! [X6,X0,X5,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | s(X0,X5) = s(X0,X6)) ) | $spl14),
  inference(subsumption_resolution,[],[f4778,f322])).
fof(f4778,plain,(
  ( ! [X6,X0,X5,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | s(X0,X5) = s(X0,X6)) ) | $spl14),
  inference(backward_demodulation,[],[f4571,f368])).
fof(f4751,plain,(
  ( ! [X0,X1] : (~sP3(happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X0,X1)),'F_0','F_0')) ) | $spl14),
  inference(backward_demodulation,[],[f4571,f2429])).
fof(f4744,plain,(
  ( ! [X0,X1] : (~sP3('F_0',happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X0,X1)),'F_0')) ) | $spl14),
  inference(backward_demodulation,[],[f4571,f2364])).
fof(f4941,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl14),
  inference(subsumption_resolution,[],[f4715,f4571])).
fof(f4715,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl14),
  inference(backward_demodulation,[],[f4571,f340])).
fof(f4199,plain,(
  ( ! [X416,X418,X417,X419] : (p(s(bool,happ(s(fun(X418,bool),X416),s(X418,happ(s(fun(bool,X418),X417),s(bool,'T_0')))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(bool,X418),X417),s(fun(X418,bool),X416)))) | s(bool,X419) = s(bool,sK35(X416,X417,bool,X418)) | s(bool,'T_0') = s(bool,X419)) )),
  inference(superposition,[],[f439,f836])).
fof(f4198,plain,(
  ( ! [X412,X414,X413,X415] : (~p(s(bool,happ(s(fun(bool,bool),X412),s(bool,'T_0')))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X414,bool),X413),s(fun(bool,bool),X412)))) | s(bool,X415) = s(bool,sK34(X412,X413,X414,bool)) | s(bool,'T_0') = s(bool,X415)) )),
  inference(superposition,[],[f436,f836])).
fof(f4197,plain,(
  ( ! [X408,X410,X409,X411] : (~p(s(bool,happ(s(fun(bool,bool),X408),s(bool,'T_0')))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X410,bool),X409),s(fun(bool,bool),X408)))) | s(bool,X411) = s(bool,sK32(X408,X409,X410,bool)) | s(bool,'T_0') = s(bool,X411)) )),
  inference(superposition,[],[f433,f836])).
fof(f4196,plain,(
  ( ! [X405,X407,X403,X404,X406] : (s(bool,'T_0') != s(bool,happ(s(fun(X405,bool),X404),s(X405,X406))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X405,bool),X404),s(fun(bool,bool),X403)))) | s(bool,X407) = s(bool,sK32(X403,X404,X405,bool)) | s(bool,'T_0') = s(bool,X407)) )),
  inference(superposition,[],[f434,f836])).
fof(f4195,plain,(
  ( ! [X399,X401,X400,X402] : (p(s(bool,happ(s(fun(bool,bool),X399),s(bool,'T_0')))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X401,bool),X400),s(fun(bool,bool),X399)))) | s(bool,X402) = s(bool,sK30(X399,X400,X401,bool)) | s(bool,'T_0') = s(bool,X402)) )),
  inference(superposition,[],[f430,f836])).
fof(f4194,plain,(
  ( ! [X397,X395,X396,X398] : (p(s(bool,happ(s(fun(bool,bool),X395),s(bool,'T_0')))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X397,bool),X396),s(fun(bool,bool),X395)))) | s(bool,X398) = s(bool,sK28(X395,X396,X397,bool)) | s(bool,'T_0') = s(bool,X398)) )),
  inference(superposition,[],[f427,f836])).
fof(f4193,plain,(
  ( ! [X391,X393,X390,X392,X394] : (s(bool,'T_0') != s(bool,happ(s(fun(X392,bool),X391),s(X392,X393))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X392,bool),X391),s(fun(bool,bool),X390)))) | s(bool,X394) = s(bool,sK28(X390,X391,X392,bool)) | s(bool,'T_0') = s(bool,X394)) )),
  inference(superposition,[],[f428,f836])).
fof(f4192,plain,(
  ( ! [X389,X387,X388,X386] : (~p(s(bool,happ(s(fun(X388,bool),X386),s(X388,happ(s(fun(bool,X388),X387),s(bool,'T_0')))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(bool,X388),X387),s(fun(X388,bool),X386)))) | s(bool,X389) = s(bool,sK26(X386,X387,bool,X388)) | s(bool,'T_0') = s(bool,X389)) )),
  inference(superposition,[],[f425,f836])).
fof(f4191,plain,(
  ( ! [X383,X385,X382,X384] : (s(X384,f(s(bool,'T_0'))) != s(X384,g(s(bool,'T_0'))) | s(fun(bool,X384),X382) = s(fun(bool,X384),X383) | s(bool,X385) = s(bool,sK25(X382,X383,X384,bool)) | s(bool,'T_0') = s(bool,X385)) )),
  inference(superposition,[],[f423,f836])).
fof(f4190,plain,(
  ( ! [X379,X380,X378,X381] : (~p(s(bool,happ(s(fun(bool,bool),X378),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X378),s(bool,X380)))) | p(s(bool,X379)) | s(bool,X381) = s(bool,sK24(X378,X379,bool)) | s(bool,'T_0') = s(bool,X381)) )),
  inference(superposition,[],[f382,f836])).
fof(f4189,plain,(
  ( ! [X377,X374,X376,X375] : (~p(s(bool,happ(s(fun(bool,bool),X374),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X374),s(bool,X376)))) | p(s(bool,X375)) | s(bool,X377) = s(bool,sK23(X374,X375,bool)) | s(bool,'T_0') = s(bool,X377)) )),
  inference(superposition,[],[f380,f836])).
fof(f4188,plain,(
  ( ! [X372,X370,X373,X371] : (~p(s(bool,happ(s(fun(bool,bool),X370),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X370),s(bool,X372)))) | ~p(s(bool,X371)) | s(bool,X373) = s(bool,sK22(X370,X371,bool)) | s(bool,'T_0') = s(bool,X373)) )),
  inference(superposition,[],[f379,f836])).
fof(f4187,plain,(
  ( ! [X366,X368,X367,X369] : (~p(s(bool,happ(s(fun(bool,bool),X366),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X366),s(bool,X368)))) | ~p(s(bool,X367)) | s(bool,X369) = s(bool,sK21(X366,X367,bool)) | s(bool,'T_0') = s(bool,X369)) )),
  inference(superposition,[],[f377,f836])).
fof(f4186,plain,(
  ( ! [X364,X365,X363] : (s(bool,'T_0') != s(bool,sK19(X363,bool)) | p(s(bool,'const.bool.?|33|_1'(s(fun(bool,bool),X363)))) | ~p(s(bool,happ(s(fun(bool,bool),X363),s(bool,X364)))) | s(bool,X365) = s(bool,sK20(X363,bool)) | s(bool,'T_0') = s(bool,X365)) )),
  inference(superposition,[],[f371,f836])).
fof(f4185,plain,(
  ( ! [X360,X361] : (s(bool,'|dollar|exists_1'(s(fun(bool,bool),X360))) = s(bool,happ(s(fun(bool,bool),X360),s(bool,'T_0'))) | s(bool,X361) = s(bool,'const.min.@_1'(s(fun(bool,bool),X360))) | s(bool,'T_0') = s(bool,X361)) )),
  inference(superposition,[],[f364,f836])).
fof(f4184,plain,(
  ( ! [X356,X358,X357,X359] : (p(s(bool,happ(s(fun(bool,bool),X358),s(bool,'T_0')))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X356,bool),X357),s(fun(bool,bool),X358)))) | s(bool,X359) = s(bool,happ(s(fun(X356,bool),X357),s(X356,sK35(X358,X357,X356,bool)))) | s(bool,'T_0') = s(bool,X359)) )),
  inference(superposition,[],[f439,f836])).
fof(f4183,plain,(
  ( ! [X352,X354,X353,X355] : (~p(s(bool,happ(s(fun(bool,bool),X354),s(bool,'T_0')))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X352,bool),X353),s(fun(bool,bool),X354)))) | s(bool,X355) = s(bool,happ(s(fun(X352,bool),X353),s(X352,sK26(X354,X353,X352,bool)))) | s(bool,'T_0') = s(bool,X355)) )),
  inference(superposition,[],[f425,f836])).
fof(f4182,plain,(
  ( ! [X350,X349,X351,X347,X348] : (s(bool,'T_0') != s(bool,sK28(X350,X348,X347,bool)) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X347,bool),X348),s(fun(bool,bool),X350)))) | s(bool,X351) = s(bool,happ(s(fun(X347,bool),X348),s(X347,X349))) | s(bool,'T_0') = s(bool,X351)) )),
  inference(superposition,[],[f428,f836])).
fof(f4181,plain,(
  ( ! [X344,X346,X343,X345,X342] : (s(bool,'T_0') != s(bool,sK32(X345,X343,X342,bool)) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X342,bool),X343),s(fun(bool,bool),X345)))) | s(bool,X346) = s(bool,happ(s(fun(X342,bool),X343),s(X342,X344))) | s(bool,'T_0') = s(bool,X346)) )),
  inference(superposition,[],[f434,f836])).
fof(f4180,plain,(
  ( ! [X341,X337,X339,X340,X338] : (~p(s(bool,happ(s(fun(bool,bool),X340),s(bool,'T_0')))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X337,bool),X338),s(fun(bool,bool),X340)))) | s(bool,X341) = s(bool,happ(s(fun(X337,bool),X338),s(X337,X339))) | s(bool,'T_0') = s(bool,X341)) )),
  inference(superposition,[],[f431,f836])).
fof(f4179,plain,(
  ( ! [X333,X335,X332,X334,X336] : (p(s(bool,happ(s(fun(bool,bool),X335),s(bool,'T_0')))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X332,bool),X333),s(fun(bool,bool),X335)))) | s(bool,X336) = s(bool,happ(s(fun(X332,bool),X333),s(X332,X334))) | s(bool,'T_0') = s(bool,X336)) )),
  inference(superposition,[],[f437,f836])).
fof(f4177,plain,(
  ( ! [X325,X327,X324,X326] : (s(bool,'T_0') != s(bool,f(s(X324,sK25(X325,X326,bool,X324)))) | s(fun(X324,bool),X325) = s(fun(X324,bool),X326) | s(bool,X327) = s(bool,g(s(X324,sK25(X325,X326,bool,X324)))) | s(bool,'T_0') = s(bool,X327)) )),
  inference(superposition,[],[f423,f836])).
fof(f4671,plain,(
  ( ! [X269,X271,X273,X270,X272,X274] : (p(s(bool,happ(s(fun(X270,bool),X272),s(X270,X273)))) | s(X270,X273) = s(X270,happ(s(fun(X269,X270),X271),s(X269,sK31(X273,X272,X271,X269,X270)))) | s(bool,X274) = s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X269,X270),X271),s(fun(X270,bool),X272))) | s(bool,'T_0') = s(bool,X274)) )),
  inference(subsumption_resolution,[],[f4164,f322])).
fof(f4164,plain,(
  ( ! [X269,X271,X273,X270,X272,X274] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X270,bool),X272),s(X270,X273)))) | s(X270,X273) = s(X270,happ(s(fun(X269,X270),X271),s(X269,sK31(X273,X272,X271,X269,X270)))) | s(bool,X274) = s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X269,X270),X271),s(fun(X270,bool),X272))) | s(bool,'T_0') = s(bool,X274)) )),
  inference(superposition,[],[f432,f836])).
fof(f4670,plain,(
  ( ! [X263,X265,X267,X268,X264,X266] : (~p(s(bool,happ(s(fun(X264,bool),X266),s(X264,X267)))) | s(X264,X267) = s(X264,happ(s(fun(X263,X264),X265),s(X263,sK27(X267,X266,X265,X263,X264)))) | s(bool,X268) = s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X263,X264),X265),s(fun(X264,bool),X266))) | s(bool,'T_0') = s(bool,X268)) )),
  inference(subsumption_resolution,[],[f4163,f322])).
fof(f4163,plain,(
  ( ! [X263,X265,X267,X268,X264,X266] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X264,bool),X266),s(X264,X267)))) | s(X264,X267) = s(X264,happ(s(fun(X263,X264),X265),s(X263,sK27(X267,X266,X265,X263,X264)))) | s(bool,X268) = s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X263,X264),X265),s(fun(X264,bool),X266))) | s(bool,'T_0') = s(bool,X268)) )),
  inference(superposition,[],[f426,f836])).
fof(f4669,plain,(
  ( ! [X261,X257,X259,X260,X262,X258] : (~p(s(bool,happ(s(fun(X258,bool),X260),s(X258,happ(s(fun(X257,X258),X259),s(X257,X261)))))) | s(bool,X262) = s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X257,X258),X259),s(fun(X258,bool),X260))) | s(bool,'T_0') = s(bool,X262)) )),
  inference(subsumption_resolution,[],[f4162,f322])).
fof(f4162,plain,(
  ( ! [X261,X257,X259,X260,X262,X258] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X258,bool),X260),s(X258,happ(s(fun(X257,X258),X259),s(X257,X261)))))) | s(bool,X262) = s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X257,X258),X259),s(fun(X258,bool),X260))) | s(bool,'T_0') = s(bool,X262)) )),
  inference(superposition,[],[f438,f836])).
fof(f4668,plain,(
  ( ! [X255,X253,X251,X254,X252,X256] : (p(s(bool,happ(s(fun(X252,bool),X254),s(X252,X255)))) | ~p(s(bool,happ(s(fun(X252,bool),X254),s(X252,happ(s(fun(X251,X252),X253),s(X251,sK33(X255,X254,X253,X251,X252))))))) | s(bool,X256) = s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X251,X252),X253),s(fun(X252,bool),X254))) | s(bool,'T_0') = s(bool,X256)) )),
  inference(subsumption_resolution,[],[f4161,f322])).
fof(f4161,plain,(
  ( ! [X255,X253,X251,X254,X252,X256] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X252,bool),X254),s(X252,X255)))) | ~p(s(bool,happ(s(fun(X252,bool),X254),s(X252,happ(s(fun(X251,X252),X253),s(X251,sK33(X255,X254,X253,X251,X252))))))) | s(bool,X256) = s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X251,X252),X253),s(fun(X252,bool),X254))) | s(bool,'T_0') = s(bool,X256)) )),
  inference(superposition,[],[f435,f836])).
fof(f4667,plain,(
  ( ! [X249,X246,X250,X248,X247,X245] : (~p(s(bool,happ(s(fun(X246,bool),X248),s(X246,X249)))) | p(s(bool,happ(s(fun(X246,bool),X248),s(X246,happ(s(fun(X245,X246),X247),s(X245,sK29(X249,X248,X247,X245,X246))))))) | s(bool,X250) = s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X245,X246),X247),s(fun(X246,bool),X248))) | s(bool,'T_0') = s(bool,X250)) )),
  inference(subsumption_resolution,[],[f4160,f322])).
fof(f4160,plain,(
  ( ! [X249,X246,X250,X248,X247,X245] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X246,bool),X248),s(X246,X249)))) | p(s(bool,happ(s(fun(X246,bool),X248),s(X246,happ(s(fun(X245,X246),X247),s(X245,sK29(X249,X248,X247,X245,X246))))))) | s(bool,X250) = s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X245,X246),X247),s(fun(X246,bool),X248))) | s(bool,'T_0') = s(bool,X250)) )),
  inference(superposition,[],[f429,f836])).
fof(f4666,plain,(
  ( ! [X239,X237,X235,X238,X236] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X238,X235),X237),s(fun(X235,bool),X236)))) | s(bool,X239) = s(bool,happ(s(fun(X235,bool),X236),s(X235,sK34(X236,X237,X238,X235)))) | s(bool,'T_0') = s(bool,X239)) )),
  inference(subsumption_resolution,[],[f4156,f322])).
fof(f4156,plain,(
  ( ! [X239,X237,X235,X238,X236] : (~p(s(bool,'T_0')) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X238,X235),X237),s(fun(X235,bool),X236)))) | s(bool,X239) = s(bool,happ(s(fun(X235,bool),X236),s(X235,sK34(X236,X237,X238,X235)))) | s(bool,'T_0') = s(bool,X239)) )),
  inference(superposition,[],[f436,f836])).
fof(f4665,plain,(
  ( ! [X231,X233,X230,X234,X232] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X233,X230),X232),s(fun(X230,bool),X231)))) | s(bool,X234) = s(bool,happ(s(fun(X230,bool),X231),s(X230,sK32(X231,X232,X233,X230)))) | s(bool,'T_0') = s(bool,X234)) )),
  inference(subsumption_resolution,[],[f4155,f322])).
fof(f4155,plain,(
  ( ! [X231,X233,X230,X234,X232] : (~p(s(bool,'T_0')) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X233,X230),X232),s(fun(X230,bool),X231)))) | s(bool,X234) = s(bool,happ(s(fun(X230,bool),X231),s(X230,sK32(X231,X232,X233,X230)))) | s(bool,'T_0') = s(bool,X234)) )),
  inference(superposition,[],[f433,f836])).
fof(f4664,plain,(
  ( ! [X218,X216,X215,X219,X217] : (p(s(bool,happ(s(fun(X215,bool),X216),s(X215,X218)))) | p(s(bool,X217)) | s(bool,X219) = s(bool,happ(s(fun(X215,bool),X216),s(X215,sK24(X216,X217,X215)))) | s(bool,'T_0') = s(bool,X219)) )),
  inference(subsumption_resolution,[],[f4152,f322])).
fof(f4152,plain,(
  ( ! [X218,X216,X215,X219,X217] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X215,bool),X216),s(X215,X218)))) | p(s(bool,X217)) | s(bool,X219) = s(bool,happ(s(fun(X215,bool),X216),s(X215,sK24(X216,X217,X215)))) | s(bool,'T_0') = s(bool,X219)) )),
  inference(superposition,[],[f382,f836])).
fof(f4663,plain,(
  ( ! [X212,X210,X213,X211,X214] : (p(s(bool,happ(s(fun(X210,bool),X211),s(X210,X213)))) | p(s(bool,X212)) | s(bool,X214) = s(bool,happ(s(fun(X210,bool),X211),s(X210,sK23(X211,X212,X210)))) | s(bool,'T_0') = s(bool,X214)) )),
  inference(subsumption_resolution,[],[f4151,f322])).
fof(f4151,plain,(
  ( ! [X212,X210,X213,X211,X214] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X210,bool),X211),s(X210,X213)))) | p(s(bool,X212)) | s(bool,X214) = s(bool,happ(s(fun(X210,bool),X211),s(X210,sK23(X211,X212,X210)))) | s(bool,'T_0') = s(bool,X214)) )),
  inference(superposition,[],[f380,f836])).
fof(f4662,plain,(
  ( ! [X208,X206,X209,X207,X205] : (p(s(bool,happ(s(fun(X205,bool),X206),s(X205,X208)))) | ~p(s(bool,X207)) | s(bool,X209) = s(bool,happ(s(fun(X205,bool),X206),s(X205,sK22(X206,X207,X205)))) | s(bool,'T_0') = s(bool,X209)) )),
  inference(subsumption_resolution,[],[f4150,f322])).
fof(f4150,plain,(
  ( ! [X208,X206,X209,X207,X205] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X205,bool),X206),s(X205,X208)))) | ~p(s(bool,X207)) | s(bool,X209) = s(bool,happ(s(fun(X205,bool),X206),s(X205,sK22(X206,X207,X205)))) | s(bool,'T_0') = s(bool,X209)) )),
  inference(superposition,[],[f379,f836])).
fof(f4661,plain,(
  ( ! [X204,X202,X200,X203,X201] : (p(s(bool,happ(s(fun(X200,bool),X201),s(X200,X203)))) | ~p(s(bool,X202)) | s(bool,X204) = s(bool,happ(s(fun(X200,bool),X201),s(X200,sK21(X201,X202,X200)))) | s(bool,'T_0') = s(bool,X204)) )),
  inference(subsumption_resolution,[],[f4149,f322])).
fof(f4149,plain,(
  ( ! [X204,X202,X200,X203,X201] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X200,bool),X201),s(X200,X203)))) | ~p(s(bool,X202)) | s(bool,X204) = s(bool,happ(s(fun(X200,bool),X201),s(X200,sK21(X201,X202,X200)))) | s(bool,'T_0') = s(bool,X204)) )),
  inference(superposition,[],[f377,f836])).
fof(f4660,plain,(
  ( ! [X191,X189,X187,X190,X188] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X189,X187),X190),s(fun(X187,bool),X188)))) | s(bool,X191) = s(bool,happ(s(fun(X187,bool),X188),s(X187,happ(s(fun(X189,X187),X190),s(X189,sK26(X188,X190,X189,X187)))))) | s(bool,'T_0') = s(bool,X191)) )),
  inference(subsumption_resolution,[],[f4146,f322])).
fof(f4146,plain,(
  ( ! [X191,X189,X187,X190,X188] : (~p(s(bool,'T_0')) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X189,X187),X190),s(fun(X187,bool),X188)))) | s(bool,X191) = s(bool,happ(s(fun(X187,bool),X188),s(X187,happ(s(fun(X189,X187),X190),s(X189,sK26(X188,X190,X189,X187)))))) | s(bool,'T_0') = s(bool,X191)) )),
  inference(superposition,[],[f425,f836])).
fof(f4659,plain,(
  ( ! [X185,X182,X186,X184,X183,X181] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X183,X181),X184),s(fun(X181,bool),X182)))) | s(bool,X186) = s(bool,happ(s(fun(X181,bool),X182),s(X181,happ(s(fun(X183,X181),X184),s(X183,X185))))) | s(bool,'T_0') = s(bool,X186)) )),
  inference(subsumption_resolution,[],[f4145,f322])).
fof(f4145,plain,(
  ( ! [X185,X182,X186,X184,X183,X181] : (~p(s(bool,'T_0')) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X183,X181),X184),s(fun(X181,bool),X182)))) | s(bool,X186) = s(bool,happ(s(fun(X181,bool),X182),s(X181,happ(s(fun(X183,X181),X184),s(X183,X185))))) | s(bool,'T_0') = s(bool,X186)) )),
  inference(superposition,[],[f431,f836])).
fof(f4658,plain,(
  ( ! [X173,X171,X174,X172] : (p(s(bool,happ(s(fun(X171,bool),X172),s(X171,sK19(X172,X171))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X171,bool),X172)))) | s(bool,X174) = s(bool,happ(s(fun(X171,bool),X172),s(X171,X173))) | s(bool,'T_0') = s(bool,X174)) )),
  inference(subsumption_resolution,[],[f4143,f322])).
fof(f4143,plain,(
  ( ! [X173,X171,X174,X172] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X171,bool),X172),s(X171,sK19(X172,X171))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X171,bool),X172)))) | s(bool,X174) = s(bool,happ(s(fun(X171,bool),X172),s(X171,X173))) | s(bool,'T_0') = s(bool,X174)) )),
  inference(superposition,[],[f369,f836])).
fof(f4657,plain,(
  ( ! [X167,X169,X170,X168] : (p(s(bool,happ(s(fun(X167,bool),X168),s(X167,sK20(X168,X167))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X167,bool),X168)))) | s(bool,X170) = s(bool,happ(s(fun(X167,bool),X168),s(X167,X169))) | s(bool,'T_0') = s(bool,X170)) )),
  inference(subsumption_resolution,[],[f4142,f322])).
fof(f4142,plain,(
  ( ! [X167,X169,X170,X168] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X167,bool),X168),s(X167,sK20(X168,X167))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X167,bool),X168)))) | s(bool,X170) = s(bool,happ(s(fun(X167,bool),X168),s(X167,X169))) | s(bool,'T_0') = s(bool,X170)) )),
  inference(superposition,[],[f370,f836])).
fof(f4655,plain,(
  ( ! [X163,X164,X162] : (p(s(bool,happ(s(fun(X162,bool),X163),s(X162,sK18(X163,X162))))) | s(bool,X164) = s(bool,'const.bool.?|33|_1'(s(fun(X162,bool),X163))) | s(bool,'T_0') = s(bool,X164)) )),
  inference(subsumption_resolution,[],[f4140,f322])).
fof(f4140,plain,(
  ( ! [X163,X164,X162] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X162,bool),X163),s(X162,sK18(X163,X162))))) | s(bool,X164) = s(bool,'const.bool.?|33|_1'(s(fun(X162,bool),X163))) | s(bool,'T_0') = s(bool,X164)) )),
  inference(superposition,[],[f367,f836])).
fof(f4654,plain,(
  ( ! [X158,X161,X159,X157,X160] : (~p(s(bool,happ(s(fun(X157,bool),X158),s(X157,X159)))) | ~p(s(bool,happ(s(fun(X157,bool),X158),s(X157,X160)))) | s(X157,X159) = s(X157,X160) | s(bool,X161) = s(bool,'const.bool.?|33|_1'(s(fun(X157,bool),X158))) | s(bool,'T_0') = s(bool,X161)) )),
  inference(subsumption_resolution,[],[f4139,f322])).
fof(f4139,plain,(
  ( ! [X158,X161,X159,X157,X160] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X157,bool),X158),s(X157,X159)))) | ~p(s(bool,happ(s(fun(X157,bool),X158),s(X157,X160)))) | s(X157,X159) = s(X157,X160) | s(bool,X161) = s(bool,'const.bool.?|33|_1'(s(fun(X157,bool),X158))) | s(bool,'T_0') = s(bool,X161)) )),
  inference(superposition,[],[f368,f836])).
fof(f4653,plain,(
  ( ! [X154,X152,X151,X149,X153,X150] : (p(s(bool,happ(s(fun(X150,bool),X152),s(X150,happ(s(fun(X149,X150),X151),s(X149,X153)))))) | s(bool,X154) = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X149,X150),X151),s(fun(X150,bool),X152))) | s(bool,'T_0') = s(bool,X154)) )),
  inference(subsumption_resolution,[],[f4136,f322])).
fof(f4136,plain,(
  ( ! [X154,X152,X151,X149,X153,X150] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X150,bool),X152),s(X150,happ(s(fun(X149,X150),X151),s(X149,X153)))))) | s(bool,X154) = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X149,X150),X151),s(fun(X150,bool),X152))) | s(bool,'T_0') = s(bool,X154)) )),
  inference(superposition,[],[f424,f836])).
fof(f4652,plain,(
  ( ! [X134,X138,X136,X135,X137] : (sP0(happ(s(fun(X135,bool),'_dst_x0_1'(s(bool,X134))),s(X135,X136)),X137) | s(bool,X134) = s(bool,X138) | s(bool,'T_0') = s(bool,X138)) )),
  inference(subsumption_resolution,[],[f4129,f322])).
fof(f4129,plain,(
  ( ! [X134,X138,X136,X135,X137] : (~p(s(bool,'T_0')) | sP0(happ(s(fun(X135,bool),'_dst_x0_1'(s(bool,X134))),s(X135,X136)),X137) | s(bool,X134) = s(bool,X138) | s(bool,'T_0') = s(bool,X138)) )),
  inference(superposition,[],[f1335,f836])).
fof(f4651,plain,(
  ( ! [X132,X130,X133,X131,X129] : (sP0(X130,happ(s(fun(X131,bool),'_dst_x0_1'(s(bool,X129))),s(X131,X132))) | s(bool,X129) = s(bool,X133) | s(bool,'T_0') = s(bool,X133)) )),
  inference(subsumption_resolution,[],[f4128,f322])).
fof(f4128,plain,(
  ( ! [X132,X130,X133,X131,X129] : (~p(s(bool,'T_0')) | sP0(X130,happ(s(fun(X131,bool),'_dst_x0_1'(s(bool,X129))),s(X131,X132))) | s(bool,X129) = s(bool,X133) | s(bool,'T_0') = s(bool,X133)) )),
  inference(superposition,[],[f1334,f836])).
fof(f4127,plain,(
  ( ! [X127,X125,X128,X126] : (sP6('T_0',X125,happ(s(fun(X126,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X126,X127))) | s(bool,X125) = s(bool,X128) | s(bool,'T_0') = s(bool,X128)) )),
  inference(superposition,[],[f1314,f836])).
fof(f4126,plain,(
  ( ! [X123,X121,X124,X122] : (~sP6('F_0',X121,happ(s(fun(X122,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X122,X123))) | s(bool,X121) = s(bool,X124) | s(bool,'T_0') = s(bool,X124)) )),
  inference(superposition,[],[f1313,f836])).
fof(f4125,plain,(
  ( ! [X118,X120,X119,X117] : (sP6('T_0',happ(s(fun(X118,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X118,X119)),X117) | s(bool,X117) = s(bool,X120) | s(bool,'T_0') = s(bool,X120)) )),
  inference(superposition,[],[f1311,f836])).
fof(f4124,plain,(
  ( ! [X116,X114,X115,X113] : (~sP6('F_0',happ(s(fun(X114,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X114,X115)),X113) | s(bool,X113) = s(bool,X116) | s(bool,'T_0') = s(bool,X116)) )),
  inference(superposition,[],[f1310,f836])).
fof(f4650,plain,(
  ( ! [X109,X107,X108,X106] : (~sP9(X106,X107,X108) | sP10(X106,X107,X108) | s(bool,X106) = s(bool,X109) | s(bool,'T_0') = s(bool,X109)) )),
  inference(subsumption_resolution,[],[f4122,f322])).
fof(f4122,plain,(
  ( ! [X109,X107,X108,X106] : (~p(s(bool,'T_0')) | ~sP9(X106,X107,X108) | sP10(X106,X107,X108) | s(bool,X106) = s(bool,X109) | s(bool,'T_0') = s(bool,X109)) )),
  inference(superposition,[],[f557,f836])).
fof(f4649,plain,(
  ( ! [X103,X105,X102,X104] : (~sP11(X102,X103,X104) | sP12(X102,X103,X104) | s(bool,X102) = s(bool,X105) | s(bool,'T_0') = s(bool,X105)) )),
  inference(subsumption_resolution,[],[f4121,f322])).
fof(f4121,plain,(
  ( ! [X103,X105,X102,X104] : (~p(s(bool,'T_0')) | ~sP11(X102,X103,X104) | sP12(X102,X103,X104) | s(bool,X102) = s(bool,X105) | s(bool,'T_0') = s(bool,X105)) )),
  inference(superposition,[],[f556,f836])).
fof(f4648,plain,(
  ( ! [X97,X95,X98,X96] : (~p(s(bool,X96)) | p(s(bool,X97)) | ~sP12(X95,X97,X96) | s(bool,X95) = s(bool,X98) | s(bool,'T_0') = s(bool,X98)) )),
  inference(subsumption_resolution,[],[f4119,f322])).
fof(f4119,plain,(
  ( ! [X97,X95,X98,X96] : (~p(s(bool,'T_0')) | ~p(s(bool,X96)) | p(s(bool,X97)) | ~sP12(X95,X97,X96) | s(bool,X95) = s(bool,X98) | s(bool,'T_0') = s(bool,X98)) )),
  inference(superposition,[],[f533,f836])).
fof(f4647,plain,(
  ( ! [X94,X92,X93,X91] : (~p(s(bool,X92)) | sP12(X91,X92,X93) | s(bool,X91) = s(bool,X94) | s(bool,'T_0') = s(bool,X94)) )),
  inference(subsumption_resolution,[],[f4118,f322])).
fof(f4118,plain,(
  ( ! [X94,X92,X93,X91] : (~p(s(bool,'T_0')) | ~p(s(bool,X92)) | sP12(X91,X92,X93) | s(bool,X91) = s(bool,X94) | s(bool,'T_0') = s(bool,X94)) )),
  inference(superposition,[],[f532,f836])).
fof(f4646,plain,(
  ( ! [X90,X88,X87,X89] : (p(s(bool,X88)) | sP12(X87,X89,X88) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(subsumption_resolution,[],[f4117,f322])).
fof(f4117,plain,(
  ( ! [X90,X88,X87,X89] : (~p(s(bool,'T_0')) | p(s(bool,X88)) | sP12(X87,X89,X88) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(superposition,[],[f531,f836])).
fof(f4645,plain,(
  ( ! [X85,X83,X86,X84] : (p(s(bool,X84)) | ~p(s(bool,X85)) | ~sP11(X85,X84,X83) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(subsumption_resolution,[],[f4116,f322])).
fof(f4116,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,'T_0')) | p(s(bool,X84)) | ~p(s(bool,X85)) | ~sP11(X85,X84,X83) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(superposition,[],[f526,f836])).
fof(f4644,plain,(
  ( ! [X80,X81,X79,X82] : (p(s(bool,X80)) | p(s(bool,X81)) | ~sP10(X79,X81,X80) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(subsumption_resolution,[],[f4115,f322])).
fof(f4115,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,'T_0')) | p(s(bool,X80)) | p(s(bool,X81)) | ~sP10(X79,X81,X80) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(superposition,[],[f516,f836])).
fof(f4643,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,X76)) | sP10(X75,X76,X77) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(subsumption_resolution,[],[f4114,f322])).
fof(f4114,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,'T_0')) | ~p(s(bool,X76)) | sP10(X75,X76,X77) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f515,f836])).
fof(f4642,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,X72)) | sP10(X71,X73,X72) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(subsumption_resolution,[],[f4113,f322])).
fof(f4113,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,'T_0')) | ~p(s(bool,X72)) | sP10(X71,X73,X72) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(superposition,[],[f514,f836])).
fof(f4641,plain,(
  ( ! [X70,X68,X69,X67] : (p(s(bool,X68)) | p(s(bool,X69)) | ~sP9(X67,X68,X69) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(subsumption_resolution,[],[f4112,f322])).
fof(f4112,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,'T_0')) | p(s(bool,X68)) | p(s(bool,X69)) | ~sP9(X67,X68,X69) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f509,f836])).
fof(f4640,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,X64)) | p(s(bool,X65)) | ~sP8(X65,X64,X63) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(subsumption_resolution,[],[f4111,f322])).
fof(f4111,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,'T_0')) | ~p(s(bool,X64)) | p(s(bool,X65)) | ~sP8(X65,X64,X63) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f501,f836])).
fof(f4639,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X60)) | p(s(bool,X61)) | ~sP7(X60,X59,X61) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(subsumption_resolution,[],[f4110,f322])).
fof(f4110,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,'T_0')) | ~p(s(bool,X60)) | p(s(bool,X61)) | ~sP7(X60,X59,X61) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f492,f836])).
fof(f4638,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,X56)) | p(s(bool,X57)) | ~sP4(X56,X55,X57) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(subsumption_resolution,[],[f4109,f322])).
fof(f4109,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,'T_0')) | ~p(s(bool,X56)) | p(s(bool,X57)) | ~sP4(X56,X55,X57) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(superposition,[],[f473,f836])).
fof(f4637,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X52)) | p(s(bool,X53)) | ~sP5(X52,X51,X53) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(subsumption_resolution,[],[f4108,f322])).
fof(f4108,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,'T_0')) | ~p(s(bool,X52)) | p(s(bool,X53)) | ~sP5(X52,X51,X53) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(superposition,[],[f469,f836])).
fof(f4636,plain,(
  ( ! [X47,X50,X48,X49] : (sP6(X48,X49,X47) | ~sP5(X48,X47,X49) | ~sP4(X49,X47,X48) | ~sP3(X47,X49,X48) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(subsumption_resolution,[],[f4107,f322])).
fof(f4107,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,'T_0')) | sP6(X48,X49,X47) | ~sP5(X48,X47,X49) | ~sP4(X49,X47,X48) | ~sP3(X47,X49,X48) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f466,f836])).
fof(f4635,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X44)) | p(s(bool,X45)) | ~sP6(X44,X43,X45) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f4106,f322])).
fof(f4106,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,'T_0')) | ~p(s(bool,X44)) | p(s(bool,X45)) | ~sP6(X44,X43,X45) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f465,f836])).
fof(f4634,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X40)) | sP2(X40,X41,X39) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(subsumption_resolution,[],[f4105,f322])).
fof(f4105,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,X40)) | sP2(X40,X41,X39) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f451,f836])).
fof(f4633,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X36)) | sP2(X36,X35,X37) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(subsumption_resolution,[],[f4104,f322])).
fof(f4104,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'T_0')) | ~p(s(bool,X36)) | sP2(X36,X35,X37) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f450,f836])).
fof(f4632,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X29)) | sP1(X29,X28) | s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(subsumption_resolution,[],[f4102,f322])).
fof(f4102,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'T_0')) | p(s(bool,X29)) | sP1(X29,X28) | s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f412,f836])).
fof(f4631,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X26)) | sP1(X25,X26) | s(bool,X25) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(subsumption_resolution,[],[f4101,f322])).
fof(f4101,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'T_0')) | p(s(bool,X26)) | sP1(X25,X26) | s(bool,X25) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f409,f836])).
fof(f4630,plain,(
  ( ! [X24,X23,X22] : (sP0(X22,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f4100,f322])).
fof(f4100,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | sP0(X22,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f394,f836])).
fof(f4629,plain,(
  ( ! [X21,X19,X20] : (sP0(X20,X19) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(subsumption_resolution,[],[f4099,f322])).
fof(f4099,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | sP0(X20,X19) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f393,f836])).
fof(f4628,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f4098,f322])).
fof(f4098,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f387,f836])).
fof(f4627,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f4094,f322])).
fof(f4094,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f340,f836])).
fof(f4075,plain,(
  ( ! [X134,X132,X133,X131] : (sP6('T_0','T_0',happ(s(fun(X132,bool),'_dst_x0_1'(s(bool,X131))),s(X132,X133))) | s(bool,X131) = s(bool,X134) | s(bool,'T_0') = s(bool,X134)) )),
  inference(superposition,[],[f1314,f836])).
fof(f4074,plain,(
  ( ! [X127,X130,X128,X129] : (~sP6('F_0','T_0',happ(s(fun(X128,bool),'_dst_x0_1'(s(bool,X127))),s(X128,X129))) | s(bool,X127) = s(bool,X130) | s(bool,'T_0') = s(bool,X130)) )),
  inference(superposition,[],[f1313,f836])).
fof(f4073,plain,(
  ( ! [X125,X123,X126,X124] : (sP6('T_0',happ(s(fun(X124,bool),'_dst_x0_1'(s(bool,X123))),s(X124,X125)),'T_0') | s(bool,X123) = s(bool,X126) | s(bool,'T_0') = s(bool,X126)) )),
  inference(superposition,[],[f1311,f836])).
fof(f4072,plain,(
  ( ! [X121,X122,X120,X119] : (~sP6('F_0',happ(s(fun(X120,bool),'_dst_x0_1'(s(bool,X119))),s(X120,X121)),'T_0') | s(bool,X119) = s(bool,X122) | s(bool,'T_0') = s(bool,X122)) )),
  inference(superposition,[],[f1310,f836])).
fof(f4039,plain,(
  ( ! [X4,X2,X3] : (s(bool,X2) != s(bool,X3) | p(s(bool,X3)) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f336,f836])).
fof(f3927,plain,(
  ( ! [X422,X424,X423,X425] : (p(s(bool,happ(s(fun(X424,bool),X422),s(X424,happ(s(fun(bool,X424),X423),s(bool,X425)))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(bool,X424),X423),s(fun(X424,bool),X422)))) | s(bool,'T_0') = s(bool,X425) | s(bool,'T_0') = s(bool,sK35(X422,X423,bool,X424))) )),
  inference(superposition,[],[f439,f836])).
fof(f3926,plain,(
  ( ! [X420,X418,X421,X419] : (~p(s(bool,happ(s(fun(bool,bool),X418),s(bool,X421)))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X420,bool),X419),s(fun(bool,bool),X418)))) | s(bool,'T_0') = s(bool,X421) | s(bool,'T_0') = s(bool,sK34(X418,X419,X420,bool))) )),
  inference(superposition,[],[f436,f836])).
fof(f3925,plain,(
  ( ! [X414,X416,X415,X417] : (~p(s(bool,happ(s(fun(bool,bool),X414),s(bool,X417)))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X416,bool),X415),s(fun(bool,bool),X414)))) | s(bool,'T_0') = s(bool,X417) | s(bool,'T_0') = s(bool,sK32(X414,X415,X416,bool))) )),
  inference(superposition,[],[f433,f836])).
fof(f3924,plain,(
  ( ! [X412,X410,X413,X409,X411] : (s(bool,X412) != s(bool,happ(s(fun(X411,bool),X410),s(X411,X413))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X411,bool),X410),s(fun(bool,bool),X409)))) | s(bool,'T_0') = s(bool,X412) | s(bool,'T_0') = s(bool,sK32(X409,X410,X411,bool))) )),
  inference(superposition,[],[f434,f836])).
fof(f3923,plain,(
  ( ! [X408,X405,X407,X406] : (p(s(bool,happ(s(fun(bool,bool),X405),s(bool,X408)))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X407,bool),X406),s(fun(bool,bool),X405)))) | s(bool,'T_0') = s(bool,X408) | s(bool,'T_0') = s(bool,sK30(X405,X406,X407,bool))) )),
  inference(superposition,[],[f430,f836])).
fof(f3922,plain,(
  ( ! [X401,X403,X404,X402] : (p(s(bool,happ(s(fun(bool,bool),X401),s(bool,X404)))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X403,bool),X402),s(fun(bool,bool),X401)))) | s(bool,'T_0') = s(bool,X404) | s(bool,'T_0') = s(bool,sK28(X401,X402,X403,bool))) )),
  inference(superposition,[],[f427,f836])).
fof(f3921,plain,(
  ( ! [X397,X399,X396,X398,X400] : (s(bool,X399) != s(bool,happ(s(fun(X398,bool),X397),s(X398,X400))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X398,bool),X397),s(fun(bool,bool),X396)))) | s(bool,'T_0') = s(bool,X399) | s(bool,'T_0') = s(bool,sK28(X396,X397,X398,bool))) )),
  inference(superposition,[],[f428,f836])).
fof(f3920,plain,(
  ( ! [X393,X395,X392,X394] : (~p(s(bool,happ(s(fun(X394,bool),X392),s(X394,happ(s(fun(bool,X394),X393),s(bool,X395)))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(bool,X394),X393),s(fun(X394,bool),X392)))) | s(bool,'T_0') = s(bool,X395) | s(bool,'T_0') = s(bool,sK26(X392,X393,bool,X394))) )),
  inference(superposition,[],[f425,f836])).
fof(f3919,plain,(
  ( ! [X389,X391,X388,X390] : (s(X390,f(s(bool,X391))) != s(X390,g(s(bool,X391))) | s(fun(bool,X390),X388) = s(fun(bool,X390),X389) | s(bool,'T_0') = s(bool,X391) | s(bool,'T_0') = s(bool,sK25(X388,X389,X390,bool))) )),
  inference(superposition,[],[f423,f836])).
fof(f3918,plain,(
  ( ! [X385,X387,X384,X386] : (~p(s(bool,happ(s(fun(bool,bool),X384),s(bool,X386)))) | p(s(bool,happ(s(fun(bool,bool),X384),s(bool,X387)))) | p(s(bool,X385)) | s(bool,'T_0') = s(bool,X386) | s(bool,'T_0') = s(bool,sK24(X384,X385,bool))) )),
  inference(superposition,[],[f382,f836])).
fof(f3917,plain,(
  ( ! [X383,X380,X382,X381] : (~p(s(bool,happ(s(fun(bool,bool),X380),s(bool,X382)))) | p(s(bool,happ(s(fun(bool,bool),X380),s(bool,X383)))) | p(s(bool,X381)) | s(bool,'T_0') = s(bool,X382) | s(bool,'T_0') = s(bool,sK23(X380,X381,bool))) )),
  inference(superposition,[],[f380,f836])).
fof(f3916,plain,(
  ( ! [X377,X379,X376,X378] : (~p(s(bool,happ(s(fun(bool,bool),X376),s(bool,X378)))) | p(s(bool,happ(s(fun(bool,bool),X376),s(bool,X379)))) | ~p(s(bool,X377)) | s(bool,'T_0') = s(bool,X378) | s(bool,'T_0') = s(bool,sK22(X376,X377,bool))) )),
  inference(superposition,[],[f379,f836])).
fof(f3915,plain,(
  ( ! [X372,X374,X373,X375] : (~p(s(bool,happ(s(fun(bool,bool),X372),s(bool,X374)))) | p(s(bool,happ(s(fun(bool,bool),X372),s(bool,X375)))) | ~p(s(bool,X373)) | s(bool,'T_0') = s(bool,X374) | s(bool,'T_0') = s(bool,sK21(X372,X373,bool))) )),
  inference(superposition,[],[f377,f836])).
fof(f3914,plain,(
  ( ! [X370,X369,X371] : (s(bool,X370) != s(bool,sK19(X369,bool)) | p(s(bool,'const.bool.?|33|_1'(s(fun(bool,bool),X369)))) | ~p(s(bool,happ(s(fun(bool,bool),X369),s(bool,X371)))) | s(bool,'T_0') = s(bool,X370) | s(bool,'T_0') = s(bool,sK20(X369,bool))) )),
  inference(superposition,[],[f371,f836])).
fof(f3913,plain,(
  ( ! [X366,X365] : (s(bool,'|dollar|exists_1'(s(fun(bool,bool),X365))) = s(bool,happ(s(fun(bool,bool),X365),s(bool,X366))) | s(bool,'T_0') = s(bool,X366) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X365)))) )),
  inference(superposition,[],[f364,f836])).
fof(f3912,plain,(
  ( ! [X364,X362,X361,X363] : (p(s(bool,happ(s(fun(bool,bool),X363),s(bool,X364)))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X361,bool),X362),s(fun(bool,bool),X363)))) | s(bool,'T_0') = s(bool,X364) | s(bool,'T_0') = s(bool,happ(s(fun(X361,bool),X362),s(X361,sK35(X363,X362,X361,bool))))) )),
  inference(superposition,[],[f439,f836])).
fof(f3911,plain,(
  ( ! [X358,X360,X357,X359] : (~p(s(bool,happ(s(fun(bool,bool),X359),s(bool,X360)))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X357,bool),X358),s(fun(bool,bool),X359)))) | s(bool,'T_0') = s(bool,X360) | s(bool,'T_0') = s(bool,happ(s(fun(X357,bool),X358),s(X357,sK26(X359,X358,X357,bool))))) )),
  inference(superposition,[],[f425,f836])).
fof(f3910,plain,(
  ( ! [X356,X352,X354,X353,X355] : (s(bool,X355) != s(bool,sK28(X356,X353,X352,bool)) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X352,bool),X353),s(fun(bool,bool),X356)))) | s(bool,'T_0') = s(bool,X355) | s(bool,'T_0') = s(bool,happ(s(fun(X352,bool),X353),s(X352,X354)))) )),
  inference(superposition,[],[f428,f836])).
fof(f3909,plain,(
  ( ! [X350,X349,X351,X347,X348] : (s(bool,X350) != s(bool,sK32(X351,X348,X347,bool)) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X347,bool),X348),s(fun(bool,bool),X351)))) | s(bool,'T_0') = s(bool,X350) | s(bool,'T_0') = s(bool,happ(s(fun(X347,bool),X348),s(X347,X349)))) )),
  inference(superposition,[],[f434,f836])).
fof(f3908,plain,(
  ( ! [X344,X346,X343,X345,X342] : (~p(s(bool,happ(s(fun(bool,bool),X346),s(bool,X345)))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X342,bool),X343),s(fun(bool,bool),X346)))) | s(bool,'T_0') = s(bool,X345) | s(bool,'T_0') = s(bool,happ(s(fun(X342,bool),X343),s(X342,X344)))) )),
  inference(superposition,[],[f431,f836])).
fof(f3907,plain,(
  ( ! [X341,X337,X339,X340,X338] : (p(s(bool,happ(s(fun(bool,bool),X341),s(bool,X340)))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X337,bool),X338),s(fun(bool,bool),X341)))) | s(bool,'T_0') = s(bool,X340) | s(bool,'T_0') = s(bool,happ(s(fun(X337,bool),X338),s(X337,X339)))) )),
  inference(superposition,[],[f437,f836])).
fof(f3905,plain,(
  ( ! [X329,X331,X332,X330] : (s(bool,X332) != s(bool,f(s(X329,sK25(X330,X331,bool,X329)))) | s(fun(X329,bool),X330) = s(fun(X329,bool),X331) | s(bool,'T_0') = s(bool,X332) | s(bool,'T_0') = s(bool,g(s(X329,sK25(X330,X331,bool,X329))))) )),
  inference(superposition,[],[f423,f836])).
fof(f3854,plain,(
  ( ! [X127,X125,X128,X126] : (sP6('T_0',X125,happ(s(fun(X127,bool),'_dst_x0_1'(s(bool,X126))),s(X127,X128))) | s(bool,'T_0') = s(bool,X126) | s(bool,'T_0') = s(bool,X125)) )),
  inference(superposition,[],[f1314,f836])).
fof(f3853,plain,(
  ( ! [X123,X121,X124,X122] : (~sP6('F_0',X121,happ(s(fun(X123,bool),'_dst_x0_1'(s(bool,X122))),s(X123,X124))) | s(bool,'T_0') = s(bool,X122) | s(bool,'T_0') = s(bool,X121)) )),
  inference(superposition,[],[f1313,f836])).
fof(f3852,plain,(
  ( ! [X118,X120,X119,X117] : (sP6('T_0',happ(s(fun(X119,bool),'_dst_x0_1'(s(bool,X118))),s(X119,X120)),X117) | s(bool,'T_0') = s(bool,X118) | s(bool,'T_0') = s(bool,X117)) )),
  inference(superposition,[],[f1311,f836])).
fof(f3851,plain,(
  ( ! [X116,X114,X115,X113] : (~sP6('F_0',happ(s(fun(X115,bool),'_dst_x0_1'(s(bool,X114))),s(X115,X116)),X113) | s(bool,'T_0') = s(bool,X114) | s(bool,'T_0') = s(bool,X113)) )),
  inference(superposition,[],[f1310,f836])).
fof(f3802,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,'T_0') = s(bool,X23) | s(bool,X25) = s(bool,happ(s(fun(X22,bool),'_dst_x0_1'(s(bool,X23))),s(X22,X24))) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f836,f421])).
fof(f3800,plain,(
  ( ! [X19,X20,X18] : (s(bool,'T_0') = s(bool,X19) | s(bool,X20) = s(bool,'|dollar|exists_1'(s(fun(X18,bool),'_dst_x0_1'(s(bool,X19))))) | s(bool,'T_0') = s(bool,X20)) )),
  inference(superposition,[],[f836,f1519])).
fof(f3799,plain,(
  ( ! [X17,X15,X16] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X15,bool),X16))) | s(bool,X17) = s(bool,happ(s(fun(X15,bool),X16),s(X15,'const.min.@_1'(s(fun(X15,bool),X16))))) | s(bool,'T_0') = s(bool,X17)) )),
  inference(superposition,[],[f836,f364])).
fof(f3788,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f836,f836])).
fof(f3773,plain,(
  ( ! [X30,X31,X29,X32] : (s(bool,X30) = s(bool,X32) | s(bool,'T_0') = s(bool,X32) | s(bool,'T_0') = s(bool,happ(s(fun(X29,bool),'_dst_x0_1'(s(bool,X30))),s(X29,X31)))) )),
  inference(superposition,[],[f836,f421])).
fof(f3771,plain,(
  ( ! [X26,X27,X25] : (s(bool,X26) = s(bool,X27) | s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X25,bool),'_dst_x0_1'(s(bool,X26)))))) )),
  inference(superposition,[],[f836,f1519])).
fof(f3770,plain,(
  ( ! [X24,X23,X22] : (s(bool,X24) = s(bool,'|dollar|exists_1'(s(fun(X22,bool),X23))) | s(bool,'T_0') = s(bool,X24) | s(bool,'T_0') = s(bool,happ(s(fun(X22,bool),X23),s(X22,'const.min.@_1'(s(fun(X22,bool),X23)))))) )),
  inference(superposition,[],[f836,f364])).
fof(f3726,plain,(
  ( ! [X2,X0,X1] : (sP0(happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X0,X1)),X2)) )),
  inference(unit_resulting_resolution,[],[f322,f1335])).
fof(f3708,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X1,X2)))) )),
  inference(unit_resulting_resolution,[],[f322,f1334])).
fof(f3361,plain,(
  ~sP1(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))),
  inference(unit_resulting_resolution,[],[f2082,f2082,f408])).
fof(f3384,plain,(
  ( ! [X0] : (sP2(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X0,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) )),
  inference(unit_resulting_resolution,[],[f2082,f2082,f451])).
fof(f3394,plain,(
  sP6('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))),
  inference(unit_resulting_resolution,[],[f588,f608,f577,f2082,f466])).
fof(f3395,plain,(
  sP6('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))),
  inference(unit_resulting_resolution,[],[f610,f582,f607,f2082,f466])).
fof(f3396,plain,(
  sP6('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))),
  inference(unit_resulting_resolution,[],[f610,f608,f607,f2082,f466])).
fof(f3405,plain,(
  ~sP5(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0')),
  inference(unit_resulting_resolution,[],[f2082,f573,f2082,f469])).
fof(f3420,plain,(
  ~sP4(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0')),
  inference(unit_resulting_resolution,[],[f2082,f573,f2082,f473])).
fof(f3430,plain,(
  ( ! [X0,X1] : (sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'|dollar|exists_1'(s(fun(X0,bool),'_dst_x0_1'(s(bool,X1)))),X1)) )),
  inference(unit_resulting_resolution,[],[f1519,f2082,f482])).
fof(f3431,plain,(
  sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'T_0')),
  inference(unit_resulting_resolution,[],[f621,f2082,f482])).
fof(f3434,plain,(
  ( ! [X0,X1] : (sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))),'|dollar|exists_1'(s(fun(X0,bool),X1)))) )),
  inference(unit_resulting_resolution,[],[f364,f2082,f482])).
fof(f3435,plain,(
  ( ! [X2,X0,X1] : (sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,X1))),s(X0,X2)),X1)) )),
  inference(unit_resulting_resolution,[],[f421,f2082,f482])).
fof(f3436,plain,(
  ( ! [X0,X1] : (sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X0,'|dollar|exists_1'(s(fun(X1,bool),'_dst_x0_1'(s(bool,X0)))))) )),
  inference(unit_resulting_resolution,[],[f1519,f2082,f482])).
fof(f3437,plain,(
  sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))),
  inference(unit_resulting_resolution,[],[f621,f2082,f482])).
fof(f3440,plain,(
  ( ! [X0,X1] : (sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'|dollar|exists_1'(s(fun(X0,bool),X1)),happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f364,f2082,f482])).
fof(f3441,plain,(
  ( ! [X2,X0,X1] : (sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X0,happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,X0))),s(X1,X2)))) )),
  inference(unit_resulting_resolution,[],[f421,f2082,f482])).
fof(f3452,plain,(
  sP8('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'T_0')),
  inference(unit_resulting_resolution,[],[f322,f612,f2082,f488])).
fof(f3458,plain,(
  sP8('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))),
  inference(unit_resulting_resolution,[],[f322,f612,f2082,f488])).
fof(f3459,plain,(
  sP8('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))),
  inference(unit_resulting_resolution,[],[f2082,f612,f2082,f488])).
fof(f3470,plain,(
  ~sP7(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0')),
  inference(unit_resulting_resolution,[],[f2082,f573,f2082,f492])).
fof(f3485,plain,(
  sP8('T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))),
  inference(unit_resulting_resolution,[],[f322,f2082,f2082,f498])).
fof(f3490,plain,(
  sP8(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'T_0')),
  inference(unit_resulting_resolution,[],[f2082,f322,f2082,f498])).
fof(f3491,plain,(
  sP8(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))),
  inference(unit_resulting_resolution,[],[f322,f2082,f2082,f498])).
fof(f3492,plain,(
  sP8(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))),
  inference(unit_resulting_resolution,[],[f2082,f2082,f2082,f498])).
fof(f3499,plain,(
  ~sP8('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))),
  inference(unit_resulting_resolution,[],[f2082,f573,f2082,f501])).
fof(f3527,plain,(
  ( ! [X0] : (sP10(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X0,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) )),
  inference(unit_resulting_resolution,[],[f2082,f2082,f514])).
fof(f3533,plain,(
  ( ! [X0] : (sP10(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X0)) )),
  inference(unit_resulting_resolution,[],[f2082,f2082,f515])).
fof(f3569,plain,(
  ~sP11(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))),
  inference(unit_resulting_resolution,[],[f573,f2082,f2082,f526])).
fof(f3585,plain,(
  ( ! [X0] : (sP12(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X0)) )),
  inference(unit_resulting_resolution,[],[f2082,f2082,f532])).
fof(f3591,plain,(
  ~sP12(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))),
  inference(unit_resulting_resolution,[],[f2082,f573,f2082,f533])).
fof(f3592,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))),
  inference(unit_resulting_resolution,[],[f573,f2082,f555])).
fof(f3599,plain,(
  ( ! [X0] : (sP12(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X0)) )),
  inference(unit_resulting_resolution,[],[f615,f2082,f556])).
fof(f3601,plain,(
  ( ! [X0] : (sP10(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X0,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) )),
  inference(unit_resulting_resolution,[],[f613,f2082,f557])).
fof(f3603,plain,(
  ( ! [X0] : (sP10(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X0)) )),
  inference(unit_resulting_resolution,[],[f614,f2082,f557])).
fof(f3605,plain,(
  ( ! [X0] : (sP2(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X0,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) )),
  inference(unit_resulting_resolution,[],[f2082,f707])).
fof(f3606,plain,(
  ( ! [X0] : (sP10('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X0,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) )),
  inference(unit_resulting_resolution,[],[f2082,f736])).
fof(f3607,plain,(
  ( ! [X0] : (sP10('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X0)) )),
  inference(unit_resulting_resolution,[],[f2082,f746])).
fof(f3608,plain,(
  ( ! [X0] : (sP12('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X0)) )),
  inference(unit_resulting_resolution,[],[f2082,f778])).
fof(f3610,plain,(
  ~sP5(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'F_0')),
  inference(unit_resulting_resolution,[],[f573,f2082,f1011])).
fof(f3613,plain,(
  ~sP4(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'F_0')),
  inference(unit_resulting_resolution,[],[f573,f2082,f1032])).
fof(f3616,plain,(
  ~sP7(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'F_0')),
  inference(unit_resulting_resolution,[],[f573,f2082,f1076])).
fof(f3619,plain,(
  ~sP8('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))),
  inference(unit_resulting_resolution,[],[f573,f2082,f1113])).
fof(f3620,plain,(
  ~sP11(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))),
  inference(unit_resulting_resolution,[],[f573,f2082,f1257])).
fof(f2997,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) )),
  inference(unit_resulting_resolution,[],[f322,f2142,f387])).
fof(f2998,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3))))) = s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))))) )),
  inference(unit_resulting_resolution,[],[f2142,f2142,f387])).
fof(f2999,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP0(X0,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X2),X3),s(fun(X2,bool),'_dst_x0_1'(s(bool,X4)))))) )),
  inference(unit_resulting_resolution,[],[f2142,f393])).
fof(f3000,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP0('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4)) )),
  inference(unit_resulting_resolution,[],[f2142,f394])).
fof(f3001,plain,(
  ( ! [X2,X0,X3,X1] : (~sP1('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f322,f2142,f408])).
fof(f3003,plain,(
  ( ! [X2,X0,X3,X1] : (~sP1('T_0','const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) )),
  inference(unit_resulting_resolution,[],[f322,f2142,f408])).
fof(f3004,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~sP1('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))))) )),
  inference(unit_resulting_resolution,[],[f2142,f2142,f408])).
fof(f3005,plain,(
  ( ! [X2,X0,X3,X1] : (sP1('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f573,f2142,f409])).
fof(f3006,plain,(
  ( ! [X2,X0,X3,X1] : (sP1('F_0','const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) )),
  inference(unit_resulting_resolution,[],[f573,f2142,f412])).
fof(f3021,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP2('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4,'T_0')) )),
  inference(unit_resulting_resolution,[],[f322,f2142,f451])).
fof(f3026,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP2('T_0',X0,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X2),X3),s(fun(X2,bool),'_dst_x0_1'(s(bool,X4)))))) )),
  inference(unit_resulting_resolution,[],[f322,f2142,f451])).
fof(f3027,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (sP2('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X5,X6),X7),s(fun(X6,bool),'_dst_x0_1'(s(bool,X8)))))) )),
  inference(unit_resulting_resolution,[],[f2142,f2142,f451])).
fof(f3033,plain,(
  ( ! [X2,X0,X3,X1] : (~sP6('T_0','const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f322,f573,f2142,f465])).
fof(f3034,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~sP6('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f2142,f573,f2142,f465])).
fof(f3037,plain,(
  ( ! [X2,X0,X3,X1] : (sP6('T_0','T_0','const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) )),
  inference(unit_resulting_resolution,[],[f587,f582,f577,f2142,f466])).
fof(f3039,plain,(
  ( ! [X2,X0,X3,X1] : (sP6('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'T_0','const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) )),
  inference(unit_resulting_resolution,[],[f609,f608,f577,f2142,f466])).
fof(f3041,plain,(
  ( ! [X2,X0,X3,X1] : (sP6('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) )),
  inference(unit_resulting_resolution,[],[f587,f582,f607,f2142,f466])).
fof(f3043,plain,(
  ( ! [X2,X0,X3,X1] : (sP6('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) )),
  inference(unit_resulting_resolution,[],[f609,f608,f607,f2142,f466])).
fof(f3044,plain,(
  ( ! [X2,X0,X3,X1] : (~sP5('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f322,f573,f2142,f469])).
fof(f3051,plain,(
  ( ! [X2,X0,X3,X1] : (~sP5('T_0','const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f322,f573,f2142,f469])).
fof(f3052,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~sP5('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f2142,f573,f2142,f469])).
fof(f3058,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP5(X0,X1,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X2,X3),X4),s(fun(X3,bool),'_dst_x0_1'(s(bool,X5)))))) )),
  inference(unit_resulting_resolution,[],[f2142,f470])).
fof(f3059,plain,(
  ( ! [X2,X0,X3,X1] : (~sP4('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f322,f573,f2142,f473])).
fof(f3066,plain,(
  ( ! [X2,X0,X3,X1] : (~sP4('T_0','const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f322,f573,f2142,f473])).
fof(f3067,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~sP4('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f2142,f573,f2142,f473])).
fof(f3073,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP4(X0,X1,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X2,X3),X4),s(fun(X3,bool),'_dst_x0_1'(s(bool,X5)))))) )),
  inference(unit_resulting_resolution,[],[f2142,f474])).
fof(f3074,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP3(X0,X1,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X2,X3),X4),s(fun(X3,bool),'_dst_x0_1'(s(bool,X5)))))) )),
  inference(unit_resulting_resolution,[],[f2142,f478])).
fof(f3075,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP3(X0,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X2),X3),s(fun(X2,bool),'_dst_x0_1'(s(bool,X4)))),X5)) )),
  inference(unit_resulting_resolution,[],[f2142,f479])).
fof(f3076,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP3('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4,X5)) )),
  inference(unit_resulting_resolution,[],[f2142,f480])).
fof(f3077,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP6('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'|dollar|exists_1'(s(fun(X4,bool),'_dst_x0_1'(s(bool,X5)))),X5)) )),
  inference(unit_resulting_resolution,[],[f1519,f2142,f482])).
fof(f3078,plain,(
  ( ! [X2,X0,X3,X1] : (sP6('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f621,f2142,f482])).
fof(f3081,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP6('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),happ(s(fun(X4,bool),X5),s(X4,'const.min.@_1'(s(fun(X4,bool),X5)))),'|dollar|exists_1'(s(fun(X4,bool),X5)))) )),
  inference(unit_resulting_resolution,[],[f364,f2142,f482])).
fof(f3082,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (sP6('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),happ(s(fun(X4,bool),'_dst_x0_1'(s(bool,X5))),s(X4,X6)),X5)) )),
  inference(unit_resulting_resolution,[],[f421,f2142,f482])).
fof(f3083,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP6('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4,'|dollar|exists_1'(s(fun(X5,bool),'_dst_x0_1'(s(bool,X4)))))) )),
  inference(unit_resulting_resolution,[],[f1519,f2142,f482])).
fof(f3084,plain,(
  ( ! [X2,X0,X3,X1] : (sP6('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) )),
  inference(unit_resulting_resolution,[],[f621,f2142,f482])).
fof(f3087,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP6('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'|dollar|exists_1'(s(fun(X4,bool),X5)),happ(s(fun(X4,bool),X5),s(X4,'const.min.@_1'(s(fun(X4,bool),X5)))))) )),
  inference(unit_resulting_resolution,[],[f364,f2142,f482])).
fof(f3088,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (sP6('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4,happ(s(fun(X5,bool),'_dst_x0_1'(s(bool,X4))),s(X5,X6)))) )),
  inference(unit_resulting_resolution,[],[f421,f2142,f482])).
fof(f3090,plain,(
  ( ! [X2,X0,X3,X1] : (~sP6('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f617,f2142,f483])).
fof(f3093,plain,(
  ( ! [X2,X0,X3,X1] : (~sP6('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f617,f2142,f483])).
fof(f3095,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP8('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4,'F_0')) )),
  inference(unit_resulting_resolution,[],[f573,f2142,f486])).
fof(f3096,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP8('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'F_0',X4)) )),
  inference(unit_resulting_resolution,[],[f573,f2142,f487])).
fof(f3099,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f322,f612,f2142,f488])).
fof(f3105,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'T_0','const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) )),
  inference(unit_resulting_resolution,[],[f322,f612,f2142,f488])).
fof(f3106,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (sP8('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))))) )),
  inference(unit_resulting_resolution,[],[f2142,f612,f2142,f488])).
fof(f3109,plain,(
  ( ! [X2,X0,X3,X1] : (~sP7('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f322,f573,f2142,f492])).
fof(f3116,plain,(
  ( ! [X2,X0,X3,X1] : (~sP7('T_0','const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f322,f573,f2142,f492])).
fof(f3117,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~sP7('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f2142,f573,f2142,f492])).
fof(f3123,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP7(X0,X1,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X2,X3),X4),s(fun(X3,bool),'_dst_x0_1'(s(bool,X5)))))) )),
  inference(unit_resulting_resolution,[],[f2142,f493])).
fof(f3124,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('T_0','const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f322,f322,f2142,f498])).
fof(f3130,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('T_0','T_0','const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) )),
  inference(unit_resulting_resolution,[],[f322,f322,f2142,f498])).
fof(f3132,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (sP8('T_0','const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))))) )),
  inference(unit_resulting_resolution,[],[f322,f2142,f2142,f498])).
fof(f3136,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'T_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f322,f322,f2142,f498])).
fof(f3137,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (sP8('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f2142,f322,f2142,f498])).
fof(f3138,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (sP8('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'T_0','const.quantHeuristics.GUESS_FORALL_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))))) )),
  inference(unit_resulting_resolution,[],[f322,f2142,f2142,f498])).
fof(f3139,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (sP8('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))),'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X8,X9),X10),s(fun(X9,bool),'_dst_x0_1'(s(bool,X11)))))) )),
  inference(unit_resulting_resolution,[],[f2142,f2142,f2142,f498])).
fof(f3140,plain,(
  ( ! [X2,X0,X3,X1] : (~sP8('F_0','const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f322,f573,f2142,f501])).
fof(f3145,plain,(
  ( ! [X2,X0,X3,X1] : (~sP8('F_0','T_0','const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) )),
  inference(unit_resulting_resolution,[],[f322,f573,f2142,f501])).
fof(f3146,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~sP8('F_0','const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))))) )),
  inference(unit_resulting_resolution,[],[f2142,f573,f2142,f501])).
fof(f3150,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP10('F_0',X0,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X2),X3),s(fun(X2,bool),'_dst_x0_1'(s(bool,X4)))))) )),
  inference(unit_resulting_resolution,[],[f573,f2142,f502])).
fof(f3154,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP10('F_0','const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4)) )),
  inference(unit_resulting_resolution,[],[f573,f2142,f503])).
fof(f3162,plain,(
  ( ! [X2,X0,X3,X1] : (~sP9('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f573,f573,f2142,f509])).
fof(f3167,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP9(X0,X1,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X2,X3),X4),s(fun(X3,bool),'_dst_x0_1'(s(bool,X5)))))) )),
  inference(unit_resulting_resolution,[],[f2142,f510])).
fof(f3168,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP9(X0,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X2),X3),s(fun(X2,bool),'_dst_x0_1'(s(bool,X4)))),X5)) )),
  inference(unit_resulting_resolution,[],[f2142,f511])).
fof(f3169,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10('T_0',X0,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X2),X3),s(fun(X2,bool),'_dst_x0_1'(s(bool,X4)))))) )),
  inference(unit_resulting_resolution,[],[f322,f2142,f514])).
fof(f3174,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (sP10('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X5,X6),X7),s(fun(X6,bool),'_dst_x0_1'(s(bool,X8)))))) )),
  inference(unit_resulting_resolution,[],[f2142,f2142,f514])).
fof(f3175,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10('T_0','const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4)) )),
  inference(unit_resulting_resolution,[],[f322,f2142,f515])).
fof(f3180,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (sP10('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))),X8)) )),
  inference(unit_resulting_resolution,[],[f2142,f2142,f515])).
fof(f3181,plain,(
  ( ! [X2,X0,X3,X1] : (~sP10('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f573,f573,f2142,f516])).
fof(f3182,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP12('F_0','const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4)) )),
  inference(unit_resulting_resolution,[],[f573,f2142,f520])).
fof(f3208,plain,(
  ( ! [X2,X0,X3,X1] : (~sP11('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f573,f322,f2142,f526])).
fof(f3215,plain,(
  ( ! [X2,X0,X3,X1] : (~sP11('T_0','F_0','const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) )),
  inference(unit_resulting_resolution,[],[f573,f322,f2142,f526])).
fof(f3216,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~sP11('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'F_0','const.quantHeuristics.GUESS_FORALL_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))))) )),
  inference(unit_resulting_resolution,[],[f573,f2142,f2142,f526])).
fof(f3222,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP11(X0,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X2),X3),s(fun(X2,bool),'_dst_x0_1'(s(bool,X4)))),X5)) )),
  inference(unit_resulting_resolution,[],[f2142,f528])).
fof(f3223,plain,(
  ( ! [X2,X0,X3,X1] : (sP12('F_0','F_0','const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) )),
  inference(unit_resulting_resolution,[],[f573,f573,f2142,f530])).
fof(f3226,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12('T_0','const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4)) )),
  inference(unit_resulting_resolution,[],[f322,f2142,f532])).
fof(f3232,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (sP12('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))),X8)) )),
  inference(unit_resulting_resolution,[],[f2142,f2142,f532])).
fof(f3233,plain,(
  ( ! [X2,X0,X3,X1] : (~sP12('T_0','F_0','const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) )),
  inference(unit_resulting_resolution,[],[f322,f573,f2142,f533])).
fof(f3237,plain,(
  ( ! [X2,X0,X3,X1] : (~sP12('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f322,f573,f2142,f533])).
fof(f3238,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~sP12('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'F_0','const.quantHeuristics.GUESS_FORALL_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))))) )),
  inference(unit_resulting_resolution,[],[f2142,f573,f2142,f533])).
fof(f3239,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) )),
  inference(unit_resulting_resolution,[],[f573,f2142,f555])).
fof(f3242,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4,'F_0')) )),
  inference(unit_resulting_resolution,[],[f599,f2142,f556])).
fof(f3245,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'T_0',X4)) )),
  inference(unit_resulting_resolution,[],[f601,f2142,f556])).
fof(f3246,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X4)) )),
  inference(unit_resulting_resolution,[],[f615,f2142,f556])).
fof(f3247,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4,'T_0')) )),
  inference(unit_resulting_resolution,[],[f595,f2142,f557])).
fof(f3248,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) )),
  inference(unit_resulting_resolution,[],[f613,f2142,f557])).
fof(f3249,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'T_0',X4)) )),
  inference(unit_resulting_resolution,[],[f596,f2142,f557])).
fof(f3250,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X4)) )),
  inference(unit_resulting_resolution,[],[f614,f2142,f557])).
fof(f3251,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP2(X0,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X2),X3),s(fun(X2,bool),'_dst_x0_1'(s(bool,X4)))),X5)) )),
  inference(unit_resulting_resolution,[],[f2142,f559])).
fof(f3252,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP2('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) )),
  inference(unit_resulting_resolution,[],[f2142,f707])).
fof(f3253,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X0,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X2),X3),s(fun(X2,bool),'_dst_x0_1'(s(bool,X4)))))) )),
  inference(unit_resulting_resolution,[],[f2142,f736])).
fof(f3254,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4)) )),
  inference(unit_resulting_resolution,[],[f2142,f746])).
fof(f3255,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4)) )),
  inference(unit_resulting_resolution,[],[f2142,f778])).
fof(f3256,plain,(
  ( ! [X2,X0,X3,X1] : (~sP6('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f573,f2142,f992])).
fof(f3257,plain,(
  ( ! [X2,X0,X3,X1] : (~sP5('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f573,f2142,f1011])).
fof(f3260,plain,(
  ( ! [X2,X0,X3,X1] : (~sP4('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f573,f2142,f1032])).
fof(f3263,plain,(
  ( ! [X2,X0,X3,X1] : (~sP7('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f573,f2142,f1076])).
fof(f3266,plain,(
  ( ! [X2,X0,X3,X1] : (~sP8('F_0','const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) )),
  inference(unit_resulting_resolution,[],[f573,f2142,f1113])).
fof(f3267,plain,(
  ( ! [X2,X0,X3,X1] : (~sP11('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) )),
  inference(unit_resulting_resolution,[],[f573,f2142,f1257])).
fof(f2565,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) )),
  inference(unit_resulting_resolution,[],[f322,f2111,f387])).
fof(f2566,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3))))) = s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))))) )),
  inference(unit_resulting_resolution,[],[f2111,f2111,f387])).
fof(f2567,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP0(X0,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X2),X3),s(fun(X2,bool),'_dst_x0_1'(s(bool,X4)))))) )),
  inference(unit_resulting_resolution,[],[f2111,f393])).
fof(f2568,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP0('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4)) )),
  inference(unit_resulting_resolution,[],[f2111,f394])).
fof(f2569,plain,(
  ( ! [X2,X0,X3,X1] : (~sP1('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f322,f2111,f408])).
fof(f2571,plain,(
  ( ! [X2,X0,X3,X1] : (~sP1('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) )),
  inference(unit_resulting_resolution,[],[f322,f2111,f408])).
fof(f2572,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~sP1('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))))) )),
  inference(unit_resulting_resolution,[],[f2111,f2111,f408])).
fof(f2573,plain,(
  ( ! [X2,X0,X3,X1] : (sP1('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f573,f2111,f409])).
fof(f2574,plain,(
  ( ! [X2,X0,X3,X1] : (sP1('F_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) )),
  inference(unit_resulting_resolution,[],[f573,f2111,f412])).
fof(f2589,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP2('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4,'T_0')) )),
  inference(unit_resulting_resolution,[],[f322,f2111,f451])).
fof(f2594,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP2('T_0',X0,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X2),X3),s(fun(X2,bool),'_dst_x0_1'(s(bool,X4)))))) )),
  inference(unit_resulting_resolution,[],[f322,f2111,f451])).
fof(f2595,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (sP2('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X5,X6),X7),s(fun(X6,bool),'_dst_x0_1'(s(bool,X8)))))) )),
  inference(unit_resulting_resolution,[],[f2111,f2111,f451])).
fof(f2601,plain,(
  ( ! [X2,X0,X3,X1] : (~sP6('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f322,f573,f2111,f465])).
fof(f2602,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~sP6('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f2111,f573,f2111,f465])).
fof(f2605,plain,(
  ( ! [X2,X0,X3,X1] : (sP6('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) )),
  inference(unit_resulting_resolution,[],[f587,f582,f577,f2111,f466])).
fof(f2607,plain,(
  ( ! [X2,X0,X3,X1] : (sP6('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) )),
  inference(unit_resulting_resolution,[],[f609,f608,f577,f2111,f466])).
fof(f2609,plain,(
  ( ! [X2,X0,X3,X1] : (sP6('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) )),
  inference(unit_resulting_resolution,[],[f587,f582,f607,f2111,f466])).
fof(f2611,plain,(
  ( ! [X2,X0,X3,X1] : (sP6('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) )),
  inference(unit_resulting_resolution,[],[f609,f608,f607,f2111,f466])).
fof(f2612,plain,(
  ( ! [X2,X0,X3,X1] : (~sP5('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f322,f573,f2111,f469])).
fof(f2619,plain,(
  ( ! [X2,X0,X3,X1] : (~sP5('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f322,f573,f2111,f469])).
fof(f2620,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~sP5('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f2111,f573,f2111,f469])).
fof(f2626,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP5(X0,X1,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X2,X3),X4),s(fun(X3,bool),'_dst_x0_1'(s(bool,X5)))))) )),
  inference(unit_resulting_resolution,[],[f2111,f470])).
fof(f2627,plain,(
  ( ! [X2,X0,X3,X1] : (~sP4('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f322,f573,f2111,f473])).
fof(f2634,plain,(
  ( ! [X2,X0,X3,X1] : (~sP4('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f322,f573,f2111,f473])).
fof(f2635,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~sP4('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f2111,f573,f2111,f473])).
fof(f2641,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP4(X0,X1,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X2,X3),X4),s(fun(X3,bool),'_dst_x0_1'(s(bool,X5)))))) )),
  inference(unit_resulting_resolution,[],[f2111,f474])).
fof(f2642,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP3(X0,X1,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X2,X3),X4),s(fun(X3,bool),'_dst_x0_1'(s(bool,X5)))))) )),
  inference(unit_resulting_resolution,[],[f2111,f478])).
fof(f2643,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP3(X0,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X2),X3),s(fun(X2,bool),'_dst_x0_1'(s(bool,X4)))),X5)) )),
  inference(unit_resulting_resolution,[],[f2111,f479])).
fof(f2644,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP3('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4,X5)) )),
  inference(unit_resulting_resolution,[],[f2111,f480])).
fof(f2645,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP6('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'|dollar|exists_1'(s(fun(X4,bool),'_dst_x0_1'(s(bool,X5)))),X5)) )),
  inference(unit_resulting_resolution,[],[f1519,f2111,f482])).
fof(f2646,plain,(
  ( ! [X2,X0,X3,X1] : (sP6('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f621,f2111,f482])).
fof(f2648,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP6('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),happ(s(fun(X4,bool),X5),s(X4,'const.min.@_1'(s(fun(X4,bool),X5)))),'|dollar|exists_1'(s(fun(X4,bool),X5)))) )),
  inference(unit_resulting_resolution,[],[f364,f2111,f482])).
fof(f2649,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (sP6('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),happ(s(fun(X4,bool),'_dst_x0_1'(s(bool,X5))),s(X4,X6)),X5)) )),
  inference(unit_resulting_resolution,[],[f421,f2111,f482])).
fof(f2650,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP6('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4,'|dollar|exists_1'(s(fun(X5,bool),'_dst_x0_1'(s(bool,X4)))))) )),
  inference(unit_resulting_resolution,[],[f1519,f2111,f482])).
fof(f2651,plain,(
  ( ! [X2,X0,X3,X1] : (sP6('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) )),
  inference(unit_resulting_resolution,[],[f621,f2111,f482])).
fof(f2653,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP6('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'|dollar|exists_1'(s(fun(X4,bool),X5)),happ(s(fun(X4,bool),X5),s(X4,'const.min.@_1'(s(fun(X4,bool),X5)))))) )),
  inference(unit_resulting_resolution,[],[f364,f2111,f482])).
fof(f2654,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (sP6('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4,happ(s(fun(X5,bool),'_dst_x0_1'(s(bool,X4))),s(X5,X6)))) )),
  inference(unit_resulting_resolution,[],[f421,f2111,f482])).
fof(f2656,plain,(
  ( ! [X2,X0,X3,X1] : (~sP6('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f617,f2111,f483])).
fof(f2658,plain,(
  ( ! [X2,X0,X3,X1] : (~sP6('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f617,f2111,f483])).
fof(f2659,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4,'F_0')) )),
  inference(unit_resulting_resolution,[],[f573,f2111,f486])).
fof(f2660,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'F_0',X4)) )),
  inference(unit_resulting_resolution,[],[f573,f2111,f487])).
fof(f2663,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f322,f612,f2111,f488])).
fof(f2669,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) )),
  inference(unit_resulting_resolution,[],[f322,f612,f2111,f488])).
fof(f2670,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (sP8('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))))) )),
  inference(unit_resulting_resolution,[],[f2111,f612,f2111,f488])).
fof(f2673,plain,(
  ( ! [X2,X0,X3,X1] : (~sP7('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f322,f573,f2111,f492])).
fof(f2680,plain,(
  ( ! [X2,X0,X3,X1] : (~sP7('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f322,f573,f2111,f492])).
fof(f2681,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~sP7('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f2111,f573,f2111,f492])).
fof(f2687,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP7(X0,X1,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X2,X3),X4),s(fun(X3,bool),'_dst_x0_1'(s(bool,X5)))))) )),
  inference(unit_resulting_resolution,[],[f2111,f493])).
fof(f2688,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f322,f322,f2111,f498])).
fof(f2694,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) )),
  inference(unit_resulting_resolution,[],[f322,f322,f2111,f498])).
fof(f2696,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (sP8('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))))) )),
  inference(unit_resulting_resolution,[],[f322,f2111,f2111,f498])).
fof(f2700,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'T_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f322,f322,f2111,f498])).
fof(f2701,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f2111,f322,f2111,f498])).
fof(f2702,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))))) )),
  inference(unit_resulting_resolution,[],[f322,f2111,f2111,f498])).
fof(f2703,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))),'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X8,X9),X10),s(fun(X9,bool),'_dst_x0_1'(s(bool,X11)))))) )),
  inference(unit_resulting_resolution,[],[f2111,f2111,f2111,f498])).
fof(f2704,plain,(
  ( ! [X2,X0,X3,X1] : (~sP8('F_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f322,f573,f2111,f501])).
fof(f2709,plain,(
  ( ! [X2,X0,X3,X1] : (~sP8('F_0','T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) )),
  inference(unit_resulting_resolution,[],[f322,f573,f2111,f501])).
fof(f2710,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~sP8('F_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))))) )),
  inference(unit_resulting_resolution,[],[f2111,f573,f2111,f501])).
fof(f2714,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP10('F_0',X0,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X2),X3),s(fun(X2,bool),'_dst_x0_1'(s(bool,X4)))))) )),
  inference(unit_resulting_resolution,[],[f573,f2111,f502])).
fof(f2718,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP10('F_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4)) )),
  inference(unit_resulting_resolution,[],[f573,f2111,f503])).
fof(f2726,plain,(
  ( ! [X2,X0,X3,X1] : (~sP9('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f573,f573,f2111,f509])).
fof(f2731,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP9(X0,X1,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X2,X3),X4),s(fun(X3,bool),'_dst_x0_1'(s(bool,X5)))))) )),
  inference(unit_resulting_resolution,[],[f2111,f510])).
fof(f2732,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP9(X0,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X2),X3),s(fun(X2,bool),'_dst_x0_1'(s(bool,X4)))),X5)) )),
  inference(unit_resulting_resolution,[],[f2111,f511])).
fof(f2733,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10('T_0',X0,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X2),X3),s(fun(X2,bool),'_dst_x0_1'(s(bool,X4)))))) )),
  inference(unit_resulting_resolution,[],[f322,f2111,f514])).
fof(f2738,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (sP10('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X5,X6),X7),s(fun(X6,bool),'_dst_x0_1'(s(bool,X8)))))) )),
  inference(unit_resulting_resolution,[],[f2111,f2111,f514])).
fof(f2739,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4)) )),
  inference(unit_resulting_resolution,[],[f322,f2111,f515])).
fof(f2744,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (sP10('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))),X8)) )),
  inference(unit_resulting_resolution,[],[f2111,f2111,f515])).
fof(f2745,plain,(
  ( ! [X2,X0,X3,X1] : (~sP10('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f573,f573,f2111,f516])).
fof(f2746,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP12('F_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4)) )),
  inference(unit_resulting_resolution,[],[f573,f2111,f520])).
fof(f2772,plain,(
  ( ! [X2,X0,X3,X1] : (~sP11('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f573,f322,f2111,f526])).
fof(f2779,plain,(
  ( ! [X2,X0,X3,X1] : (~sP11('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) )),
  inference(unit_resulting_resolution,[],[f573,f322,f2111,f526])).
fof(f2780,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~sP11('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'F_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))))) )),
  inference(unit_resulting_resolution,[],[f573,f2111,f2111,f526])).
fof(f2786,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP11(X0,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X2),X3),s(fun(X2,bool),'_dst_x0_1'(s(bool,X4)))),X5)) )),
  inference(unit_resulting_resolution,[],[f2111,f528])).
fof(f2787,plain,(
  ( ! [X2,X0,X3,X1] : (sP12('F_0','F_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) )),
  inference(unit_resulting_resolution,[],[f573,f573,f2111,f530])).
fof(f2790,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4)) )),
  inference(unit_resulting_resolution,[],[f322,f2111,f532])).
fof(f2796,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (sP12('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))),X8)) )),
  inference(unit_resulting_resolution,[],[f2111,f2111,f532])).
fof(f2797,plain,(
  ( ! [X2,X0,X3,X1] : (~sP12('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) )),
  inference(unit_resulting_resolution,[],[f322,f573,f2111,f533])).
fof(f2801,plain,(
  ( ! [X2,X0,X3,X1] : (~sP12('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f322,f573,f2111,f533])).
fof(f2802,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~sP12('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'F_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))))) )),
  inference(unit_resulting_resolution,[],[f2111,f573,f2111,f533])).
fof(f2803,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) )),
  inference(unit_resulting_resolution,[],[f573,f2111,f555])).
fof(f2806,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4,'F_0')) )),
  inference(unit_resulting_resolution,[],[f599,f2111,f556])).
fof(f2809,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'T_0',X4)) )),
  inference(unit_resulting_resolution,[],[f601,f2111,f556])).
fof(f2810,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X4)) )),
  inference(unit_resulting_resolution,[],[f615,f2111,f556])).
fof(f2811,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4,'T_0')) )),
  inference(unit_resulting_resolution,[],[f595,f2111,f557])).
fof(f2812,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) )),
  inference(unit_resulting_resolution,[],[f613,f2111,f557])).
fof(f2813,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'T_0',X4)) )),
  inference(unit_resulting_resolution,[],[f596,f2111,f557])).
fof(f2814,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X4)) )),
  inference(unit_resulting_resolution,[],[f614,f2111,f557])).
fof(f2815,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP2(X0,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X2),X3),s(fun(X2,bool),'_dst_x0_1'(s(bool,X4)))),X5)) )),
  inference(unit_resulting_resolution,[],[f2111,f559])).
fof(f2816,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP2('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) )),
  inference(unit_resulting_resolution,[],[f2111,f707])).
fof(f2817,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X0,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X2),X3),s(fun(X2,bool),'_dst_x0_1'(s(bool,X4)))))) )),
  inference(unit_resulting_resolution,[],[f2111,f736])).
fof(f2818,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4)) )),
  inference(unit_resulting_resolution,[],[f2111,f746])).
fof(f2819,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4)) )),
  inference(unit_resulting_resolution,[],[f2111,f778])).
fof(f2820,plain,(
  ( ! [X2,X0,X3,X1] : (~sP6('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f573,f2111,f992])).
fof(f2821,plain,(
  ( ! [X2,X0,X3,X1] : (~sP5('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f573,f2111,f1011])).
fof(f2824,plain,(
  ( ! [X2,X0,X3,X1] : (~sP4('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f573,f2111,f1032])).
fof(f2827,plain,(
  ( ! [X2,X0,X3,X1] : (~sP7('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f573,f2111,f1076])).
fof(f2830,plain,(
  ( ! [X2,X0,X3,X1] : (~sP8('F_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) )),
  inference(unit_resulting_resolution,[],[f573,f2111,f1113])).
fof(f2831,plain,(
  ( ! [X2,X0,X3,X1] : (~sP11('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) )),
  inference(unit_resulting_resolution,[],[f573,f2111,f1257])).
fof(f2533,plain,(
  ( ! [X125,X123,X124,X122] : (sP6('T_0','|dollar|exists_1'(s(fun(X122,bool),'_dst_x0_1'(s(bool,X123)))),happ(s(fun(X124,bool),'_dst_x0_1'(s(bool,X123))),s(X124,X125)))) )),
  inference(superposition,[],[f1314,f1519])).
fof(f2532,plain,(
  ( ! [X121,X118,X120,X119] : (~sP6('F_0','|dollar|exists_1'(s(fun(X118,bool),'_dst_x0_1'(s(bool,X119)))),happ(s(fun(X120,bool),'_dst_x0_1'(s(bool,X119))),s(X120,X121)))) )),
  inference(superposition,[],[f1313,f1519])).
fof(f2531,plain,(
  ( ! [X116,X114,X117,X115] : (sP6('T_0',happ(s(fun(X116,bool),'_dst_x0_1'(s(bool,X115))),s(X116,X117)),'|dollar|exists_1'(s(fun(X114,bool),'_dst_x0_1'(s(bool,X115)))))) )),
  inference(superposition,[],[f1311,f1519])).
fof(f2530,plain,(
  ( ! [X111,X112,X110,X113] : (~sP6('F_0',happ(s(fun(X112,bool),'_dst_x0_1'(s(bool,X111))),s(X112,X113)),'|dollar|exists_1'(s(fun(X110,bool),'_dst_x0_1'(s(bool,X111)))))) )),
  inference(superposition,[],[f1310,f1519])).
fof(f2529,plain,(
  ( ! [X109,X107,X108,X106] : (~p(s(bool,X107)) | ~sP9('|dollar|exists_1'(s(fun(X106,bool),'_dst_x0_1'(s(bool,X107)))),X108,X109) | sP10('|dollar|exists_1'(s(fun(X106,bool),'_dst_x0_1'(s(bool,X107)))),X108,X109)) )),
  inference(superposition,[],[f557,f1519])).
fof(f2528,plain,(
  ( ! [X103,X105,X102,X104] : (~p(s(bool,X103)) | ~sP11('|dollar|exists_1'(s(fun(X102,bool),'_dst_x0_1'(s(bool,X103)))),X104,X105) | sP12('|dollar|exists_1'(s(fun(X102,bool),'_dst_x0_1'(s(bool,X103)))),X104,X105)) )),
  inference(superposition,[],[f556,f1519])).
fof(f2526,plain,(
  ( ! [X97,X95,X98,X96] : (~p(s(bool,X96)) | ~p(s(bool,X97)) | p(s(bool,X98)) | ~sP12('|dollar|exists_1'(s(fun(X95,bool),'_dst_x0_1'(s(bool,X96)))),X98,X97)) )),
  inference(superposition,[],[f533,f1519])).
fof(f2525,plain,(
  ( ! [X94,X92,X93,X91] : (~p(s(bool,X92)) | ~p(s(bool,X93)) | sP12('|dollar|exists_1'(s(fun(X91,bool),'_dst_x0_1'(s(bool,X92)))),X93,X94)) )),
  inference(superposition,[],[f532,f1519])).
fof(f2524,plain,(
  ( ! [X90,X88,X87,X89] : (~p(s(bool,X88)) | p(s(bool,X89)) | sP12('|dollar|exists_1'(s(fun(X87,bool),'_dst_x0_1'(s(bool,X88)))),X90,X89)) )),
  inference(superposition,[],[f531,f1519])).
fof(f2523,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,X84)) | p(s(bool,X85)) | ~p(s(bool,X86)) | ~sP11(X86,X85,'|dollar|exists_1'(s(fun(X83,bool),'_dst_x0_1'(s(bool,X84)))))) )),
  inference(superposition,[],[f526,f1519])).
fof(f2522,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,X80)) | p(s(bool,X81)) | p(s(bool,X82)) | ~sP10('|dollar|exists_1'(s(fun(X79,bool),'_dst_x0_1'(s(bool,X80)))),X82,X81)) )),
  inference(superposition,[],[f516,f1519])).
fof(f2521,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,X76)) | ~p(s(bool,X77)) | sP10('|dollar|exists_1'(s(fun(X75,bool),'_dst_x0_1'(s(bool,X76)))),X77,X78)) )),
  inference(superposition,[],[f515,f1519])).
fof(f2520,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,X72)) | ~p(s(bool,X73)) | sP10('|dollar|exists_1'(s(fun(X71,bool),'_dst_x0_1'(s(bool,X72)))),X74,X73)) )),
  inference(superposition,[],[f514,f1519])).
fof(f2519,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,X68)) | p(s(bool,X69)) | p(s(bool,X70)) | ~sP9('|dollar|exists_1'(s(fun(X67,bool),'_dst_x0_1'(s(bool,X68)))),X69,X70)) )),
  inference(superposition,[],[f509,f1519])).
fof(f2518,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,X64)) | ~p(s(bool,X65)) | p(s(bool,X66)) | ~sP8(X66,X65,'|dollar|exists_1'(s(fun(X63,bool),'_dst_x0_1'(s(bool,X64)))))) )),
  inference(superposition,[],[f501,f1519])).
fof(f2517,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X60)) | ~p(s(bool,X61)) | p(s(bool,X62)) | ~sP7(X61,'|dollar|exists_1'(s(fun(X59,bool),'_dst_x0_1'(s(bool,X60)))),X62)) )),
  inference(superposition,[],[f492,f1519])).
fof(f2516,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,X56)) | ~p(s(bool,X57)) | p(s(bool,X58)) | ~sP4(X57,'|dollar|exists_1'(s(fun(X55,bool),'_dst_x0_1'(s(bool,X56)))),X58)) )),
  inference(superposition,[],[f473,f1519])).
fof(f2515,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X52)) | ~p(s(bool,X53)) | p(s(bool,X54)) | ~sP5(X53,'|dollar|exists_1'(s(fun(X51,bool),'_dst_x0_1'(s(bool,X52)))),X54)) )),
  inference(superposition,[],[f469,f1519])).
fof(f2514,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X48)) | sP6(X49,X50,'|dollar|exists_1'(s(fun(X47,bool),'_dst_x0_1'(s(bool,X48))))) | ~sP5(X49,'|dollar|exists_1'(s(fun(X47,bool),'_dst_x0_1'(s(bool,X48)))),X50) | ~sP4(X50,'|dollar|exists_1'(s(fun(X47,bool),'_dst_x0_1'(s(bool,X48)))),X49) | ~sP3('|dollar|exists_1'(s(fun(X47,bool),'_dst_x0_1'(s(bool,X48)))),X50,X49)) )),
  inference(superposition,[],[f466,f1519])).
fof(f2513,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X44)) | ~p(s(bool,X45)) | p(s(bool,X46)) | ~sP6(X45,'|dollar|exists_1'(s(fun(X43,bool),'_dst_x0_1'(s(bool,X44)))),X46)) )),
  inference(superposition,[],[f465,f1519])).
fof(f2512,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X40)) | ~p(s(bool,X41)) | sP2(X41,X42,'|dollar|exists_1'(s(fun(X39,bool),'_dst_x0_1'(s(bool,X40)))))) )),
  inference(superposition,[],[f451,f1519])).
fof(f2511,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X36)) | ~p(s(bool,X37)) | sP2(X37,'|dollar|exists_1'(s(fun(X35,bool),'_dst_x0_1'(s(bool,X36)))),X38)) )),
  inference(superposition,[],[f450,f1519])).
fof(f2509,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,X29)) | p(s(bool,X30)) | sP1(X30,'|dollar|exists_1'(s(fun(X28,bool),'_dst_x0_1'(s(bool,X29)))))) )),
  inference(superposition,[],[f412,f1519])).
fof(f2508,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,X26)) | p(s(bool,X27)) | sP1('|dollar|exists_1'(s(fun(X25,bool),'_dst_x0_1'(s(bool,X26)))),X27)) )),
  inference(superposition,[],[f409,f1519])).
fof(f2507,plain,(
  ( ! [X24,X23,X22] : (sP0('|dollar|exists_1'(s(fun(X22,bool),'_dst_x0_1'(s(bool,X23)))),X24) | ~p(s(bool,X23))) )),
  inference(superposition,[],[f394,f1519])).
fof(f2506,plain,(
  ( ! [X21,X19,X20] : (sP0(X21,'|dollar|exists_1'(s(fun(X19,bool),'_dst_x0_1'(s(bool,X20))))) | ~p(s(bool,X20))) )),
  inference(superposition,[],[f393,f1519])).
fof(f2497,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') = s(bool,X3) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X2,bool),'_dst_x0_1'(s(bool,X3)))))) )),
  inference(superposition,[],[f1519,f329])).
fof(f2485,plain,(
  ( ! [X0,X1] : (~sP6('F_0','|dollar|exists_1'(s(fun(X0,bool),'_dst_x0_1'(s(bool,X1)))),X1)) )),
  inference(unit_resulting_resolution,[],[f573,f1519,f484])).
fof(f2486,plain,(
  ( ! [X0,X1] : (sP6('T_0','|dollar|exists_1'(s(fun(X0,bool),'_dst_x0_1'(s(bool,X1)))),X1)) )),
  inference(unit_resulting_resolution,[],[f322,f1519,f482])).
fof(f2488,plain,(
  ( ! [X0,X1] : (~sP6('F_0',X0,'|dollar|exists_1'(s(fun(X1,bool),'_dst_x0_1'(s(bool,X0)))))) )),
  inference(unit_resulting_resolution,[],[f573,f1519,f484])).
fof(f2489,plain,(
  ( ! [X0,X1] : (sP6('T_0',X0,'|dollar|exists_1'(s(fun(X1,bool),'_dst_x0_1'(s(bool,X0)))))) )),
  inference(unit_resulting_resolution,[],[f322,f1519,f482])).
fof(f2464,plain,(
  ( ! [X12,X10,X13,X11] : (sP6('T_0',happ(s(fun(X10,bool),X11),s(X10,'const.min.@_1'(s(fun(X10,bool),X11)))),happ(s(fun(X12,bool),'_dst_x0_1'(s(bool,'|dollar|exists_1'(s(fun(X10,bool),X11))))),s(X12,X13)))) )),
  inference(superposition,[],[f1314,f364])).
fof(f2462,plain,(
  ( ! [X6,X7] : (sP6('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),happ(s(fun(X6,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X6,X7)))) )),
  inference(superposition,[],[f1314,f621])).
fof(f2461,plain,(
  ( ! [X4,X5,X3] : (sP6('T_0','F_0',happ(s(fun(X4,bool),'_dst_x0_1'(s(bool,X3))),s(X4,X5))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1314,f329])).
fof(f2460,plain,(
  ( ! [X2,X0,X1] : (sP6('T_0',X0,happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X1,X2))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1314,f329])).
fof(f2456,plain,(
  ( ! [X2,X0,X1] : (sP5('T_0',happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,X1))),s(X0,X2)),X1)) )),
  inference(unit_resulting_resolution,[],[f1314,f464])).
fof(f2436,plain,(
  ( ! [X12,X10,X13,X11] : (~sP6('F_0',happ(s(fun(X10,bool),X11),s(X10,'const.min.@_1'(s(fun(X10,bool),X11)))),happ(s(fun(X12,bool),'_dst_x0_1'(s(bool,'|dollar|exists_1'(s(fun(X10,bool),X11))))),s(X12,X13)))) )),
  inference(superposition,[],[f1313,f364])).
fof(f2434,plain,(
  ( ! [X6,X7] : (~sP6('F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),happ(s(fun(X6,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X6,X7)))) )),
  inference(superposition,[],[f1313,f621])).
fof(f2433,plain,(
  ( ! [X4,X5,X3] : (~sP6('F_0','F_0',happ(s(fun(X4,bool),'_dst_x0_1'(s(bool,X3))),s(X4,X5))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1313,f329])).
fof(f2432,plain,(
  ( ! [X2,X0,X1] : (~sP6('F_0',X0,happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X1,X2))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1313,f329])).
fof(f2427,plain,(
  ( ! [X0,X1] : (~sP4('T_0',happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X0,X1)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f573,f588,f577,f1313,f468])).
fof(f2444,plain,(
  ( ! [X0,X1] : (~sP4('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X0,X1)),'F_0')) )),
  inference(forward_demodulation,[],[f2428,f621])).
fof(f2428,plain,(
  ( ! [X0,X1] : (~sP4('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16))))),s(X0,X1)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f573,f610,f607,f1313,f468])).
fof(f2429,plain,(
  ( ! [X0,X1] : (~sP3(happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X0,X1)),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f573,f585,f580,f1313,f467])).
fof(f2408,plain,(
  ( ! [X14,X17,X15,X18,X16] : (sP6('T_0',happ(s(fun(X17,bool),'_dst_x0_1'(s(bool,X15))),s(X17,X18)),happ(s(fun(X14,bool),'_dst_x0_1'(s(bool,X15))),s(X14,X16)))) )),
  inference(superposition,[],[f1311,f421])).
fof(f2407,plain,(
  ( ! [X12,X10,X13,X11] : (sP6('T_0',happ(s(fun(X12,bool),'_dst_x0_1'(s(bool,'|dollar|exists_1'(s(fun(X10,bool),X11))))),s(X12,X13)),happ(s(fun(X10,bool),X11),s(X10,'const.min.@_1'(s(fun(X10,bool),X11)))))) )),
  inference(superposition,[],[f1311,f364])).
fof(f2405,plain,(
  ( ! [X6,X7] : (sP6('T_0',happ(s(fun(X6,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X6,X7)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) )),
  inference(superposition,[],[f1311,f621])).
fof(f2404,plain,(
  ( ! [X4,X5,X3] : (sP6('T_0',happ(s(fun(X4,bool),'_dst_x0_1'(s(bool,X3))),s(X4,X5)),'F_0') | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1311,f329])).
fof(f2403,plain,(
  ( ! [X2,X0,X1] : (sP6('T_0',happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X1,X2)),X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1311,f329])).
fof(f2400,plain,(
  ( ! [X2,X0,X1] : (sP5('T_0',X0,happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,X0))),s(X1,X2)))) )),
  inference(unit_resulting_resolution,[],[f1311,f464])).
fof(f2375,plain,(
  ( ! [X14,X17,X15,X18,X16] : (~sP6('F_0',happ(s(fun(X17,bool),'_dst_x0_1'(s(bool,X15))),s(X17,X18)),happ(s(fun(X14,bool),'_dst_x0_1'(s(bool,X15))),s(X14,X16)))) )),
  inference(superposition,[],[f1310,f421])).
fof(f2374,plain,(
  ( ! [X12,X10,X13,X11] : (~sP6('F_0',happ(s(fun(X12,bool),'_dst_x0_1'(s(bool,'|dollar|exists_1'(s(fun(X10,bool),X11))))),s(X12,X13)),happ(s(fun(X10,bool),X11),s(X10,'const.min.@_1'(s(fun(X10,bool),X11)))))) )),
  inference(superposition,[],[f1310,f364])).
fof(f2372,plain,(
  ( ! [X6,X7] : (~sP6('F_0',happ(s(fun(X6,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X6,X7)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) )),
  inference(superposition,[],[f1310,f621])).
fof(f2371,plain,(
  ( ! [X4,X5,X3] : (~sP6('F_0',happ(s(fun(X4,bool),'_dst_x0_1'(s(bool,X3))),s(X4,X5)),'F_0') | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1310,f329])).
fof(f2370,plain,(
  ( ! [X2,X0,X1] : (~sP6('F_0',happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X1,X2)),X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1310,f329])).
fof(f2386,plain,(
  ( ! [X0,X1] : (~sP4(happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X0,X1)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'F_0')) )),
  inference(forward_demodulation,[],[f2360,f621])).
fof(f2360,plain,(
  ( ! [X0,X1] : (~sP4(happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16))))),s(X0,X1)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f573,f611,f580,f1310,f468])).
fof(f2364,plain,(
  ( ! [X0,X1] : (~sP3('F_0',happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X0,X1)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f573,f583,f578,f1310,f468])).
fof(f2367,plain,(
  ( ! [X0,X1] : (~sP4(happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X0,X1)),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f322,f589,f580,f1310,f466])).
fof(f2349,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X4,X0),X5),s(fun(X0,bool),X2)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,happ(s(fun(X4,X0),X5),s(X4,sK33(X6,happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X5,X4,X0)))))))) )),
  inference(forward_demodulation,[],[f2348,f421])).
fof(f2348,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))) | ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X4,X0),X5),s(fun(X0,bool),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3))),s(X0,happ(s(fun(X4,X0),X5),s(X4,sK33(X6,happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X5,X4,X0)))))))) )),
  inference(forward_demodulation,[],[f2345,f421])).
fof(f2345,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X4,X0),X5),s(fun(X0,bool),X2)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3))),s(X0,X6)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3))),s(X0,happ(s(fun(X4,X0),X5),s(X4,sK33(X6,happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X5,X4,X0)))))))) )),
  inference(superposition,[],[f435,f421])).
fof(f2347,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X3),s(fun(X1,bool),X5)))) | ~p(s(bool,happ(s(fun(X1,bool),X5),s(X1,happ(s(fun(X0,X1),X3),s(X0,sK33(X6,X5,happ(s(fun(X2,fun(X0,X1)),'_dst_x0_1'(s(fun(X0,X1),X3))),s(X2,X4)),X0,X1))))))) | p(s(bool,happ(s(fun(X1,bool),X5),s(X1,X6))))) )),
  inference(forward_demodulation,[],[f2344,f421])).
fof(f2344,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X3),s(fun(X1,bool),X5)))) | p(s(bool,happ(s(fun(X1,bool),X5),s(X1,X6)))) | ~p(s(bool,happ(s(fun(X1,bool),X5),s(X1,happ(s(fun(X0,X1),happ(s(fun(X2,fun(X0,X1)),'_dst_x0_1'(s(fun(X0,X1),X3))),s(X2,X4))),s(X0,sK33(X6,X5,happ(s(fun(X2,fun(X0,X1)),'_dst_x0_1'(s(fun(X0,X1),X3))),s(X2,X4)),X0,X1)))))))) )),
  inference(superposition,[],[f435,f421])).
fof(f2336,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X4,X0),X5),s(fun(X0,bool),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,happ(s(fun(X4,X0),X5),s(X4,sK29(X6,happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X5,X4,X0)))))))) )),
  inference(forward_demodulation,[],[f2335,f421])).
fof(f2335,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))) | ~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X4,X0),X5),s(fun(X0,bool),X2)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3))),s(X0,happ(s(fun(X4,X0),X5),s(X4,sK29(X6,happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X5,X4,X0)))))))) )),
  inference(forward_demodulation,[],[f2332,f421])).
fof(f2332,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X4,X0),X5),s(fun(X0,bool),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3))),s(X0,X6)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3))),s(X0,happ(s(fun(X4,X0),X5),s(X4,sK29(X6,happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X5,X4,X0)))))))) )),
  inference(superposition,[],[f429,f421])).
fof(f2334,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X3),s(fun(X1,bool),X5)))) | p(s(bool,happ(s(fun(X1,bool),X5),s(X1,happ(s(fun(X0,X1),X3),s(X0,sK29(X6,X5,happ(s(fun(X2,fun(X0,X1)),'_dst_x0_1'(s(fun(X0,X1),X3))),s(X2,X4)),X0,X1))))))) | ~p(s(bool,happ(s(fun(X1,bool),X5),s(X1,X6))))) )),
  inference(forward_demodulation,[],[f2331,f421])).
fof(f2331,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X3),s(fun(X1,bool),X5)))) | ~p(s(bool,happ(s(fun(X1,bool),X5),s(X1,X6)))) | p(s(bool,happ(s(fun(X1,bool),X5),s(X1,happ(s(fun(X0,X1),happ(s(fun(X2,fun(X0,X1)),'_dst_x0_1'(s(fun(X0,X1),X3))),s(X2,X4))),s(X0,sK29(X6,X5,happ(s(fun(X2,fun(X0,X1)),'_dst_x0_1'(s(fun(X0,X1),X3))),s(X2,X4)),X0,X1)))))))) )),
  inference(superposition,[],[f429,f421])).
fof(f2306,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X4,X0),X5),s(fun(X0,bool),X2)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))) | s(X0,X6) = s(X0,happ(s(fun(X4,X0),X5),s(X4,sK31(X6,happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X5,X4,X0))))) )),
  inference(forward_demodulation,[],[f2303,f421])).
fof(f2303,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X4,X0),X5),s(fun(X0,bool),X2)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3))),s(X0,X6)))) | s(X0,X6) = s(X0,happ(s(fun(X4,X0),X5),s(X4,sK31(X6,happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X5,X4,X0))))) )),
  inference(superposition,[],[f432,f421])).
fof(f2305,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X0,X1),X3),s(fun(X1,bool),X5)))) | s(X1,X6) = s(X1,happ(s(fun(X0,X1),X3),s(X0,sK31(X6,X5,happ(s(fun(X2,fun(X0,X1)),'_dst_x0_1'(s(fun(X0,X1),X3))),s(X2,X4)),X0,X1)))) | p(s(bool,happ(s(fun(X1,bool),X5),s(X1,X6))))) )),
  inference(forward_demodulation,[],[f2302,f421])).
fof(f2302,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X0,X1),X3),s(fun(X1,bool),X5)))) | p(s(bool,happ(s(fun(X1,bool),X5),s(X1,X6)))) | s(X1,X6) = s(X1,happ(s(fun(X0,X1),happ(s(fun(X2,fun(X0,X1)),'_dst_x0_1'(s(fun(X0,X1),X3))),s(X2,X4))),s(X0,sK31(X6,X5,happ(s(fun(X2,fun(X0,X1)),'_dst_x0_1'(s(fun(X0,X1),X3))),s(X2,X4)),X0,X1))))) )),
  inference(superposition,[],[f432,f421])).
fof(f2287,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X4,X0),X5),s(fun(X0,bool),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))) | s(X0,X6) = s(X0,happ(s(fun(X4,X0),X5),s(X4,sK27(X6,happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X5,X4,X0))))) )),
  inference(forward_demodulation,[],[f2284,f421])).
fof(f2284,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X4,X0),X5),s(fun(X0,bool),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3))),s(X0,X6)))) | s(X0,X6) = s(X0,happ(s(fun(X4,X0),X5),s(X4,sK27(X6,happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X5,X4,X0))))) )),
  inference(superposition,[],[f426,f421])).
fof(f2286,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X0,X1),X3),s(fun(X1,bool),X5)))) | s(X1,X6) = s(X1,happ(s(fun(X0,X1),X3),s(X0,sK27(X6,X5,happ(s(fun(X2,fun(X0,X1)),'_dst_x0_1'(s(fun(X0,X1),X3))),s(X2,X4)),X0,X1)))) | ~p(s(bool,happ(s(fun(X1,bool),X5),s(X1,X6))))) )),
  inference(forward_demodulation,[],[f2283,f421])).
fof(f2283,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X0,X1),X3),s(fun(X1,bool),X5)))) | ~p(s(bool,happ(s(fun(X1,bool),X5),s(X1,X6)))) | s(X1,X6) = s(X1,happ(s(fun(X0,X1),happ(s(fun(X2,fun(X0,X1)),'_dst_x0_1'(s(fun(X0,X1),X3))),s(X2,X4))),s(X0,sK27(X6,X5,happ(s(fun(X2,fun(X0,X1)),'_dst_x0_1'(s(fun(X0,X1),X3))),s(X2,X4)),X0,X1))))) )),
  inference(superposition,[],[f426,f421])).
fof(f2248,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X6,X4),X7),s(fun(X4,bool),'_dst_x0_1'(s(bool,X5)))))) | p(s(bool,X5))) )),
  inference(superposition,[],[f439,f421])).
fof(f2246,plain,(
  ( ! [X6,X4,X5,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X4,X3),'_dst_x0_1'(s(X3,X5))),s(fun(X3,bool),X6)))) | p(s(bool,happ(s(fun(X3,bool),X6),s(X3,X5))))) )),
  inference(superposition,[],[f439,f421])).
fof(f2245,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X0,bool),X1),s(fun(bool,bool),X2)))) | p(s(bool,happ(s(fun(bool,bool),X2),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK35(X2,X1,X0,bool))))) )),
  inference(superposition,[],[f439,f329])).
fof(f2244,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X2,bool),X0),s(X2,happ(s(fun(bool,X2),X1),s(bool,'F_0')))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(bool,X2),X1),s(fun(X2,bool),X0)))) | s(bool,'T_0') = s(bool,sK35(X0,X1,bool,X2))) )),
  inference(superposition,[],[f439,f329])).
fof(f2250,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),X5),s(X1,happ(s(fun(X0,X1),X3),s(X0,sK35(X5,happ(s(fun(X2,fun(X0,X1)),'_dst_x0_1'(s(fun(X0,X1),X3))),s(X2,X4)),X0,X1))))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X0,X1),X3),s(fun(X1,bool),X5))))) )),
  inference(forward_demodulation,[],[f2243,f421])).
fof(f2243,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),X5),s(X1,happ(s(fun(X0,X1),X3),s(X0,sK35(X5,happ(s(fun(X2,fun(X0,X1)),'_dst_x0_1'(s(fun(X0,X1),X3))),s(X2,X4)),X0,X1))))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X0,X1),happ(s(fun(X2,fun(X0,X1)),'_dst_x0_1'(s(fun(X0,X1),X3))),s(X2,X4))),s(fun(X1,bool),X5))))) )),
  inference(superposition,[],[f439,f421])).
fof(f2249,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,happ(s(fun(X4,X0),X5),s(X4,sK35(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X5,X4,X0))))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X4,X0),X5),s(fun(X0,bool),X2))))) )),
  inference(forward_demodulation,[],[f2242,f421])).
fof(f2242,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,happ(s(fun(X4,X0),X5),s(X4,sK35(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X5,X4,X0))))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X4,X0),X5),s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3))))))) )),
  inference(superposition,[],[f439,f421])).
fof(f2232,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,X5)) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X6,X4),X7),s(fun(X4,bool),'_dst_x0_1'(s(bool,X5))))))) )),
  inference(superposition,[],[f425,f421])).
fof(f2229,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X2),s(bool,'F_0')))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X0,bool),X1),s(fun(bool,bool),X2)))) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK26(X2,X1,X0,bool))))) )),
  inference(superposition,[],[f425,f329])).
fof(f2228,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X2,bool),X0),s(X2,happ(s(fun(bool,X2),X1),s(bool,'F_0')))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(bool,X2),X1),s(fun(X2,bool),X0)))) | s(bool,'T_0') = s(bool,sK26(X0,X1,bool,X2))) )),
  inference(superposition,[],[f425,f329])).
fof(f2234,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X1,bool),X5),s(X1,happ(s(fun(X0,X1),X3),s(X0,sK26(X5,happ(s(fun(X2,fun(X0,X1)),'_dst_x0_1'(s(fun(X0,X1),X3))),s(X2,X4)),X0,X1))))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X0,X1),X3),s(fun(X1,bool),X5))))) )),
  inference(forward_demodulation,[],[f2227,f421])).
fof(f2227,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X1,bool),X5),s(X1,happ(s(fun(X0,X1),X3),s(X0,sK26(X5,happ(s(fun(X2,fun(X0,X1)),'_dst_x0_1'(s(fun(X0,X1),X3))),s(X2,X4)),X0,X1))))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X0,X1),happ(s(fun(X2,fun(X0,X1)),'_dst_x0_1'(s(fun(X0,X1),X3))),s(X2,X4))),s(fun(X1,bool),X5))))) )),
  inference(superposition,[],[f425,f421])).
fof(f2233,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,happ(s(fun(X4,X0),X5),s(X4,sK26(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X5,X4,X0))))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X4,X0),X5),s(fun(X0,bool),X2))))) )),
  inference(forward_demodulation,[],[f2226,f421])).
fof(f2226,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,happ(s(fun(X4,X0),X5),s(X4,sK26(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X5,X4,X0))))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X4,X0),X5),s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3))))))) )),
  inference(superposition,[],[f425,f421])).
fof(f2205,plain,(
  ( ! [X0,X1] : (~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK20(X1,X0))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1))))) )),
  inference(superposition,[],[f370,f364])).
fof(f2209,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK20(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X0))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X2))))) )),
  inference(forward_demodulation,[],[f2208,f421])).
fof(f2208,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK20(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3))))))) )),
  inference(forward_demodulation,[],[f2198,f421])).
fof(f2198,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3))),s(X0,sK20(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X0))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3))))))) )),
  inference(superposition,[],[f370,f421])).
fof(f2182,plain,(
  ( ! [X0,X1] : (~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK19(X1,X0))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1))))) )),
  inference(superposition,[],[f369,f364])).
fof(f2186,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK19(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X0))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X2))))) )),
  inference(forward_demodulation,[],[f2185,f421])).
fof(f2185,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK19(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3))))))) )),
  inference(forward_demodulation,[],[f2175,f421])).
fof(f2175,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3))),s(X0,sK19(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X0))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3))))))) )),
  inference(superposition,[],[f369,f421])).
fof(f2142,plain,(
  ( ! [X6,X8,X7,X5] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X7,X5),X8),s(fun(X5,bool),'_dst_x0_1'(s(bool,X6))))))) )),
  inference(subsumption_resolution,[],[f2138,f1899])).
fof(f2138,plain,(
  ( ! [X6,X8,X7,X5] : (p(s(bool,X6)) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X7,X5),X8),s(fun(X5,bool),'_dst_x0_1'(s(bool,X6))))))) )),
  inference(superposition,[],[f437,f421])).
fof(f2136,plain,(
  ( ! [X8,X7,X11,X9] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X8,X7),'_dst_x0_1'(s(X7,X9))),s(fun(X7,bool),X11)))) | p(s(bool,happ(s(fun(X7,bool),X11),s(X7,X9))))) )),
  inference(superposition,[],[f437,f421])).
fof(f2135,plain,(
  ( ! [X6,X4,X5,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X3,bool),X4),s(fun(bool,bool),X6)))) | p(s(bool,happ(s(fun(bool,bool),X6),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,happ(s(fun(X3,bool),X4),s(X3,X5)))) )),
  inference(superposition,[],[f437,f329])).
fof(f2134,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(bool,bool),X2),s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,bool),X1),s(fun(bool,bool),X2))))) )),
  inference(superposition,[],[f437,f364])).
fof(f2111,plain,(
  ( ! [X6,X8,X7,X5] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X7,X5),X8),s(fun(X5,bool),'_dst_x0_1'(s(bool,X6))))))) )),
  inference(subsumption_resolution,[],[f2108,f1875])).
fof(f2108,plain,(
  ( ! [X6,X8,X7,X5] : (~p(s(bool,X6)) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X7,X5),X8),s(fun(X5,bool),'_dst_x0_1'(s(bool,X6))))))) )),
  inference(superposition,[],[f431,f421])).
fof(f2105,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,happ(s(fun(bool,bool),X6),s(bool,'F_0')))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X3,bool),X4),s(fun(bool,bool),X6)))) | s(bool,'T_0') = s(bool,happ(s(fun(X3,bool),X4),s(X3,X5)))) )),
  inference(superposition,[],[f431,f329])).
fof(f2104,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X2),s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,bool),X1),s(fun(bool,bool),X2))))) )),
  inference(superposition,[],[f431,f364])).
fof(f2061,plain,(
  ( ! [X8,X7,X11,X9] : (s(X7,X9) != s(X7,sK32(X11,'_dst_x0_1'(s(X7,X9)),X8,X7)) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X8,X7),'_dst_x0_1'(s(X7,X9))),s(fun(X7,bool),X11))))) )),
  inference(superposition,[],[f434,f421])).
fof(f2060,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,'F_0') != s(bool,sK32(X6,X4,X3,bool)) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X3,bool),X4),s(fun(bool,bool),X6)))) | s(bool,'T_0') = s(bool,happ(s(fun(X3,bool),X4),s(X3,X5)))) )),
  inference(superposition,[],[f434,f329])).
fof(f2059,plain,(
  ( ! [X2,X0,X1] : (s(bool,sK32(X2,X1,X0,bool)) != s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X0,bool),X1),s(fun(bool,bool),X2))))) )),
  inference(superposition,[],[f434,f364])).
fof(f2051,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') != s(bool,happ(s(fun(X2,bool),X1),s(X2,X3))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X2,bool),X1),s(fun(bool,bool),X0)))) | s(bool,'T_0') = s(bool,sK32(X0,X1,X2,bool))) )),
  inference(superposition,[],[f434,f329])).
fof(f2037,plain,(
  ( ! [X8,X7,X11,X9] : (s(X7,X9) != s(X7,sK28(X11,'_dst_x0_1'(s(X7,X9)),X8,X7)) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X8,X7),'_dst_x0_1'(s(X7,X9))),s(fun(X7,bool),X11))))) )),
  inference(superposition,[],[f428,f421])).
fof(f2036,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,'F_0') != s(bool,sK28(X6,X4,X3,bool)) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X3,bool),X4),s(fun(bool,bool),X6)))) | s(bool,'T_0') = s(bool,happ(s(fun(X3,bool),X4),s(X3,X5)))) )),
  inference(superposition,[],[f428,f329])).
fof(f2035,plain,(
  ( ! [X2,X0,X1] : (s(bool,sK28(X2,X1,X0,bool)) != s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X0,bool),X1),s(fun(bool,bool),X2))))) )),
  inference(superposition,[],[f428,f364])).
fof(f2027,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') != s(bool,happ(s(fun(X2,bool),X1),s(X2,X3))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X2,bool),X1),s(fun(bool,bool),X0)))) | s(bool,'T_0') = s(bool,sK28(X0,X1,X2,bool))) )),
  inference(superposition,[],[f428,f329])).
fof(f2000,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,sK19(X0,bool)) | p(s(bool,'const.bool.?|33|_1'(s(fun(bool,bool),X0)))) | ~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X1)))) | s(bool,'T_0') = s(bool,sK20(X0,bool))) )),
  inference(superposition,[],[f371,f329])).
fof(f1973,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X2)))) | p(s(bool,X1)) | s(bool,'T_0') = s(bool,sK24(X0,X1,bool))) )),
  inference(superposition,[],[f382,f329])).
fof(f1977,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK24(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X4,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) | p(s(bool,X4))) )),
  inference(forward_demodulation,[],[f1972,f421])).
fof(f1972,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK24(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X4,X0))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3))),s(X0,X5)))) | p(s(bool,X4))) )),
  inference(superposition,[],[f382,f421])).
fof(f1948,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X2)))) | p(s(bool,X1)) | s(bool,'T_0') = s(bool,sK23(X0,X1,bool))) )),
  inference(superposition,[],[f380,f329])).
fof(f1952,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK23(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X4,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) | p(s(bool,X4))) )),
  inference(forward_demodulation,[],[f1947,f421])).
fof(f1947,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK23(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X4,X0))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3))),s(X0,X5)))) | p(s(bool,X4))) )),
  inference(superposition,[],[f380,f421])).
fof(f1930,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X2)))) | ~p(s(bool,X1)) | s(bool,'T_0') = s(bool,sK22(X0,X1,bool))) )),
  inference(superposition,[],[f379,f329])).
fof(f1934,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK22(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X4,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) | ~p(s(bool,X4))) )),
  inference(forward_demodulation,[],[f1929,f421])).
fof(f1929,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK22(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X4,X0))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3))),s(X0,X5)))) | ~p(s(bool,X4))) )),
  inference(superposition,[],[f379,f421])).
fof(f1912,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X2)))) | ~p(s(bool,X1)) | s(bool,'T_0') = s(bool,sK21(X0,X1,bool))) )),
  inference(superposition,[],[f377,f329])).
fof(f1916,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK21(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X4,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) | ~p(s(bool,X4))) )),
  inference(forward_demodulation,[],[f1911,f421])).
fof(f1911,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK21(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X4,X0))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3))),s(X0,X5)))) | ~p(s(bool,X4))) )),
  inference(superposition,[],[f377,f421])).
fof(f1899,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,X5)) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X7,X4),X6),s(fun(X4,bool),'_dst_x0_1'(s(bool,X5))))))) )),
  inference(superposition,[],[f436,f421])).
fof(f1897,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X2,bool),X1),s(fun(bool,bool),X0)))) | s(bool,'T_0') = s(bool,sK34(X0,X1,X2,bool))) )),
  inference(superposition,[],[f436,f329])).
fof(f1900,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK34(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X4,X5,X0))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X5,X0),X4),s(fun(X0,bool),X2))))) )),
  inference(forward_demodulation,[],[f1896,f421])).
fof(f1896,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK34(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X4,X5,X0))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X5,X0),X4),s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3))))))) )),
  inference(superposition,[],[f436,f421])).
fof(f1886,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,X5)) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X7,X4),X6),s(fun(X4,bool),'_dst_x0_1'(s(bool,X5))))))) )),
  inference(superposition,[],[f433,f421])).
fof(f1884,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X2,bool),X1),s(fun(bool,bool),X0)))) | s(bool,'T_0') = s(bool,sK32(X0,X1,X2,bool))) )),
  inference(superposition,[],[f433,f329])).
fof(f1887,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK32(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X4,X5,X0))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X5,X0),X4),s(fun(X0,bool),X2))))) )),
  inference(forward_demodulation,[],[f1883,f421])).
fof(f1883,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK32(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X4,X5,X0))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X5,X0),X4),s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3))))))) )),
  inference(superposition,[],[f433,f421])).
fof(f1875,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,X5)) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X7,X4),X6),s(fun(X4,bool),'_dst_x0_1'(s(bool,X5))))))) )),
  inference(superposition,[],[f430,f421])).
fof(f1873,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X2,bool),X1),s(fun(bool,bool),X0)))) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK30(X0,X1,X2,bool))) )),
  inference(superposition,[],[f430,f329])).
fof(f1876,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK30(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X4,X5,X0))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X5,X0),X4),s(fun(X0,bool),X2))))) )),
  inference(forward_demodulation,[],[f1872,f421])).
fof(f1872,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK30(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X4,X5,X0))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X5,X0),X4),s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3))))))) )),
  inference(superposition,[],[f430,f421])).
fof(f1865,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X7,X4),X6),s(fun(X4,bool),'_dst_x0_1'(s(bool,X5)))))) | p(s(bool,X5))) )),
  inference(superposition,[],[f427,f421])).
fof(f1863,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X2,bool),X1),s(fun(bool,bool),X0)))) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK28(X0,X1,X2,bool))) )),
  inference(superposition,[],[f427,f329])).
fof(f1866,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK28(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X4,X5,X0))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X5,X0),X4),s(fun(X0,bool),X2))))) )),
  inference(forward_demodulation,[],[f1862,f421])).
fof(f1862,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK28(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X4,X5,X0))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X5,X0),X4),s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3))))))) )),
  inference(superposition,[],[f427,f421])).
fof(f1552,plain,(
  ( ! [X109,X107,X108,X106] : (sP10(happ(s(fun(X106,bool),X107),s(X106,'const.min.@_1'(s(fun(X106,bool),X107)))),X108,X109) | ~sP9(happ(s(fun(X106,bool),X107),s(X106,'const.min.@_1'(s(fun(X106,bool),X107)))),X108,X109) | ~p(s(bool,'|dollar|exists_1'(s(fun(X106,bool),X107))))) )),
  inference(superposition,[],[f557,f364])).
fof(f1551,plain,(
  ( ! [X103,X105,X102,X104] : (sP12(happ(s(fun(X102,bool),X103),s(X102,'const.min.@_1'(s(fun(X102,bool),X103)))),X104,X105) | ~sP11(happ(s(fun(X102,bool),X103),s(X102,'const.min.@_1'(s(fun(X102,bool),X103)))),X104,X105) | ~p(s(bool,'|dollar|exists_1'(s(fun(X102,bool),X103))))) )),
  inference(superposition,[],[f556,f364])).
fof(f1549,plain,(
  ( ! [X97,X95,X98,X96] : (~sP12(happ(s(fun(X95,bool),X96),s(X95,'const.min.@_1'(s(fun(X95,bool),X96)))),X98,X97) | ~p(s(bool,X97)) | p(s(bool,X98)) | ~p(s(bool,'|dollar|exists_1'(s(fun(X95,bool),X96))))) )),
  inference(superposition,[],[f533,f364])).
fof(f1548,plain,(
  ( ! [X94,X92,X93,X91] : (sP12(happ(s(fun(X91,bool),X92),s(X91,'const.min.@_1'(s(fun(X91,bool),X92)))),X93,X94) | ~p(s(bool,X93)) | ~p(s(bool,'|dollar|exists_1'(s(fun(X91,bool),X92))))) )),
  inference(superposition,[],[f532,f364])).
fof(f1547,plain,(
  ( ! [X90,X88,X87,X89] : (sP12(happ(s(fun(X87,bool),X88),s(X87,'const.min.@_1'(s(fun(X87,bool),X88)))),X90,X89) | p(s(bool,X89)) | ~p(s(bool,'|dollar|exists_1'(s(fun(X87,bool),X88))))) )),
  inference(superposition,[],[f531,f364])).
fof(f1546,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,'|dollar|exists_1'(s(fun(X83,bool),X84)))) | p(s(bool,X85)) | ~p(s(bool,X86)) | ~sP11(X86,X85,happ(s(fun(X83,bool),X84),s(X83,'const.min.@_1'(s(fun(X83,bool),X84)))))) )),
  inference(superposition,[],[f526,f364])).
fof(f1545,plain,(
  ( ! [X80,X81,X79,X82] : (~sP10(happ(s(fun(X79,bool),X80),s(X79,'const.min.@_1'(s(fun(X79,bool),X80)))),X82,X81) | p(s(bool,X81)) | p(s(bool,X82)) | ~p(s(bool,'|dollar|exists_1'(s(fun(X79,bool),X80))))) )),
  inference(superposition,[],[f516,f364])).
fof(f1544,plain,(
  ( ! [X78,X76,X77,X75] : (sP10(happ(s(fun(X75,bool),X76),s(X75,'const.min.@_1'(s(fun(X75,bool),X76)))),X77,X78) | ~p(s(bool,X77)) | ~p(s(bool,'|dollar|exists_1'(s(fun(X75,bool),X76))))) )),
  inference(superposition,[],[f515,f364])).
fof(f1543,plain,(
  ( ! [X74,X72,X71,X73] : (sP10(happ(s(fun(X71,bool),X72),s(X71,'const.min.@_1'(s(fun(X71,bool),X72)))),X74,X73) | ~p(s(bool,X73)) | ~p(s(bool,'|dollar|exists_1'(s(fun(X71,bool),X72))))) )),
  inference(superposition,[],[f514,f364])).
fof(f1542,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,'|dollar|exists_1'(s(fun(X67,bool),X68)))) | p(s(bool,X69)) | p(s(bool,X70)) | ~sP9(happ(s(fun(X67,bool),X68),s(X67,'const.min.@_1'(s(fun(X67,bool),X68)))),X69,X70)) )),
  inference(superposition,[],[f509,f364])).
fof(f1541,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,'|dollar|exists_1'(s(fun(X63,bool),X64)))) | ~p(s(bool,X65)) | p(s(bool,X66)) | ~sP8(X66,X65,happ(s(fun(X63,bool),X64),s(X63,'const.min.@_1'(s(fun(X63,bool),X64)))))) )),
  inference(superposition,[],[f501,f364])).
fof(f1540,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,'|dollar|exists_1'(s(fun(X59,bool),X60)))) | ~p(s(bool,X61)) | p(s(bool,X62)) | ~sP7(X61,happ(s(fun(X59,bool),X60),s(X59,'const.min.@_1'(s(fun(X59,bool),X60)))),X62)) )),
  inference(superposition,[],[f492,f364])).
fof(f1539,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,'|dollar|exists_1'(s(fun(X55,bool),X56)))) | ~p(s(bool,X57)) | p(s(bool,X58)) | ~sP4(X57,happ(s(fun(X55,bool),X56),s(X55,'const.min.@_1'(s(fun(X55,bool),X56)))),X58)) )),
  inference(superposition,[],[f473,f364])).
fof(f1538,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,'|dollar|exists_1'(s(fun(X51,bool),X52)))) | ~p(s(bool,X53)) | p(s(bool,X54)) | ~sP5(X53,happ(s(fun(X51,bool),X52),s(X51,'const.min.@_1'(s(fun(X51,bool),X52)))),X54)) )),
  inference(superposition,[],[f469,f364])).
fof(f1537,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,'|dollar|exists_1'(s(fun(X47,bool),X48)))) | sP6(X49,X50,happ(s(fun(X47,bool),X48),s(X47,'const.min.@_1'(s(fun(X47,bool),X48))))) | ~sP5(X49,happ(s(fun(X47,bool),X48),s(X47,'const.min.@_1'(s(fun(X47,bool),X48)))),X50) | ~sP4(X50,happ(s(fun(X47,bool),X48),s(X47,'const.min.@_1'(s(fun(X47,bool),X48)))),X49) | ~sP3(happ(s(fun(X47,bool),X48),s(X47,'const.min.@_1'(s(fun(X47,bool),X48)))),X50,X49)) )),
  inference(superposition,[],[f466,f364])).
fof(f1536,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,'|dollar|exists_1'(s(fun(X43,bool),X44)))) | ~p(s(bool,X45)) | p(s(bool,X46)) | ~sP6(X45,happ(s(fun(X43,bool),X44),s(X43,'const.min.@_1'(s(fun(X43,bool),X44)))),X46)) )),
  inference(superposition,[],[f465,f364])).
fof(f1534,plain,(
  ( ! [X37,X35,X38,X36] : (sP2(X37,happ(s(fun(X35,bool),X36),s(X35,'const.min.@_1'(s(fun(X35,bool),X36)))),X38) | ~p(s(bool,X37)) | ~p(s(bool,'|dollar|exists_1'(s(fun(X35,bool),X36))))) )),
  inference(superposition,[],[f450,f364])).
fof(f1532,plain,(
  ( ! [X30,X28,X29] : (sP1(X30,happ(s(fun(X28,bool),X29),s(X28,'const.min.@_1'(s(fun(X28,bool),X29))))) | p(s(bool,X30)) | ~p(s(bool,'|dollar|exists_1'(s(fun(X28,bool),X29))))) )),
  inference(superposition,[],[f412,f364])).
fof(f1531,plain,(
  ( ! [X26,X27,X25] : (sP1(happ(s(fun(X25,bool),X26),s(X25,'const.min.@_1'(s(fun(X25,bool),X26)))),X27) | p(s(bool,X27)) | ~p(s(bool,'|dollar|exists_1'(s(fun(X25,bool),X26))))) )),
  inference(superposition,[],[f409,f364])).
fof(f1530,plain,(
  ( ! [X24,X23,X22] : (sP0(happ(s(fun(X22,bool),X23),s(X22,'const.min.@_1'(s(fun(X22,bool),X23)))),X24) | ~p(s(bool,'|dollar|exists_1'(s(fun(X22,bool),X23))))) )),
  inference(superposition,[],[f394,f364])).
fof(f1529,plain,(
  ( ! [X21,X19,X20] : (sP0(X21,happ(s(fun(X19,bool),X20),s(X19,'const.min.@_1'(s(fun(X19,bool),X20))))) | ~p(s(bool,'|dollar|exists_1'(s(fun(X19,bool),X20))))) )),
  inference(superposition,[],[f393,f364])).
fof(f1517,plain,(
  ( ! [X0] : (s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0))) = s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X0)))) )),
  inference(superposition,[],[f364,f329])).
fof(f1513,plain,(
  ( ! [X0,X1] : (sP6('T_0',happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))),'|dollar|exists_1'(s(fun(X0,bool),X1)))) )),
  inference(unit_resulting_resolution,[],[f322,f364,f482])).
fof(f1515,plain,(
  ( ! [X0,X1] : (sP6('T_0','|dollar|exists_1'(s(fun(X0,bool),X1)),happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f322,f364,f482])).
fof(f1503,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X2)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK18(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X0)))))) )),
  inference(forward_demodulation,[],[f1500,f421])).
fof(f1500,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X2)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3))),s(X0,sK18(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X0)))))) )),
  inference(superposition,[],[f367,f421])).
fof(f1439,plain,(
  sP6('F_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))),
  inference(unit_resulting_resolution,[],[f611,f585,f580,f573,f467])).
fof(f1402,plain,(
  ( ! [X14,X12,X10,X13,X11] : (sP6(X13,X14,happ(s(fun(X10,bool),'_dst_x0_1'(s(bool,X11))),s(X10,X12))) | ~p(s(bool,X11)) | ~sP5(X13,happ(s(fun(X10,bool),'_dst_x0_1'(s(bool,X11))),s(X10,X12)),X14) | ~sP4(X14,happ(s(fun(X10,bool),'_dst_x0_1'(s(bool,X11))),s(X10,X12)),X13) | ~sP3(happ(s(fun(X10,bool),'_dst_x0_1'(s(bool,X11))),s(X10,X12)),X14,X13)) )),
  inference(superposition,[],[f466,f421])).
fof(f1405,plain,(
  ( ! [X6,X7] : (sP6(X6,X7,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16))) | ~sP5(X6,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X7) | ~sP4(X7,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X6)) )),
  inference(subsumption_resolution,[],[f1404,f611])).
fof(f1404,plain,(
  ( ! [X6,X7] : (sP6(X6,X7,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16))) | ~sP5(X6,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X7) | ~sP4(X7,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X6) | ~sP3('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X7,X6)) )),
  inference(subsumption_resolution,[],[f1400,f322])).
fof(f1400,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | sP6(X6,X7,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16))) | ~sP5(X6,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X7) | ~sP4(X7,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X6) | ~sP3('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X7,X6)) )),
  inference(superposition,[],[f466,f621])).
fof(f1386,plain,(
  sP6('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f589,f582,f577,f322,f466])).
fof(f1387,plain,(
  sP6('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'T_0','T_0')),
  inference(unit_resulting_resolution,[],[f589,f608,f577,f322,f466])).
fof(f1389,plain,(
  sP6('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'T_0')),
  inference(unit_resulting_resolution,[],[f589,f608,f607,f322,f466])).
fof(f1397,plain,(
  ~sP4('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'T_0','F_0')),
  inference(unit_resulting_resolution,[],[f589,f580,f949,f322,f466])).
fof(f1381,plain,(
  ~sP12('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))),
  inference(unit_resulting_resolution,[],[f1371,f521])).
fof(f1371,plain,(
  ~sP11('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))),
  inference(unit_resulting_resolution,[],[f573,f322,f1257])).
fof(f1335,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,X32)) | sP0(happ(s(fun(X31,bool),'_dst_x0_1'(s(bool,X32))),s(X31,X33)),X34)) )),
  inference(superposition,[],[f394,f421])).
fof(f1334,plain,(
  ( ! [X30,X28,X29,X27] : (~p(s(bool,X28)) | sP0(X30,happ(s(fun(X27,bool),'_dst_x0_1'(s(bool,X28))),s(X27,X29)))) )),
  inference(superposition,[],[f393,f421])).
fof(f1314,plain,(
  ( ! [X2,X0,X1] : (sP6('T_0',X0,happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,X0))),s(X1,X2)))) )),
  inference(unit_resulting_resolution,[],[f322,f421,f482])).
fof(f1298,plain,(
  ( ! [X6,X7] : (~sP12('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X7,X6) | p(s(bool,X7)) | ~p(s(bool,X6))) )),
  inference(subsumption_resolution,[],[f1293,f322])).
fof(f1293,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | ~p(s(bool,X6)) | p(s(bool,X7)) | ~sP12('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X7,X6)) )),
  inference(superposition,[],[f533,f621])).
fof(f1197,plain,(
  ~sP4('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'F_0')),
  inference(unit_resulting_resolution,[],[f322,f573,f1032])).
fof(f1196,plain,(
  ( ! [X6,X7] : (~sP10('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X7,X6) | p(s(bool,X7)) | p(s(bool,X6))) )),
  inference(subsumption_resolution,[],[f1191,f322])).
fof(f1191,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | p(s(bool,X6)) | p(s(bool,X7)) | ~sP10('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X7,X6)) )),
  inference(superposition,[],[f516,f621])).
fof(f1055,plain,(
  sP8('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f590,f322,f322,f488])).
fof(f1056,plain,(
  sP8('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'T_0','T_0')),
  inference(unit_resulting_resolution,[],[f612,f322,f322,f488])).
fof(f938,plain,(
  ~sP6('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f322,f617,f483])).
fof(f927,plain,(
  sP6('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'T_0')),
  inference(unit_resulting_resolution,[],[f322,f621,f482])).
fof(f929,plain,(
  sP6('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))),
  inference(unit_resulting_resolution,[],[f322,f621,f482])).
fof(f818,plain,(
  ( ! [X2,X3] : (~sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X2,X3) | sP10('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X2,X3)) )),
  inference(subsumption_resolution,[],[f817,f322])).
fof(f817,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | ~sP9('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X2,X3) | sP10('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X2,X3)) )),
  inference(superposition,[],[f557,f621])).
fof(f800,plain,(
  ( ! [X2,X3] : (~sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X2,X3) | sP12('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X2,X3)) )),
  inference(subsumption_resolution,[],[f799,f322])).
fof(f799,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | ~sP11('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X2,X3) | sP12('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X2,X3)) )),
  inference(superposition,[],[f556,f621])).
fof(f797,plain,(
  ( ! [X0] : (sP12('T_0','const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X0)) )),
  inference(unit_resulting_resolution,[],[f615,f322,f556])).
fof(f775,plain,(
  ( ! [X0] : (sP12('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f322,f322,f532])).
fof(f756,plain,(
  ( ! [X0] : (sP10('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f322,f746])).
fof(f743,plain,(
  ( ! [X0] : (sP10('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f322,f322,f515])).
fof(f740,plain,(
  ( ! [X0] : (sP10('const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f322,f736])).
fof(f733,plain,(
  ( ! [X0] : (sP10('T_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f322,f322,f514])).
fof(f704,plain,(
  ( ! [X0] : (sP2('T_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f322,f322,f451])).
fof(f640,plain,(
  ~sP1('T_0','T_0')),
  inference(unit_resulting_resolution,[],[f322,f322,f408])).
fof(f612,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) )),
  inference(unit_resulting_resolution,[],[f318,f493])).
fof(f604,plain,(
  ( ! [X0,X1] : (sP2(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f322,f559])).
fof(f576,plain,(
  ( ! [X0] : (sP0('T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f322,f394])).
fof(f575,plain,(
  ( ! [X0] : (sP0(X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f322,f393])).
fof(f435,plain,(
  ( ! [X6,X2,X0,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK33(X6,X3,X2,X1,X0)))))))) )),
  inference(cnf_transformation,[],[f268])).
fof(f436,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,sK34(X3,X2,X1,X0))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f268])).
fof(f432,plain,(
  ( ! [X6,X2,X0,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | s(X0,X6) = s(X0,happ(s(fun(X1,X0),X2),s(X1,sK31(X6,X3,X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f265])).
fof(f433,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,sK32(X3,X2,X1,X0))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f265])).
fof(f429,plain,(
  ( ! [X6,X2,X0,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK29(X6,X3,X2,X1,X0)))))))) )),
  inference(cnf_transformation,[],[f262])).
fof(f426,plain,(
  ( ! [X6,X2,X0,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | s(X0,X6) = s(X0,happ(s(fun(X1,X0),X2),s(X1,sK27(X6,X3,X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f259])).
fof(f380,plain,(
  ( ! [X2,X0,X5,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK23(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f240])).
fof(f240,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK23(X2,X1,X0))))) & ~p(s(bool,X1))) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) | p(s(bool,X1))) & ((~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK24(X2,X1,X0))))) & ~p(s(bool,X1))) | ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | p(s(bool,X1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK23,sK24])],[f239])).
fof(f239,plain,(
  ! [X0,X1,X2] : ((? [X6] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))) & ~p(s(bool,X1))) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) | p(s(bool,X1))) & ((? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) & ~p(s(bool,X1))) | ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | p(s(bool,X1)))))),
  inference(rectify,[],[f238])).
fof(f238,plain,(
  ! [X0,X1,X2] : ((? [X3] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) & ~p(s(bool,X1))) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,X1))) & ((? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) & ~p(s(bool,X1))) | ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | p(s(bool,X1)))))),
  inference(flattening,[],[f237])).
fof(f237,plain,(
  ! [X0,X1,X2] : ((? [X3] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) & ~p(s(bool,X1))) | (! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,X1)))) & ((? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) & ~p(s(bool,X1))) | ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | p(s(bool,X1)))))),
  inference(nnf_transformation,[],[f98])).
fof(f98,plain,(
  ! [X0,X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | p(s(bool,X1))) <=> (! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,X1))))),
  inference(flattening,[],[f97])).
fof(f97,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | p(s(bool,X1))) <=> (! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,X1))))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X7] : ! [X17,X9] : (! [X14] : (p(s(bool,happ(s(fun(X7,bool),X9),s(X7,X14)))) | p(s(bool,X17))) <=> (! [X14] : p(s(bool,happ(s(fun(X7,bool),X9),s(X7,X14)))) | p(s(bool,X17))))),
  file('thm.quantHeuristics.GUESSES_UEXISTS_THM3','thm.bool.LEFT_FORALL_OR_THM')).
fof(f382,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK24(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f240])).
fof(f377,plain,(
  ( ! [X2,X0,X5,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK21(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f236])).
fof(f236,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK21(X2,X1,X0))))) | ! [X5] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))))) & ((~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK22(X2,X1,X0)))))) | (p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK21,sK22])],[f235])).
fof(f235,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,X1)) | ? [X6] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))) | ! [X5] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))))) & (? [X4] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(rectify,[],[f234])).
fof(f234,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,X1)) | ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | ! [X4] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (? [X4] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(flattening,[],[f233])).
fof(f233,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) | ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ! [X4] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (? [X4] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f96])).
fof(f96,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> ! [X4] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f95])).
fof(f95,plain,(
  ! [X0] : ! [X1,X2] : ((p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> ! [X4] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X7] : ! [X9,X17] : ((p(s(bool,X9)) & ! [X14] : p(s(bool,happ(s(fun(X7,bool),X17),s(X7,X14))))) <=> ! [X14] : (p(s(bool,X9)) & p(s(bool,happ(s(fun(X7,bool),X17),s(X7,X14))))))),
  file('thm.quantHeuristics.GUESSES_UEXISTS_THM3','thm.bool.RIGHT_AND_FORALL_THM')).
fof(f379,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK22(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f236])).
fof(f5059,plain,(
  $false | $spl14),
  inference(global_subsumption,[],[f321,f320,f319,f318,f322,f573,f329,f337,f336,f341,f340,f364,f371,f370,f369,f368,f367,f379,f377,f382,f380,f387,f384,f394,f393,f574,f412,f409,f408,f407,f421,f423,f425,f424,f428,f427,f426,f431,f430,f429,f434,f433,f432,f437,f436,f435,f439,f438,f451,f450,f449,f559,f447,f446,f468,f467,f466,f465,f464,f463,f462,f472,f471,f470,f469,f476,f475,f474,f473,f480,f479,f478,f477,f484,f483,f482,f481,f558,f488,f487,f486,f485,f495,f494,f493,f492,f501,f498,f497,f496,f508,f507,f506,f557,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f525,f524,f523,f556,f521,f520,f519,f529,f528,f527,f526,f533,f532,f531,f530,f555,f575,f576,f577,f578,f580,f582,f583,f585,f587,f588,f589,f590,f591,f593,f595,f596,f597,f599,f601,f602,f604,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f617,f621,f639,f640,f643,f649,f657,f664,f681,f682,f685,f688,f695,f704,f707,f708,f711,f712,f716,f725,f729,f730,f733,f736,f740,f743,f746,f755,f756,f757,f765,f768,f772,f775,f778,f797,f800,f815,f813,f818,f836,f884,f902,f925,f923,f929,f927,f941,f938,f951,f949,f992,f1003,f1011,f1024,f1032,f1036,f1042,f1051,f1056,f1055,f1068,f1076,f1080,f1113,f1124,f1131,f1140,f1152,f1160,f1164,f1178,f1196,f1197,f1212,f1222,f1237,f1249,f1257,f1261,f1272,f1282,f1298,f1314,f1313,f1311,f1310,f1325,f1334,f1335,f1336,f1337,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1347,f1348,f1349,f1350,f1351,f1352,f1353,f1355,f1356,f1371,f1381,f1397,f1389,f1387,f1386,f1405,f1402,f1439,f1429,f1428,f1484,f1503,f1515,f1514,f1513,f1512,f1517,f1518,f1519,f1529,f1530,f1531,f1532,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1543,f1544,f1545,f1546,f1547,f1548,f1549,f1551,f1552,f1866,f1863,f1867,f1865,f1876,f1873,f1877,f1875,f1887,f1884,f1886,f1900,f1897,f1899,f1916,f1912,f1934,f1930,f1952,f1948,f1977,f1973,f2000,f2015,f2016,f2027,f2038,f2035,f2036,f2037,f2051,f2062,f2059,f2060,f2061,f2082,f2104,f2105,f2106,f2111,f2134,f2135,f2136,f2141,f2142,f2186,f2182,f2209,f2205,f2233,f2234,f2228,f2229,f2230,f2232,f2249,f2250,f2244,f2245,f2246,f2251,f2248,f2286,f2287,f2305,f2306,f2334,f2336,f2347,f2349,f2367,f2364,f2386,f2370,f2371,f2372,f2374,f2375,f2400,f2403,f2404,f2405,f2407,f2408,f2429,f2444,f2427,f2432,f2433,f2434,f2436,f2456,f2460,f2461,f2462,f2464,f2489,f2488,f2486,f2485,f2497,f2506,f2507,f2508,f2509,f2511,f2512,f2513,f2514,f2515,f2516,f2517,f2518,f2519,f2520,f2521,f2522,f2523,f2524,f2525,f2526,f2528,f2529,f2530,f2531,f2532,f2533,f2831,f2830,f2827,f2824,f2821,f2820,f2819,f2818,f2817,f2816,f2815,f2814,f2813,f2812,f2811,f2810,f2809,f2806,f2803,f2802,f2801,f2797,f2796,f2790,f2787,f2786,f2780,f2779,f2772,f2746,f2745,f2744,f2739,f2738,f2733,f2732,f2731,f2726,f2718,f2714,f2710,f2709,f2704,f2703,f2702,f2701,f2700,f2696,f2694,f2688,f2687,f2681,f2680,f2673,f2670,f2669,f2663,f2660,f2659,f2658,f2656,f2654,f2653,f2651,f2650,f2649,f2648,f2646,f2645,f2644,f2643,f2642,f2641,f2635,f2634,f2627,f2626,f2620,f2619,f2612,f2611,f2609,f2607,f2605,f2602,f2601,f2595,f2594,f2589,f2574,f2573,f2572,f2571,f2569,f2568,f2567,f2566,f2565,f3267,f3266,f3263,f3260,f3257,f3256,f3255,f3254,f3253,f3252,f3251,f3250,f3249,f3248,f3247,f3246,f3245,f3242,f3239,f3238,f3237,f3233,f3232,f3226,f3223,f3222,f3216,f3215,f3208,f3182,f3181,f3180,f3175,f3174,f3169,f3168,f3167,f3162,f3154,f3150,f3146,f3145,f3140,f3139,f3138,f3137,f3136,f3132,f3130,f3124,f3123,f3117,f3116,f3109,f3106,f3105,f3099,f3096,f3095,f3093,f3090,f3088,f3087,f3084,f3083,f3082,f3081,f3078,f3077,f3076,f3075,f3074,f3073,f3067,f3066,f3059,f3058,f3052,f3051,f3044,f3043,f3041,f3039,f3037,f3034,f3033,f3027,f3026,f3021,f3006,f3005,f3004,f3003,f3001,f3000,f2999,f2998,f2997,f3620,f3619,f3616,f3613,f3610,f3609,f3608,f3607,f3606,f3605,f3604,f3603,f3602,f3601,f3600,f3599,f3598,f3595,f3592,f3591,f3590,f3586,f3585,f3579,f3576,f3575,f3569,f3568,f3561,f3535,f3534,f3533,f3528,f3527,f3522,f3521,f3520,f3515,f3507,f3503,f3499,f3498,f3493,f3492,f3491,f3490,f3489,f3485,f3483,f3477,f3476,f3470,f3469,f3462,f3459,f3458,f3452,f3449,f3448,f3446,f3443,f3441,f3440,f3437,f3436,f3435,f3434,f3431,f3430,f3429,f3428,f3427,f3426,f3420,f3419,f3412,f3411,f3405,f3404,f3397,f3396,f3395,f3394,f3393,f3391,f3390,f3384,f3383,f3378,f3363,f3362,f3361,f3360,f3358,f3357,f3356,f3354,f3708,f3726,f3770,f3771,f3773,f3788,f3799,f3800,f3802,f3851,f3852,f3853,f3854,f3905,f3907,f3908,f3909,f3910,f3911,f3912,f3913,f3914,f3915,f3916,f3917,f3918,f3919,f3920,f3921,f3922,f3923,f3924,f3925,f3926,f3927,f4039,f4569,f4573,f4576,f4579,f4583,f4586,f4589,f4592,f4595,f4598,f4601,f4604,f4607,f4610,f4613,f4616,f4619,f4622,f4625,f4072,f4073,f4074,f4075,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4646,f4647,f4648,f4649,f4650,f4124,f4125,f4126,f4127,f4651,f4652,f4653,f4654,f4655,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4667,f4668,f4669,f4670,f4671,f4177,f4179,f4180,f4181,f4182,f4183,f4184,f4185,f4186,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4677,f4679,f4681,f4683,f4686,f4688,f4690,f4692,f4694,f4696,f4698,f4700,f4702,f4704,f4706,f4708,f4710,f4712,f4714,f4571,f4941,f4744,f4751,f4966,f4972,f4982,f4986,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f5037,f5039,f5041,f5043,f5045,f5047,f5049,f5051,f4899,f4900,f4901])).
fof(f5058,plain,(
  $false | $spl14),
  inference(global_subsumption,[],[f321,f320,f319,f318,f322,f573,f329,f337,f336,f341,f340,f364,f371,f370,f369,f368,f367,f379,f377,f382,f380,f387,f384,f394,f393,f574,f412,f409,f408,f407,f421,f423,f425,f424,f428,f427,f426,f431,f430,f429,f434,f433,f432,f437,f436,f435,f439,f438,f451,f450,f449,f559,f447,f446,f468,f467,f466,f465,f464,f463,f462,f472,f471,f470,f469,f476,f475,f474,f473,f480,f479,f478,f477,f484,f483,f482,f481,f558,f488,f487,f486,f485,f495,f494,f493,f492,f501,f498,f497,f496,f508,f507,f506,f557,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f525,f524,f523,f556,f521,f520,f519,f529,f528,f527,f526,f533,f532,f531,f530,f555,f575,f576,f577,f578,f580,f582,f583,f585,f587,f588,f589,f590,f591,f593,f595,f596,f597,f599,f601,f602,f604,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f617,f621,f639,f640,f643,f649,f657,f664,f681,f682,f685,f688,f695,f704,f707,f708,f711,f712,f716,f725,f729,f730,f733,f736,f740,f743,f746,f755,f756,f757,f765,f768,f772,f775,f778,f797,f800,f815,f813,f818,f836,f884,f902,f925,f923,f929,f927,f941,f938,f951,f949,f992,f1003,f1011,f1024,f1032,f1036,f1042,f1051,f1056,f1055,f1068,f1076,f1080,f1113,f1124,f1131,f1140,f1152,f1160,f1164,f1178,f1196,f1197,f1212,f1222,f1237,f1249,f1257,f1261,f1272,f1282,f1298,f1314,f1313,f1311,f1310,f1325,f1334,f1335,f1336,f1337,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1347,f1348,f1349,f1350,f1351,f1352,f1353,f1355,f1356,f1371,f1381,f1397,f1389,f1387,f1386,f1405,f1402,f1439,f1429,f1428,f1484,f1503,f1515,f1514,f1513,f1512,f1517,f1518,f1519,f1529,f1530,f1531,f1532,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1543,f1544,f1545,f1546,f1547,f1548,f1549,f1551,f1552,f1866,f1863,f1867,f1865,f1876,f1873,f1877,f1875,f1887,f1884,f1886,f1900,f1897,f1899,f1916,f1912,f1934,f1930,f1952,f1948,f1977,f1973,f2000,f2015,f2016,f2027,f2038,f2035,f2036,f2037,f2051,f2062,f2059,f2060,f2061,f2082,f2104,f2105,f2106,f2111,f2134,f2135,f2136,f2141,f2142,f2186,f2182,f2209,f2205,f2233,f2234,f2228,f2229,f2230,f2232,f2249,f2250,f2244,f2245,f2246,f2251,f2248,f2286,f2287,f2305,f2306,f2334,f2336,f2347,f2349,f2367,f2364,f2386,f2370,f2371,f2372,f2374,f2375,f2400,f2403,f2404,f2405,f2407,f2408,f2429,f2444,f2427,f2432,f2433,f2434,f2436,f2456,f2460,f2461,f2462,f2464,f2489,f2488,f2486,f2485,f2497,f2506,f2507,f2508,f2509,f2511,f2512,f2513,f2514,f2515,f2516,f2517,f2518,f2519,f2520,f2521,f2522,f2523,f2524,f2525,f2526,f2528,f2529,f2530,f2531,f2532,f2533,f2831,f2830,f2827,f2824,f2821,f2820,f2819,f2818,f2817,f2816,f2815,f2814,f2813,f2812,f2811,f2810,f2809,f2806,f2803,f2802,f2801,f2797,f2796,f2790,f2787,f2786,f2780,f2779,f2772,f2746,f2745,f2744,f2739,f2738,f2733,f2732,f2731,f2726,f2718,f2714,f2710,f2709,f2704,f2703,f2702,f2701,f2700,f2696,f2694,f2688,f2687,f2681,f2680,f2673,f2670,f2669,f2663,f2660,f2659,f2658,f2656,f2654,f2653,f2651,f2650,f2649,f2648,f2646,f2645,f2644,f2643,f2642,f2641,f2635,f2634,f2627,f2626,f2620,f2619,f2612,f2611,f2609,f2607,f2605,f2602,f2601,f2595,f2594,f2589,f2574,f2573,f2572,f2571,f2569,f2568,f2567,f2566,f2565,f3267,f3266,f3263,f3260,f3257,f3256,f3255,f3254,f3253,f3252,f3251,f3250,f3249,f3248,f3247,f3246,f3245,f3242,f3239,f3238,f3237,f3233,f3232,f3226,f3223,f3222,f3216,f3215,f3208,f3182,f3181,f3180,f3175,f3174,f3169,f3168,f3167,f3162,f3154,f3150,f3146,f3145,f3140,f3139,f3138,f3137,f3136,f3132,f3130,f3124,f3123,f3117,f3116,f3109,f3106,f3105,f3099,f3096,f3095,f3093,f3090,f3088,f3087,f3084,f3083,f3082,f3081,f3078,f3077,f3076,f3075,f3074,f3073,f3067,f3066,f3059,f3058,f3052,f3051,f3044,f3043,f3041,f3039,f3037,f3034,f3033,f3027,f3026,f3021,f3006,f3005,f3004,f3003,f3001,f3000,f2999,f2998,f2997,f3620,f3619,f3616,f3613,f3610,f3609,f3608,f3607,f3606,f3605,f3604,f3603,f3602,f3601,f3600,f3599,f3598,f3595,f3592,f3591,f3590,f3586,f3585,f3579,f3576,f3575,f3569,f3568,f3561,f3535,f3534,f3533,f3528,f3527,f3522,f3521,f3520,f3515,f3507,f3503,f3499,f3498,f3493,f3492,f3491,f3490,f3489,f3485,f3483,f3477,f3476,f3470,f3469,f3462,f3459,f3458,f3452,f3449,f3448,f3446,f3443,f3441,f3440,f3437,f3436,f3435,f3434,f3431,f3430,f3429,f3428,f3427,f3426,f3420,f3419,f3412,f3411,f3405,f3404,f3397,f3396,f3395,f3394,f3393,f3391,f3390,f3384,f3383,f3378,f3363,f3362,f3361,f3360,f3358,f3357,f3356,f3354,f3708,f3726,f3770,f3771,f3773,f3788,f3799,f3800,f3802,f3851,f3852,f3853,f3854,f3905,f3907,f3908,f3909,f3910,f3911,f3912,f3913,f3914,f3915,f3916,f3917,f3918,f3919,f3920,f3921,f3922,f3923,f3924,f3925,f3926,f3927,f4039,f4569,f4573,f4576,f4579,f4583,f4586,f4589,f4592,f4595,f4598,f4601,f4604,f4607,f4610,f4613,f4616,f4619,f4622,f4625,f4072,f4073,f4074,f4075,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4646,f4647,f4648,f4649,f4650,f4124,f4125,f4126,f4127,f4651,f4652,f4653,f4654,f4655,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4667,f4668,f4669,f4670,f4671,f4177,f4179,f4180,f4181,f4182,f4183,f4184,f4185,f4186,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4677,f4679,f4681,f4683,f4686,f4688,f4690,f4692,f4694,f4696,f4698,f4700,f4702,f4704,f4706,f4708,f4710,f4712,f4714,f4571,f4941,f4744,f4751,f4966,f4972,f4982,f4986,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f5037,f5039,f5041,f5043,f5045,f5047,f5049,f5051,f4899,f4900])).
fof(f5057,plain,(
  $false | $spl14),
  inference(global_subsumption,[],[f321,f320,f319,f318,f322,f573,f329,f337,f336,f341,f340,f364,f371,f370,f369,f368,f367,f379,f377,f382,f380,f387,f384,f394,f393,f574,f412,f409,f408,f407,f421,f423,f425,f424,f428,f427,f426,f431,f430,f429,f434,f433,f432,f437,f436,f435,f439,f438,f451,f450,f449,f559,f447,f446,f468,f467,f466,f465,f464,f463,f462,f472,f471,f470,f469,f476,f475,f474,f473,f480,f479,f478,f477,f484,f483,f482,f481,f558,f488,f487,f486,f485,f495,f494,f493,f492,f501,f498,f497,f496,f508,f507,f506,f557,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f525,f524,f523,f556,f521,f520,f519,f529,f528,f527,f526,f533,f532,f531,f530,f555,f575,f576,f577,f578,f580,f582,f583,f585,f587,f588,f589,f590,f591,f593,f595,f596,f597,f599,f601,f602,f604,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f617,f621,f639,f640,f643,f649,f657,f664,f681,f682,f685,f688,f695,f704,f707,f708,f711,f712,f716,f725,f729,f730,f733,f736,f740,f743,f746,f755,f756,f757,f765,f768,f772,f775,f778,f797,f800,f815,f813,f818,f836,f884,f902,f925,f923,f929,f927,f941,f938,f951,f949,f992,f1003,f1011,f1024,f1032,f1036,f1042,f1051,f1056,f1055,f1068,f1076,f1080,f1113,f1124,f1131,f1140,f1152,f1160,f1164,f1178,f1196,f1197,f1212,f1222,f1237,f1249,f1257,f1261,f1272,f1282,f1298,f1314,f1313,f1311,f1310,f1325,f1334,f1335,f1336,f1337,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1347,f1348,f1349,f1350,f1351,f1352,f1353,f1355,f1356,f1371,f1381,f1397,f1389,f1387,f1386,f1405,f1402,f1439,f1429,f1428,f1484,f1503,f1515,f1514,f1513,f1512,f1517,f1518,f1519,f1529,f1530,f1531,f1532,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1543,f1544,f1545,f1546,f1547,f1548,f1549,f1551,f1552,f1866,f1863,f1867,f1865,f1876,f1873,f1877,f1875,f1887,f1884,f1886,f1900,f1897,f1899,f1916,f1912,f1934,f1930,f1952,f1948,f1977,f1973,f2000,f2015,f2016,f2027,f2038,f2035,f2036,f2037,f2051,f2062,f2059,f2060,f2061,f2082,f2104,f2105,f2106,f2111,f2134,f2135,f2136,f2141,f2142,f2186,f2182,f2209,f2205,f2233,f2234,f2228,f2229,f2230,f2232,f2249,f2250,f2244,f2245,f2246,f2251,f2248,f2286,f2287,f2305,f2306,f2334,f2336,f2347,f2349,f2367,f2364,f2386,f2370,f2371,f2372,f2374,f2375,f2400,f2403,f2404,f2405,f2407,f2408,f2429,f2444,f2427,f2432,f2433,f2434,f2436,f2456,f2460,f2461,f2462,f2464,f2489,f2488,f2486,f2485,f2497,f2506,f2507,f2508,f2509,f2511,f2512,f2513,f2514,f2515,f2516,f2517,f2518,f2519,f2520,f2521,f2522,f2523,f2524,f2525,f2526,f2528,f2529,f2530,f2531,f2532,f2533,f2831,f2830,f2827,f2824,f2821,f2820,f2819,f2818,f2817,f2816,f2815,f2814,f2813,f2812,f2811,f2810,f2809,f2806,f2803,f2802,f2801,f2797,f2796,f2790,f2787,f2786,f2780,f2779,f2772,f2746,f2745,f2744,f2739,f2738,f2733,f2732,f2731,f2726,f2718,f2714,f2710,f2709,f2704,f2703,f2702,f2701,f2700,f2696,f2694,f2688,f2687,f2681,f2680,f2673,f2670,f2669,f2663,f2660,f2659,f2658,f2656,f2654,f2653,f2651,f2650,f2649,f2648,f2646,f2645,f2644,f2643,f2642,f2641,f2635,f2634,f2627,f2626,f2620,f2619,f2612,f2611,f2609,f2607,f2605,f2602,f2601,f2595,f2594,f2589,f2574,f2573,f2572,f2571,f2569,f2568,f2567,f2566,f2565,f3267,f3266,f3263,f3260,f3257,f3256,f3255,f3254,f3253,f3252,f3251,f3250,f3249,f3248,f3247,f3246,f3245,f3242,f3239,f3238,f3237,f3233,f3232,f3226,f3223,f3222,f3216,f3215,f3208,f3182,f3181,f3180,f3175,f3174,f3169,f3168,f3167,f3162,f3154,f3150,f3146,f3145,f3140,f3139,f3138,f3137,f3136,f3132,f3130,f3124,f3123,f3117,f3116,f3109,f3106,f3105,f3099,f3096,f3095,f3093,f3090,f3088,f3087,f3084,f3083,f3082,f3081,f3078,f3077,f3076,f3075,f3074,f3073,f3067,f3066,f3059,f3058,f3052,f3051,f3044,f3043,f3041,f3039,f3037,f3034,f3033,f3027,f3026,f3021,f3006,f3005,f3004,f3003,f3001,f3000,f2999,f2998,f2997,f3620,f3619,f3616,f3613,f3610,f3609,f3608,f3607,f3606,f3605,f3604,f3603,f3602,f3601,f3600,f3599,f3598,f3595,f3592,f3591,f3590,f3586,f3585,f3579,f3576,f3575,f3569,f3568,f3561,f3535,f3534,f3533,f3528,f3527,f3522,f3521,f3520,f3515,f3507,f3503,f3499,f3498,f3493,f3492,f3491,f3490,f3489,f3485,f3483,f3477,f3476,f3470,f3469,f3462,f3459,f3458,f3452,f3449,f3448,f3446,f3443,f3441,f3440,f3437,f3436,f3435,f3434,f3431,f3430,f3429,f3428,f3427,f3426,f3420,f3419,f3412,f3411,f3405,f3404,f3397,f3396,f3395,f3394,f3393,f3391,f3390,f3384,f3383,f3378,f3363,f3362,f3361,f3360,f3358,f3357,f3356,f3354,f3708,f3726,f3770,f3771,f3773,f3788,f3799,f3800,f3802,f3851,f3852,f3853,f3854,f3905,f3907,f3908,f3909,f3910,f3911,f3912,f3913,f3914,f3915,f3916,f3917,f3918,f3919,f3920,f3921,f3922,f3923,f3924,f3925,f3926,f3927,f4039,f4569,f4573,f4576,f4579,f4583,f4586,f4589,f4592,f4595,f4598,f4601,f4604,f4607,f4610,f4613,f4616,f4619,f4622,f4625,f4072,f4073,f4074,f4075,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4646,f4647,f4648,f4649,f4650,f4124,f4125,f4126,f4127,f4651,f4652,f4653,f4654,f4655,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4667,f4668,f4669,f4670,f4671,f4177,f4179,f4180,f4181,f4182,f4183,f4184,f4185,f4186,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4677,f4679,f4681,f4683,f4686,f4688,f4690,f4692,f4694,f4696,f4698,f4700,f4702,f4704,f4706,f4708,f4710,f4712,f4714,f4571,f4941,f4744,f4751,f4966,f4972,f4982,f4986,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f5037,f5039,f5041,f5043,f5045,f5047,f5049,f5051,f4899])).
fof(f5052,plain,(
  $false | $spl14),
  inference(global_subsumption,[],[f321,f320,f319,f318,f322,f573,f329,f337,f336,f341,f340,f364,f371,f370,f369,f368,f367,f379,f377,f382,f380,f387,f384,f394,f393,f574,f412,f409,f408,f407,f421,f423,f425,f424,f428,f427,f426,f431,f430,f429,f434,f433,f432,f437,f436,f435,f439,f438,f451,f450,f449,f559,f447,f446,f468,f467,f466,f465,f464,f463,f462,f472,f471,f470,f469,f476,f475,f474,f473,f480,f479,f478,f477,f484,f483,f482,f481,f558,f488,f487,f486,f485,f495,f494,f493,f492,f501,f498,f497,f496,f508,f507,f506,f557,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f525,f524,f523,f556,f521,f520,f519,f529,f528,f527,f526,f533,f532,f531,f530,f555,f575,f576,f577,f578,f580,f582,f583,f585,f587,f588,f589,f590,f591,f593,f595,f596,f597,f599,f601,f602,f604,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f617,f621,f639,f640,f643,f649,f657,f664,f681,f682,f685,f688,f695,f704,f707,f708,f711,f712,f716,f725,f729,f730,f733,f736,f740,f743,f746,f755,f756,f757,f765,f768,f772,f775,f778,f797,f800,f815,f813,f818,f836,f884,f902,f925,f923,f929,f927,f941,f938,f951,f949,f992,f1003,f1011,f1024,f1032,f1036,f1042,f1051,f1056,f1055,f1068,f1076,f1080,f1113,f1124,f1131,f1140,f1152,f1160,f1164,f1178,f1196,f1197,f1212,f1222,f1237,f1249,f1257,f1261,f1272,f1282,f1298,f1314,f1313,f1311,f1310,f1325,f1334,f1335,f1336,f1337,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1347,f1348,f1349,f1350,f1351,f1352,f1353,f1355,f1356,f1371,f1381,f1397,f1389,f1387,f1386,f1405,f1402,f1439,f1429,f1428,f1484,f1503,f1515,f1514,f1513,f1512,f1517,f1518,f1519,f1529,f1530,f1531,f1532,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1543,f1544,f1545,f1546,f1547,f1548,f1549,f1551,f1552,f1866,f1863,f1867,f1865,f1876,f1873,f1877,f1875,f1887,f1884,f1886,f1900,f1897,f1899,f1916,f1912,f1934,f1930,f1952,f1948,f1977,f1973,f2000,f2015,f2016,f2027,f2038,f2035,f2036,f2037,f2051,f2062,f2059,f2060,f2061,f2082,f2104,f2105,f2106,f2111,f2134,f2135,f2136,f2141,f2142,f2186,f2182,f2209,f2205,f2233,f2234,f2228,f2229,f2230,f2232,f2249,f2250,f2244,f2245,f2246,f2251,f2248,f2286,f2287,f2305,f2306,f2334,f2336,f2347,f2349,f2367,f2364,f2386,f2370,f2371,f2372,f2374,f2375,f2400,f2403,f2404,f2405,f2407,f2408,f2429,f2444,f2427,f2432,f2433,f2434,f2436,f2456,f2460,f2461,f2462,f2464,f2489,f2488,f2486,f2485,f2497,f2506,f2507,f2508,f2509,f2511,f2512,f2513,f2514,f2515,f2516,f2517,f2518,f2519,f2520,f2521,f2522,f2523,f2524,f2525,f2526,f2528,f2529,f2530,f2531,f2532,f2533,f2831,f2830,f2827,f2824,f2821,f2820,f2819,f2818,f2817,f2816,f2815,f2814,f2813,f2812,f2811,f2810,f2809,f2806,f2803,f2802,f2801,f2797,f2796,f2790,f2787,f2786,f2780,f2779,f2772,f2746,f2745,f2744,f2739,f2738,f2733,f2732,f2731,f2726,f2718,f2714,f2710,f2709,f2704,f2703,f2702,f2701,f2700,f2696,f2694,f2688,f2687,f2681,f2680,f2673,f2670,f2669,f2663,f2660,f2659,f2658,f2656,f2654,f2653,f2651,f2650,f2649,f2648,f2646,f2645,f2644,f2643,f2642,f2641,f2635,f2634,f2627,f2626,f2620,f2619,f2612,f2611,f2609,f2607,f2605,f2602,f2601,f2595,f2594,f2589,f2574,f2573,f2572,f2571,f2569,f2568,f2567,f2566,f2565,f3267,f3266,f3263,f3260,f3257,f3256,f3255,f3254,f3253,f3252,f3251,f3250,f3249,f3248,f3247,f3246,f3245,f3242,f3239,f3238,f3237,f3233,f3232,f3226,f3223,f3222,f3216,f3215,f3208,f3182,f3181,f3180,f3175,f3174,f3169,f3168,f3167,f3162,f3154,f3150,f3146,f3145,f3140,f3139,f3138,f3137,f3136,f3132,f3130,f3124,f3123,f3117,f3116,f3109,f3106,f3105,f3099,f3096,f3095,f3093,f3090,f3088,f3087,f3084,f3083,f3082,f3081,f3078,f3077,f3076,f3075,f3074,f3073,f3067,f3066,f3059,f3058,f3052,f3051,f3044,f3043,f3041,f3039,f3037,f3034,f3033,f3027,f3026,f3021,f3006,f3005,f3004,f3003,f3001,f3000,f2999,f2998,f2997,f3620,f3619,f3616,f3613,f3610,f3609,f3608,f3607,f3606,f3605,f3604,f3603,f3602,f3601,f3600,f3599,f3598,f3595,f3592,f3591,f3590,f3586,f3585,f3579,f3576,f3575,f3569,f3568,f3561,f3535,f3534,f3533,f3528,f3527,f3522,f3521,f3520,f3515,f3507,f3503,f3499,f3498,f3493,f3492,f3491,f3490,f3489,f3485,f3483,f3477,f3476,f3470,f3469,f3462,f3459,f3458,f3452,f3449,f3448,f3446,f3443,f3441,f3440,f3437,f3436,f3435,f3434,f3431,f3430,f3429,f3428,f3427,f3426,f3420,f3419,f3412,f3411,f3405,f3404,f3397,f3396,f3395,f3394,f3393,f3391,f3390,f3384,f3383,f3378,f3363,f3362,f3361,f3360,f3358,f3357,f3356,f3354,f3708,f3726,f3770,f3771,f3773,f3788,f3799,f3800,f3802,f3851,f3852,f3853,f3854,f3905,f3907,f3908,f3909,f3910,f3911,f3912,f3913,f3914,f3915,f3916,f3917,f3918,f3919,f3920,f3921,f3922,f3923,f3924,f3925,f3926,f3927,f4039,f4569,f4573,f4576,f4579,f4583,f4586,f4589,f4592,f4595,f4598,f4601,f4604,f4607,f4610,f4613,f4616,f4619,f4622,f4625,f4072,f4073,f4074,f4075,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4646,f4647,f4648,f4649,f4650,f4124,f4125,f4126,f4127,f4651,f4652,f4653,f4654,f4655,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4667,f4668,f4669,f4670,f4671,f4177,f4179,f4180,f4181,f4182,f4183,f4184,f4185,f4186,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4677,f4679,f4681,f4683,f4686,f4688,f4690,f4692,f4694,f4696,f4698,f4700,f4702,f4704,f4706,f4708,f4710,f4712,f4714,f4571,f4941,f4744,f4751,f4966,f4972,f4982,f4986,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f5037,f5039,f5041,f5043,f5045,f5047,f5049,f5051])).
fof(f5050,plain,(
  $false | $spl14),
  inference(global_subsumption,[],[f321,f320,f319,f318,f322,f573,f329,f337,f336,f341,f340,f364,f371,f370,f369,f368,f367,f379,f377,f382,f380,f387,f384,f394,f393,f574,f412,f409,f408,f407,f421,f423,f425,f424,f428,f427,f426,f431,f430,f429,f434,f433,f432,f437,f436,f435,f439,f438,f451,f450,f449,f559,f447,f446,f468,f467,f466,f465,f464,f463,f462,f472,f471,f470,f469,f476,f475,f474,f473,f480,f479,f478,f477,f484,f483,f482,f481,f558,f488,f487,f486,f485,f495,f494,f493,f492,f501,f498,f497,f496,f508,f507,f506,f557,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f525,f524,f523,f556,f521,f520,f519,f529,f528,f527,f526,f533,f532,f531,f530,f555,f575,f576,f577,f578,f580,f582,f583,f585,f587,f588,f589,f590,f591,f593,f595,f596,f597,f599,f601,f602,f604,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f617,f621,f639,f640,f643,f649,f657,f664,f681,f682,f685,f688,f695,f704,f707,f708,f711,f712,f716,f725,f729,f730,f733,f736,f740,f743,f746,f755,f756,f757,f765,f768,f772,f775,f778,f797,f800,f815,f813,f818,f836,f884,f902,f925,f923,f929,f927,f941,f938,f951,f949,f992,f1003,f1011,f1024,f1032,f1036,f1042,f1051,f1056,f1055,f1068,f1076,f1080,f1113,f1124,f1131,f1140,f1152,f1160,f1164,f1178,f1196,f1197,f1212,f1222,f1237,f1249,f1257,f1261,f1272,f1282,f1298,f1314,f1313,f1311,f1310,f1325,f1334,f1335,f1336,f1337,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1347,f1348,f1349,f1350,f1351,f1352,f1353,f1355,f1356,f1371,f1381,f1397,f1389,f1387,f1386,f1405,f1402,f1439,f1429,f1428,f1484,f1503,f1515,f1514,f1513,f1512,f1517,f1518,f1519,f1529,f1530,f1531,f1532,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1543,f1544,f1545,f1546,f1547,f1548,f1549,f1551,f1552,f1866,f1863,f1867,f1865,f1876,f1873,f1877,f1875,f1887,f1884,f1886,f1900,f1897,f1899,f1916,f1912,f1934,f1930,f1952,f1948,f1977,f1973,f2000,f2015,f2016,f2027,f2038,f2035,f2036,f2037,f2051,f2062,f2059,f2060,f2061,f2082,f2104,f2105,f2106,f2111,f2134,f2135,f2136,f2141,f2142,f2186,f2182,f2209,f2205,f2233,f2234,f2228,f2229,f2230,f2232,f2249,f2250,f2244,f2245,f2246,f2251,f2248,f2286,f2287,f2305,f2306,f2334,f2336,f2347,f2349,f2367,f2364,f2386,f2370,f2371,f2372,f2374,f2375,f2400,f2403,f2404,f2405,f2407,f2408,f2429,f2444,f2427,f2432,f2433,f2434,f2436,f2456,f2460,f2461,f2462,f2464,f2489,f2488,f2486,f2485,f2497,f2506,f2507,f2508,f2509,f2511,f2512,f2513,f2514,f2515,f2516,f2517,f2518,f2519,f2520,f2521,f2522,f2523,f2524,f2525,f2526,f2528,f2529,f2530,f2531,f2532,f2533,f2831,f2830,f2827,f2824,f2821,f2820,f2819,f2818,f2817,f2816,f2815,f2814,f2813,f2812,f2811,f2810,f2809,f2806,f2803,f2802,f2801,f2797,f2796,f2790,f2787,f2786,f2780,f2779,f2772,f2746,f2745,f2744,f2739,f2738,f2733,f2732,f2731,f2726,f2718,f2714,f2710,f2709,f2704,f2703,f2702,f2701,f2700,f2696,f2694,f2688,f2687,f2681,f2680,f2673,f2670,f2669,f2663,f2660,f2659,f2658,f2656,f2654,f2653,f2651,f2650,f2649,f2648,f2646,f2645,f2644,f2643,f2642,f2641,f2635,f2634,f2627,f2626,f2620,f2619,f2612,f2611,f2609,f2607,f2605,f2602,f2601,f2595,f2594,f2589,f2574,f2573,f2572,f2571,f2569,f2568,f2567,f2566,f2565,f3267,f3266,f3263,f3260,f3257,f3256,f3255,f3254,f3253,f3252,f3251,f3250,f3249,f3248,f3247,f3246,f3245,f3242,f3239,f3238,f3237,f3233,f3232,f3226,f3223,f3222,f3216,f3215,f3208,f3182,f3181,f3180,f3175,f3174,f3169,f3168,f3167,f3162,f3154,f3150,f3146,f3145,f3140,f3139,f3138,f3137,f3136,f3132,f3130,f3124,f3123,f3117,f3116,f3109,f3106,f3105,f3099,f3096,f3095,f3093,f3090,f3088,f3087,f3084,f3083,f3082,f3081,f3078,f3077,f3076,f3075,f3074,f3073,f3067,f3066,f3059,f3058,f3052,f3051,f3044,f3043,f3041,f3039,f3037,f3034,f3033,f3027,f3026,f3021,f3006,f3005,f3004,f3003,f3001,f3000,f2999,f2998,f2997,f3620,f3619,f3616,f3613,f3610,f3609,f3608,f3607,f3606,f3605,f3604,f3603,f3602,f3601,f3600,f3599,f3598,f3595,f3592,f3591,f3590,f3586,f3585,f3579,f3576,f3575,f3569,f3568,f3561,f3535,f3534,f3533,f3528,f3527,f3522,f3521,f3520,f3515,f3507,f3503,f3499,f3498,f3493,f3492,f3491,f3490,f3489,f3485,f3483,f3477,f3476,f3470,f3469,f3462,f3459,f3458,f3452,f3449,f3448,f3446,f3443,f3441,f3440,f3437,f3436,f3435,f3434,f3431,f3430,f3429,f3428,f3427,f3426,f3420,f3419,f3412,f3411,f3405,f3404,f3397,f3396,f3395,f3394,f3393,f3391,f3390,f3384,f3383,f3378,f3363,f3362,f3361,f3360,f3358,f3357,f3356,f3354,f3708,f3726,f3770,f3771,f3773,f3788,f3799,f3800,f3802,f3851,f3852,f3853,f3854,f3905,f3907,f3908,f3909,f3910,f3911,f3912,f3913,f3914,f3915,f3916,f3917,f3918,f3919,f3920,f3921,f3922,f3923,f3924,f3925,f3926,f3927,f4039,f4569,f4573,f4576,f4579,f4583,f4586,f4589,f4592,f4595,f4598,f4601,f4604,f4607,f4610,f4613,f4616,f4619,f4622,f4625,f4072,f4073,f4074,f4075,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4646,f4647,f4648,f4649,f4650,f4124,f4125,f4126,f4127,f4651,f4652,f4653,f4654,f4655,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4667,f4668,f4669,f4670,f4671,f4177,f4179,f4180,f4181,f4182,f4183,f4184,f4185,f4186,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4677,f4679,f4681,f4683,f4686,f4688,f4690,f4692,f4694,f4696,f4698,f4700,f4702,f4704,f4706,f4708,f4710,f4712,f4714,f4571,f4941,f4744,f4751,f4966,f4972,f4982,f4986,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f5037,f5039,f5041,f5043,f5045,f5047,f5049])).
fof(f5048,plain,(
  $false | $spl14),
  inference(global_subsumption,[],[f321,f320,f319,f318,f322,f573,f329,f337,f336,f341,f340,f364,f371,f370,f369,f368,f367,f379,f377,f382,f380,f387,f384,f394,f393,f574,f412,f409,f408,f407,f421,f423,f425,f424,f428,f427,f426,f431,f430,f429,f434,f433,f432,f437,f436,f435,f439,f438,f451,f450,f449,f559,f447,f446,f468,f467,f466,f465,f464,f463,f462,f472,f471,f470,f469,f476,f475,f474,f473,f480,f479,f478,f477,f484,f483,f482,f481,f558,f488,f487,f486,f485,f495,f494,f493,f492,f501,f498,f497,f496,f508,f507,f506,f557,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f525,f524,f523,f556,f521,f520,f519,f529,f528,f527,f526,f533,f532,f531,f530,f555,f575,f576,f577,f578,f580,f582,f583,f585,f587,f588,f589,f590,f591,f593,f595,f596,f597,f599,f601,f602,f604,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f617,f621,f639,f640,f643,f649,f657,f664,f681,f682,f685,f688,f695,f704,f707,f708,f711,f712,f716,f725,f729,f730,f733,f736,f740,f743,f746,f755,f756,f757,f765,f768,f772,f775,f778,f797,f800,f815,f813,f818,f836,f884,f902,f925,f923,f929,f927,f941,f938,f951,f949,f992,f1003,f1011,f1024,f1032,f1036,f1042,f1051,f1056,f1055,f1068,f1076,f1080,f1113,f1124,f1131,f1140,f1152,f1160,f1164,f1178,f1196,f1197,f1212,f1222,f1237,f1249,f1257,f1261,f1272,f1282,f1298,f1314,f1313,f1311,f1310,f1325,f1334,f1335,f1336,f1337,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1347,f1348,f1349,f1350,f1351,f1352,f1353,f1355,f1356,f1371,f1381,f1397,f1389,f1387,f1386,f1405,f1402,f1439,f1429,f1428,f1484,f1503,f1515,f1514,f1513,f1512,f1517,f1518,f1519,f1529,f1530,f1531,f1532,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1543,f1544,f1545,f1546,f1547,f1548,f1549,f1551,f1552,f1866,f1863,f1867,f1865,f1876,f1873,f1877,f1875,f1887,f1884,f1886,f1900,f1897,f1899,f1916,f1912,f1934,f1930,f1952,f1948,f1977,f1973,f2000,f2015,f2016,f2027,f2038,f2035,f2036,f2037,f2051,f2062,f2059,f2060,f2061,f2082,f2104,f2105,f2106,f2111,f2134,f2135,f2136,f2141,f2142,f2186,f2182,f2209,f2205,f2233,f2234,f2228,f2229,f2230,f2232,f2249,f2250,f2244,f2245,f2246,f2251,f2248,f2286,f2287,f2305,f2306,f2334,f2336,f2347,f2349,f2367,f2364,f2386,f2370,f2371,f2372,f2374,f2375,f2400,f2403,f2404,f2405,f2407,f2408,f2429,f2444,f2427,f2432,f2433,f2434,f2436,f2456,f2460,f2461,f2462,f2464,f2489,f2488,f2486,f2485,f2497,f2506,f2507,f2508,f2509,f2511,f2512,f2513,f2514,f2515,f2516,f2517,f2518,f2519,f2520,f2521,f2522,f2523,f2524,f2525,f2526,f2528,f2529,f2530,f2531,f2532,f2533,f2831,f2830,f2827,f2824,f2821,f2820,f2819,f2818,f2817,f2816,f2815,f2814,f2813,f2812,f2811,f2810,f2809,f2806,f2803,f2802,f2801,f2797,f2796,f2790,f2787,f2786,f2780,f2779,f2772,f2746,f2745,f2744,f2739,f2738,f2733,f2732,f2731,f2726,f2718,f2714,f2710,f2709,f2704,f2703,f2702,f2701,f2700,f2696,f2694,f2688,f2687,f2681,f2680,f2673,f2670,f2669,f2663,f2660,f2659,f2658,f2656,f2654,f2653,f2651,f2650,f2649,f2648,f2646,f2645,f2644,f2643,f2642,f2641,f2635,f2634,f2627,f2626,f2620,f2619,f2612,f2611,f2609,f2607,f2605,f2602,f2601,f2595,f2594,f2589,f2574,f2573,f2572,f2571,f2569,f2568,f2567,f2566,f2565,f3267,f3266,f3263,f3260,f3257,f3256,f3255,f3254,f3253,f3252,f3251,f3250,f3249,f3248,f3247,f3246,f3245,f3242,f3239,f3238,f3237,f3233,f3232,f3226,f3223,f3222,f3216,f3215,f3208,f3182,f3181,f3180,f3175,f3174,f3169,f3168,f3167,f3162,f3154,f3150,f3146,f3145,f3140,f3139,f3138,f3137,f3136,f3132,f3130,f3124,f3123,f3117,f3116,f3109,f3106,f3105,f3099,f3096,f3095,f3093,f3090,f3088,f3087,f3084,f3083,f3082,f3081,f3078,f3077,f3076,f3075,f3074,f3073,f3067,f3066,f3059,f3058,f3052,f3051,f3044,f3043,f3041,f3039,f3037,f3034,f3033,f3027,f3026,f3021,f3006,f3005,f3004,f3003,f3001,f3000,f2999,f2998,f2997,f3620,f3619,f3616,f3613,f3610,f3609,f3608,f3607,f3606,f3605,f3604,f3603,f3602,f3601,f3600,f3599,f3598,f3595,f3592,f3591,f3590,f3586,f3585,f3579,f3576,f3575,f3569,f3568,f3561,f3535,f3534,f3533,f3528,f3527,f3522,f3521,f3520,f3515,f3507,f3503,f3499,f3498,f3493,f3492,f3491,f3490,f3489,f3485,f3483,f3477,f3476,f3470,f3469,f3462,f3459,f3458,f3452,f3449,f3448,f3446,f3443,f3441,f3440,f3437,f3436,f3435,f3434,f3431,f3430,f3429,f3428,f3427,f3426,f3420,f3419,f3412,f3411,f3405,f3404,f3397,f3396,f3395,f3394,f3393,f3391,f3390,f3384,f3383,f3378,f3363,f3362,f3361,f3360,f3358,f3357,f3356,f3354,f3708,f3726,f3770,f3771,f3773,f3788,f3799,f3800,f3802,f3851,f3852,f3853,f3854,f3905,f3907,f3908,f3909,f3910,f3911,f3912,f3913,f3914,f3915,f3916,f3917,f3918,f3919,f3920,f3921,f3922,f3923,f3924,f3925,f3926,f3927,f4039,f4569,f4573,f4576,f4579,f4583,f4586,f4589,f4592,f4595,f4598,f4601,f4604,f4607,f4610,f4613,f4616,f4619,f4622,f4625,f4072,f4073,f4074,f4075,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4646,f4647,f4648,f4649,f4650,f4124,f4125,f4126,f4127,f4651,f4652,f4653,f4654,f4655,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4667,f4668,f4669,f4670,f4671,f4177,f4179,f4180,f4181,f4182,f4183,f4184,f4185,f4186,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4677,f4679,f4681,f4683,f4686,f4688,f4690,f4692,f4694,f4696,f4698,f4700,f4702,f4704,f4706,f4708,f4710,f4712,f4714,f4571,f4941,f4744,f4751,f4966,f4972,f4982,f4986,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f5037,f5039,f5041,f5043,f5045,f5047])).
fof(f5046,plain,(
  $false | $spl14),
  inference(global_subsumption,[],[f321,f320,f319,f318,f322,f573,f329,f337,f336,f341,f340,f364,f371,f370,f369,f368,f367,f379,f377,f382,f380,f387,f384,f394,f393,f574,f412,f409,f408,f407,f421,f423,f425,f424,f428,f427,f426,f431,f430,f429,f434,f433,f432,f437,f436,f435,f439,f438,f451,f450,f449,f559,f447,f446,f468,f467,f466,f465,f464,f463,f462,f472,f471,f470,f469,f476,f475,f474,f473,f480,f479,f478,f477,f484,f483,f482,f481,f558,f488,f487,f486,f485,f495,f494,f493,f492,f501,f498,f497,f496,f508,f507,f506,f557,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f525,f524,f523,f556,f521,f520,f519,f529,f528,f527,f526,f533,f532,f531,f530,f555,f575,f576,f577,f578,f580,f582,f583,f585,f587,f588,f589,f590,f591,f593,f595,f596,f597,f599,f601,f602,f604,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f617,f621,f639,f640,f643,f649,f657,f664,f681,f682,f685,f688,f695,f704,f707,f708,f711,f712,f716,f725,f729,f730,f733,f736,f740,f743,f746,f755,f756,f757,f765,f768,f772,f775,f778,f797,f800,f815,f813,f818,f836,f884,f902,f925,f923,f929,f927,f941,f938,f951,f949,f992,f1003,f1011,f1024,f1032,f1036,f1042,f1051,f1056,f1055,f1068,f1076,f1080,f1113,f1124,f1131,f1140,f1152,f1160,f1164,f1178,f1196,f1197,f1212,f1222,f1237,f1249,f1257,f1261,f1272,f1282,f1298,f1314,f1313,f1311,f1310,f1325,f1334,f1335,f1336,f1337,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1347,f1348,f1349,f1350,f1351,f1352,f1353,f1355,f1356,f1371,f1381,f1397,f1389,f1387,f1386,f1405,f1402,f1439,f1429,f1428,f1484,f1503,f1515,f1514,f1513,f1512,f1517,f1518,f1519,f1529,f1530,f1531,f1532,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1543,f1544,f1545,f1546,f1547,f1548,f1549,f1551,f1552,f1866,f1863,f1867,f1865,f1876,f1873,f1877,f1875,f1887,f1884,f1886,f1900,f1897,f1899,f1916,f1912,f1934,f1930,f1952,f1948,f1977,f1973,f2000,f2015,f2016,f2027,f2038,f2035,f2036,f2037,f2051,f2062,f2059,f2060,f2061,f2082,f2104,f2105,f2106,f2111,f2134,f2135,f2136,f2141,f2142,f2186,f2182,f2209,f2205,f2233,f2234,f2228,f2229,f2230,f2232,f2249,f2250,f2244,f2245,f2246,f2251,f2248,f2286,f2287,f2305,f2306,f2334,f2336,f2347,f2349,f2367,f2364,f2386,f2370,f2371,f2372,f2374,f2375,f2400,f2403,f2404,f2405,f2407,f2408,f2429,f2444,f2427,f2432,f2433,f2434,f2436,f2456,f2460,f2461,f2462,f2464,f2489,f2488,f2486,f2485,f2497,f2506,f2507,f2508,f2509,f2511,f2512,f2513,f2514,f2515,f2516,f2517,f2518,f2519,f2520,f2521,f2522,f2523,f2524,f2525,f2526,f2528,f2529,f2530,f2531,f2532,f2533,f2831,f2830,f2827,f2824,f2821,f2820,f2819,f2818,f2817,f2816,f2815,f2814,f2813,f2812,f2811,f2810,f2809,f2806,f2803,f2802,f2801,f2797,f2796,f2790,f2787,f2786,f2780,f2779,f2772,f2746,f2745,f2744,f2739,f2738,f2733,f2732,f2731,f2726,f2718,f2714,f2710,f2709,f2704,f2703,f2702,f2701,f2700,f2696,f2694,f2688,f2687,f2681,f2680,f2673,f2670,f2669,f2663,f2660,f2659,f2658,f2656,f2654,f2653,f2651,f2650,f2649,f2648,f2646,f2645,f2644,f2643,f2642,f2641,f2635,f2634,f2627,f2626,f2620,f2619,f2612,f2611,f2609,f2607,f2605,f2602,f2601,f2595,f2594,f2589,f2574,f2573,f2572,f2571,f2569,f2568,f2567,f2566,f2565,f3267,f3266,f3263,f3260,f3257,f3256,f3255,f3254,f3253,f3252,f3251,f3250,f3249,f3248,f3247,f3246,f3245,f3242,f3239,f3238,f3237,f3233,f3232,f3226,f3223,f3222,f3216,f3215,f3208,f3182,f3181,f3180,f3175,f3174,f3169,f3168,f3167,f3162,f3154,f3150,f3146,f3145,f3140,f3139,f3138,f3137,f3136,f3132,f3130,f3124,f3123,f3117,f3116,f3109,f3106,f3105,f3099,f3096,f3095,f3093,f3090,f3088,f3087,f3084,f3083,f3082,f3081,f3078,f3077,f3076,f3075,f3074,f3073,f3067,f3066,f3059,f3058,f3052,f3051,f3044,f3043,f3041,f3039,f3037,f3034,f3033,f3027,f3026,f3021,f3006,f3005,f3004,f3003,f3001,f3000,f2999,f2998,f2997,f3620,f3619,f3616,f3613,f3610,f3609,f3608,f3607,f3606,f3605,f3604,f3603,f3602,f3601,f3600,f3599,f3598,f3595,f3592,f3591,f3590,f3586,f3585,f3579,f3576,f3575,f3569,f3568,f3561,f3535,f3534,f3533,f3528,f3527,f3522,f3521,f3520,f3515,f3507,f3503,f3499,f3498,f3493,f3492,f3491,f3490,f3489,f3485,f3483,f3477,f3476,f3470,f3469,f3462,f3459,f3458,f3452,f3449,f3448,f3446,f3443,f3441,f3440,f3437,f3436,f3435,f3434,f3431,f3430,f3429,f3428,f3427,f3426,f3420,f3419,f3412,f3411,f3405,f3404,f3397,f3396,f3395,f3394,f3393,f3391,f3390,f3384,f3383,f3378,f3363,f3362,f3361,f3360,f3358,f3357,f3356,f3354,f3708,f3726,f3770,f3771,f3773,f3788,f3799,f3800,f3802,f3851,f3852,f3853,f3854,f3905,f3907,f3908,f3909,f3910,f3911,f3912,f3913,f3914,f3915,f3916,f3917,f3918,f3919,f3920,f3921,f3922,f3923,f3924,f3925,f3926,f3927,f4039,f4569,f4573,f4576,f4579,f4583,f4586,f4589,f4592,f4595,f4598,f4601,f4604,f4607,f4610,f4613,f4616,f4619,f4622,f4625,f4072,f4073,f4074,f4075,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4646,f4647,f4648,f4649,f4650,f4124,f4125,f4126,f4127,f4651,f4652,f4653,f4654,f4655,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4667,f4668,f4669,f4670,f4671,f4177,f4179,f4180,f4181,f4182,f4183,f4184,f4185,f4186,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4677,f4679,f4681,f4683,f4686,f4688,f4690,f4692,f4694,f4696,f4698,f4700,f4702,f4704,f4706,f4708,f4710,f4712,f4714,f4571,f4941,f4744,f4751,f4966,f4972,f4982,f4986,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f5037,f5039,f5041,f5043,f5045])).
fof(f5044,plain,(
  $false | $spl14),
  inference(global_subsumption,[],[f321,f320,f319,f318,f322,f573,f329,f337,f336,f341,f340,f364,f371,f370,f369,f368,f367,f379,f377,f382,f380,f387,f384,f394,f393,f574,f412,f409,f408,f407,f421,f423,f425,f424,f428,f427,f426,f431,f430,f429,f434,f433,f432,f437,f436,f435,f439,f438,f451,f450,f449,f559,f447,f446,f468,f467,f466,f465,f464,f463,f462,f472,f471,f470,f469,f476,f475,f474,f473,f480,f479,f478,f477,f484,f483,f482,f481,f558,f488,f487,f486,f485,f495,f494,f493,f492,f501,f498,f497,f496,f508,f507,f506,f557,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f525,f524,f523,f556,f521,f520,f519,f529,f528,f527,f526,f533,f532,f531,f530,f555,f575,f576,f577,f578,f580,f582,f583,f585,f587,f588,f589,f590,f591,f593,f595,f596,f597,f599,f601,f602,f604,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f617,f621,f639,f640,f643,f649,f657,f664,f681,f682,f685,f688,f695,f704,f707,f708,f711,f712,f716,f725,f729,f730,f733,f736,f740,f743,f746,f755,f756,f757,f765,f768,f772,f775,f778,f797,f800,f815,f813,f818,f836,f884,f902,f925,f923,f929,f927,f941,f938,f951,f949,f992,f1003,f1011,f1024,f1032,f1036,f1042,f1051,f1056,f1055,f1068,f1076,f1080,f1113,f1124,f1131,f1140,f1152,f1160,f1164,f1178,f1196,f1197,f1212,f1222,f1237,f1249,f1257,f1261,f1272,f1282,f1298,f1314,f1313,f1311,f1310,f1325,f1334,f1335,f1336,f1337,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1347,f1348,f1349,f1350,f1351,f1352,f1353,f1355,f1356,f1371,f1381,f1397,f1389,f1387,f1386,f1405,f1402,f1439,f1429,f1428,f1484,f1503,f1515,f1514,f1513,f1512,f1517,f1518,f1519,f1529,f1530,f1531,f1532,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1543,f1544,f1545,f1546,f1547,f1548,f1549,f1551,f1552,f1866,f1863,f1867,f1865,f1876,f1873,f1877,f1875,f1887,f1884,f1886,f1900,f1897,f1899,f1916,f1912,f1934,f1930,f1952,f1948,f1977,f1973,f2000,f2015,f2016,f2027,f2038,f2035,f2036,f2037,f2051,f2062,f2059,f2060,f2061,f2082,f2104,f2105,f2106,f2111,f2134,f2135,f2136,f2141,f2142,f2186,f2182,f2209,f2205,f2233,f2234,f2228,f2229,f2230,f2232,f2249,f2250,f2244,f2245,f2246,f2251,f2248,f2286,f2287,f2305,f2306,f2334,f2336,f2347,f2349,f2367,f2364,f2386,f2370,f2371,f2372,f2374,f2375,f2400,f2403,f2404,f2405,f2407,f2408,f2429,f2444,f2427,f2432,f2433,f2434,f2436,f2456,f2460,f2461,f2462,f2464,f2489,f2488,f2486,f2485,f2497,f2506,f2507,f2508,f2509,f2511,f2512,f2513,f2514,f2515,f2516,f2517,f2518,f2519,f2520,f2521,f2522,f2523,f2524,f2525,f2526,f2528,f2529,f2530,f2531,f2532,f2533,f2831,f2830,f2827,f2824,f2821,f2820,f2819,f2818,f2817,f2816,f2815,f2814,f2813,f2812,f2811,f2810,f2809,f2806,f2803,f2802,f2801,f2797,f2796,f2790,f2787,f2786,f2780,f2779,f2772,f2746,f2745,f2744,f2739,f2738,f2733,f2732,f2731,f2726,f2718,f2714,f2710,f2709,f2704,f2703,f2702,f2701,f2700,f2696,f2694,f2688,f2687,f2681,f2680,f2673,f2670,f2669,f2663,f2660,f2659,f2658,f2656,f2654,f2653,f2651,f2650,f2649,f2648,f2646,f2645,f2644,f2643,f2642,f2641,f2635,f2634,f2627,f2626,f2620,f2619,f2612,f2611,f2609,f2607,f2605,f2602,f2601,f2595,f2594,f2589,f2574,f2573,f2572,f2571,f2569,f2568,f2567,f2566,f2565,f3267,f3266,f3263,f3260,f3257,f3256,f3255,f3254,f3253,f3252,f3251,f3250,f3249,f3248,f3247,f3246,f3245,f3242,f3239,f3238,f3237,f3233,f3232,f3226,f3223,f3222,f3216,f3215,f3208,f3182,f3181,f3180,f3175,f3174,f3169,f3168,f3167,f3162,f3154,f3150,f3146,f3145,f3140,f3139,f3138,f3137,f3136,f3132,f3130,f3124,f3123,f3117,f3116,f3109,f3106,f3105,f3099,f3096,f3095,f3093,f3090,f3088,f3087,f3084,f3083,f3082,f3081,f3078,f3077,f3076,f3075,f3074,f3073,f3067,f3066,f3059,f3058,f3052,f3051,f3044,f3043,f3041,f3039,f3037,f3034,f3033,f3027,f3026,f3021,f3006,f3005,f3004,f3003,f3001,f3000,f2999,f2998,f2997,f3620,f3619,f3616,f3613,f3610,f3609,f3608,f3607,f3606,f3605,f3604,f3603,f3602,f3601,f3600,f3599,f3598,f3595,f3592,f3591,f3590,f3586,f3585,f3579,f3576,f3575,f3569,f3568,f3561,f3535,f3534,f3533,f3528,f3527,f3522,f3521,f3520,f3515,f3507,f3503,f3499,f3498,f3493,f3492,f3491,f3490,f3489,f3485,f3483,f3477,f3476,f3470,f3469,f3462,f3459,f3458,f3452,f3449,f3448,f3446,f3443,f3441,f3440,f3437,f3436,f3435,f3434,f3431,f3430,f3429,f3428,f3427,f3426,f3420,f3419,f3412,f3411,f3405,f3404,f3397,f3396,f3395,f3394,f3393,f3391,f3390,f3384,f3383,f3378,f3363,f3362,f3361,f3360,f3358,f3357,f3356,f3354,f3708,f3726,f3770,f3771,f3773,f3788,f3799,f3800,f3802,f3851,f3852,f3853,f3854,f3905,f3907,f3908,f3909,f3910,f3911,f3912,f3913,f3914,f3915,f3916,f3917,f3918,f3919,f3920,f3921,f3922,f3923,f3924,f3925,f3926,f3927,f4039,f4569,f4573,f4576,f4579,f4583,f4586,f4589,f4592,f4595,f4598,f4601,f4604,f4607,f4610,f4613,f4616,f4619,f4622,f4625,f4072,f4073,f4074,f4075,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4646,f4647,f4648,f4649,f4650,f4124,f4125,f4126,f4127,f4651,f4652,f4653,f4654,f4655,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4667,f4668,f4669,f4670,f4671,f4177,f4179,f4180,f4181,f4182,f4183,f4184,f4185,f4186,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4677,f4679,f4681,f4683,f4686,f4688,f4690,f4692,f4694,f4696,f4698,f4700,f4702,f4704,f4706,f4708,f4710,f4712,f4714,f4571,f4941,f4744,f4751,f4966,f4972,f4982,f4986,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f5037,f5039,f5041,f5043])).
fof(f5042,plain,(
  $false | $spl14),
  inference(global_subsumption,[],[f321,f320,f319,f318,f322,f573,f329,f337,f336,f341,f340,f364,f371,f370,f369,f368,f367,f379,f377,f382,f380,f387,f384,f394,f393,f574,f412,f409,f408,f407,f421,f423,f425,f424,f428,f427,f426,f431,f430,f429,f434,f433,f432,f437,f436,f435,f439,f438,f451,f450,f449,f559,f447,f446,f468,f467,f466,f465,f464,f463,f462,f472,f471,f470,f469,f476,f475,f474,f473,f480,f479,f478,f477,f484,f483,f482,f481,f558,f488,f487,f486,f485,f495,f494,f493,f492,f501,f498,f497,f496,f508,f507,f506,f557,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f525,f524,f523,f556,f521,f520,f519,f529,f528,f527,f526,f533,f532,f531,f530,f555,f575,f576,f577,f578,f580,f582,f583,f585,f587,f588,f589,f590,f591,f593,f595,f596,f597,f599,f601,f602,f604,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f617,f621,f639,f640,f643,f649,f657,f664,f681,f682,f685,f688,f695,f704,f707,f708,f711,f712,f716,f725,f729,f730,f733,f736,f740,f743,f746,f755,f756,f757,f765,f768,f772,f775,f778,f797,f800,f815,f813,f818,f836,f884,f902,f925,f923,f929,f927,f941,f938,f951,f949,f992,f1003,f1011,f1024,f1032,f1036,f1042,f1051,f1056,f1055,f1068,f1076,f1080,f1113,f1124,f1131,f1140,f1152,f1160,f1164,f1178,f1196,f1197,f1212,f1222,f1237,f1249,f1257,f1261,f1272,f1282,f1298,f1314,f1313,f1311,f1310,f1325,f1334,f1335,f1336,f1337,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1347,f1348,f1349,f1350,f1351,f1352,f1353,f1355,f1356,f1371,f1381,f1397,f1389,f1387,f1386,f1405,f1402,f1439,f1429,f1428,f1484,f1503,f1515,f1514,f1513,f1512,f1517,f1518,f1519,f1529,f1530,f1531,f1532,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1543,f1544,f1545,f1546,f1547,f1548,f1549,f1551,f1552,f1866,f1863,f1867,f1865,f1876,f1873,f1877,f1875,f1887,f1884,f1886,f1900,f1897,f1899,f1916,f1912,f1934,f1930,f1952,f1948,f1977,f1973,f2000,f2015,f2016,f2027,f2038,f2035,f2036,f2037,f2051,f2062,f2059,f2060,f2061,f2082,f2104,f2105,f2106,f2111,f2134,f2135,f2136,f2141,f2142,f2186,f2182,f2209,f2205,f2233,f2234,f2228,f2229,f2230,f2232,f2249,f2250,f2244,f2245,f2246,f2251,f2248,f2286,f2287,f2305,f2306,f2334,f2336,f2347,f2349,f2367,f2364,f2386,f2370,f2371,f2372,f2374,f2375,f2400,f2403,f2404,f2405,f2407,f2408,f2429,f2444,f2427,f2432,f2433,f2434,f2436,f2456,f2460,f2461,f2462,f2464,f2489,f2488,f2486,f2485,f2497,f2506,f2507,f2508,f2509,f2511,f2512,f2513,f2514,f2515,f2516,f2517,f2518,f2519,f2520,f2521,f2522,f2523,f2524,f2525,f2526,f2528,f2529,f2530,f2531,f2532,f2533,f2831,f2830,f2827,f2824,f2821,f2820,f2819,f2818,f2817,f2816,f2815,f2814,f2813,f2812,f2811,f2810,f2809,f2806,f2803,f2802,f2801,f2797,f2796,f2790,f2787,f2786,f2780,f2779,f2772,f2746,f2745,f2744,f2739,f2738,f2733,f2732,f2731,f2726,f2718,f2714,f2710,f2709,f2704,f2703,f2702,f2701,f2700,f2696,f2694,f2688,f2687,f2681,f2680,f2673,f2670,f2669,f2663,f2660,f2659,f2658,f2656,f2654,f2653,f2651,f2650,f2649,f2648,f2646,f2645,f2644,f2643,f2642,f2641,f2635,f2634,f2627,f2626,f2620,f2619,f2612,f2611,f2609,f2607,f2605,f2602,f2601,f2595,f2594,f2589,f2574,f2573,f2572,f2571,f2569,f2568,f2567,f2566,f2565,f3267,f3266,f3263,f3260,f3257,f3256,f3255,f3254,f3253,f3252,f3251,f3250,f3249,f3248,f3247,f3246,f3245,f3242,f3239,f3238,f3237,f3233,f3232,f3226,f3223,f3222,f3216,f3215,f3208,f3182,f3181,f3180,f3175,f3174,f3169,f3168,f3167,f3162,f3154,f3150,f3146,f3145,f3140,f3139,f3138,f3137,f3136,f3132,f3130,f3124,f3123,f3117,f3116,f3109,f3106,f3105,f3099,f3096,f3095,f3093,f3090,f3088,f3087,f3084,f3083,f3082,f3081,f3078,f3077,f3076,f3075,f3074,f3073,f3067,f3066,f3059,f3058,f3052,f3051,f3044,f3043,f3041,f3039,f3037,f3034,f3033,f3027,f3026,f3021,f3006,f3005,f3004,f3003,f3001,f3000,f2999,f2998,f2997,f3620,f3619,f3616,f3613,f3610,f3609,f3608,f3607,f3606,f3605,f3604,f3603,f3602,f3601,f3600,f3599,f3598,f3595,f3592,f3591,f3590,f3586,f3585,f3579,f3576,f3575,f3569,f3568,f3561,f3535,f3534,f3533,f3528,f3527,f3522,f3521,f3520,f3515,f3507,f3503,f3499,f3498,f3493,f3492,f3491,f3490,f3489,f3485,f3483,f3477,f3476,f3470,f3469,f3462,f3459,f3458,f3452,f3449,f3448,f3446,f3443,f3441,f3440,f3437,f3436,f3435,f3434,f3431,f3430,f3429,f3428,f3427,f3426,f3420,f3419,f3412,f3411,f3405,f3404,f3397,f3396,f3395,f3394,f3393,f3391,f3390,f3384,f3383,f3378,f3363,f3362,f3361,f3360,f3358,f3357,f3356,f3354,f3708,f3726,f3770,f3771,f3773,f3788,f3799,f3800,f3802,f3851,f3852,f3853,f3854,f3905,f3907,f3908,f3909,f3910,f3911,f3912,f3913,f3914,f3915,f3916,f3917,f3918,f3919,f3920,f3921,f3922,f3923,f3924,f3925,f3926,f3927,f4039,f4569,f4573,f4576,f4579,f4583,f4586,f4589,f4592,f4595,f4598,f4601,f4604,f4607,f4610,f4613,f4616,f4619,f4622,f4625,f4072,f4073,f4074,f4075,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4646,f4647,f4648,f4649,f4650,f4124,f4125,f4126,f4127,f4651,f4652,f4653,f4654,f4655,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4667,f4668,f4669,f4670,f4671,f4177,f4179,f4180,f4181,f4182,f4183,f4184,f4185,f4186,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4677,f4679,f4681,f4683,f4686,f4688,f4690,f4692,f4694,f4696,f4698,f4700,f4702,f4704,f4706,f4708,f4710,f4712,f4714,f4571,f4941,f4744,f4751,f4966,f4972,f4982,f4986,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f5037,f5039,f5041])).
fof(f5040,plain,(
  $false | $spl14),
  inference(global_subsumption,[],[f321,f320,f319,f318,f322,f573,f329,f337,f336,f341,f340,f364,f371,f370,f369,f368,f367,f379,f377,f382,f380,f387,f384,f394,f393,f574,f412,f409,f408,f407,f421,f423,f425,f424,f428,f427,f426,f431,f430,f429,f434,f433,f432,f437,f436,f435,f439,f438,f451,f450,f449,f559,f447,f446,f468,f467,f466,f465,f464,f463,f462,f472,f471,f470,f469,f476,f475,f474,f473,f480,f479,f478,f477,f484,f483,f482,f481,f558,f488,f487,f486,f485,f495,f494,f493,f492,f501,f498,f497,f496,f508,f507,f506,f557,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f525,f524,f523,f556,f521,f520,f519,f529,f528,f527,f526,f533,f532,f531,f530,f555,f575,f576,f577,f578,f580,f582,f583,f585,f587,f588,f589,f590,f591,f593,f595,f596,f597,f599,f601,f602,f604,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f617,f621,f639,f640,f643,f649,f657,f664,f681,f682,f685,f688,f695,f704,f707,f708,f711,f712,f716,f725,f729,f730,f733,f736,f740,f743,f746,f755,f756,f757,f765,f768,f772,f775,f778,f797,f800,f815,f813,f818,f836,f884,f902,f925,f923,f929,f927,f941,f938,f951,f949,f992,f1003,f1011,f1024,f1032,f1036,f1042,f1051,f1056,f1055,f1068,f1076,f1080,f1113,f1124,f1131,f1140,f1152,f1160,f1164,f1178,f1196,f1197,f1212,f1222,f1237,f1249,f1257,f1261,f1272,f1282,f1298,f1314,f1313,f1311,f1310,f1325,f1334,f1335,f1336,f1337,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1347,f1348,f1349,f1350,f1351,f1352,f1353,f1355,f1356,f1371,f1381,f1397,f1389,f1387,f1386,f1405,f1402,f1439,f1429,f1428,f1484,f1503,f1515,f1514,f1513,f1512,f1517,f1518,f1519,f1529,f1530,f1531,f1532,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1543,f1544,f1545,f1546,f1547,f1548,f1549,f1551,f1552,f1866,f1863,f1867,f1865,f1876,f1873,f1877,f1875,f1887,f1884,f1886,f1900,f1897,f1899,f1916,f1912,f1934,f1930,f1952,f1948,f1977,f1973,f2000,f2015,f2016,f2027,f2038,f2035,f2036,f2037,f2051,f2062,f2059,f2060,f2061,f2082,f2104,f2105,f2106,f2111,f2134,f2135,f2136,f2141,f2142,f2186,f2182,f2209,f2205,f2233,f2234,f2228,f2229,f2230,f2232,f2249,f2250,f2244,f2245,f2246,f2251,f2248,f2286,f2287,f2305,f2306,f2334,f2336,f2347,f2349,f2367,f2364,f2386,f2370,f2371,f2372,f2374,f2375,f2400,f2403,f2404,f2405,f2407,f2408,f2429,f2444,f2427,f2432,f2433,f2434,f2436,f2456,f2460,f2461,f2462,f2464,f2489,f2488,f2486,f2485,f2497,f2506,f2507,f2508,f2509,f2511,f2512,f2513,f2514,f2515,f2516,f2517,f2518,f2519,f2520,f2521,f2522,f2523,f2524,f2525,f2526,f2528,f2529,f2530,f2531,f2532,f2533,f2831,f2830,f2827,f2824,f2821,f2820,f2819,f2818,f2817,f2816,f2815,f2814,f2813,f2812,f2811,f2810,f2809,f2806,f2803,f2802,f2801,f2797,f2796,f2790,f2787,f2786,f2780,f2779,f2772,f2746,f2745,f2744,f2739,f2738,f2733,f2732,f2731,f2726,f2718,f2714,f2710,f2709,f2704,f2703,f2702,f2701,f2700,f2696,f2694,f2688,f2687,f2681,f2680,f2673,f2670,f2669,f2663,f2660,f2659,f2658,f2656,f2654,f2653,f2651,f2650,f2649,f2648,f2646,f2645,f2644,f2643,f2642,f2641,f2635,f2634,f2627,f2626,f2620,f2619,f2612,f2611,f2609,f2607,f2605,f2602,f2601,f2595,f2594,f2589,f2574,f2573,f2572,f2571,f2569,f2568,f2567,f2566,f2565,f3267,f3266,f3263,f3260,f3257,f3256,f3255,f3254,f3253,f3252,f3251,f3250,f3249,f3248,f3247,f3246,f3245,f3242,f3239,f3238,f3237,f3233,f3232,f3226,f3223,f3222,f3216,f3215,f3208,f3182,f3181,f3180,f3175,f3174,f3169,f3168,f3167,f3162,f3154,f3150,f3146,f3145,f3140,f3139,f3138,f3137,f3136,f3132,f3130,f3124,f3123,f3117,f3116,f3109,f3106,f3105,f3099,f3096,f3095,f3093,f3090,f3088,f3087,f3084,f3083,f3082,f3081,f3078,f3077,f3076,f3075,f3074,f3073,f3067,f3066,f3059,f3058,f3052,f3051,f3044,f3043,f3041,f3039,f3037,f3034,f3033,f3027,f3026,f3021,f3006,f3005,f3004,f3003,f3001,f3000,f2999,f2998,f2997,f3620,f3619,f3616,f3613,f3610,f3609,f3608,f3607,f3606,f3605,f3604,f3603,f3602,f3601,f3600,f3599,f3598,f3595,f3592,f3591,f3590,f3586,f3585,f3579,f3576,f3575,f3569,f3568,f3561,f3535,f3534,f3533,f3528,f3527,f3522,f3521,f3520,f3515,f3507,f3503,f3499,f3498,f3493,f3492,f3491,f3490,f3489,f3485,f3483,f3477,f3476,f3470,f3469,f3462,f3459,f3458,f3452,f3449,f3448,f3446,f3443,f3441,f3440,f3437,f3436,f3435,f3434,f3431,f3430,f3429,f3428,f3427,f3426,f3420,f3419,f3412,f3411,f3405,f3404,f3397,f3396,f3395,f3394,f3393,f3391,f3390,f3384,f3383,f3378,f3363,f3362,f3361,f3360,f3358,f3357,f3356,f3354,f3708,f3726,f3770,f3771,f3773,f3788,f3799,f3800,f3802,f3851,f3852,f3853,f3854,f3905,f3907,f3908,f3909,f3910,f3911,f3912,f3913,f3914,f3915,f3916,f3917,f3918,f3919,f3920,f3921,f3922,f3923,f3924,f3925,f3926,f3927,f4039,f4569,f4573,f4576,f4579,f4583,f4586,f4589,f4592,f4595,f4598,f4601,f4604,f4607,f4610,f4613,f4616,f4619,f4622,f4625,f4072,f4073,f4074,f4075,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4646,f4647,f4648,f4649,f4650,f4124,f4125,f4126,f4127,f4651,f4652,f4653,f4654,f4655,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4667,f4668,f4669,f4670,f4671,f4177,f4179,f4180,f4181,f4182,f4183,f4184,f4185,f4186,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4677,f4679,f4681,f4683,f4686,f4688,f4690,f4692,f4694,f4696,f4698,f4700,f4702,f4704,f4706,f4708,f4710,f4712,f4714,f4571,f4941,f4744,f4751,f4966,f4972,f4982,f4986,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f5037,f5039])).
fof(f5038,plain,(
  $false | $spl14),
  inference(global_subsumption,[],[f321,f320,f319,f318,f322,f573,f329,f337,f336,f341,f340,f364,f371,f370,f369,f368,f367,f379,f377,f382,f380,f387,f384,f394,f393,f574,f412,f409,f408,f407,f421,f423,f425,f424,f428,f427,f426,f431,f430,f429,f434,f433,f432,f437,f436,f435,f439,f438,f451,f450,f449,f559,f447,f446,f468,f467,f466,f465,f464,f463,f462,f472,f471,f470,f469,f476,f475,f474,f473,f480,f479,f478,f477,f484,f483,f482,f481,f558,f488,f487,f486,f485,f495,f494,f493,f492,f501,f498,f497,f496,f508,f507,f506,f557,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f525,f524,f523,f556,f521,f520,f519,f529,f528,f527,f526,f533,f532,f531,f530,f555,f575,f576,f577,f578,f580,f582,f583,f585,f587,f588,f589,f590,f591,f593,f595,f596,f597,f599,f601,f602,f604,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f617,f621,f639,f640,f643,f649,f657,f664,f681,f682,f685,f688,f695,f704,f707,f708,f711,f712,f716,f725,f729,f730,f733,f736,f740,f743,f746,f755,f756,f757,f765,f768,f772,f775,f778,f797,f800,f815,f813,f818,f836,f884,f902,f925,f923,f929,f927,f941,f938,f951,f949,f992,f1003,f1011,f1024,f1032,f1036,f1042,f1051,f1056,f1055,f1068,f1076,f1080,f1113,f1124,f1131,f1140,f1152,f1160,f1164,f1178,f1196,f1197,f1212,f1222,f1237,f1249,f1257,f1261,f1272,f1282,f1298,f1314,f1313,f1311,f1310,f1325,f1334,f1335,f1336,f1337,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1347,f1348,f1349,f1350,f1351,f1352,f1353,f1355,f1356,f1371,f1381,f1397,f1389,f1387,f1386,f1405,f1402,f1439,f1429,f1428,f1484,f1503,f1515,f1514,f1513,f1512,f1517,f1518,f1519,f1529,f1530,f1531,f1532,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1543,f1544,f1545,f1546,f1547,f1548,f1549,f1551,f1552,f1866,f1863,f1867,f1865,f1876,f1873,f1877,f1875,f1887,f1884,f1886,f1900,f1897,f1899,f1916,f1912,f1934,f1930,f1952,f1948,f1977,f1973,f2000,f2015,f2016,f2027,f2038,f2035,f2036,f2037,f2051,f2062,f2059,f2060,f2061,f2082,f2104,f2105,f2106,f2111,f2134,f2135,f2136,f2141,f2142,f2186,f2182,f2209,f2205,f2233,f2234,f2228,f2229,f2230,f2232,f2249,f2250,f2244,f2245,f2246,f2251,f2248,f2286,f2287,f2305,f2306,f2334,f2336,f2347,f2349,f2367,f2364,f2386,f2370,f2371,f2372,f2374,f2375,f2400,f2403,f2404,f2405,f2407,f2408,f2429,f2444,f2427,f2432,f2433,f2434,f2436,f2456,f2460,f2461,f2462,f2464,f2489,f2488,f2486,f2485,f2497,f2506,f2507,f2508,f2509,f2511,f2512,f2513,f2514,f2515,f2516,f2517,f2518,f2519,f2520,f2521,f2522,f2523,f2524,f2525,f2526,f2528,f2529,f2530,f2531,f2532,f2533,f2831,f2830,f2827,f2824,f2821,f2820,f2819,f2818,f2817,f2816,f2815,f2814,f2813,f2812,f2811,f2810,f2809,f2806,f2803,f2802,f2801,f2797,f2796,f2790,f2787,f2786,f2780,f2779,f2772,f2746,f2745,f2744,f2739,f2738,f2733,f2732,f2731,f2726,f2718,f2714,f2710,f2709,f2704,f2703,f2702,f2701,f2700,f2696,f2694,f2688,f2687,f2681,f2680,f2673,f2670,f2669,f2663,f2660,f2659,f2658,f2656,f2654,f2653,f2651,f2650,f2649,f2648,f2646,f2645,f2644,f2643,f2642,f2641,f2635,f2634,f2627,f2626,f2620,f2619,f2612,f2611,f2609,f2607,f2605,f2602,f2601,f2595,f2594,f2589,f2574,f2573,f2572,f2571,f2569,f2568,f2567,f2566,f2565,f3267,f3266,f3263,f3260,f3257,f3256,f3255,f3254,f3253,f3252,f3251,f3250,f3249,f3248,f3247,f3246,f3245,f3242,f3239,f3238,f3237,f3233,f3232,f3226,f3223,f3222,f3216,f3215,f3208,f3182,f3181,f3180,f3175,f3174,f3169,f3168,f3167,f3162,f3154,f3150,f3146,f3145,f3140,f3139,f3138,f3137,f3136,f3132,f3130,f3124,f3123,f3117,f3116,f3109,f3106,f3105,f3099,f3096,f3095,f3093,f3090,f3088,f3087,f3084,f3083,f3082,f3081,f3078,f3077,f3076,f3075,f3074,f3073,f3067,f3066,f3059,f3058,f3052,f3051,f3044,f3043,f3041,f3039,f3037,f3034,f3033,f3027,f3026,f3021,f3006,f3005,f3004,f3003,f3001,f3000,f2999,f2998,f2997,f3620,f3619,f3616,f3613,f3610,f3609,f3608,f3607,f3606,f3605,f3604,f3603,f3602,f3601,f3600,f3599,f3598,f3595,f3592,f3591,f3590,f3586,f3585,f3579,f3576,f3575,f3569,f3568,f3561,f3535,f3534,f3533,f3528,f3527,f3522,f3521,f3520,f3515,f3507,f3503,f3499,f3498,f3493,f3492,f3491,f3490,f3489,f3485,f3483,f3477,f3476,f3470,f3469,f3462,f3459,f3458,f3452,f3449,f3448,f3446,f3443,f3441,f3440,f3437,f3436,f3435,f3434,f3431,f3430,f3429,f3428,f3427,f3426,f3420,f3419,f3412,f3411,f3405,f3404,f3397,f3396,f3395,f3394,f3393,f3391,f3390,f3384,f3383,f3378,f3363,f3362,f3361,f3360,f3358,f3357,f3356,f3354,f3708,f3726,f3770,f3771,f3773,f3788,f3799,f3800,f3802,f3851,f3852,f3853,f3854,f3905,f3907,f3908,f3909,f3910,f3911,f3912,f3913,f3914,f3915,f3916,f3917,f3918,f3919,f3920,f3921,f3922,f3923,f3924,f3925,f3926,f3927,f4039,f4569,f4573,f4576,f4579,f4583,f4586,f4589,f4592,f4595,f4598,f4601,f4604,f4607,f4610,f4613,f4616,f4619,f4622,f4625,f4072,f4073,f4074,f4075,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4646,f4647,f4648,f4649,f4650,f4124,f4125,f4126,f4127,f4651,f4652,f4653,f4654,f4655,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4667,f4668,f4669,f4670,f4671,f4177,f4179,f4180,f4181,f4182,f4183,f4184,f4185,f4186,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4677,f4679,f4681,f4683,f4686,f4688,f4690,f4692,f4694,f4696,f4698,f4700,f4702,f4704,f4706,f4708,f4710,f4712,f4714,f4571,f4941,f4744,f4751,f4966,f4972,f4982,f4986,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f5037])).
fof(f5036,plain,(
  $false | $spl14),
  inference(global_subsumption,[],[f321,f320,f319,f318,f322,f573,f329,f337,f336,f341,f340,f364,f371,f370,f369,f368,f367,f379,f377,f382,f380,f387,f384,f394,f393,f574,f412,f409,f408,f407,f421,f423,f425,f424,f428,f427,f426,f431,f430,f429,f434,f433,f432,f437,f436,f435,f439,f438,f451,f450,f449,f559,f447,f446,f468,f467,f466,f465,f464,f463,f462,f472,f471,f470,f469,f476,f475,f474,f473,f480,f479,f478,f477,f484,f483,f482,f481,f558,f488,f487,f486,f485,f495,f494,f493,f492,f501,f498,f497,f496,f508,f507,f506,f557,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f525,f524,f523,f556,f521,f520,f519,f529,f528,f527,f526,f533,f532,f531,f530,f555,f575,f576,f577,f578,f580,f582,f583,f585,f587,f588,f589,f590,f591,f593,f595,f596,f597,f599,f601,f602,f604,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f617,f621,f639,f640,f643,f649,f657,f664,f681,f682,f685,f688,f695,f704,f707,f708,f711,f712,f716,f725,f729,f730,f733,f736,f740,f743,f746,f755,f756,f757,f765,f768,f772,f775,f778,f797,f800,f815,f813,f818,f836,f884,f902,f925,f923,f929,f927,f941,f938,f951,f949,f992,f1003,f1011,f1024,f1032,f1036,f1042,f1051,f1056,f1055,f1068,f1076,f1080,f1113,f1124,f1131,f1140,f1152,f1160,f1164,f1178,f1196,f1197,f1212,f1222,f1237,f1249,f1257,f1261,f1272,f1282,f1298,f1314,f1313,f1311,f1310,f1325,f1334,f1335,f1336,f1337,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1347,f1348,f1349,f1350,f1351,f1352,f1353,f1355,f1356,f1371,f1381,f1397,f1389,f1387,f1386,f1405,f1402,f1439,f1429,f1428,f1484,f1503,f1515,f1514,f1513,f1512,f1517,f1518,f1519,f1529,f1530,f1531,f1532,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1543,f1544,f1545,f1546,f1547,f1548,f1549,f1551,f1552,f1866,f1863,f1867,f1865,f1876,f1873,f1877,f1875,f1887,f1884,f1886,f1900,f1897,f1899,f1916,f1912,f1934,f1930,f1952,f1948,f1977,f1973,f2000,f2015,f2016,f2027,f2038,f2035,f2036,f2037,f2051,f2062,f2059,f2060,f2061,f2082,f2104,f2105,f2106,f2111,f2134,f2135,f2136,f2141,f2142,f2186,f2182,f2209,f2205,f2233,f2234,f2228,f2229,f2230,f2232,f2249,f2250,f2244,f2245,f2246,f2251,f2248,f2286,f2287,f2305,f2306,f2334,f2336,f2347,f2349,f2367,f2364,f2386,f2370,f2371,f2372,f2374,f2375,f2400,f2403,f2404,f2405,f2407,f2408,f2429,f2444,f2427,f2432,f2433,f2434,f2436,f2456,f2460,f2461,f2462,f2464,f2489,f2488,f2486,f2485,f2497,f2506,f2507,f2508,f2509,f2511,f2512,f2513,f2514,f2515,f2516,f2517,f2518,f2519,f2520,f2521,f2522,f2523,f2524,f2525,f2526,f2528,f2529,f2530,f2531,f2532,f2533,f2831,f2830,f2827,f2824,f2821,f2820,f2819,f2818,f2817,f2816,f2815,f2814,f2813,f2812,f2811,f2810,f2809,f2806,f2803,f2802,f2801,f2797,f2796,f2790,f2787,f2786,f2780,f2779,f2772,f2746,f2745,f2744,f2739,f2738,f2733,f2732,f2731,f2726,f2718,f2714,f2710,f2709,f2704,f2703,f2702,f2701,f2700,f2696,f2694,f2688,f2687,f2681,f2680,f2673,f2670,f2669,f2663,f2660,f2659,f2658,f2656,f2654,f2653,f2651,f2650,f2649,f2648,f2646,f2645,f2644,f2643,f2642,f2641,f2635,f2634,f2627,f2626,f2620,f2619,f2612,f2611,f2609,f2607,f2605,f2602,f2601,f2595,f2594,f2589,f2574,f2573,f2572,f2571,f2569,f2568,f2567,f2566,f2565,f3267,f3266,f3263,f3260,f3257,f3256,f3255,f3254,f3253,f3252,f3251,f3250,f3249,f3248,f3247,f3246,f3245,f3242,f3239,f3238,f3237,f3233,f3232,f3226,f3223,f3222,f3216,f3215,f3208,f3182,f3181,f3180,f3175,f3174,f3169,f3168,f3167,f3162,f3154,f3150,f3146,f3145,f3140,f3139,f3138,f3137,f3136,f3132,f3130,f3124,f3123,f3117,f3116,f3109,f3106,f3105,f3099,f3096,f3095,f3093,f3090,f3088,f3087,f3084,f3083,f3082,f3081,f3078,f3077,f3076,f3075,f3074,f3073,f3067,f3066,f3059,f3058,f3052,f3051,f3044,f3043,f3041,f3039,f3037,f3034,f3033,f3027,f3026,f3021,f3006,f3005,f3004,f3003,f3001,f3000,f2999,f2998,f2997,f3620,f3619,f3616,f3613,f3610,f3609,f3608,f3607,f3606,f3605,f3604,f3603,f3602,f3601,f3600,f3599,f3598,f3595,f3592,f3591,f3590,f3586,f3585,f3579,f3576,f3575,f3569,f3568,f3561,f3535,f3534,f3533,f3528,f3527,f3522,f3521,f3520,f3515,f3507,f3503,f3499,f3498,f3493,f3492,f3491,f3490,f3489,f3485,f3483,f3477,f3476,f3470,f3469,f3462,f3459,f3458,f3452,f3449,f3448,f3446,f3443,f3441,f3440,f3437,f3436,f3435,f3434,f3431,f3430,f3429,f3428,f3427,f3426,f3420,f3419,f3412,f3411,f3405,f3404,f3397,f3396,f3395,f3394,f3393,f3391,f3390,f3384,f3383,f3378,f3363,f3362,f3361,f3360,f3358,f3357,f3356,f3354,f3708,f3726,f3770,f3771,f3773,f3788,f3799,f3800,f3802,f3851,f3852,f3853,f3854,f3905,f3907,f3908,f3909,f3910,f3911,f3912,f3913,f3914,f3915,f3916,f3917,f3918,f3919,f3920,f3921,f3922,f3923,f3924,f3925,f3926,f3927,f4039,f4569,f4573,f4576,f4579,f4583,f4586,f4589,f4592,f4595,f4598,f4601,f4604,f4607,f4610,f4613,f4616,f4619,f4622,f4625,f4072,f4073,f4074,f4075,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4646,f4647,f4648,f4649,f4650,f4124,f4125,f4126,f4127,f4651,f4652,f4653,f4654,f4655,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4667,f4668,f4669,f4670,f4671,f4177,f4179,f4180,f4181,f4182,f4183,f4184,f4185,f4186,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4677,f4679,f4681,f4683,f4686,f4688,f4690,f4692,f4694,f4696,f4698,f4700,f4702,f4704,f4706,f4708,f4710,f4712,f4714,f4571,f4941,f4744,f4751,f4966,f4972,f4982,f4986,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035])).
fof(f5034,plain,(
  $false | $spl14),
  inference(global_subsumption,[],[f321,f320,f319,f318,f322,f573,f329,f337,f336,f341,f340,f364,f371,f370,f369,f368,f367,f379,f377,f382,f380,f387,f384,f394,f393,f574,f412,f409,f408,f407,f421,f423,f425,f424,f428,f427,f426,f431,f430,f429,f434,f433,f432,f437,f436,f435,f439,f438,f451,f450,f449,f559,f447,f446,f468,f467,f466,f465,f464,f463,f462,f472,f471,f470,f469,f476,f475,f474,f473,f480,f479,f478,f477,f484,f483,f482,f481,f558,f488,f487,f486,f485,f495,f494,f493,f492,f501,f498,f497,f496,f508,f507,f506,f557,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f525,f524,f523,f556,f521,f520,f519,f529,f528,f527,f526,f533,f532,f531,f530,f555,f575,f576,f577,f578,f580,f582,f583,f585,f587,f588,f589,f590,f591,f593,f595,f596,f597,f599,f601,f602,f604,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f617,f621,f639,f640,f643,f649,f657,f664,f681,f682,f685,f688,f695,f704,f707,f708,f711,f712,f716,f725,f729,f730,f733,f736,f740,f743,f746,f755,f756,f757,f765,f768,f772,f775,f778,f797,f800,f815,f813,f818,f836,f884,f902,f925,f923,f929,f927,f941,f938,f951,f949,f992,f1003,f1011,f1024,f1032,f1036,f1042,f1051,f1056,f1055,f1068,f1076,f1080,f1113,f1124,f1131,f1140,f1152,f1160,f1164,f1178,f1196,f1197,f1212,f1222,f1237,f1249,f1257,f1261,f1272,f1282,f1298,f1314,f1313,f1311,f1310,f1325,f1334,f1335,f1336,f1337,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1347,f1348,f1349,f1350,f1351,f1352,f1353,f1355,f1356,f1371,f1381,f1397,f1389,f1387,f1386,f1405,f1402,f1439,f1429,f1428,f1484,f1503,f1515,f1514,f1513,f1512,f1517,f1518,f1519,f1529,f1530,f1531,f1532,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1543,f1544,f1545,f1546,f1547,f1548,f1549,f1551,f1552,f1866,f1863,f1867,f1865,f1876,f1873,f1877,f1875,f1887,f1884,f1886,f1900,f1897,f1899,f1916,f1912,f1934,f1930,f1952,f1948,f1977,f1973,f2000,f2015,f2016,f2027,f2038,f2035,f2036,f2037,f2051,f2062,f2059,f2060,f2061,f2082,f2104,f2105,f2106,f2111,f2134,f2135,f2136,f2141,f2142,f2186,f2182,f2209,f2205,f2233,f2234,f2228,f2229,f2230,f2232,f2249,f2250,f2244,f2245,f2246,f2251,f2248,f2286,f2287,f2305,f2306,f2334,f2336,f2347,f2349,f2367,f2364,f2386,f2370,f2371,f2372,f2374,f2375,f2400,f2403,f2404,f2405,f2407,f2408,f2429,f2444,f2427,f2432,f2433,f2434,f2436,f2456,f2460,f2461,f2462,f2464,f2489,f2488,f2486,f2485,f2497,f2506,f2507,f2508,f2509,f2511,f2512,f2513,f2514,f2515,f2516,f2517,f2518,f2519,f2520,f2521,f2522,f2523,f2524,f2525,f2526,f2528,f2529,f2530,f2531,f2532,f2533,f2831,f2830,f2827,f2824,f2821,f2820,f2819,f2818,f2817,f2816,f2815,f2814,f2813,f2812,f2811,f2810,f2809,f2806,f2803,f2802,f2801,f2797,f2796,f2790,f2787,f2786,f2780,f2779,f2772,f2746,f2745,f2744,f2739,f2738,f2733,f2732,f2731,f2726,f2718,f2714,f2710,f2709,f2704,f2703,f2702,f2701,f2700,f2696,f2694,f2688,f2687,f2681,f2680,f2673,f2670,f2669,f2663,f2660,f2659,f2658,f2656,f2654,f2653,f2651,f2650,f2649,f2648,f2646,f2645,f2644,f2643,f2642,f2641,f2635,f2634,f2627,f2626,f2620,f2619,f2612,f2611,f2609,f2607,f2605,f2602,f2601,f2595,f2594,f2589,f2574,f2573,f2572,f2571,f2569,f2568,f2567,f2566,f2565,f3267,f3266,f3263,f3260,f3257,f3256,f3255,f3254,f3253,f3252,f3251,f3250,f3249,f3248,f3247,f3246,f3245,f3242,f3239,f3238,f3237,f3233,f3232,f3226,f3223,f3222,f3216,f3215,f3208,f3182,f3181,f3180,f3175,f3174,f3169,f3168,f3167,f3162,f3154,f3150,f3146,f3145,f3140,f3139,f3138,f3137,f3136,f3132,f3130,f3124,f3123,f3117,f3116,f3109,f3106,f3105,f3099,f3096,f3095,f3093,f3090,f3088,f3087,f3084,f3083,f3082,f3081,f3078,f3077,f3076,f3075,f3074,f3073,f3067,f3066,f3059,f3058,f3052,f3051,f3044,f3043,f3041,f3039,f3037,f3034,f3033,f3027,f3026,f3021,f3006,f3005,f3004,f3003,f3001,f3000,f2999,f2998,f2997,f3620,f3619,f3616,f3613,f3610,f3609,f3608,f3607,f3606,f3605,f3604,f3603,f3602,f3601,f3600,f3599,f3598,f3595,f3592,f3591,f3590,f3586,f3585,f3579,f3576,f3575,f3569,f3568,f3561,f3535,f3534,f3533,f3528,f3527,f3522,f3521,f3520,f3515,f3507,f3503,f3499,f3498,f3493,f3492,f3491,f3490,f3489,f3485,f3483,f3477,f3476,f3470,f3469,f3462,f3459,f3458,f3452,f3449,f3448,f3446,f3443,f3441,f3440,f3437,f3436,f3435,f3434,f3431,f3430,f3429,f3428,f3427,f3426,f3420,f3419,f3412,f3411,f3405,f3404,f3397,f3396,f3395,f3394,f3393,f3391,f3390,f3384,f3383,f3378,f3363,f3362,f3361,f3360,f3358,f3357,f3356,f3354,f3708,f3726,f3770,f3771,f3773,f3788,f3799,f3800,f3802,f3851,f3852,f3853,f3854,f3905,f3907,f3908,f3909,f3910,f3911,f3912,f3913,f3914,f3915,f3916,f3917,f3918,f3919,f3920,f3921,f3922,f3923,f3924,f3925,f3926,f3927,f4039,f4569,f4573,f4576,f4579,f4583,f4586,f4589,f4592,f4595,f4598,f4601,f4604,f4607,f4610,f4613,f4616,f4619,f4622,f4625,f4072,f4073,f4074,f4075,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4646,f4647,f4648,f4649,f4650,f4124,f4125,f4126,f4127,f4651,f4652,f4653,f4654,f4655,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4667,f4668,f4669,f4670,f4671,f4177,f4179,f4180,f4181,f4182,f4183,f4184,f4185,f4186,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4677,f4679,f4681,f4683,f4686,f4688,f4690,f4692,f4694,f4696,f4698,f4700,f4702,f4704,f4706,f4708,f4710,f4712,f4714,f4571,f4941,f4744,f4751,f4966,f4972,f4982,f4986,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033])).
fof(f5032,plain,(
  $false | $spl14),
  inference(global_subsumption,[],[f321,f320,f319,f318,f322,f573,f329,f337,f336,f341,f340,f364,f371,f370,f369,f368,f367,f379,f377,f382,f380,f387,f384,f394,f393,f574,f412,f409,f408,f407,f421,f423,f425,f424,f428,f427,f426,f431,f430,f429,f434,f433,f432,f437,f436,f435,f439,f438,f451,f450,f449,f559,f447,f446,f468,f467,f466,f465,f464,f463,f462,f472,f471,f470,f469,f476,f475,f474,f473,f480,f479,f478,f477,f484,f483,f482,f481,f558,f488,f487,f486,f485,f495,f494,f493,f492,f501,f498,f497,f496,f508,f507,f506,f557,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f525,f524,f523,f556,f521,f520,f519,f529,f528,f527,f526,f533,f532,f531,f530,f555,f575,f576,f577,f578,f580,f582,f583,f585,f587,f588,f589,f590,f591,f593,f595,f596,f597,f599,f601,f602,f604,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f617,f621,f639,f640,f643,f649,f657,f664,f681,f682,f685,f688,f695,f704,f707,f708,f711,f712,f716,f725,f729,f730,f733,f736,f740,f743,f746,f755,f756,f757,f765,f768,f772,f775,f778,f797,f800,f815,f813,f818,f836,f884,f902,f925,f923,f929,f927,f941,f938,f951,f949,f992,f1003,f1011,f1024,f1032,f1036,f1042,f1051,f1056,f1055,f1068,f1076,f1080,f1113,f1124,f1131,f1140,f1152,f1160,f1164,f1178,f1196,f1197,f1212,f1222,f1237,f1249,f1257,f1261,f1272,f1282,f1298,f1314,f1313,f1311,f1310,f1325,f1334,f1335,f1336,f1337,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1347,f1348,f1349,f1350,f1351,f1352,f1353,f1355,f1356,f1371,f1381,f1397,f1389,f1387,f1386,f1405,f1402,f1439,f1429,f1428,f1484,f1503,f1515,f1514,f1513,f1512,f1517,f1518,f1519,f1529,f1530,f1531,f1532,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1543,f1544,f1545,f1546,f1547,f1548,f1549,f1551,f1552,f1866,f1863,f1867,f1865,f1876,f1873,f1877,f1875,f1887,f1884,f1886,f1900,f1897,f1899,f1916,f1912,f1934,f1930,f1952,f1948,f1977,f1973,f2000,f2015,f2016,f2027,f2038,f2035,f2036,f2037,f2051,f2062,f2059,f2060,f2061,f2082,f2104,f2105,f2106,f2111,f2134,f2135,f2136,f2141,f2142,f2186,f2182,f2209,f2205,f2233,f2234,f2228,f2229,f2230,f2232,f2249,f2250,f2244,f2245,f2246,f2251,f2248,f2286,f2287,f2305,f2306,f2334,f2336,f2347,f2349,f2367,f2364,f2386,f2370,f2371,f2372,f2374,f2375,f2400,f2403,f2404,f2405,f2407,f2408,f2429,f2444,f2427,f2432,f2433,f2434,f2436,f2456,f2460,f2461,f2462,f2464,f2489,f2488,f2486,f2485,f2497,f2506,f2507,f2508,f2509,f2511,f2512,f2513,f2514,f2515,f2516,f2517,f2518,f2519,f2520,f2521,f2522,f2523,f2524,f2525,f2526,f2528,f2529,f2530,f2531,f2532,f2533,f2831,f2830,f2827,f2824,f2821,f2820,f2819,f2818,f2817,f2816,f2815,f2814,f2813,f2812,f2811,f2810,f2809,f2806,f2803,f2802,f2801,f2797,f2796,f2790,f2787,f2786,f2780,f2779,f2772,f2746,f2745,f2744,f2739,f2738,f2733,f2732,f2731,f2726,f2718,f2714,f2710,f2709,f2704,f2703,f2702,f2701,f2700,f2696,f2694,f2688,f2687,f2681,f2680,f2673,f2670,f2669,f2663,f2660,f2659,f2658,f2656,f2654,f2653,f2651,f2650,f2649,f2648,f2646,f2645,f2644,f2643,f2642,f2641,f2635,f2634,f2627,f2626,f2620,f2619,f2612,f2611,f2609,f2607,f2605,f2602,f2601,f2595,f2594,f2589,f2574,f2573,f2572,f2571,f2569,f2568,f2567,f2566,f2565,f3267,f3266,f3263,f3260,f3257,f3256,f3255,f3254,f3253,f3252,f3251,f3250,f3249,f3248,f3247,f3246,f3245,f3242,f3239,f3238,f3237,f3233,f3232,f3226,f3223,f3222,f3216,f3215,f3208,f3182,f3181,f3180,f3175,f3174,f3169,f3168,f3167,f3162,f3154,f3150,f3146,f3145,f3140,f3139,f3138,f3137,f3136,f3132,f3130,f3124,f3123,f3117,f3116,f3109,f3106,f3105,f3099,f3096,f3095,f3093,f3090,f3088,f3087,f3084,f3083,f3082,f3081,f3078,f3077,f3076,f3075,f3074,f3073,f3067,f3066,f3059,f3058,f3052,f3051,f3044,f3043,f3041,f3039,f3037,f3034,f3033,f3027,f3026,f3021,f3006,f3005,f3004,f3003,f3001,f3000,f2999,f2998,f2997,f3620,f3619,f3616,f3613,f3610,f3609,f3608,f3607,f3606,f3605,f3604,f3603,f3602,f3601,f3600,f3599,f3598,f3595,f3592,f3591,f3590,f3586,f3585,f3579,f3576,f3575,f3569,f3568,f3561,f3535,f3534,f3533,f3528,f3527,f3522,f3521,f3520,f3515,f3507,f3503,f3499,f3498,f3493,f3492,f3491,f3490,f3489,f3485,f3483,f3477,f3476,f3470,f3469,f3462,f3459,f3458,f3452,f3449,f3448,f3446,f3443,f3441,f3440,f3437,f3436,f3435,f3434,f3431,f3430,f3429,f3428,f3427,f3426,f3420,f3419,f3412,f3411,f3405,f3404,f3397,f3396,f3395,f3394,f3393,f3391,f3390,f3384,f3383,f3378,f3363,f3362,f3361,f3360,f3358,f3357,f3356,f3354,f3708,f3726,f3770,f3771,f3773,f3788,f3799,f3800,f3802,f3851,f3852,f3853,f3854,f3905,f3907,f3908,f3909,f3910,f3911,f3912,f3913,f3914,f3915,f3916,f3917,f3918,f3919,f3920,f3921,f3922,f3923,f3924,f3925,f3926,f3927,f4039,f4569,f4573,f4576,f4579,f4583,f4586,f4589,f4592,f4595,f4598,f4601,f4604,f4607,f4610,f4613,f4616,f4619,f4622,f4625,f4072,f4073,f4074,f4075,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4646,f4647,f4648,f4649,f4650,f4124,f4125,f4126,f4127,f4651,f4652,f4653,f4654,f4655,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4667,f4668,f4669,f4670,f4671,f4177,f4179,f4180,f4181,f4182,f4183,f4184,f4185,f4186,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4677,f4679,f4681,f4683,f4686,f4688,f4690,f4692,f4694,f4696,f4698,f4700,f4702,f4704,f4706,f4708,f4710,f4712,f4714,f4571,f4941,f4744,f4751,f4966,f4972,f4982,f4986,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031])).
fof(f5030,plain,(
  $false | $spl14),
  inference(global_subsumption,[],[f321,f320,f319,f318,f322,f573,f329,f337,f336,f341,f340,f364,f371,f370,f369,f368,f367,f379,f377,f382,f380,f387,f384,f394,f393,f574,f412,f409,f408,f407,f421,f423,f425,f424,f428,f427,f426,f431,f430,f429,f434,f433,f432,f437,f436,f435,f439,f438,f451,f450,f449,f559,f447,f446,f468,f467,f466,f465,f464,f463,f462,f472,f471,f470,f469,f476,f475,f474,f473,f480,f479,f478,f477,f484,f483,f482,f481,f558,f488,f487,f486,f485,f495,f494,f493,f492,f501,f498,f497,f496,f508,f507,f506,f557,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f525,f524,f523,f556,f521,f520,f519,f529,f528,f527,f526,f533,f532,f531,f530,f555,f575,f576,f577,f578,f580,f582,f583,f585,f587,f588,f589,f590,f591,f593,f595,f596,f597,f599,f601,f602,f604,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f617,f621,f639,f640,f643,f649,f657,f664,f681,f682,f685,f688,f695,f704,f707,f708,f711,f712,f716,f725,f729,f730,f733,f736,f740,f743,f746,f755,f756,f757,f765,f768,f772,f775,f778,f797,f800,f815,f813,f818,f836,f884,f902,f925,f923,f929,f927,f941,f938,f951,f949,f992,f1003,f1011,f1024,f1032,f1036,f1042,f1051,f1056,f1055,f1068,f1076,f1080,f1113,f1124,f1131,f1140,f1152,f1160,f1164,f1178,f1196,f1197,f1212,f1222,f1237,f1249,f1257,f1261,f1272,f1282,f1298,f1314,f1313,f1311,f1310,f1325,f1334,f1335,f1336,f1337,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1347,f1348,f1349,f1350,f1351,f1352,f1353,f1355,f1356,f1371,f1381,f1397,f1389,f1387,f1386,f1405,f1402,f1439,f1429,f1428,f1484,f1503,f1515,f1514,f1513,f1512,f1517,f1518,f1519,f1529,f1530,f1531,f1532,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1543,f1544,f1545,f1546,f1547,f1548,f1549,f1551,f1552,f1866,f1863,f1867,f1865,f1876,f1873,f1877,f1875,f1887,f1884,f1886,f1900,f1897,f1899,f1916,f1912,f1934,f1930,f1952,f1948,f1977,f1973,f2000,f2015,f2016,f2027,f2038,f2035,f2036,f2037,f2051,f2062,f2059,f2060,f2061,f2082,f2104,f2105,f2106,f2111,f2134,f2135,f2136,f2141,f2142,f2186,f2182,f2209,f2205,f2233,f2234,f2228,f2229,f2230,f2232,f2249,f2250,f2244,f2245,f2246,f2251,f2248,f2286,f2287,f2305,f2306,f2334,f2336,f2347,f2349,f2367,f2364,f2386,f2370,f2371,f2372,f2374,f2375,f2400,f2403,f2404,f2405,f2407,f2408,f2429,f2444,f2427,f2432,f2433,f2434,f2436,f2456,f2460,f2461,f2462,f2464,f2489,f2488,f2486,f2485,f2497,f2506,f2507,f2508,f2509,f2511,f2512,f2513,f2514,f2515,f2516,f2517,f2518,f2519,f2520,f2521,f2522,f2523,f2524,f2525,f2526,f2528,f2529,f2530,f2531,f2532,f2533,f2831,f2830,f2827,f2824,f2821,f2820,f2819,f2818,f2817,f2816,f2815,f2814,f2813,f2812,f2811,f2810,f2809,f2806,f2803,f2802,f2801,f2797,f2796,f2790,f2787,f2786,f2780,f2779,f2772,f2746,f2745,f2744,f2739,f2738,f2733,f2732,f2731,f2726,f2718,f2714,f2710,f2709,f2704,f2703,f2702,f2701,f2700,f2696,f2694,f2688,f2687,f2681,f2680,f2673,f2670,f2669,f2663,f2660,f2659,f2658,f2656,f2654,f2653,f2651,f2650,f2649,f2648,f2646,f2645,f2644,f2643,f2642,f2641,f2635,f2634,f2627,f2626,f2620,f2619,f2612,f2611,f2609,f2607,f2605,f2602,f2601,f2595,f2594,f2589,f2574,f2573,f2572,f2571,f2569,f2568,f2567,f2566,f2565,f3267,f3266,f3263,f3260,f3257,f3256,f3255,f3254,f3253,f3252,f3251,f3250,f3249,f3248,f3247,f3246,f3245,f3242,f3239,f3238,f3237,f3233,f3232,f3226,f3223,f3222,f3216,f3215,f3208,f3182,f3181,f3180,f3175,f3174,f3169,f3168,f3167,f3162,f3154,f3150,f3146,f3145,f3140,f3139,f3138,f3137,f3136,f3132,f3130,f3124,f3123,f3117,f3116,f3109,f3106,f3105,f3099,f3096,f3095,f3093,f3090,f3088,f3087,f3084,f3083,f3082,f3081,f3078,f3077,f3076,f3075,f3074,f3073,f3067,f3066,f3059,f3058,f3052,f3051,f3044,f3043,f3041,f3039,f3037,f3034,f3033,f3027,f3026,f3021,f3006,f3005,f3004,f3003,f3001,f3000,f2999,f2998,f2997,f3620,f3619,f3616,f3613,f3610,f3609,f3608,f3607,f3606,f3605,f3604,f3603,f3602,f3601,f3600,f3599,f3598,f3595,f3592,f3591,f3590,f3586,f3585,f3579,f3576,f3575,f3569,f3568,f3561,f3535,f3534,f3533,f3528,f3527,f3522,f3521,f3520,f3515,f3507,f3503,f3499,f3498,f3493,f3492,f3491,f3490,f3489,f3485,f3483,f3477,f3476,f3470,f3469,f3462,f3459,f3458,f3452,f3449,f3448,f3446,f3443,f3441,f3440,f3437,f3436,f3435,f3434,f3431,f3430,f3429,f3428,f3427,f3426,f3420,f3419,f3412,f3411,f3405,f3404,f3397,f3396,f3395,f3394,f3393,f3391,f3390,f3384,f3383,f3378,f3363,f3362,f3361,f3360,f3358,f3357,f3356,f3354,f3708,f3726,f3770,f3771,f3773,f3788,f3799,f3800,f3802,f3851,f3852,f3853,f3854,f3905,f3907,f3908,f3909,f3910,f3911,f3912,f3913,f3914,f3915,f3916,f3917,f3918,f3919,f3920,f3921,f3922,f3923,f3924,f3925,f3926,f3927,f4039,f4569,f4573,f4576,f4579,f4583,f4586,f4589,f4592,f4595,f4598,f4601,f4604,f4607,f4610,f4613,f4616,f4619,f4622,f4625,f4072,f4073,f4074,f4075,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4646,f4647,f4648,f4649,f4650,f4124,f4125,f4126,f4127,f4651,f4652,f4653,f4654,f4655,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4667,f4668,f4669,f4670,f4671,f4177,f4179,f4180,f4181,f4182,f4183,f4184,f4185,f4186,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4677,f4679,f4681,f4683,f4686,f4688,f4690,f4692,f4694,f4696,f4698,f4700,f4702,f4704,f4706,f4708,f4710,f4712,f4714,f4571,f4941,f4744,f4751,f4966,f4972,f4982,f4986,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029])).
fof(f5028,plain,(
  $false | $spl14),
  inference(global_subsumption,[],[f321,f320,f319,f318,f322,f573,f329,f337,f336,f341,f340,f364,f371,f370,f369,f368,f367,f379,f377,f382,f380,f387,f384,f394,f393,f574,f412,f409,f408,f407,f421,f423,f425,f424,f428,f427,f426,f431,f430,f429,f434,f433,f432,f437,f436,f435,f439,f438,f451,f450,f449,f559,f447,f446,f468,f467,f466,f465,f464,f463,f462,f472,f471,f470,f469,f476,f475,f474,f473,f480,f479,f478,f477,f484,f483,f482,f481,f558,f488,f487,f486,f485,f495,f494,f493,f492,f501,f498,f497,f496,f508,f507,f506,f557,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f525,f524,f523,f556,f521,f520,f519,f529,f528,f527,f526,f533,f532,f531,f530,f555,f575,f576,f577,f578,f580,f582,f583,f585,f587,f588,f589,f590,f591,f593,f595,f596,f597,f599,f601,f602,f604,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f617,f621,f639,f640,f643,f649,f657,f664,f681,f682,f685,f688,f695,f704,f707,f708,f711,f712,f716,f725,f729,f730,f733,f736,f740,f743,f746,f755,f756,f757,f765,f768,f772,f775,f778,f797,f800,f815,f813,f818,f836,f884,f902,f925,f923,f929,f927,f941,f938,f951,f949,f992,f1003,f1011,f1024,f1032,f1036,f1042,f1051,f1056,f1055,f1068,f1076,f1080,f1113,f1124,f1131,f1140,f1152,f1160,f1164,f1178,f1196,f1197,f1212,f1222,f1237,f1249,f1257,f1261,f1272,f1282,f1298,f1314,f1313,f1311,f1310,f1325,f1334,f1335,f1336,f1337,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1347,f1348,f1349,f1350,f1351,f1352,f1353,f1355,f1356,f1371,f1381,f1397,f1389,f1387,f1386,f1405,f1402,f1439,f1429,f1428,f1484,f1503,f1515,f1514,f1513,f1512,f1517,f1518,f1519,f1529,f1530,f1531,f1532,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1543,f1544,f1545,f1546,f1547,f1548,f1549,f1551,f1552,f1866,f1863,f1867,f1865,f1876,f1873,f1877,f1875,f1887,f1884,f1886,f1900,f1897,f1899,f1916,f1912,f1934,f1930,f1952,f1948,f1977,f1973,f2000,f2015,f2016,f2027,f2038,f2035,f2036,f2037,f2051,f2062,f2059,f2060,f2061,f2082,f2104,f2105,f2106,f2111,f2134,f2135,f2136,f2141,f2142,f2186,f2182,f2209,f2205,f2233,f2234,f2228,f2229,f2230,f2232,f2249,f2250,f2244,f2245,f2246,f2251,f2248,f2286,f2287,f2305,f2306,f2334,f2336,f2347,f2349,f2367,f2364,f2386,f2370,f2371,f2372,f2374,f2375,f2400,f2403,f2404,f2405,f2407,f2408,f2429,f2444,f2427,f2432,f2433,f2434,f2436,f2456,f2460,f2461,f2462,f2464,f2489,f2488,f2486,f2485,f2497,f2506,f2507,f2508,f2509,f2511,f2512,f2513,f2514,f2515,f2516,f2517,f2518,f2519,f2520,f2521,f2522,f2523,f2524,f2525,f2526,f2528,f2529,f2530,f2531,f2532,f2533,f2831,f2830,f2827,f2824,f2821,f2820,f2819,f2818,f2817,f2816,f2815,f2814,f2813,f2812,f2811,f2810,f2809,f2806,f2803,f2802,f2801,f2797,f2796,f2790,f2787,f2786,f2780,f2779,f2772,f2746,f2745,f2744,f2739,f2738,f2733,f2732,f2731,f2726,f2718,f2714,f2710,f2709,f2704,f2703,f2702,f2701,f2700,f2696,f2694,f2688,f2687,f2681,f2680,f2673,f2670,f2669,f2663,f2660,f2659,f2658,f2656,f2654,f2653,f2651,f2650,f2649,f2648,f2646,f2645,f2644,f2643,f2642,f2641,f2635,f2634,f2627,f2626,f2620,f2619,f2612,f2611,f2609,f2607,f2605,f2602,f2601,f2595,f2594,f2589,f2574,f2573,f2572,f2571,f2569,f2568,f2567,f2566,f2565,f3267,f3266,f3263,f3260,f3257,f3256,f3255,f3254,f3253,f3252,f3251,f3250,f3249,f3248,f3247,f3246,f3245,f3242,f3239,f3238,f3237,f3233,f3232,f3226,f3223,f3222,f3216,f3215,f3208,f3182,f3181,f3180,f3175,f3174,f3169,f3168,f3167,f3162,f3154,f3150,f3146,f3145,f3140,f3139,f3138,f3137,f3136,f3132,f3130,f3124,f3123,f3117,f3116,f3109,f3106,f3105,f3099,f3096,f3095,f3093,f3090,f3088,f3087,f3084,f3083,f3082,f3081,f3078,f3077,f3076,f3075,f3074,f3073,f3067,f3066,f3059,f3058,f3052,f3051,f3044,f3043,f3041,f3039,f3037,f3034,f3033,f3027,f3026,f3021,f3006,f3005,f3004,f3003,f3001,f3000,f2999,f2998,f2997,f3620,f3619,f3616,f3613,f3610,f3609,f3608,f3607,f3606,f3605,f3604,f3603,f3602,f3601,f3600,f3599,f3598,f3595,f3592,f3591,f3590,f3586,f3585,f3579,f3576,f3575,f3569,f3568,f3561,f3535,f3534,f3533,f3528,f3527,f3522,f3521,f3520,f3515,f3507,f3503,f3499,f3498,f3493,f3492,f3491,f3490,f3489,f3485,f3483,f3477,f3476,f3470,f3469,f3462,f3459,f3458,f3452,f3449,f3448,f3446,f3443,f3441,f3440,f3437,f3436,f3435,f3434,f3431,f3430,f3429,f3428,f3427,f3426,f3420,f3419,f3412,f3411,f3405,f3404,f3397,f3396,f3395,f3394,f3393,f3391,f3390,f3384,f3383,f3378,f3363,f3362,f3361,f3360,f3358,f3357,f3356,f3354,f3708,f3726,f3770,f3771,f3773,f3788,f3799,f3800,f3802,f3851,f3852,f3853,f3854,f3905,f3907,f3908,f3909,f3910,f3911,f3912,f3913,f3914,f3915,f3916,f3917,f3918,f3919,f3920,f3921,f3922,f3923,f3924,f3925,f3926,f3927,f4039,f4569,f4573,f4576,f4579,f4583,f4586,f4589,f4592,f4595,f4598,f4601,f4604,f4607,f4610,f4613,f4616,f4619,f4622,f4625,f4072,f4073,f4074,f4075,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4646,f4647,f4648,f4649,f4650,f4124,f4125,f4126,f4127,f4651,f4652,f4653,f4654,f4655,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4667,f4668,f4669,f4670,f4671,f4177,f4179,f4180,f4181,f4182,f4183,f4184,f4185,f4186,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4677,f4679,f4681,f4683,f4686,f4688,f4690,f4692,f4694,f4696,f4698,f4700,f4702,f4704,f4706,f4708,f4710,f4712,f4714,f4571,f4941,f4744,f4751,f4966,f4972,f4982,f4986,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027])).
fof(f5026,plain,(
  $false | $spl14),
  inference(global_subsumption,[],[f321,f320,f319,f318,f322,f573,f329,f337,f336,f341,f340,f364,f371,f370,f369,f368,f367,f379,f377,f382,f380,f387,f384,f394,f393,f574,f412,f409,f408,f407,f421,f423,f425,f424,f428,f427,f426,f431,f430,f429,f434,f433,f432,f437,f436,f435,f439,f438,f451,f450,f449,f559,f447,f446,f468,f467,f466,f465,f464,f463,f462,f472,f471,f470,f469,f476,f475,f474,f473,f480,f479,f478,f477,f484,f483,f482,f481,f558,f488,f487,f486,f485,f495,f494,f493,f492,f501,f498,f497,f496,f508,f507,f506,f557,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f525,f524,f523,f556,f521,f520,f519,f529,f528,f527,f526,f533,f532,f531,f530,f555,f575,f576,f577,f578,f580,f582,f583,f585,f587,f588,f589,f590,f591,f593,f595,f596,f597,f599,f601,f602,f604,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f617,f621,f639,f640,f643,f649,f657,f664,f681,f682,f685,f688,f695,f704,f707,f708,f711,f712,f716,f725,f729,f730,f733,f736,f740,f743,f746,f755,f756,f757,f765,f768,f772,f775,f778,f797,f800,f815,f813,f818,f836,f884,f902,f925,f923,f929,f927,f941,f938,f951,f949,f992,f1003,f1011,f1024,f1032,f1036,f1042,f1051,f1056,f1055,f1068,f1076,f1080,f1113,f1124,f1131,f1140,f1152,f1160,f1164,f1178,f1196,f1197,f1212,f1222,f1237,f1249,f1257,f1261,f1272,f1282,f1298,f1314,f1313,f1311,f1310,f1325,f1334,f1335,f1336,f1337,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1347,f1348,f1349,f1350,f1351,f1352,f1353,f1355,f1356,f1371,f1381,f1397,f1389,f1387,f1386,f1405,f1402,f1439,f1429,f1428,f1484,f1503,f1515,f1514,f1513,f1512,f1517,f1518,f1519,f1529,f1530,f1531,f1532,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1543,f1544,f1545,f1546,f1547,f1548,f1549,f1551,f1552,f1866,f1863,f1867,f1865,f1876,f1873,f1877,f1875,f1887,f1884,f1886,f1900,f1897,f1899,f1916,f1912,f1934,f1930,f1952,f1948,f1977,f1973,f2000,f2015,f2016,f2027,f2038,f2035,f2036,f2037,f2051,f2062,f2059,f2060,f2061,f2082,f2104,f2105,f2106,f2111,f2134,f2135,f2136,f2141,f2142,f2186,f2182,f2209,f2205,f2233,f2234,f2228,f2229,f2230,f2232,f2249,f2250,f2244,f2245,f2246,f2251,f2248,f2286,f2287,f2305,f2306,f2334,f2336,f2347,f2349,f2367,f2364,f2386,f2370,f2371,f2372,f2374,f2375,f2400,f2403,f2404,f2405,f2407,f2408,f2429,f2444,f2427,f2432,f2433,f2434,f2436,f2456,f2460,f2461,f2462,f2464,f2489,f2488,f2486,f2485,f2497,f2506,f2507,f2508,f2509,f2511,f2512,f2513,f2514,f2515,f2516,f2517,f2518,f2519,f2520,f2521,f2522,f2523,f2524,f2525,f2526,f2528,f2529,f2530,f2531,f2532,f2533,f2831,f2830,f2827,f2824,f2821,f2820,f2819,f2818,f2817,f2816,f2815,f2814,f2813,f2812,f2811,f2810,f2809,f2806,f2803,f2802,f2801,f2797,f2796,f2790,f2787,f2786,f2780,f2779,f2772,f2746,f2745,f2744,f2739,f2738,f2733,f2732,f2731,f2726,f2718,f2714,f2710,f2709,f2704,f2703,f2702,f2701,f2700,f2696,f2694,f2688,f2687,f2681,f2680,f2673,f2670,f2669,f2663,f2660,f2659,f2658,f2656,f2654,f2653,f2651,f2650,f2649,f2648,f2646,f2645,f2644,f2643,f2642,f2641,f2635,f2634,f2627,f2626,f2620,f2619,f2612,f2611,f2609,f2607,f2605,f2602,f2601,f2595,f2594,f2589,f2574,f2573,f2572,f2571,f2569,f2568,f2567,f2566,f2565,f3267,f3266,f3263,f3260,f3257,f3256,f3255,f3254,f3253,f3252,f3251,f3250,f3249,f3248,f3247,f3246,f3245,f3242,f3239,f3238,f3237,f3233,f3232,f3226,f3223,f3222,f3216,f3215,f3208,f3182,f3181,f3180,f3175,f3174,f3169,f3168,f3167,f3162,f3154,f3150,f3146,f3145,f3140,f3139,f3138,f3137,f3136,f3132,f3130,f3124,f3123,f3117,f3116,f3109,f3106,f3105,f3099,f3096,f3095,f3093,f3090,f3088,f3087,f3084,f3083,f3082,f3081,f3078,f3077,f3076,f3075,f3074,f3073,f3067,f3066,f3059,f3058,f3052,f3051,f3044,f3043,f3041,f3039,f3037,f3034,f3033,f3027,f3026,f3021,f3006,f3005,f3004,f3003,f3001,f3000,f2999,f2998,f2997,f3620,f3619,f3616,f3613,f3610,f3609,f3608,f3607,f3606,f3605,f3604,f3603,f3602,f3601,f3600,f3599,f3598,f3595,f3592,f3591,f3590,f3586,f3585,f3579,f3576,f3575,f3569,f3568,f3561,f3535,f3534,f3533,f3528,f3527,f3522,f3521,f3520,f3515,f3507,f3503,f3499,f3498,f3493,f3492,f3491,f3490,f3489,f3485,f3483,f3477,f3476,f3470,f3469,f3462,f3459,f3458,f3452,f3449,f3448,f3446,f3443,f3441,f3440,f3437,f3436,f3435,f3434,f3431,f3430,f3429,f3428,f3427,f3426,f3420,f3419,f3412,f3411,f3405,f3404,f3397,f3396,f3395,f3394,f3393,f3391,f3390,f3384,f3383,f3378,f3363,f3362,f3361,f3360,f3358,f3357,f3356,f3354,f3708,f3726,f3770,f3771,f3773,f3788,f3799,f3800,f3802,f3851,f3852,f3853,f3854,f3905,f3907,f3908,f3909,f3910,f3911,f3912,f3913,f3914,f3915,f3916,f3917,f3918,f3919,f3920,f3921,f3922,f3923,f3924,f3925,f3926,f3927,f4039,f4569,f4573,f4576,f4579,f4583,f4586,f4589,f4592,f4595,f4598,f4601,f4604,f4607,f4610,f4613,f4616,f4619,f4622,f4625,f4072,f4073,f4074,f4075,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4646,f4647,f4648,f4649,f4650,f4124,f4125,f4126,f4127,f4651,f4652,f4653,f4654,f4655,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4667,f4668,f4669,f4670,f4671,f4177,f4179,f4180,f4181,f4182,f4183,f4184,f4185,f4186,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4677,f4679,f4681,f4683,f4686,f4688,f4690,f4692,f4694,f4696,f4698,f4700,f4702,f4704,f4706,f4708,f4710,f4712,f4714,f4571,f4941,f4744,f4751,f4966,f4972,f4982,f4986,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025])).
fof(f5024,plain,(
  $false | $spl14),
  inference(global_subsumption,[],[f321,f320,f319,f318,f322,f573,f329,f337,f336,f341,f340,f364,f371,f370,f369,f368,f367,f379,f377,f382,f380,f387,f384,f394,f393,f574,f412,f409,f408,f407,f421,f423,f425,f424,f428,f427,f426,f431,f430,f429,f434,f433,f432,f437,f436,f435,f439,f438,f451,f450,f449,f559,f447,f446,f468,f467,f466,f465,f464,f463,f462,f472,f471,f470,f469,f476,f475,f474,f473,f480,f479,f478,f477,f484,f483,f482,f481,f558,f488,f487,f486,f485,f495,f494,f493,f492,f501,f498,f497,f496,f508,f507,f506,f557,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f525,f524,f523,f556,f521,f520,f519,f529,f528,f527,f526,f533,f532,f531,f530,f555,f575,f576,f577,f578,f580,f582,f583,f585,f587,f588,f589,f590,f591,f593,f595,f596,f597,f599,f601,f602,f604,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f617,f621,f639,f640,f643,f649,f657,f664,f681,f682,f685,f688,f695,f704,f707,f708,f711,f712,f716,f725,f729,f730,f733,f736,f740,f743,f746,f755,f756,f757,f765,f768,f772,f775,f778,f797,f800,f815,f813,f818,f836,f884,f902,f925,f923,f929,f927,f941,f938,f951,f949,f992,f1003,f1011,f1024,f1032,f1036,f1042,f1051,f1056,f1055,f1068,f1076,f1080,f1113,f1124,f1131,f1140,f1152,f1160,f1164,f1178,f1196,f1197,f1212,f1222,f1237,f1249,f1257,f1261,f1272,f1282,f1298,f1314,f1313,f1311,f1310,f1325,f1334,f1335,f1336,f1337,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1347,f1348,f1349,f1350,f1351,f1352,f1353,f1355,f1356,f1371,f1381,f1397,f1389,f1387,f1386,f1405,f1402,f1439,f1429,f1428,f1484,f1503,f1515,f1514,f1513,f1512,f1517,f1518,f1519,f1529,f1530,f1531,f1532,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1543,f1544,f1545,f1546,f1547,f1548,f1549,f1551,f1552,f1866,f1863,f1867,f1865,f1876,f1873,f1877,f1875,f1887,f1884,f1886,f1900,f1897,f1899,f1916,f1912,f1934,f1930,f1952,f1948,f1977,f1973,f2000,f2015,f2016,f2027,f2038,f2035,f2036,f2037,f2051,f2062,f2059,f2060,f2061,f2082,f2104,f2105,f2106,f2111,f2134,f2135,f2136,f2141,f2142,f2186,f2182,f2209,f2205,f2233,f2234,f2228,f2229,f2230,f2232,f2249,f2250,f2244,f2245,f2246,f2251,f2248,f2286,f2287,f2305,f2306,f2334,f2336,f2347,f2349,f2367,f2364,f2386,f2370,f2371,f2372,f2374,f2375,f2400,f2403,f2404,f2405,f2407,f2408,f2429,f2444,f2427,f2432,f2433,f2434,f2436,f2456,f2460,f2461,f2462,f2464,f2489,f2488,f2486,f2485,f2497,f2506,f2507,f2508,f2509,f2511,f2512,f2513,f2514,f2515,f2516,f2517,f2518,f2519,f2520,f2521,f2522,f2523,f2524,f2525,f2526,f2528,f2529,f2530,f2531,f2532,f2533,f2831,f2830,f2827,f2824,f2821,f2820,f2819,f2818,f2817,f2816,f2815,f2814,f2813,f2812,f2811,f2810,f2809,f2806,f2803,f2802,f2801,f2797,f2796,f2790,f2787,f2786,f2780,f2779,f2772,f2746,f2745,f2744,f2739,f2738,f2733,f2732,f2731,f2726,f2718,f2714,f2710,f2709,f2704,f2703,f2702,f2701,f2700,f2696,f2694,f2688,f2687,f2681,f2680,f2673,f2670,f2669,f2663,f2660,f2659,f2658,f2656,f2654,f2653,f2651,f2650,f2649,f2648,f2646,f2645,f2644,f2643,f2642,f2641,f2635,f2634,f2627,f2626,f2620,f2619,f2612,f2611,f2609,f2607,f2605,f2602,f2601,f2595,f2594,f2589,f2574,f2573,f2572,f2571,f2569,f2568,f2567,f2566,f2565,f3267,f3266,f3263,f3260,f3257,f3256,f3255,f3254,f3253,f3252,f3251,f3250,f3249,f3248,f3247,f3246,f3245,f3242,f3239,f3238,f3237,f3233,f3232,f3226,f3223,f3222,f3216,f3215,f3208,f3182,f3181,f3180,f3175,f3174,f3169,f3168,f3167,f3162,f3154,f3150,f3146,f3145,f3140,f3139,f3138,f3137,f3136,f3132,f3130,f3124,f3123,f3117,f3116,f3109,f3106,f3105,f3099,f3096,f3095,f3093,f3090,f3088,f3087,f3084,f3083,f3082,f3081,f3078,f3077,f3076,f3075,f3074,f3073,f3067,f3066,f3059,f3058,f3052,f3051,f3044,f3043,f3041,f3039,f3037,f3034,f3033,f3027,f3026,f3021,f3006,f3005,f3004,f3003,f3001,f3000,f2999,f2998,f2997,f3620,f3619,f3616,f3613,f3610,f3609,f3608,f3607,f3606,f3605,f3604,f3603,f3602,f3601,f3600,f3599,f3598,f3595,f3592,f3591,f3590,f3586,f3585,f3579,f3576,f3575,f3569,f3568,f3561,f3535,f3534,f3533,f3528,f3527,f3522,f3521,f3520,f3515,f3507,f3503,f3499,f3498,f3493,f3492,f3491,f3490,f3489,f3485,f3483,f3477,f3476,f3470,f3469,f3462,f3459,f3458,f3452,f3449,f3448,f3446,f3443,f3441,f3440,f3437,f3436,f3435,f3434,f3431,f3430,f3429,f3428,f3427,f3426,f3420,f3419,f3412,f3411,f3405,f3404,f3397,f3396,f3395,f3394,f3393,f3391,f3390,f3384,f3383,f3378,f3363,f3362,f3361,f3360,f3358,f3357,f3356,f3354,f3708,f3726,f3770,f3771,f3773,f3788,f3799,f3800,f3802,f3851,f3852,f3853,f3854,f3905,f3907,f3908,f3909,f3910,f3911,f3912,f3913,f3914,f3915,f3916,f3917,f3918,f3919,f3920,f3921,f3922,f3923,f3924,f3925,f3926,f3927,f4039,f4569,f4573,f4576,f4579,f4583,f4586,f4589,f4592,f4595,f4598,f4601,f4604,f4607,f4610,f4613,f4616,f4619,f4622,f4625,f4072,f4073,f4074,f4075,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4646,f4647,f4648,f4649,f4650,f4124,f4125,f4126,f4127,f4651,f4652,f4653,f4654,f4655,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4667,f4668,f4669,f4670,f4671,f4177,f4179,f4180,f4181,f4182,f4183,f4184,f4185,f4186,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4677,f4679,f4681,f4683,f4686,f4688,f4690,f4692,f4694,f4696,f4698,f4700,f4702,f4704,f4706,f4708,f4710,f4712,f4714,f4571,f4941,f4744,f4751,f4966,f4972,f4982,f4986,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023])).
fof(f5022,plain,(
  $false | $spl14),
  inference(global_subsumption,[],[f321,f320,f319,f318,f322,f573,f329,f337,f336,f341,f340,f364,f371,f370,f369,f368,f367,f379,f377,f382,f380,f387,f384,f394,f393,f574,f412,f409,f408,f407,f421,f423,f425,f424,f428,f427,f426,f431,f430,f429,f434,f433,f432,f437,f436,f435,f439,f438,f451,f450,f449,f559,f447,f446,f468,f467,f466,f465,f464,f463,f462,f472,f471,f470,f469,f476,f475,f474,f473,f480,f479,f478,f477,f484,f483,f482,f481,f558,f488,f487,f486,f485,f495,f494,f493,f492,f501,f498,f497,f496,f508,f507,f506,f557,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f525,f524,f523,f556,f521,f520,f519,f529,f528,f527,f526,f533,f532,f531,f530,f555,f575,f576,f577,f578,f580,f582,f583,f585,f587,f588,f589,f590,f591,f593,f595,f596,f597,f599,f601,f602,f604,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f617,f621,f639,f640,f643,f649,f657,f664,f681,f682,f685,f688,f695,f704,f707,f708,f711,f712,f716,f725,f729,f730,f733,f736,f740,f743,f746,f755,f756,f757,f765,f768,f772,f775,f778,f797,f800,f815,f813,f818,f836,f884,f902,f925,f923,f929,f927,f941,f938,f951,f949,f992,f1003,f1011,f1024,f1032,f1036,f1042,f1051,f1056,f1055,f1068,f1076,f1080,f1113,f1124,f1131,f1140,f1152,f1160,f1164,f1178,f1196,f1197,f1212,f1222,f1237,f1249,f1257,f1261,f1272,f1282,f1298,f1314,f1313,f1311,f1310,f1325,f1334,f1335,f1336,f1337,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1347,f1348,f1349,f1350,f1351,f1352,f1353,f1355,f1356,f1371,f1381,f1397,f1389,f1387,f1386,f1405,f1402,f1439,f1429,f1428,f1484,f1503,f1515,f1514,f1513,f1512,f1517,f1518,f1519,f1529,f1530,f1531,f1532,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1543,f1544,f1545,f1546,f1547,f1548,f1549,f1551,f1552,f1866,f1863,f1867,f1865,f1876,f1873,f1877,f1875,f1887,f1884,f1886,f1900,f1897,f1899,f1916,f1912,f1934,f1930,f1952,f1948,f1977,f1973,f2000,f2015,f2016,f2027,f2038,f2035,f2036,f2037,f2051,f2062,f2059,f2060,f2061,f2082,f2104,f2105,f2106,f2111,f2134,f2135,f2136,f2141,f2142,f2186,f2182,f2209,f2205,f2233,f2234,f2228,f2229,f2230,f2232,f2249,f2250,f2244,f2245,f2246,f2251,f2248,f2286,f2287,f2305,f2306,f2334,f2336,f2347,f2349,f2367,f2364,f2386,f2370,f2371,f2372,f2374,f2375,f2400,f2403,f2404,f2405,f2407,f2408,f2429,f2444,f2427,f2432,f2433,f2434,f2436,f2456,f2460,f2461,f2462,f2464,f2489,f2488,f2486,f2485,f2497,f2506,f2507,f2508,f2509,f2511,f2512,f2513,f2514,f2515,f2516,f2517,f2518,f2519,f2520,f2521,f2522,f2523,f2524,f2525,f2526,f2528,f2529,f2530,f2531,f2532,f2533,f2831,f2830,f2827,f2824,f2821,f2820,f2819,f2818,f2817,f2816,f2815,f2814,f2813,f2812,f2811,f2810,f2809,f2806,f2803,f2802,f2801,f2797,f2796,f2790,f2787,f2786,f2780,f2779,f2772,f2746,f2745,f2744,f2739,f2738,f2733,f2732,f2731,f2726,f2718,f2714,f2710,f2709,f2704,f2703,f2702,f2701,f2700,f2696,f2694,f2688,f2687,f2681,f2680,f2673,f2670,f2669,f2663,f2660,f2659,f2658,f2656,f2654,f2653,f2651,f2650,f2649,f2648,f2646,f2645,f2644,f2643,f2642,f2641,f2635,f2634,f2627,f2626,f2620,f2619,f2612,f2611,f2609,f2607,f2605,f2602,f2601,f2595,f2594,f2589,f2574,f2573,f2572,f2571,f2569,f2568,f2567,f2566,f2565,f3267,f3266,f3263,f3260,f3257,f3256,f3255,f3254,f3253,f3252,f3251,f3250,f3249,f3248,f3247,f3246,f3245,f3242,f3239,f3238,f3237,f3233,f3232,f3226,f3223,f3222,f3216,f3215,f3208,f3182,f3181,f3180,f3175,f3174,f3169,f3168,f3167,f3162,f3154,f3150,f3146,f3145,f3140,f3139,f3138,f3137,f3136,f3132,f3130,f3124,f3123,f3117,f3116,f3109,f3106,f3105,f3099,f3096,f3095,f3093,f3090,f3088,f3087,f3084,f3083,f3082,f3081,f3078,f3077,f3076,f3075,f3074,f3073,f3067,f3066,f3059,f3058,f3052,f3051,f3044,f3043,f3041,f3039,f3037,f3034,f3033,f3027,f3026,f3021,f3006,f3005,f3004,f3003,f3001,f3000,f2999,f2998,f2997,f3620,f3619,f3616,f3613,f3610,f3609,f3608,f3607,f3606,f3605,f3604,f3603,f3602,f3601,f3600,f3599,f3598,f3595,f3592,f3591,f3590,f3586,f3585,f3579,f3576,f3575,f3569,f3568,f3561,f3535,f3534,f3533,f3528,f3527,f3522,f3521,f3520,f3515,f3507,f3503,f3499,f3498,f3493,f3492,f3491,f3490,f3489,f3485,f3483,f3477,f3476,f3470,f3469,f3462,f3459,f3458,f3452,f3449,f3448,f3446,f3443,f3441,f3440,f3437,f3436,f3435,f3434,f3431,f3430,f3429,f3428,f3427,f3426,f3420,f3419,f3412,f3411,f3405,f3404,f3397,f3396,f3395,f3394,f3393,f3391,f3390,f3384,f3383,f3378,f3363,f3362,f3361,f3360,f3358,f3357,f3356,f3354,f3708,f3726,f3770,f3771,f3773,f3788,f3799,f3800,f3802,f3851,f3852,f3853,f3854,f3905,f3907,f3908,f3909,f3910,f3911,f3912,f3913,f3914,f3915,f3916,f3917,f3918,f3919,f3920,f3921,f3922,f3923,f3924,f3925,f3926,f3927,f4039,f4569,f4573,f4576,f4579,f4583,f4586,f4589,f4592,f4595,f4598,f4601,f4604,f4607,f4610,f4613,f4616,f4619,f4622,f4625,f4072,f4073,f4074,f4075,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4646,f4647,f4648,f4649,f4650,f4124,f4125,f4126,f4127,f4651,f4652,f4653,f4654,f4655,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4667,f4668,f4669,f4670,f4671,f4177,f4179,f4180,f4181,f4182,f4183,f4184,f4185,f4186,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4677,f4679,f4681,f4683,f4686,f4688,f4690,f4692,f4694,f4696,f4698,f4700,f4702,f4704,f4706,f4708,f4710,f4712,f4714,f4571,f4941,f4744,f4751,f4966,f4972,f4982,f4986,f5009,f5011,f5013,f5015,f5017,f5019,f5021])).
fof(f5020,plain,(
  $false | $spl14),
  inference(global_subsumption,[],[f321,f320,f319,f318,f322,f573,f329,f337,f336,f341,f340,f364,f371,f370,f369,f368,f367,f379,f377,f382,f380,f387,f384,f394,f393,f574,f412,f409,f408,f407,f421,f423,f425,f424,f428,f427,f426,f431,f430,f429,f434,f433,f432,f437,f436,f435,f439,f438,f451,f450,f449,f559,f447,f446,f468,f467,f466,f465,f464,f463,f462,f472,f471,f470,f469,f476,f475,f474,f473,f480,f479,f478,f477,f484,f483,f482,f481,f558,f488,f487,f486,f485,f495,f494,f493,f492,f501,f498,f497,f496,f508,f507,f506,f557,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f525,f524,f523,f556,f521,f520,f519,f529,f528,f527,f526,f533,f532,f531,f530,f555,f575,f576,f577,f578,f580,f582,f583,f585,f587,f588,f589,f590,f591,f593,f595,f596,f597,f599,f601,f602,f604,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f617,f621,f639,f640,f643,f649,f657,f664,f681,f682,f685,f688,f695,f704,f707,f708,f711,f712,f716,f725,f729,f730,f733,f736,f740,f743,f746,f755,f756,f757,f765,f768,f772,f775,f778,f797,f800,f815,f813,f818,f836,f884,f902,f925,f923,f929,f927,f941,f938,f951,f949,f992,f1003,f1011,f1024,f1032,f1036,f1042,f1051,f1056,f1055,f1068,f1076,f1080,f1113,f1124,f1131,f1140,f1152,f1160,f1164,f1178,f1196,f1197,f1212,f1222,f1237,f1249,f1257,f1261,f1272,f1282,f1298,f1314,f1313,f1311,f1310,f1325,f1334,f1335,f1336,f1337,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1347,f1348,f1349,f1350,f1351,f1352,f1353,f1355,f1356,f1371,f1381,f1397,f1389,f1387,f1386,f1405,f1402,f1439,f1429,f1428,f1484,f1503,f1515,f1514,f1513,f1512,f1517,f1518,f1519,f1529,f1530,f1531,f1532,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1543,f1544,f1545,f1546,f1547,f1548,f1549,f1551,f1552,f1866,f1863,f1867,f1865,f1876,f1873,f1877,f1875,f1887,f1884,f1886,f1900,f1897,f1899,f1916,f1912,f1934,f1930,f1952,f1948,f1977,f1973,f2000,f2015,f2016,f2027,f2038,f2035,f2036,f2037,f2051,f2062,f2059,f2060,f2061,f2082,f2104,f2105,f2106,f2111,f2134,f2135,f2136,f2141,f2142,f2186,f2182,f2209,f2205,f2233,f2234,f2228,f2229,f2230,f2232,f2249,f2250,f2244,f2245,f2246,f2251,f2248,f2286,f2287,f2305,f2306,f2334,f2336,f2347,f2349,f2367,f2364,f2386,f2370,f2371,f2372,f2374,f2375,f2400,f2403,f2404,f2405,f2407,f2408,f2429,f2444,f2427,f2432,f2433,f2434,f2436,f2456,f2460,f2461,f2462,f2464,f2489,f2488,f2486,f2485,f2497,f2506,f2507,f2508,f2509,f2511,f2512,f2513,f2514,f2515,f2516,f2517,f2518,f2519,f2520,f2521,f2522,f2523,f2524,f2525,f2526,f2528,f2529,f2530,f2531,f2532,f2533,f2831,f2830,f2827,f2824,f2821,f2820,f2819,f2818,f2817,f2816,f2815,f2814,f2813,f2812,f2811,f2810,f2809,f2806,f2803,f2802,f2801,f2797,f2796,f2790,f2787,f2786,f2780,f2779,f2772,f2746,f2745,f2744,f2739,f2738,f2733,f2732,f2731,f2726,f2718,f2714,f2710,f2709,f2704,f2703,f2702,f2701,f2700,f2696,f2694,f2688,f2687,f2681,f2680,f2673,f2670,f2669,f2663,f2660,f2659,f2658,f2656,f2654,f2653,f2651,f2650,f2649,f2648,f2646,f2645,f2644,f2643,f2642,f2641,f2635,f2634,f2627,f2626,f2620,f2619,f2612,f2611,f2609,f2607,f2605,f2602,f2601,f2595,f2594,f2589,f2574,f2573,f2572,f2571,f2569,f2568,f2567,f2566,f2565,f3267,f3266,f3263,f3260,f3257,f3256,f3255,f3254,f3253,f3252,f3251,f3250,f3249,f3248,f3247,f3246,f3245,f3242,f3239,f3238,f3237,f3233,f3232,f3226,f3223,f3222,f3216,f3215,f3208,f3182,f3181,f3180,f3175,f3174,f3169,f3168,f3167,f3162,f3154,f3150,f3146,f3145,f3140,f3139,f3138,f3137,f3136,f3132,f3130,f3124,f3123,f3117,f3116,f3109,f3106,f3105,f3099,f3096,f3095,f3093,f3090,f3088,f3087,f3084,f3083,f3082,f3081,f3078,f3077,f3076,f3075,f3074,f3073,f3067,f3066,f3059,f3058,f3052,f3051,f3044,f3043,f3041,f3039,f3037,f3034,f3033,f3027,f3026,f3021,f3006,f3005,f3004,f3003,f3001,f3000,f2999,f2998,f2997,f3620,f3619,f3616,f3613,f3610,f3609,f3608,f3607,f3606,f3605,f3604,f3603,f3602,f3601,f3600,f3599,f3598,f3595,f3592,f3591,f3590,f3586,f3585,f3579,f3576,f3575,f3569,f3568,f3561,f3535,f3534,f3533,f3528,f3527,f3522,f3521,f3520,f3515,f3507,f3503,f3499,f3498,f3493,f3492,f3491,f3490,f3489,f3485,f3483,f3477,f3476,f3470,f3469,f3462,f3459,f3458,f3452,f3449,f3448,f3446,f3443,f3441,f3440,f3437,f3436,f3435,f3434,f3431,f3430,f3429,f3428,f3427,f3426,f3420,f3419,f3412,f3411,f3405,f3404,f3397,f3396,f3395,f3394,f3393,f3391,f3390,f3384,f3383,f3378,f3363,f3362,f3361,f3360,f3358,f3357,f3356,f3354,f3708,f3726,f3770,f3771,f3773,f3788,f3799,f3800,f3802,f3851,f3852,f3853,f3854,f3905,f3907,f3908,f3909,f3910,f3911,f3912,f3913,f3914,f3915,f3916,f3917,f3918,f3919,f3920,f3921,f3922,f3923,f3924,f3925,f3926,f3927,f4039,f4569,f4573,f4576,f4579,f4583,f4586,f4589,f4592,f4595,f4598,f4601,f4604,f4607,f4610,f4613,f4616,f4619,f4622,f4625,f4072,f4073,f4074,f4075,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4646,f4647,f4648,f4649,f4650,f4124,f4125,f4126,f4127,f4651,f4652,f4653,f4654,f4655,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4667,f4668,f4669,f4670,f4671,f4177,f4179,f4180,f4181,f4182,f4183,f4184,f4185,f4186,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4677,f4679,f4681,f4683,f4686,f4688,f4690,f4692,f4694,f4696,f4698,f4700,f4702,f4704,f4706,f4708,f4710,f4712,f4714,f4571,f4941,f4744,f4751,f4966,f4972,f4982,f4986,f5009,f5011,f5013,f5015,f5017,f5019])).
fof(f5018,plain,(
  $false | $spl14),
  inference(global_subsumption,[],[f321,f320,f319,f318,f322,f573,f329,f337,f336,f341,f340,f364,f371,f370,f369,f368,f367,f379,f377,f382,f380,f387,f384,f394,f393,f574,f412,f409,f408,f407,f421,f423,f425,f424,f428,f427,f426,f431,f430,f429,f434,f433,f432,f437,f436,f435,f439,f438,f451,f450,f449,f559,f447,f446,f468,f467,f466,f465,f464,f463,f462,f472,f471,f470,f469,f476,f475,f474,f473,f480,f479,f478,f477,f484,f483,f482,f481,f558,f488,f487,f486,f485,f495,f494,f493,f492,f501,f498,f497,f496,f508,f507,f506,f557,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f525,f524,f523,f556,f521,f520,f519,f529,f528,f527,f526,f533,f532,f531,f530,f555,f575,f576,f577,f578,f580,f582,f583,f585,f587,f588,f589,f590,f591,f593,f595,f596,f597,f599,f601,f602,f604,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f617,f621,f639,f640,f643,f649,f657,f664,f681,f682,f685,f688,f695,f704,f707,f708,f711,f712,f716,f725,f729,f730,f733,f736,f740,f743,f746,f755,f756,f757,f765,f768,f772,f775,f778,f797,f800,f815,f813,f818,f836,f884,f902,f925,f923,f929,f927,f941,f938,f951,f949,f992,f1003,f1011,f1024,f1032,f1036,f1042,f1051,f1056,f1055,f1068,f1076,f1080,f1113,f1124,f1131,f1140,f1152,f1160,f1164,f1178,f1196,f1197,f1212,f1222,f1237,f1249,f1257,f1261,f1272,f1282,f1298,f1314,f1313,f1311,f1310,f1325,f1334,f1335,f1336,f1337,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1347,f1348,f1349,f1350,f1351,f1352,f1353,f1355,f1356,f1371,f1381,f1397,f1389,f1387,f1386,f1405,f1402,f1439,f1429,f1428,f1484,f1503,f1515,f1514,f1513,f1512,f1517,f1518,f1519,f1529,f1530,f1531,f1532,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1543,f1544,f1545,f1546,f1547,f1548,f1549,f1551,f1552,f1866,f1863,f1867,f1865,f1876,f1873,f1877,f1875,f1887,f1884,f1886,f1900,f1897,f1899,f1916,f1912,f1934,f1930,f1952,f1948,f1977,f1973,f2000,f2015,f2016,f2027,f2038,f2035,f2036,f2037,f2051,f2062,f2059,f2060,f2061,f2082,f2104,f2105,f2106,f2111,f2134,f2135,f2136,f2141,f2142,f2186,f2182,f2209,f2205,f2233,f2234,f2228,f2229,f2230,f2232,f2249,f2250,f2244,f2245,f2246,f2251,f2248,f2286,f2287,f2305,f2306,f2334,f2336,f2347,f2349,f2367,f2364,f2386,f2370,f2371,f2372,f2374,f2375,f2400,f2403,f2404,f2405,f2407,f2408,f2429,f2444,f2427,f2432,f2433,f2434,f2436,f2456,f2460,f2461,f2462,f2464,f2489,f2488,f2486,f2485,f2497,f2506,f2507,f2508,f2509,f2511,f2512,f2513,f2514,f2515,f2516,f2517,f2518,f2519,f2520,f2521,f2522,f2523,f2524,f2525,f2526,f2528,f2529,f2530,f2531,f2532,f2533,f2831,f2830,f2827,f2824,f2821,f2820,f2819,f2818,f2817,f2816,f2815,f2814,f2813,f2812,f2811,f2810,f2809,f2806,f2803,f2802,f2801,f2797,f2796,f2790,f2787,f2786,f2780,f2779,f2772,f2746,f2745,f2744,f2739,f2738,f2733,f2732,f2731,f2726,f2718,f2714,f2710,f2709,f2704,f2703,f2702,f2701,f2700,f2696,f2694,f2688,f2687,f2681,f2680,f2673,f2670,f2669,f2663,f2660,f2659,f2658,f2656,f2654,f2653,f2651,f2650,f2649,f2648,f2646,f2645,f2644,f2643,f2642,f2641,f2635,f2634,f2627,f2626,f2620,f2619,f2612,f2611,f2609,f2607,f2605,f2602,f2601,f2595,f2594,f2589,f2574,f2573,f2572,f2571,f2569,f2568,f2567,f2566,f2565,f3267,f3266,f3263,f3260,f3257,f3256,f3255,f3254,f3253,f3252,f3251,f3250,f3249,f3248,f3247,f3246,f3245,f3242,f3239,f3238,f3237,f3233,f3232,f3226,f3223,f3222,f3216,f3215,f3208,f3182,f3181,f3180,f3175,f3174,f3169,f3168,f3167,f3162,f3154,f3150,f3146,f3145,f3140,f3139,f3138,f3137,f3136,f3132,f3130,f3124,f3123,f3117,f3116,f3109,f3106,f3105,f3099,f3096,f3095,f3093,f3090,f3088,f3087,f3084,f3083,f3082,f3081,f3078,f3077,f3076,f3075,f3074,f3073,f3067,f3066,f3059,f3058,f3052,f3051,f3044,f3043,f3041,f3039,f3037,f3034,f3033,f3027,f3026,f3021,f3006,f3005,f3004,f3003,f3001,f3000,f2999,f2998,f2997,f3620,f3619,f3616,f3613,f3610,f3609,f3608,f3607,f3606,f3605,f3604,f3603,f3602,f3601,f3600,f3599,f3598,f3595,f3592,f3591,f3590,f3586,f3585,f3579,f3576,f3575,f3569,f3568,f3561,f3535,f3534,f3533,f3528,f3527,f3522,f3521,f3520,f3515,f3507,f3503,f3499,f3498,f3493,f3492,f3491,f3490,f3489,f3485,f3483,f3477,f3476,f3470,f3469,f3462,f3459,f3458,f3452,f3449,f3448,f3446,f3443,f3441,f3440,f3437,f3436,f3435,f3434,f3431,f3430,f3429,f3428,f3427,f3426,f3420,f3419,f3412,f3411,f3405,f3404,f3397,f3396,f3395,f3394,f3393,f3391,f3390,f3384,f3383,f3378,f3363,f3362,f3361,f3360,f3358,f3357,f3356,f3354,f3708,f3726,f3770,f3771,f3773,f3788,f3799,f3800,f3802,f3851,f3852,f3853,f3854,f3905,f3907,f3908,f3909,f3910,f3911,f3912,f3913,f3914,f3915,f3916,f3917,f3918,f3919,f3920,f3921,f3922,f3923,f3924,f3925,f3926,f3927,f4039,f4569,f4573,f4576,f4579,f4583,f4586,f4589,f4592,f4595,f4598,f4601,f4604,f4607,f4610,f4613,f4616,f4619,f4622,f4625,f4072,f4073,f4074,f4075,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4646,f4647,f4648,f4649,f4650,f4124,f4125,f4126,f4127,f4651,f4652,f4653,f4654,f4655,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4667,f4668,f4669,f4670,f4671,f4177,f4179,f4180,f4181,f4182,f4183,f4184,f4185,f4186,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4677,f4679,f4681,f4683,f4686,f4688,f4690,f4692,f4694,f4696,f4698,f4700,f4702,f4704,f4706,f4708,f4710,f4712,f4714,f4571,f4941,f4744,f4751,f4966,f4972,f4982,f4986,f5009,f5011,f5013,f5015,f5017])).
fof(f5016,plain,(
  $false | $spl14),
  inference(global_subsumption,[],[f321,f320,f319,f318,f322,f573,f329,f337,f336,f341,f340,f364,f371,f370,f369,f368,f367,f379,f377,f382,f380,f387,f384,f394,f393,f574,f412,f409,f408,f407,f421,f423,f425,f424,f428,f427,f426,f431,f430,f429,f434,f433,f432,f437,f436,f435,f439,f438,f451,f450,f449,f559,f447,f446,f468,f467,f466,f465,f464,f463,f462,f472,f471,f470,f469,f476,f475,f474,f473,f480,f479,f478,f477,f484,f483,f482,f481,f558,f488,f487,f486,f485,f495,f494,f493,f492,f501,f498,f497,f496,f508,f507,f506,f557,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f525,f524,f523,f556,f521,f520,f519,f529,f528,f527,f526,f533,f532,f531,f530,f555,f575,f576,f577,f578,f580,f582,f583,f585,f587,f588,f589,f590,f591,f593,f595,f596,f597,f599,f601,f602,f604,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f617,f621,f639,f640,f643,f649,f657,f664,f681,f682,f685,f688,f695,f704,f707,f708,f711,f712,f716,f725,f729,f730,f733,f736,f740,f743,f746,f755,f756,f757,f765,f768,f772,f775,f778,f797,f800,f815,f813,f818,f836,f884,f902,f925,f923,f929,f927,f941,f938,f951,f949,f992,f1003,f1011,f1024,f1032,f1036,f1042,f1051,f1056,f1055,f1068,f1076,f1080,f1113,f1124,f1131,f1140,f1152,f1160,f1164,f1178,f1196,f1197,f1212,f1222,f1237,f1249,f1257,f1261,f1272,f1282,f1298,f1314,f1313,f1311,f1310,f1325,f1334,f1335,f1336,f1337,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1347,f1348,f1349,f1350,f1351,f1352,f1353,f1355,f1356,f1371,f1381,f1397,f1389,f1387,f1386,f1405,f1402,f1439,f1429,f1428,f1484,f1503,f1515,f1514,f1513,f1512,f1517,f1518,f1519,f1529,f1530,f1531,f1532,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1543,f1544,f1545,f1546,f1547,f1548,f1549,f1551,f1552,f1866,f1863,f1867,f1865,f1876,f1873,f1877,f1875,f1887,f1884,f1886,f1900,f1897,f1899,f1916,f1912,f1934,f1930,f1952,f1948,f1977,f1973,f2000,f2015,f2016,f2027,f2038,f2035,f2036,f2037,f2051,f2062,f2059,f2060,f2061,f2082,f2104,f2105,f2106,f2111,f2134,f2135,f2136,f2141,f2142,f2186,f2182,f2209,f2205,f2233,f2234,f2228,f2229,f2230,f2232,f2249,f2250,f2244,f2245,f2246,f2251,f2248,f2286,f2287,f2305,f2306,f2334,f2336,f2347,f2349,f2367,f2364,f2386,f2370,f2371,f2372,f2374,f2375,f2400,f2403,f2404,f2405,f2407,f2408,f2429,f2444,f2427,f2432,f2433,f2434,f2436,f2456,f2460,f2461,f2462,f2464,f2489,f2488,f2486,f2485,f2497,f2506,f2507,f2508,f2509,f2511,f2512,f2513,f2514,f2515,f2516,f2517,f2518,f2519,f2520,f2521,f2522,f2523,f2524,f2525,f2526,f2528,f2529,f2530,f2531,f2532,f2533,f2831,f2830,f2827,f2824,f2821,f2820,f2819,f2818,f2817,f2816,f2815,f2814,f2813,f2812,f2811,f2810,f2809,f2806,f2803,f2802,f2801,f2797,f2796,f2790,f2787,f2786,f2780,f2779,f2772,f2746,f2745,f2744,f2739,f2738,f2733,f2732,f2731,f2726,f2718,f2714,f2710,f2709,f2704,f2703,f2702,f2701,f2700,f2696,f2694,f2688,f2687,f2681,f2680,f2673,f2670,f2669,f2663,f2660,f2659,f2658,f2656,f2654,f2653,f2651,f2650,f2649,f2648,f2646,f2645,f2644,f2643,f2642,f2641,f2635,f2634,f2627,f2626,f2620,f2619,f2612,f2611,f2609,f2607,f2605,f2602,f2601,f2595,f2594,f2589,f2574,f2573,f2572,f2571,f2569,f2568,f2567,f2566,f2565,f3267,f3266,f3263,f3260,f3257,f3256,f3255,f3254,f3253,f3252,f3251,f3250,f3249,f3248,f3247,f3246,f3245,f3242,f3239,f3238,f3237,f3233,f3232,f3226,f3223,f3222,f3216,f3215,f3208,f3182,f3181,f3180,f3175,f3174,f3169,f3168,f3167,f3162,f3154,f3150,f3146,f3145,f3140,f3139,f3138,f3137,f3136,f3132,f3130,f3124,f3123,f3117,f3116,f3109,f3106,f3105,f3099,f3096,f3095,f3093,f3090,f3088,f3087,f3084,f3083,f3082,f3081,f3078,f3077,f3076,f3075,f3074,f3073,f3067,f3066,f3059,f3058,f3052,f3051,f3044,f3043,f3041,f3039,f3037,f3034,f3033,f3027,f3026,f3021,f3006,f3005,f3004,f3003,f3001,f3000,f2999,f2998,f2997,f3620,f3619,f3616,f3613,f3610,f3609,f3608,f3607,f3606,f3605,f3604,f3603,f3602,f3601,f3600,f3599,f3598,f3595,f3592,f3591,f3590,f3586,f3585,f3579,f3576,f3575,f3569,f3568,f3561,f3535,f3534,f3533,f3528,f3527,f3522,f3521,f3520,f3515,f3507,f3503,f3499,f3498,f3493,f3492,f3491,f3490,f3489,f3485,f3483,f3477,f3476,f3470,f3469,f3462,f3459,f3458,f3452,f3449,f3448,f3446,f3443,f3441,f3440,f3437,f3436,f3435,f3434,f3431,f3430,f3429,f3428,f3427,f3426,f3420,f3419,f3412,f3411,f3405,f3404,f3397,f3396,f3395,f3394,f3393,f3391,f3390,f3384,f3383,f3378,f3363,f3362,f3361,f3360,f3358,f3357,f3356,f3354,f3708,f3726,f3770,f3771,f3773,f3788,f3799,f3800,f3802,f3851,f3852,f3853,f3854,f3905,f3907,f3908,f3909,f3910,f3911,f3912,f3913,f3914,f3915,f3916,f3917,f3918,f3919,f3920,f3921,f3922,f3923,f3924,f3925,f3926,f3927,f4039,f4569,f4573,f4576,f4579,f4583,f4586,f4589,f4592,f4595,f4598,f4601,f4604,f4607,f4610,f4613,f4616,f4619,f4622,f4625,f4072,f4073,f4074,f4075,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4646,f4647,f4648,f4649,f4650,f4124,f4125,f4126,f4127,f4651,f4652,f4653,f4654,f4655,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4667,f4668,f4669,f4670,f4671,f4177,f4179,f4180,f4181,f4182,f4183,f4184,f4185,f4186,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4677,f4679,f4681,f4683,f4686,f4688,f4690,f4692,f4694,f4696,f4698,f4700,f4702,f4704,f4706,f4708,f4710,f4712,f4714,f4571,f4941,f4744,f4751,f4966,f4972,f4982,f4986,f5009,f5011,f5013,f5015])).
fof(f5014,plain,(
  $false | $spl14),
  inference(global_subsumption,[],[f321,f320,f319,f318,f322,f573,f329,f337,f336,f341,f340,f364,f371,f370,f369,f368,f367,f379,f377,f382,f380,f387,f384,f394,f393,f574,f412,f409,f408,f407,f421,f423,f425,f424,f428,f427,f426,f431,f430,f429,f434,f433,f432,f437,f436,f435,f439,f438,f451,f450,f449,f559,f447,f446,f468,f467,f466,f465,f464,f463,f462,f472,f471,f470,f469,f476,f475,f474,f473,f480,f479,f478,f477,f484,f483,f482,f481,f558,f488,f487,f486,f485,f495,f494,f493,f492,f501,f498,f497,f496,f508,f507,f506,f557,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f525,f524,f523,f556,f521,f520,f519,f529,f528,f527,f526,f533,f532,f531,f530,f555,f575,f576,f577,f578,f580,f582,f583,f585,f587,f588,f589,f590,f591,f593,f595,f596,f597,f599,f601,f602,f604,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f617,f621,f639,f640,f643,f649,f657,f664,f681,f682,f685,f688,f695,f704,f707,f708,f711,f712,f716,f725,f729,f730,f733,f736,f740,f743,f746,f755,f756,f757,f765,f768,f772,f775,f778,f797,f800,f815,f813,f818,f836,f884,f902,f925,f923,f929,f927,f941,f938,f951,f949,f992,f1003,f1011,f1024,f1032,f1036,f1042,f1051,f1056,f1055,f1068,f1076,f1080,f1113,f1124,f1131,f1140,f1152,f1160,f1164,f1178,f1196,f1197,f1212,f1222,f1237,f1249,f1257,f1261,f1272,f1282,f1298,f1314,f1313,f1311,f1310,f1325,f1334,f1335,f1336,f1337,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1347,f1348,f1349,f1350,f1351,f1352,f1353,f1355,f1356,f1371,f1381,f1397,f1389,f1387,f1386,f1405,f1402,f1439,f1429,f1428,f1484,f1503,f1515,f1514,f1513,f1512,f1517,f1518,f1519,f1529,f1530,f1531,f1532,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1543,f1544,f1545,f1546,f1547,f1548,f1549,f1551,f1552,f1866,f1863,f1867,f1865,f1876,f1873,f1877,f1875,f1887,f1884,f1886,f1900,f1897,f1899,f1916,f1912,f1934,f1930,f1952,f1948,f1977,f1973,f2000,f2015,f2016,f2027,f2038,f2035,f2036,f2037,f2051,f2062,f2059,f2060,f2061,f2082,f2104,f2105,f2106,f2111,f2134,f2135,f2136,f2141,f2142,f2186,f2182,f2209,f2205,f2233,f2234,f2228,f2229,f2230,f2232,f2249,f2250,f2244,f2245,f2246,f2251,f2248,f2286,f2287,f2305,f2306,f2334,f2336,f2347,f2349,f2367,f2364,f2386,f2370,f2371,f2372,f2374,f2375,f2400,f2403,f2404,f2405,f2407,f2408,f2429,f2444,f2427,f2432,f2433,f2434,f2436,f2456,f2460,f2461,f2462,f2464,f2489,f2488,f2486,f2485,f2497,f2506,f2507,f2508,f2509,f2511,f2512,f2513,f2514,f2515,f2516,f2517,f2518,f2519,f2520,f2521,f2522,f2523,f2524,f2525,f2526,f2528,f2529,f2530,f2531,f2532,f2533,f2831,f2830,f2827,f2824,f2821,f2820,f2819,f2818,f2817,f2816,f2815,f2814,f2813,f2812,f2811,f2810,f2809,f2806,f2803,f2802,f2801,f2797,f2796,f2790,f2787,f2786,f2780,f2779,f2772,f2746,f2745,f2744,f2739,f2738,f2733,f2732,f2731,f2726,f2718,f2714,f2710,f2709,f2704,f2703,f2702,f2701,f2700,f2696,f2694,f2688,f2687,f2681,f2680,f2673,f2670,f2669,f2663,f2660,f2659,f2658,f2656,f2654,f2653,f2651,f2650,f2649,f2648,f2646,f2645,f2644,f2643,f2642,f2641,f2635,f2634,f2627,f2626,f2620,f2619,f2612,f2611,f2609,f2607,f2605,f2602,f2601,f2595,f2594,f2589,f2574,f2573,f2572,f2571,f2569,f2568,f2567,f2566,f2565,f3267,f3266,f3263,f3260,f3257,f3256,f3255,f3254,f3253,f3252,f3251,f3250,f3249,f3248,f3247,f3246,f3245,f3242,f3239,f3238,f3237,f3233,f3232,f3226,f3223,f3222,f3216,f3215,f3208,f3182,f3181,f3180,f3175,f3174,f3169,f3168,f3167,f3162,f3154,f3150,f3146,f3145,f3140,f3139,f3138,f3137,f3136,f3132,f3130,f3124,f3123,f3117,f3116,f3109,f3106,f3105,f3099,f3096,f3095,f3093,f3090,f3088,f3087,f3084,f3083,f3082,f3081,f3078,f3077,f3076,f3075,f3074,f3073,f3067,f3066,f3059,f3058,f3052,f3051,f3044,f3043,f3041,f3039,f3037,f3034,f3033,f3027,f3026,f3021,f3006,f3005,f3004,f3003,f3001,f3000,f2999,f2998,f2997,f3620,f3619,f3616,f3613,f3610,f3609,f3608,f3607,f3606,f3605,f3604,f3603,f3602,f3601,f3600,f3599,f3598,f3595,f3592,f3591,f3590,f3586,f3585,f3579,f3576,f3575,f3569,f3568,f3561,f3535,f3534,f3533,f3528,f3527,f3522,f3521,f3520,f3515,f3507,f3503,f3499,f3498,f3493,f3492,f3491,f3490,f3489,f3485,f3483,f3477,f3476,f3470,f3469,f3462,f3459,f3458,f3452,f3449,f3448,f3446,f3443,f3441,f3440,f3437,f3436,f3435,f3434,f3431,f3430,f3429,f3428,f3427,f3426,f3420,f3419,f3412,f3411,f3405,f3404,f3397,f3396,f3395,f3394,f3393,f3391,f3390,f3384,f3383,f3378,f3363,f3362,f3361,f3360,f3358,f3357,f3356,f3354,f3708,f3726,f3770,f3771,f3773,f3788,f3799,f3800,f3802,f3851,f3852,f3853,f3854,f3905,f3907,f3908,f3909,f3910,f3911,f3912,f3913,f3914,f3915,f3916,f3917,f3918,f3919,f3920,f3921,f3922,f3923,f3924,f3925,f3926,f3927,f4039,f4569,f4573,f4576,f4579,f4583,f4586,f4589,f4592,f4595,f4598,f4601,f4604,f4607,f4610,f4613,f4616,f4619,f4622,f4625,f4072,f4073,f4074,f4075,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4646,f4647,f4648,f4649,f4650,f4124,f4125,f4126,f4127,f4651,f4652,f4653,f4654,f4655,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4667,f4668,f4669,f4670,f4671,f4177,f4179,f4180,f4181,f4182,f4183,f4184,f4185,f4186,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4677,f4679,f4681,f4683,f4686,f4688,f4690,f4692,f4694,f4696,f4698,f4700,f4702,f4704,f4706,f4708,f4710,f4712,f4714,f4571,f4941,f4744,f4751,f4966,f4972,f4982,f4986,f5009,f5011,f5013])).
fof(f5012,plain,(
  $false | $spl14),
  inference(global_subsumption,[],[f321,f320,f319,f318,f322,f573,f329,f337,f336,f341,f340,f364,f371,f370,f369,f368,f367,f379,f377,f382,f380,f387,f384,f394,f393,f574,f412,f409,f408,f407,f421,f423,f425,f424,f428,f427,f426,f431,f430,f429,f434,f433,f432,f437,f436,f435,f439,f438,f451,f450,f449,f559,f447,f446,f468,f467,f466,f465,f464,f463,f462,f472,f471,f470,f469,f476,f475,f474,f473,f480,f479,f478,f477,f484,f483,f482,f481,f558,f488,f487,f486,f485,f495,f494,f493,f492,f501,f498,f497,f496,f508,f507,f506,f557,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f525,f524,f523,f556,f521,f520,f519,f529,f528,f527,f526,f533,f532,f531,f530,f555,f575,f576,f577,f578,f580,f582,f583,f585,f587,f588,f589,f590,f591,f593,f595,f596,f597,f599,f601,f602,f604,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f617,f621,f639,f640,f643,f649,f657,f664,f681,f682,f685,f688,f695,f704,f707,f708,f711,f712,f716,f725,f729,f730,f733,f736,f740,f743,f746,f755,f756,f757,f765,f768,f772,f775,f778,f797,f800,f815,f813,f818,f836,f884,f902,f925,f923,f929,f927,f941,f938,f951,f949,f992,f1003,f1011,f1024,f1032,f1036,f1042,f1051,f1056,f1055,f1068,f1076,f1080,f1113,f1124,f1131,f1140,f1152,f1160,f1164,f1178,f1196,f1197,f1212,f1222,f1237,f1249,f1257,f1261,f1272,f1282,f1298,f1314,f1313,f1311,f1310,f1325,f1334,f1335,f1336,f1337,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1347,f1348,f1349,f1350,f1351,f1352,f1353,f1355,f1356,f1371,f1381,f1397,f1389,f1387,f1386,f1405,f1402,f1439,f1429,f1428,f1484,f1503,f1515,f1514,f1513,f1512,f1517,f1518,f1519,f1529,f1530,f1531,f1532,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1543,f1544,f1545,f1546,f1547,f1548,f1549,f1551,f1552,f1866,f1863,f1867,f1865,f1876,f1873,f1877,f1875,f1887,f1884,f1886,f1900,f1897,f1899,f1916,f1912,f1934,f1930,f1952,f1948,f1977,f1973,f2000,f2015,f2016,f2027,f2038,f2035,f2036,f2037,f2051,f2062,f2059,f2060,f2061,f2082,f2104,f2105,f2106,f2111,f2134,f2135,f2136,f2141,f2142,f2186,f2182,f2209,f2205,f2233,f2234,f2228,f2229,f2230,f2232,f2249,f2250,f2244,f2245,f2246,f2251,f2248,f2286,f2287,f2305,f2306,f2334,f2336,f2347,f2349,f2367,f2364,f2386,f2370,f2371,f2372,f2374,f2375,f2400,f2403,f2404,f2405,f2407,f2408,f2429,f2444,f2427,f2432,f2433,f2434,f2436,f2456,f2460,f2461,f2462,f2464,f2489,f2488,f2486,f2485,f2497,f2506,f2507,f2508,f2509,f2511,f2512,f2513,f2514,f2515,f2516,f2517,f2518,f2519,f2520,f2521,f2522,f2523,f2524,f2525,f2526,f2528,f2529,f2530,f2531,f2532,f2533,f2831,f2830,f2827,f2824,f2821,f2820,f2819,f2818,f2817,f2816,f2815,f2814,f2813,f2812,f2811,f2810,f2809,f2806,f2803,f2802,f2801,f2797,f2796,f2790,f2787,f2786,f2780,f2779,f2772,f2746,f2745,f2744,f2739,f2738,f2733,f2732,f2731,f2726,f2718,f2714,f2710,f2709,f2704,f2703,f2702,f2701,f2700,f2696,f2694,f2688,f2687,f2681,f2680,f2673,f2670,f2669,f2663,f2660,f2659,f2658,f2656,f2654,f2653,f2651,f2650,f2649,f2648,f2646,f2645,f2644,f2643,f2642,f2641,f2635,f2634,f2627,f2626,f2620,f2619,f2612,f2611,f2609,f2607,f2605,f2602,f2601,f2595,f2594,f2589,f2574,f2573,f2572,f2571,f2569,f2568,f2567,f2566,f2565,f3267,f3266,f3263,f3260,f3257,f3256,f3255,f3254,f3253,f3252,f3251,f3250,f3249,f3248,f3247,f3246,f3245,f3242,f3239,f3238,f3237,f3233,f3232,f3226,f3223,f3222,f3216,f3215,f3208,f3182,f3181,f3180,f3175,f3174,f3169,f3168,f3167,f3162,f3154,f3150,f3146,f3145,f3140,f3139,f3138,f3137,f3136,f3132,f3130,f3124,f3123,f3117,f3116,f3109,f3106,f3105,f3099,f3096,f3095,f3093,f3090,f3088,f3087,f3084,f3083,f3082,f3081,f3078,f3077,f3076,f3075,f3074,f3073,f3067,f3066,f3059,f3058,f3052,f3051,f3044,f3043,f3041,f3039,f3037,f3034,f3033,f3027,f3026,f3021,f3006,f3005,f3004,f3003,f3001,f3000,f2999,f2998,f2997,f3620,f3619,f3616,f3613,f3610,f3609,f3608,f3607,f3606,f3605,f3604,f3603,f3602,f3601,f3600,f3599,f3598,f3595,f3592,f3591,f3590,f3586,f3585,f3579,f3576,f3575,f3569,f3568,f3561,f3535,f3534,f3533,f3528,f3527,f3522,f3521,f3520,f3515,f3507,f3503,f3499,f3498,f3493,f3492,f3491,f3490,f3489,f3485,f3483,f3477,f3476,f3470,f3469,f3462,f3459,f3458,f3452,f3449,f3448,f3446,f3443,f3441,f3440,f3437,f3436,f3435,f3434,f3431,f3430,f3429,f3428,f3427,f3426,f3420,f3419,f3412,f3411,f3405,f3404,f3397,f3396,f3395,f3394,f3393,f3391,f3390,f3384,f3383,f3378,f3363,f3362,f3361,f3360,f3358,f3357,f3356,f3354,f3708,f3726,f3770,f3771,f3773,f3788,f3799,f3800,f3802,f3851,f3852,f3853,f3854,f3905,f3907,f3908,f3909,f3910,f3911,f3912,f3913,f3914,f3915,f3916,f3917,f3918,f3919,f3920,f3921,f3922,f3923,f3924,f3925,f3926,f3927,f4039,f4569,f4573,f4576,f4579,f4583,f4586,f4589,f4592,f4595,f4598,f4601,f4604,f4607,f4610,f4613,f4616,f4619,f4622,f4625,f4072,f4073,f4074,f4075,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4646,f4647,f4648,f4649,f4650,f4124,f4125,f4126,f4127,f4651,f4652,f4653,f4654,f4655,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4667,f4668,f4669,f4670,f4671,f4177,f4179,f4180,f4181,f4182,f4183,f4184,f4185,f4186,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4677,f4679,f4681,f4683,f4686,f4688,f4690,f4692,f4694,f4696,f4698,f4700,f4702,f4704,f4706,f4708,f4710,f4712,f4714,f4571,f4941,f4744,f4751,f4966,f4972,f4982,f4986,f5009,f5011])).
fof(f5010,plain,(
  $false | $spl14),
  inference(global_subsumption,[],[f321,f320,f319,f318,f322,f573,f329,f337,f336,f341,f340,f364,f371,f370,f369,f368,f367,f379,f377,f382,f380,f387,f384,f394,f393,f574,f412,f409,f408,f407,f421,f423,f425,f424,f428,f427,f426,f431,f430,f429,f434,f433,f432,f437,f436,f435,f439,f438,f451,f450,f449,f559,f447,f446,f468,f467,f466,f465,f464,f463,f462,f472,f471,f470,f469,f476,f475,f474,f473,f480,f479,f478,f477,f484,f483,f482,f481,f558,f488,f487,f486,f485,f495,f494,f493,f492,f501,f498,f497,f496,f508,f507,f506,f557,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f525,f524,f523,f556,f521,f520,f519,f529,f528,f527,f526,f533,f532,f531,f530,f555,f575,f576,f577,f578,f580,f582,f583,f585,f587,f588,f589,f590,f591,f593,f595,f596,f597,f599,f601,f602,f604,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f617,f621,f639,f640,f643,f649,f657,f664,f681,f682,f685,f688,f695,f704,f707,f708,f711,f712,f716,f725,f729,f730,f733,f736,f740,f743,f746,f755,f756,f757,f765,f768,f772,f775,f778,f797,f800,f815,f813,f818,f836,f884,f902,f925,f923,f929,f927,f941,f938,f951,f949,f992,f1003,f1011,f1024,f1032,f1036,f1042,f1051,f1056,f1055,f1068,f1076,f1080,f1113,f1124,f1131,f1140,f1152,f1160,f1164,f1178,f1196,f1197,f1212,f1222,f1237,f1249,f1257,f1261,f1272,f1282,f1298,f1314,f1313,f1311,f1310,f1325,f1334,f1335,f1336,f1337,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1347,f1348,f1349,f1350,f1351,f1352,f1353,f1355,f1356,f1371,f1381,f1397,f1389,f1387,f1386,f1405,f1402,f1439,f1429,f1428,f1484,f1503,f1515,f1514,f1513,f1512,f1517,f1518,f1519,f1529,f1530,f1531,f1532,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1543,f1544,f1545,f1546,f1547,f1548,f1549,f1551,f1552,f1866,f1863,f1867,f1865,f1876,f1873,f1877,f1875,f1887,f1884,f1886,f1900,f1897,f1899,f1916,f1912,f1934,f1930,f1952,f1948,f1977,f1973,f2000,f2015,f2016,f2027,f2038,f2035,f2036,f2037,f2051,f2062,f2059,f2060,f2061,f2082,f2104,f2105,f2106,f2111,f2134,f2135,f2136,f2141,f2142,f2186,f2182,f2209,f2205,f2233,f2234,f2228,f2229,f2230,f2232,f2249,f2250,f2244,f2245,f2246,f2251,f2248,f2286,f2287,f2305,f2306,f2334,f2336,f2347,f2349,f2367,f2364,f2386,f2370,f2371,f2372,f2374,f2375,f2400,f2403,f2404,f2405,f2407,f2408,f2429,f2444,f2427,f2432,f2433,f2434,f2436,f2456,f2460,f2461,f2462,f2464,f2489,f2488,f2486,f2485,f2497,f2506,f2507,f2508,f2509,f2511,f2512,f2513,f2514,f2515,f2516,f2517,f2518,f2519,f2520,f2521,f2522,f2523,f2524,f2525,f2526,f2528,f2529,f2530,f2531,f2532,f2533,f2831,f2830,f2827,f2824,f2821,f2820,f2819,f2818,f2817,f2816,f2815,f2814,f2813,f2812,f2811,f2810,f2809,f2806,f2803,f2802,f2801,f2797,f2796,f2790,f2787,f2786,f2780,f2779,f2772,f2746,f2745,f2744,f2739,f2738,f2733,f2732,f2731,f2726,f2718,f2714,f2710,f2709,f2704,f2703,f2702,f2701,f2700,f2696,f2694,f2688,f2687,f2681,f2680,f2673,f2670,f2669,f2663,f2660,f2659,f2658,f2656,f2654,f2653,f2651,f2650,f2649,f2648,f2646,f2645,f2644,f2643,f2642,f2641,f2635,f2634,f2627,f2626,f2620,f2619,f2612,f2611,f2609,f2607,f2605,f2602,f2601,f2595,f2594,f2589,f2574,f2573,f2572,f2571,f2569,f2568,f2567,f2566,f2565,f3267,f3266,f3263,f3260,f3257,f3256,f3255,f3254,f3253,f3252,f3251,f3250,f3249,f3248,f3247,f3246,f3245,f3242,f3239,f3238,f3237,f3233,f3232,f3226,f3223,f3222,f3216,f3215,f3208,f3182,f3181,f3180,f3175,f3174,f3169,f3168,f3167,f3162,f3154,f3150,f3146,f3145,f3140,f3139,f3138,f3137,f3136,f3132,f3130,f3124,f3123,f3117,f3116,f3109,f3106,f3105,f3099,f3096,f3095,f3093,f3090,f3088,f3087,f3084,f3083,f3082,f3081,f3078,f3077,f3076,f3075,f3074,f3073,f3067,f3066,f3059,f3058,f3052,f3051,f3044,f3043,f3041,f3039,f3037,f3034,f3033,f3027,f3026,f3021,f3006,f3005,f3004,f3003,f3001,f3000,f2999,f2998,f2997,f3620,f3619,f3616,f3613,f3610,f3609,f3608,f3607,f3606,f3605,f3604,f3603,f3602,f3601,f3600,f3599,f3598,f3595,f3592,f3591,f3590,f3586,f3585,f3579,f3576,f3575,f3569,f3568,f3561,f3535,f3534,f3533,f3528,f3527,f3522,f3521,f3520,f3515,f3507,f3503,f3499,f3498,f3493,f3492,f3491,f3490,f3489,f3485,f3483,f3477,f3476,f3470,f3469,f3462,f3459,f3458,f3452,f3449,f3448,f3446,f3443,f3441,f3440,f3437,f3436,f3435,f3434,f3431,f3430,f3429,f3428,f3427,f3426,f3420,f3419,f3412,f3411,f3405,f3404,f3397,f3396,f3395,f3394,f3393,f3391,f3390,f3384,f3383,f3378,f3363,f3362,f3361,f3360,f3358,f3357,f3356,f3354,f3708,f3726,f3770,f3771,f3773,f3788,f3799,f3800,f3802,f3851,f3852,f3853,f3854,f3905,f3907,f3908,f3909,f3910,f3911,f3912,f3913,f3914,f3915,f3916,f3917,f3918,f3919,f3920,f3921,f3922,f3923,f3924,f3925,f3926,f3927,f4039,f4569,f4573,f4576,f4579,f4583,f4586,f4589,f4592,f4595,f4598,f4601,f4604,f4607,f4610,f4613,f4616,f4619,f4622,f4625,f4072,f4073,f4074,f4075,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4646,f4647,f4648,f4649,f4650,f4124,f4125,f4126,f4127,f4651,f4652,f4653,f4654,f4655,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4667,f4668,f4669,f4670,f4671,f4177,f4179,f4180,f4181,f4182,f4183,f4184,f4185,f4186,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4677,f4679,f4681,f4683,f4686,f4688,f4690,f4692,f4694,f4696,f4698,f4700,f4702,f4704,f4706,f4708,f4710,f4712,f4714,f4571,f4941,f4744,f4751,f4966,f4972,f4982,f4986,f5009])).
fof(f5006,plain,(
  $false | ($spl6 | $spl14)),
  inference(global_subsumption,[],[f5005,f321,f320,f319,f318,f322,f573,f329,f337,f336,f341,f340,f364,f371,f370,f369,f368,f367,f379,f377,f382,f380,f387,f384,f394,f393,f574,f412,f409,f408,f407,f421,f423,f425,f424,f428,f427,f426,f431,f430,f429,f434,f433,f432,f437,f436,f435,f439,f438,f451,f450,f449,f559,f447,f446,f468,f467,f466,f465,f464,f463,f462,f472,f471,f470,f469,f476,f475,f474,f473,f480,f479,f478,f477,f484,f483,f482,f481,f558,f488,f487,f486,f485,f495,f494,f493,f492,f501,f498,f497,f496,f508,f507,f506,f557,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f525,f524,f523,f556,f521,f520,f519,f529,f528,f527,f526,f533,f532,f531,f530,f555,f575,f576,f577,f578,f580,f582,f583,f585,f587,f588,f589,f590,f591,f593,f595,f596,f597,f599,f601,f602,f604,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f617,f621,f639,f640,f643,f649,f657,f664,f681,f682,f685,f688,f695,f704,f707,f708,f711,f712,f716,f725,f729,f730,f733,f736,f740,f743,f746,f755,f756,f757,f765,f768,f772,f775,f778,f797,f800,f815,f813,f818,f836,f884,f902,f925,f923,f929,f927,f941,f938,f951,f949,f992,f1003,f1011,f1024,f1032,f1036,f1042,f1051,f1056,f1055,f1068,f1076,f1080,f1113,f1124,f1131,f1140,f1152,f1160,f1164,f1178,f1196,f1197,f1212,f1222,f1237,f1249,f1257,f1261,f1272,f1282,f1298,f1314,f1313,f1311,f1310,f1325,f1334,f1335,f1336,f1337,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1347,f1348,f1349,f1350,f1351,f1352,f1353,f1355,f1356,f1371,f1381,f1397,f1389,f1387,f1386,f1405,f1402,f1439,f1429,f1428,f1484,f1503,f1515,f1514,f1513,f1512,f1517,f1518,f1519,f1529,f1530,f1531,f1532,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1543,f1544,f1545,f1546,f1547,f1548,f1549,f1551,f1552,f1866,f1863,f1867,f1865,f1876,f1873,f1877,f1875,f1887,f1884,f1886,f1900,f1897,f1899,f1916,f1912,f1934,f1930,f1952,f1948,f1977,f1973,f2000,f2015,f2016,f2027,f2038,f2035,f2036,f2037,f2051,f2062,f2059,f2060,f2061,f2082,f2104,f2105,f2106,f2111,f2134,f2135,f2136,f2141,f2142,f2186,f2182,f2209,f2205,f2233,f2234,f2228,f2229,f2230,f2232,f2249,f2250,f2244,f2245,f2246,f2251,f2248,f2286,f2287,f2305,f2306,f2334,f2336,f2347,f2349,f2367,f2364,f2386,f2370,f2371,f2372,f2374,f2375,f2400,f2403,f2404,f2405,f2407,f2408,f2429,f2444,f2427,f2432,f2433,f2434,f2436,f2456,f2460,f2461,f2462,f2464,f2489,f2488,f2486,f2485,f2497,f2506,f2507,f2508,f2509,f2511,f2512,f2513,f2514,f2515,f2516,f2517,f2518,f2519,f2520,f2521,f2522,f2523,f2524,f2525,f2526,f2528,f2529,f2530,f2531,f2532,f2533,f2831,f2830,f2827,f2824,f2821,f2820,f2819,f2818,f2817,f2816,f2815,f2814,f2813,f2812,f2811,f2810,f2809,f2806,f2803,f2802,f2801,f2797,f2796,f2790,f2787,f2786,f2780,f2779,f2772,f2746,f2745,f2744,f2739,f2738,f2733,f2732,f2731,f2726,f2718,f2714,f2710,f2709,f2704,f2703,f2702,f2701,f2700,f2696,f2694,f2688,f2687,f2681,f2680,f2673,f2670,f2669,f2663,f2660,f2659,f2658,f2656,f2654,f2653,f2651,f2650,f2649,f2648,f2646,f2645,f2644,f2643,f2642,f2641,f2635,f2634,f2627,f2626,f2620,f2619,f2612,f2611,f2609,f2607,f2605,f2602,f2601,f2595,f2594,f2589,f2574,f2573,f2572,f2571,f2569,f2568,f2567,f2566,f2565,f3267,f3266,f3263,f3260,f3257,f3256,f3255,f3254,f3253,f3252,f3251,f3250,f3249,f3248,f3247,f3246,f3245,f3242,f3239,f3238,f3237,f3233,f3232,f3226,f3223,f3222,f3216,f3215,f3208,f3182,f3181,f3180,f3175,f3174,f3169,f3168,f3167,f3162,f3154,f3150,f3146,f3145,f3140,f3139,f3138,f3137,f3136,f3132,f3130,f3124,f3123,f3117,f3116,f3109,f3106,f3105,f3099,f3096,f3095,f3093,f3090,f3088,f3087,f3084,f3083,f3082,f3081,f3078,f3077,f3076,f3075,f3074,f3073,f3067,f3066,f3059,f3058,f3052,f3051,f3044,f3043,f3041,f3039,f3037,f3034,f3033,f3027,f3026,f3021,f3006,f3005,f3004,f3003,f3001,f3000,f2999,f2998,f2997,f3620,f3619,f3616,f3613,f3610,f3609,f3608,f3607,f3606,f3605,f3604,f3603,f3602,f3601,f3600,f3599,f3598,f3595,f3592,f3591,f3590,f3586,f3585,f3579,f3576,f3575,f3569,f3568,f3561,f3535,f3534,f3533,f3528,f3527,f3522,f3521,f3520,f3515,f3507,f3503,f3499,f3498,f3493,f3492,f3491,f3490,f3489,f3485,f3483,f3477,f3476,f3470,f3469,f3462,f3459,f3458,f3452,f3449,f3448,f3446,f3443,f3441,f3440,f3437,f3436,f3435,f3434,f3431,f3430,f3429,f3428,f3427,f3426,f3420,f3419,f3412,f3411,f3405,f3404,f3397,f3396,f3395,f3394,f3393,f3391,f3390,f3384,f3383,f3378,f3363,f3362,f3361,f3360,f3358,f3357,f3356,f3354,f3708,f3726,f3770,f3771,f3773,f3788,f3799,f3800,f3802,f3851,f3852,f3853,f3854,f3905,f3907,f3908,f3909,f3910,f3911,f3912,f3913,f3914,f3915,f3916,f3917,f3918,f3919,f3920,f3921,f3922,f3923,f3924,f3925,f3926,f3927,f4039,f4569,f4573,f4576,f4579,f4583,f4586,f4589,f4592,f4595,f4598,f4601,f4604,f4607,f4610,f4613,f4616,f4619,f4622,f4625,f4072,f4073,f4074,f4075,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4646,f4647,f4648,f4649,f4650,f4124,f4125,f4126,f4127,f4651,f4652,f4653,f4654,f4655,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4667,f4668,f4669,f4670,f4671,f4177,f4179,f4180,f4181,f4182,f4183,f4184,f4185,f4186,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4677,f4679,f4681,f4683,f4686,f4688,f4690,f4692,f4694,f4696,f4698,f4700,f4702,f4704,f4706,f4708,f4710,f4712,f4714,f4571,f4941,f4744,f4751,f4966,f4972,f4982,f4986])).
fof(f5005,plain,(
  ( ! [X5] : (s(sK14,sK15) = s(sK14,X5)) ) | ($spl6 | $spl14)),
  inference(subsumption_resolution,[],[f4854,f322])).
fof(f4854,plain,(
  ( ! [X5] : (~p(s(bool,'T_0')) | s(sK14,sK15) = s(sK14,X5)) ) | ($spl6 | $spl14)),
  inference(backward_demodulation,[],[f4571,f572])).
fof(f4996,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f4995,f322])).
fof(f4995,plain,(
  ~p(s(bool,'T_0')) | $spl14),
  inference(forward_demodulation,[],[f4994,f4571])).
fof(f4994,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))) ) | $spl14),
  inference(subsumption_resolution,[],[f4834,f322])).
fof(f4834,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))) ) | $spl14),
  inference(backward_demodulation,[],[f4571,f438])).
fof(f4987,plain,(
  $false | $spl14),
  inference(global_subsumption,[],[f321,f320,f319,f318,f322,f573,f329,f337,f336,f341,f340,f364,f371,f370,f369,f368,f367,f379,f377,f382,f380,f387,f384,f394,f393,f574,f412,f409,f408,f407,f421,f423,f425,f424,f428,f427,f426,f431,f430,f429,f434,f433,f432,f437,f436,f435,f439,f438,f451,f450,f449,f559,f447,f446,f468,f467,f466,f465,f464,f463,f462,f472,f471,f470,f469,f476,f475,f474,f473,f480,f479,f478,f477,f484,f483,f482,f481,f558,f488,f487,f486,f485,f495,f494,f493,f492,f501,f498,f497,f496,f508,f507,f506,f557,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f525,f524,f523,f556,f521,f520,f519,f529,f528,f527,f526,f533,f532,f531,f530,f555,f575,f576,f577,f578,f580,f582,f583,f585,f587,f588,f589,f590,f591,f593,f595,f596,f597,f599,f601,f602,f604,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f617,f621,f639,f640,f643,f649,f657,f664,f681,f682,f685,f688,f695,f704,f707,f708,f711,f712,f716,f725,f729,f730,f733,f736,f740,f743,f746,f755,f756,f757,f765,f768,f772,f775,f778,f797,f800,f815,f813,f818,f836,f884,f902,f925,f923,f929,f927,f941,f938,f951,f949,f992,f1003,f1011,f1024,f1032,f1036,f1042,f1051,f1056,f1055,f1068,f1076,f1080,f1113,f1124,f1131,f1140,f1152,f1160,f1164,f1178,f1196,f1197,f1212,f1222,f1237,f1249,f1257,f1261,f1272,f1282,f1298,f1314,f1313,f1311,f1310,f1325,f1334,f1335,f1336,f1337,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1347,f1348,f1349,f1350,f1351,f1352,f1353,f1355,f1356,f1371,f1381,f1397,f1389,f1387,f1386,f1405,f1402,f1439,f1429,f1428,f1484,f1503,f1515,f1514,f1513,f1512,f1517,f1518,f1519,f1529,f1530,f1531,f1532,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1543,f1544,f1545,f1546,f1547,f1548,f1549,f1551,f1552,f1866,f1863,f1867,f1865,f1876,f1873,f1877,f1875,f1887,f1884,f1886,f1900,f1897,f1899,f1916,f1912,f1934,f1930,f1952,f1948,f1977,f1973,f2000,f2015,f2016,f2027,f2038,f2035,f2036,f2037,f2051,f2062,f2059,f2060,f2061,f2082,f2104,f2105,f2106,f2111,f2134,f2135,f2136,f2141,f2142,f2186,f2182,f2209,f2205,f2233,f2234,f2228,f2229,f2230,f2232,f2249,f2250,f2244,f2245,f2246,f2251,f2248,f2286,f2287,f2305,f2306,f2334,f2336,f2347,f2349,f2367,f2364,f2386,f2370,f2371,f2372,f2374,f2375,f2400,f2403,f2404,f2405,f2407,f2408,f2429,f2444,f2427,f2432,f2433,f2434,f2436,f2456,f2460,f2461,f2462,f2464,f2489,f2488,f2486,f2485,f2497,f2506,f2507,f2508,f2509,f2511,f2512,f2513,f2514,f2515,f2516,f2517,f2518,f2519,f2520,f2521,f2522,f2523,f2524,f2525,f2526,f2528,f2529,f2530,f2531,f2532,f2533,f2831,f2830,f2827,f2824,f2821,f2820,f2819,f2818,f2817,f2816,f2815,f2814,f2813,f2812,f2811,f2810,f2809,f2806,f2803,f2802,f2801,f2797,f2796,f2790,f2787,f2786,f2780,f2779,f2772,f2746,f2745,f2744,f2739,f2738,f2733,f2732,f2731,f2726,f2718,f2714,f2710,f2709,f2704,f2703,f2702,f2701,f2700,f2696,f2694,f2688,f2687,f2681,f2680,f2673,f2670,f2669,f2663,f2660,f2659,f2658,f2656,f2654,f2653,f2651,f2650,f2649,f2648,f2646,f2645,f2644,f2643,f2642,f2641,f2635,f2634,f2627,f2626,f2620,f2619,f2612,f2611,f2609,f2607,f2605,f2602,f2601,f2595,f2594,f2589,f2574,f2573,f2572,f2571,f2569,f2568,f2567,f2566,f2565,f3267,f3266,f3263,f3260,f3257,f3256,f3255,f3254,f3253,f3252,f3251,f3250,f3249,f3248,f3247,f3246,f3245,f3242,f3239,f3238,f3237,f3233,f3232,f3226,f3223,f3222,f3216,f3215,f3208,f3182,f3181,f3180,f3175,f3174,f3169,f3168,f3167,f3162,f3154,f3150,f3146,f3145,f3140,f3139,f3138,f3137,f3136,f3132,f3130,f3124,f3123,f3117,f3116,f3109,f3106,f3105,f3099,f3096,f3095,f3093,f3090,f3088,f3087,f3084,f3083,f3082,f3081,f3078,f3077,f3076,f3075,f3074,f3073,f3067,f3066,f3059,f3058,f3052,f3051,f3044,f3043,f3041,f3039,f3037,f3034,f3033,f3027,f3026,f3021,f3006,f3005,f3004,f3003,f3001,f3000,f2999,f2998,f2997,f3620,f3619,f3616,f3613,f3610,f3609,f3608,f3607,f3606,f3605,f3604,f3603,f3602,f3601,f3600,f3599,f3598,f3595,f3592,f3591,f3590,f3586,f3585,f3579,f3576,f3575,f3569,f3568,f3561,f3535,f3534,f3533,f3528,f3527,f3522,f3521,f3520,f3515,f3507,f3503,f3499,f3498,f3493,f3492,f3491,f3490,f3489,f3485,f3483,f3477,f3476,f3470,f3469,f3462,f3459,f3458,f3452,f3449,f3448,f3446,f3443,f3441,f3440,f3437,f3436,f3435,f3434,f3431,f3430,f3429,f3428,f3427,f3426,f3420,f3419,f3412,f3411,f3405,f3404,f3397,f3396,f3395,f3394,f3393,f3391,f3390,f3384,f3383,f3378,f3363,f3362,f3361,f3360,f3358,f3357,f3356,f3354,f3708,f3726,f3770,f3771,f3773,f3788,f3799,f3800,f3802,f3851,f3852,f3853,f3854,f3905,f3907,f3908,f3909,f3910,f3911,f3912,f3913,f3914,f3915,f3916,f3917,f3918,f3919,f3920,f3921,f3922,f3923,f3924,f3925,f3926,f3927,f4039,f4569,f4573,f4576,f4579,f4583,f4586,f4589,f4592,f4595,f4598,f4601,f4604,f4607,f4610,f4613,f4616,f4619,f4622,f4625,f4072,f4073,f4074,f4075,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4646,f4647,f4648,f4649,f4650,f4124,f4125,f4126,f4127,f4651,f4652,f4653,f4654,f4655,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4667,f4668,f4669,f4670,f4671,f4177,f4179,f4180,f4181,f4182,f4183,f4184,f4185,f4186,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4677,f4679,f4681,f4683,f4686,f4688,f4690,f4692,f4694,f4696,f4698,f4700,f4702,f4704,f4706,f4708,f4710,f4712,f4714,f4571,f4941,f4744,f4751,f4966,f4972,f4982,f4986])).
fof(f4983,plain,(
  $false | $spl14),
  inference(global_subsumption,[],[f321,f320,f319,f318,f322,f573,f329,f337,f336,f341,f340,f364,f371,f370,f369,f368,f367,f379,f377,f382,f380,f387,f384,f394,f393,f574,f412,f409,f408,f407,f421,f423,f425,f424,f428,f427,f426,f431,f430,f429,f434,f433,f432,f437,f436,f435,f439,f438,f451,f450,f449,f559,f447,f446,f468,f467,f466,f465,f464,f463,f462,f472,f471,f470,f469,f476,f475,f474,f473,f480,f479,f478,f477,f484,f483,f482,f481,f558,f488,f487,f486,f485,f495,f494,f493,f492,f501,f498,f497,f496,f508,f507,f506,f557,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f525,f524,f523,f556,f521,f520,f519,f529,f528,f527,f526,f533,f532,f531,f530,f555,f575,f576,f577,f578,f580,f582,f583,f585,f587,f588,f589,f590,f591,f593,f595,f596,f597,f599,f601,f602,f604,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f617,f621,f639,f640,f643,f649,f657,f664,f681,f682,f685,f688,f695,f704,f707,f708,f711,f712,f716,f725,f729,f730,f733,f736,f740,f743,f746,f755,f756,f757,f765,f768,f772,f775,f778,f797,f800,f815,f813,f818,f836,f884,f902,f925,f923,f929,f927,f941,f938,f951,f949,f992,f1003,f1011,f1024,f1032,f1036,f1042,f1051,f1056,f1055,f1068,f1076,f1080,f1113,f1124,f1131,f1140,f1152,f1160,f1164,f1178,f1196,f1197,f1212,f1222,f1237,f1249,f1257,f1261,f1272,f1282,f1298,f1314,f1313,f1311,f1310,f1325,f1334,f1335,f1336,f1337,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1347,f1348,f1349,f1350,f1351,f1352,f1353,f1355,f1356,f1371,f1381,f1397,f1389,f1387,f1386,f1405,f1402,f1439,f1429,f1428,f1484,f1503,f1515,f1514,f1513,f1512,f1517,f1518,f1519,f1529,f1530,f1531,f1532,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1543,f1544,f1545,f1546,f1547,f1548,f1549,f1551,f1552,f1866,f1863,f1867,f1865,f1876,f1873,f1877,f1875,f1887,f1884,f1886,f1900,f1897,f1899,f1916,f1912,f1934,f1930,f1952,f1948,f1977,f1973,f2000,f2015,f2016,f2027,f2038,f2035,f2036,f2037,f2051,f2062,f2059,f2060,f2061,f2082,f2104,f2105,f2106,f2111,f2134,f2135,f2136,f2141,f2142,f2186,f2182,f2209,f2205,f2233,f2234,f2228,f2229,f2230,f2232,f2249,f2250,f2244,f2245,f2246,f2251,f2248,f2286,f2287,f2305,f2306,f2334,f2336,f2347,f2349,f2367,f2364,f2386,f2370,f2371,f2372,f2374,f2375,f2400,f2403,f2404,f2405,f2407,f2408,f2429,f2444,f2427,f2432,f2433,f2434,f2436,f2456,f2460,f2461,f2462,f2464,f2489,f2488,f2486,f2485,f2497,f2506,f2507,f2508,f2509,f2511,f2512,f2513,f2514,f2515,f2516,f2517,f2518,f2519,f2520,f2521,f2522,f2523,f2524,f2525,f2526,f2528,f2529,f2530,f2531,f2532,f2533,f2831,f2830,f2827,f2824,f2821,f2820,f2819,f2818,f2817,f2816,f2815,f2814,f2813,f2812,f2811,f2810,f2809,f2806,f2803,f2802,f2801,f2797,f2796,f2790,f2787,f2786,f2780,f2779,f2772,f2746,f2745,f2744,f2739,f2738,f2733,f2732,f2731,f2726,f2718,f2714,f2710,f2709,f2704,f2703,f2702,f2701,f2700,f2696,f2694,f2688,f2687,f2681,f2680,f2673,f2670,f2669,f2663,f2660,f2659,f2658,f2656,f2654,f2653,f2651,f2650,f2649,f2648,f2646,f2645,f2644,f2643,f2642,f2641,f2635,f2634,f2627,f2626,f2620,f2619,f2612,f2611,f2609,f2607,f2605,f2602,f2601,f2595,f2594,f2589,f2574,f2573,f2572,f2571,f2569,f2568,f2567,f2566,f2565,f3267,f3266,f3263,f3260,f3257,f3256,f3255,f3254,f3253,f3252,f3251,f3250,f3249,f3248,f3247,f3246,f3245,f3242,f3239,f3238,f3237,f3233,f3232,f3226,f3223,f3222,f3216,f3215,f3208,f3182,f3181,f3180,f3175,f3174,f3169,f3168,f3167,f3162,f3154,f3150,f3146,f3145,f3140,f3139,f3138,f3137,f3136,f3132,f3130,f3124,f3123,f3117,f3116,f3109,f3106,f3105,f3099,f3096,f3095,f3093,f3090,f3088,f3087,f3084,f3083,f3082,f3081,f3078,f3077,f3076,f3075,f3074,f3073,f3067,f3066,f3059,f3058,f3052,f3051,f3044,f3043,f3041,f3039,f3037,f3034,f3033,f3027,f3026,f3021,f3006,f3005,f3004,f3003,f3001,f3000,f2999,f2998,f2997,f3620,f3619,f3616,f3613,f3610,f3609,f3608,f3607,f3606,f3605,f3604,f3603,f3602,f3601,f3600,f3599,f3598,f3595,f3592,f3591,f3590,f3586,f3585,f3579,f3576,f3575,f3569,f3568,f3561,f3535,f3534,f3533,f3528,f3527,f3522,f3521,f3520,f3515,f3507,f3503,f3499,f3498,f3493,f3492,f3491,f3490,f3489,f3485,f3483,f3477,f3476,f3470,f3469,f3462,f3459,f3458,f3452,f3449,f3448,f3446,f3443,f3441,f3440,f3437,f3436,f3435,f3434,f3431,f3430,f3429,f3428,f3427,f3426,f3420,f3419,f3412,f3411,f3405,f3404,f3397,f3396,f3395,f3394,f3393,f3391,f3390,f3384,f3383,f3378,f3363,f3362,f3361,f3360,f3358,f3357,f3356,f3354,f3708,f3726,f3770,f3771,f3773,f3788,f3799,f3800,f3802,f3851,f3852,f3853,f3854,f3905,f3907,f3908,f3909,f3910,f3911,f3912,f3913,f3914,f3915,f3916,f3917,f3918,f3919,f3920,f3921,f3922,f3923,f3924,f3925,f3926,f3927,f4039,f4569,f4573,f4576,f4579,f4583,f4586,f4589,f4592,f4595,f4598,f4601,f4604,f4607,f4610,f4613,f4616,f4619,f4622,f4625,f4072,f4073,f4074,f4075,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4646,f4647,f4648,f4649,f4650,f4124,f4125,f4126,f4127,f4651,f4652,f4653,f4654,f4655,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4667,f4668,f4669,f4670,f4671,f4177,f4179,f4180,f4181,f4182,f4183,f4184,f4185,f4186,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4677,f4679,f4681,f4683,f4686,f4688,f4690,f4692,f4694,f4696,f4698,f4700,f4702,f4704,f4706,f4708,f4710,f4712,f4714,f4571,f4941,f4744,f4751,f4966,f4972,f4982])).
fof(f4973,plain,(
  $false | $spl14),
  inference(global_subsumption,[],[f321,f320,f319,f318,f322,f573,f329,f337,f336,f341,f340,f364,f371,f370,f369,f368,f367,f379,f377,f382,f380,f387,f384,f394,f393,f574,f412,f409,f408,f407,f421,f423,f425,f424,f428,f427,f426,f431,f430,f429,f434,f433,f432,f437,f436,f435,f439,f438,f451,f450,f449,f559,f447,f446,f468,f467,f466,f465,f464,f463,f462,f472,f471,f470,f469,f476,f475,f474,f473,f480,f479,f478,f477,f484,f483,f482,f481,f558,f488,f487,f486,f485,f495,f494,f493,f492,f501,f498,f497,f496,f508,f507,f506,f557,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f525,f524,f523,f556,f521,f520,f519,f529,f528,f527,f526,f533,f532,f531,f530,f555,f575,f576,f577,f578,f580,f582,f583,f585,f587,f588,f589,f590,f591,f593,f595,f596,f597,f599,f601,f602,f604,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f617,f621,f639,f640,f643,f649,f657,f664,f681,f682,f685,f688,f695,f704,f707,f708,f711,f712,f716,f725,f729,f730,f733,f736,f740,f743,f746,f755,f756,f757,f765,f768,f772,f775,f778,f797,f800,f815,f813,f818,f836,f884,f902,f925,f923,f929,f927,f941,f938,f951,f949,f992,f1003,f1011,f1024,f1032,f1036,f1042,f1051,f1056,f1055,f1068,f1076,f1080,f1113,f1124,f1131,f1140,f1152,f1160,f1164,f1178,f1196,f1197,f1212,f1222,f1237,f1249,f1257,f1261,f1272,f1282,f1298,f1314,f1313,f1311,f1310,f1325,f1334,f1335,f1336,f1337,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1347,f1348,f1349,f1350,f1351,f1352,f1353,f1355,f1356,f1371,f1381,f1397,f1389,f1387,f1386,f1405,f1402,f1439,f1429,f1428,f1484,f1503,f1515,f1514,f1513,f1512,f1517,f1518,f1519,f1529,f1530,f1531,f1532,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1543,f1544,f1545,f1546,f1547,f1548,f1549,f1551,f1552,f1866,f1863,f1867,f1865,f1876,f1873,f1877,f1875,f1887,f1884,f1886,f1900,f1897,f1899,f1916,f1912,f1934,f1930,f1952,f1948,f1977,f1973,f2000,f2015,f2016,f2027,f2038,f2035,f2036,f2037,f2051,f2062,f2059,f2060,f2061,f2082,f2104,f2105,f2106,f2111,f2134,f2135,f2136,f2141,f2142,f2186,f2182,f2209,f2205,f2233,f2234,f2228,f2229,f2230,f2232,f2249,f2250,f2244,f2245,f2246,f2251,f2248,f2286,f2287,f2305,f2306,f2334,f2336,f2347,f2349,f2367,f2364,f2386,f2370,f2371,f2372,f2374,f2375,f2400,f2403,f2404,f2405,f2407,f2408,f2429,f2444,f2427,f2432,f2433,f2434,f2436,f2456,f2460,f2461,f2462,f2464,f2489,f2488,f2486,f2485,f2497,f2506,f2507,f2508,f2509,f2511,f2512,f2513,f2514,f2515,f2516,f2517,f2518,f2519,f2520,f2521,f2522,f2523,f2524,f2525,f2526,f2528,f2529,f2530,f2531,f2532,f2533,f2831,f2830,f2827,f2824,f2821,f2820,f2819,f2818,f2817,f2816,f2815,f2814,f2813,f2812,f2811,f2810,f2809,f2806,f2803,f2802,f2801,f2797,f2796,f2790,f2787,f2786,f2780,f2779,f2772,f2746,f2745,f2744,f2739,f2738,f2733,f2732,f2731,f2726,f2718,f2714,f2710,f2709,f2704,f2703,f2702,f2701,f2700,f2696,f2694,f2688,f2687,f2681,f2680,f2673,f2670,f2669,f2663,f2660,f2659,f2658,f2656,f2654,f2653,f2651,f2650,f2649,f2648,f2646,f2645,f2644,f2643,f2642,f2641,f2635,f2634,f2627,f2626,f2620,f2619,f2612,f2611,f2609,f2607,f2605,f2602,f2601,f2595,f2594,f2589,f2574,f2573,f2572,f2571,f2569,f2568,f2567,f2566,f2565,f3267,f3266,f3263,f3260,f3257,f3256,f3255,f3254,f3253,f3252,f3251,f3250,f3249,f3248,f3247,f3246,f3245,f3242,f3239,f3238,f3237,f3233,f3232,f3226,f3223,f3222,f3216,f3215,f3208,f3182,f3181,f3180,f3175,f3174,f3169,f3168,f3167,f3162,f3154,f3150,f3146,f3145,f3140,f3139,f3138,f3137,f3136,f3132,f3130,f3124,f3123,f3117,f3116,f3109,f3106,f3105,f3099,f3096,f3095,f3093,f3090,f3088,f3087,f3084,f3083,f3082,f3081,f3078,f3077,f3076,f3075,f3074,f3073,f3067,f3066,f3059,f3058,f3052,f3051,f3044,f3043,f3041,f3039,f3037,f3034,f3033,f3027,f3026,f3021,f3006,f3005,f3004,f3003,f3001,f3000,f2999,f2998,f2997,f3620,f3619,f3616,f3613,f3610,f3609,f3608,f3607,f3606,f3605,f3604,f3603,f3602,f3601,f3600,f3599,f3598,f3595,f3592,f3591,f3590,f3586,f3585,f3579,f3576,f3575,f3569,f3568,f3561,f3535,f3534,f3533,f3528,f3527,f3522,f3521,f3520,f3515,f3507,f3503,f3499,f3498,f3493,f3492,f3491,f3490,f3489,f3485,f3483,f3477,f3476,f3470,f3469,f3462,f3459,f3458,f3452,f3449,f3448,f3446,f3443,f3441,f3440,f3437,f3436,f3435,f3434,f3431,f3430,f3429,f3428,f3427,f3426,f3420,f3419,f3412,f3411,f3405,f3404,f3397,f3396,f3395,f3394,f3393,f3391,f3390,f3384,f3383,f3378,f3363,f3362,f3361,f3360,f3358,f3357,f3356,f3354,f3708,f3726,f3770,f3771,f3773,f3788,f3799,f3800,f3802,f3851,f3852,f3853,f3854,f3905,f3907,f3908,f3909,f3910,f3911,f3912,f3913,f3914,f3915,f3916,f3917,f3918,f3919,f3920,f3921,f3922,f3923,f3924,f3925,f3926,f3927,f4039,f4569,f4573,f4576,f4579,f4583,f4586,f4589,f4592,f4595,f4598,f4601,f4604,f4607,f4610,f4613,f4616,f4619,f4622,f4625,f4072,f4073,f4074,f4075,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4646,f4647,f4648,f4649,f4650,f4124,f4125,f4126,f4127,f4651,f4652,f4653,f4654,f4655,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4667,f4668,f4669,f4670,f4671,f4177,f4179,f4180,f4181,f4182,f4183,f4184,f4185,f4186,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4677,f4679,f4681,f4683,f4686,f4688,f4690,f4692,f4694,f4696,f4698,f4700,f4702,f4704,f4706,f4708,f4710,f4712,f4714,f4571,f4941,f4744,f4751,f4966,f4972])).
fof(f4967,plain,(
  $false | $spl14),
  inference(global_subsumption,[],[f321,f320,f319,f318,f322,f573,f329,f337,f336,f341,f340,f364,f371,f370,f369,f368,f367,f379,f377,f382,f380,f387,f384,f394,f393,f574,f412,f409,f408,f407,f421,f423,f425,f424,f428,f427,f426,f431,f430,f429,f434,f433,f432,f437,f436,f435,f439,f438,f451,f450,f449,f559,f447,f446,f468,f467,f466,f465,f464,f463,f462,f472,f471,f470,f469,f476,f475,f474,f473,f480,f479,f478,f477,f484,f483,f482,f481,f558,f488,f487,f486,f485,f495,f494,f493,f492,f501,f498,f497,f496,f508,f507,f506,f557,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f525,f524,f523,f556,f521,f520,f519,f529,f528,f527,f526,f533,f532,f531,f530,f555,f575,f576,f577,f578,f580,f582,f583,f585,f587,f588,f589,f590,f591,f593,f595,f596,f597,f599,f601,f602,f604,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f617,f621,f639,f640,f643,f649,f657,f664,f681,f682,f685,f688,f695,f704,f707,f708,f711,f712,f716,f725,f729,f730,f733,f736,f740,f743,f746,f755,f756,f757,f765,f768,f772,f775,f778,f797,f800,f815,f813,f818,f836,f884,f902,f925,f923,f929,f927,f941,f938,f951,f949,f992,f1003,f1011,f1024,f1032,f1036,f1042,f1051,f1056,f1055,f1068,f1076,f1080,f1113,f1124,f1131,f1140,f1152,f1160,f1164,f1178,f1196,f1197,f1212,f1222,f1237,f1249,f1257,f1261,f1272,f1282,f1298,f1314,f1313,f1311,f1310,f1325,f1334,f1335,f1336,f1337,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1347,f1348,f1349,f1350,f1351,f1352,f1353,f1355,f1356,f1371,f1381,f1397,f1389,f1387,f1386,f1405,f1402,f1439,f1429,f1428,f1484,f1503,f1515,f1514,f1513,f1512,f1517,f1518,f1519,f1529,f1530,f1531,f1532,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1543,f1544,f1545,f1546,f1547,f1548,f1549,f1551,f1552,f1866,f1863,f1867,f1865,f1876,f1873,f1877,f1875,f1887,f1884,f1886,f1900,f1897,f1899,f1916,f1912,f1934,f1930,f1952,f1948,f1977,f1973,f2000,f2015,f2016,f2027,f2038,f2035,f2036,f2037,f2051,f2062,f2059,f2060,f2061,f2082,f2104,f2105,f2106,f2111,f2134,f2135,f2136,f2141,f2142,f2186,f2182,f2209,f2205,f2233,f2234,f2228,f2229,f2230,f2232,f2249,f2250,f2244,f2245,f2246,f2251,f2248,f2286,f2287,f2305,f2306,f2334,f2336,f2347,f2349,f2367,f2364,f2386,f2370,f2371,f2372,f2374,f2375,f2400,f2403,f2404,f2405,f2407,f2408,f2429,f2444,f2427,f2432,f2433,f2434,f2436,f2456,f2460,f2461,f2462,f2464,f2489,f2488,f2486,f2485,f2497,f2506,f2507,f2508,f2509,f2511,f2512,f2513,f2514,f2515,f2516,f2517,f2518,f2519,f2520,f2521,f2522,f2523,f2524,f2525,f2526,f2528,f2529,f2530,f2531,f2532,f2533,f2831,f2830,f2827,f2824,f2821,f2820,f2819,f2818,f2817,f2816,f2815,f2814,f2813,f2812,f2811,f2810,f2809,f2806,f2803,f2802,f2801,f2797,f2796,f2790,f2787,f2786,f2780,f2779,f2772,f2746,f2745,f2744,f2739,f2738,f2733,f2732,f2731,f2726,f2718,f2714,f2710,f2709,f2704,f2703,f2702,f2701,f2700,f2696,f2694,f2688,f2687,f2681,f2680,f2673,f2670,f2669,f2663,f2660,f2659,f2658,f2656,f2654,f2653,f2651,f2650,f2649,f2648,f2646,f2645,f2644,f2643,f2642,f2641,f2635,f2634,f2627,f2626,f2620,f2619,f2612,f2611,f2609,f2607,f2605,f2602,f2601,f2595,f2594,f2589,f2574,f2573,f2572,f2571,f2569,f2568,f2567,f2566,f2565,f3267,f3266,f3263,f3260,f3257,f3256,f3255,f3254,f3253,f3252,f3251,f3250,f3249,f3248,f3247,f3246,f3245,f3242,f3239,f3238,f3237,f3233,f3232,f3226,f3223,f3222,f3216,f3215,f3208,f3182,f3181,f3180,f3175,f3174,f3169,f3168,f3167,f3162,f3154,f3150,f3146,f3145,f3140,f3139,f3138,f3137,f3136,f3132,f3130,f3124,f3123,f3117,f3116,f3109,f3106,f3105,f3099,f3096,f3095,f3093,f3090,f3088,f3087,f3084,f3083,f3082,f3081,f3078,f3077,f3076,f3075,f3074,f3073,f3067,f3066,f3059,f3058,f3052,f3051,f3044,f3043,f3041,f3039,f3037,f3034,f3033,f3027,f3026,f3021,f3006,f3005,f3004,f3003,f3001,f3000,f2999,f2998,f2997,f3620,f3619,f3616,f3613,f3610,f3609,f3608,f3607,f3606,f3605,f3604,f3603,f3602,f3601,f3600,f3599,f3598,f3595,f3592,f3591,f3590,f3586,f3585,f3579,f3576,f3575,f3569,f3568,f3561,f3535,f3534,f3533,f3528,f3527,f3522,f3521,f3520,f3515,f3507,f3503,f3499,f3498,f3493,f3492,f3491,f3490,f3489,f3485,f3483,f3477,f3476,f3470,f3469,f3462,f3459,f3458,f3452,f3449,f3448,f3446,f3443,f3441,f3440,f3437,f3436,f3435,f3434,f3431,f3430,f3429,f3428,f3427,f3426,f3420,f3419,f3412,f3411,f3405,f3404,f3397,f3396,f3395,f3394,f3393,f3391,f3390,f3384,f3383,f3378,f3363,f3362,f3361,f3360,f3358,f3357,f3356,f3354,f3708,f3726,f3770,f3771,f3773,f3788,f3799,f3800,f3802,f3851,f3852,f3853,f3854,f3905,f3907,f3908,f3909,f3910,f3911,f3912,f3913,f3914,f3915,f3916,f3917,f3918,f3919,f3920,f3921,f3922,f3923,f3924,f3925,f3926,f3927,f4039,f4569,f4573,f4576,f4579,f4583,f4586,f4589,f4592,f4595,f4598,f4601,f4604,f4607,f4610,f4613,f4616,f4619,f4622,f4625,f4072,f4073,f4074,f4075,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4646,f4647,f4648,f4649,f4650,f4124,f4125,f4126,f4127,f4651,f4652,f4653,f4654,f4655,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4667,f4668,f4669,f4670,f4671,f4177,f4179,f4180,f4181,f4182,f4183,f4184,f4185,f4186,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4677,f4679,f4681,f4683,f4686,f4688,f4690,f4692,f4694,f4696,f4698,f4700,f4702,f4704,f4706,f4708,f4710,f4712,f4714,f4571,f4941,f4744,f4751,f4966])).
fof(f4957,plain,(
  $false | ($spl11 | $spl14)),
  inference(global_subsumption,[],[f4762,f321,f320,f319,f318,f322,f573,f329,f337,f336,f341,f340,f364,f371,f370,f369,f368,f367,f379,f377,f382,f380,f387,f384,f394,f393,f574,f412,f409,f408,f407,f421,f423,f425,f424,f428,f427,f426,f431,f430,f429,f434,f433,f432,f437,f436,f435,f439,f438,f451,f450,f449,f559,f447,f446,f468,f467,f466,f465,f464,f463,f462,f472,f471,f470,f469,f476,f475,f474,f473,f480,f479,f478,f477,f484,f483,f482,f481,f558,f488,f487,f486,f485,f495,f494,f493,f492,f501,f498,f497,f496,f508,f507,f506,f557,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f525,f524,f523,f556,f521,f520,f519,f529,f528,f527,f526,f533,f532,f531,f530,f555,f575,f576,f577,f578,f580,f582,f583,f585,f587,f588,f589,f590,f591,f593,f595,f596,f597,f599,f601,f602,f604,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f617,f621,f639,f640,f643,f649,f657,f664,f681,f682,f685,f688,f695,f704,f707,f708,f711,f712,f716,f725,f729,f730,f733,f736,f740,f743,f746,f755,f756,f757,f765,f768,f772,f775,f778,f797,f800,f815,f813,f818,f836,f884,f902,f925,f923,f929,f927,f941,f938,f951,f949,f992,f1003,f1011,f1024,f1032,f1036,f1042,f1051,f1056,f1055,f1068,f1076,f1080,f1113,f1124,f1131,f1140,f1152,f1160,f1164,f1178,f1196,f1197,f1212,f1222,f1237,f1249,f1257,f1261,f1272,f1282,f1298,f1314,f1313,f1311,f1310,f1325,f1334,f1335,f1336,f1337,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1347,f1348,f1349,f1350,f1351,f1352,f1353,f1355,f1356,f1371,f1381,f1397,f1389,f1387,f1386,f1405,f1402,f1439,f1429,f1428,f1484,f1503,f1515,f1514,f1513,f1512,f1517,f1518,f1519,f1529,f1530,f1531,f1532,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1543,f1544,f1545,f1546,f1547,f1548,f1549,f1551,f1552,f1866,f1863,f1867,f1865,f1876,f1873,f1877,f1875,f1887,f1884,f1886,f1900,f1897,f1899,f1916,f1912,f1934,f1930,f1952,f1948,f1977,f1973,f2000,f2015,f2016,f2027,f2038,f2035,f2036,f2037,f2051,f2062,f2059,f2060,f2061,f2082,f2104,f2105,f2106,f2111,f2134,f2135,f2136,f2141,f2142,f2186,f2182,f2209,f2205,f2233,f2234,f2228,f2229,f2230,f2232,f2249,f2250,f2244,f2245,f2246,f2251,f2248,f2286,f2287,f2305,f2306,f2334,f2336,f2347,f2349,f2367,f2364,f2386,f2370,f2371,f2372,f2374,f2375,f2400,f2403,f2404,f2405,f2407,f2408,f2429,f2444,f2427,f2432,f2433,f2434,f2436,f2456,f2460,f2461,f2462,f2464,f2489,f2488,f2486,f2485,f2497,f2506,f2507,f2508,f2509,f2511,f2512,f2513,f2514,f2515,f2516,f2517,f2518,f2519,f2520,f2521,f2522,f2523,f2524,f2525,f2526,f2528,f2529,f2530,f2531,f2532,f2533,f2831,f2830,f2827,f2824,f2821,f2820,f2819,f2818,f2817,f2816,f2815,f2814,f2813,f2812,f2811,f2810,f2809,f2806,f2803,f2802,f2801,f2797,f2796,f2790,f2787,f2786,f2780,f2779,f2772,f2746,f2745,f2744,f2739,f2738,f2733,f2732,f2731,f2726,f2718,f2714,f2710,f2709,f2704,f2703,f2702,f2701,f2700,f2696,f2694,f2688,f2687,f2681,f2680,f2673,f2670,f2669,f2663,f2660,f2659,f2658,f2656,f2654,f2653,f2651,f2650,f2649,f2648,f2646,f2645,f2644,f2643,f2642,f2641,f2635,f2634,f2627,f2626,f2620,f2619,f2612,f2611,f2609,f2607,f2605,f2602,f2601,f2595,f2594,f2589,f2574,f2573,f2572,f2571,f2569,f2568,f2567,f2566,f2565,f3267,f3266,f3263,f3260,f3257,f3256,f3255,f3254,f3253,f3252,f3251,f3250,f3249,f3248,f3247,f3246,f3245,f3242,f3239,f3238,f3237,f3233,f3232,f3226,f3223,f3222,f3216,f3215,f3208,f3182,f3181,f3180,f3175,f3174,f3169,f3168,f3167,f3162,f3154,f3150,f3146,f3145,f3140,f3139,f3138,f3137,f3136,f3132,f3130,f3124,f3123,f3117,f3116,f3109,f3106,f3105,f3099,f3096,f3095,f3093,f3090,f3088,f3087,f3084,f3083,f3082,f3081,f3078,f3077,f3076,f3075,f3074,f3073,f3067,f3066,f3059,f3058,f3052,f3051,f3044,f3043,f3041,f3039,f3037,f3034,f3033,f3027,f3026,f3021,f3006,f3005,f3004,f3003,f3001,f3000,f2999,f2998,f2997,f3620,f3619,f3616,f3613,f3610,f3609,f3608,f3607,f3606,f3605,f3604,f3603,f3602,f3601,f3600,f3599,f3598,f3595,f3592,f3591,f3590,f3586,f3585,f3579,f3576,f3575,f3569,f3568,f3561,f3535,f3534,f3533,f3528,f3527,f3522,f3521,f3520,f3515,f3507,f3503,f3499,f3498,f3493,f3492,f3491,f3490,f3489,f3485,f3483,f3477,f3476,f3470,f3469,f3462,f3459,f3458,f3452,f3449,f3448,f3446,f3443,f3441,f3440,f3437,f3436,f3435,f3434,f3431,f3430,f3429,f3428,f3427,f3426,f3420,f3419,f3412,f3411,f3405,f3404,f3397,f3396,f3395,f3394,f3393,f3391,f3390,f3384,f3383,f3378,f3363,f3362,f3361,f3360,f3358,f3357,f3356,f3354,f3708,f3726,f3770,f3771,f3773,f3788,f3799,f3800,f3802,f3851,f3852,f3853,f3854,f3905,f3907,f3908,f3909,f3910,f3911,f3912,f3913,f3914,f3915,f3916,f3917,f3918,f3919,f3920,f3921,f3922,f3923,f3924,f3925,f3926,f3927,f4039,f4569,f4573,f4576,f4579,f4583,f4586,f4589,f4592,f4595,f4598,f4601,f4604,f4607,f4610,f4613,f4616,f4619,f4622,f4625,f4072,f4073,f4074,f4075,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4646,f4647,f4648,f4649,f4650,f4124,f4125,f4126,f4127,f4651,f4652,f4653,f4654,f4655,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4667,f4668,f4669,f4670,f4671,f4177,f4179,f4180,f4181,f4182,f4183,f4184,f4185,f4186,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4677,f4679,f4681,f4683,f4686,f4688,f4690,f4692,f4694,f4696,f4698,f4700,f4702,f4704,f4706,f4708,f4710,f4712,f4714,f4571,f4941,f4744,f4751])).
fof(f4762,plain,(
  ( ! [X52,X53] : (sP6('T_0','|dollar|exists_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X52,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X52,X53)))) ) | ($spl11 | $spl14)),
  inference(backward_demodulation,[],[f4571,f2956])).
fof(f2956,plain,(
  ( ! [X52,X53] : (sP6('T_0','|dollar|exists_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X52,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X52,X53)))) ) | $spl11),
  inference(superposition,[],[f1314,f1572])).
fof(f4956,plain,(
  $false | ($spl11 | $spl14)),
  inference(global_subsumption,[],[f4761,f321,f320,f319,f318,f322,f573,f329,f337,f336,f341,f340,f364,f371,f370,f369,f368,f367,f379,f377,f382,f380,f387,f384,f394,f393,f574,f412,f409,f408,f407,f421,f423,f425,f424,f428,f427,f426,f431,f430,f429,f434,f433,f432,f437,f436,f435,f439,f438,f451,f450,f449,f559,f447,f446,f468,f467,f466,f465,f464,f463,f462,f472,f471,f470,f469,f476,f475,f474,f473,f480,f479,f478,f477,f484,f483,f482,f481,f558,f488,f487,f486,f485,f495,f494,f493,f492,f501,f498,f497,f496,f508,f507,f506,f557,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f525,f524,f523,f556,f521,f520,f519,f529,f528,f527,f526,f533,f532,f531,f530,f555,f575,f576,f577,f578,f580,f582,f583,f585,f587,f588,f589,f590,f591,f593,f595,f596,f597,f599,f601,f602,f604,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f617,f621,f639,f640,f643,f649,f657,f664,f681,f682,f685,f688,f695,f704,f707,f708,f711,f712,f716,f725,f729,f730,f733,f736,f740,f743,f746,f755,f756,f757,f765,f768,f772,f775,f778,f797,f800,f815,f813,f818,f836,f884,f902,f925,f923,f929,f927,f941,f938,f951,f949,f992,f1003,f1011,f1024,f1032,f1036,f1042,f1051,f1056,f1055,f1068,f1076,f1080,f1113,f1124,f1131,f1140,f1152,f1160,f1164,f1178,f1196,f1197,f1212,f1222,f1237,f1249,f1257,f1261,f1272,f1282,f1298,f1314,f1313,f1311,f1310,f1325,f1334,f1335,f1336,f1337,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1347,f1348,f1349,f1350,f1351,f1352,f1353,f1355,f1356,f1371,f1381,f1397,f1389,f1387,f1386,f1405,f1402,f1439,f1429,f1428,f1484,f1503,f1515,f1514,f1513,f1512,f1517,f1518,f1519,f1529,f1530,f1531,f1532,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1543,f1544,f1545,f1546,f1547,f1548,f1549,f1551,f1552,f1866,f1863,f1867,f1865,f1876,f1873,f1877,f1875,f1887,f1884,f1886,f1900,f1897,f1899,f1916,f1912,f1934,f1930,f1952,f1948,f1977,f1973,f2000,f2015,f2016,f2027,f2038,f2035,f2036,f2037,f2051,f2062,f2059,f2060,f2061,f2082,f2104,f2105,f2106,f2111,f2134,f2135,f2136,f2141,f2142,f2186,f2182,f2209,f2205,f2233,f2234,f2228,f2229,f2230,f2232,f2249,f2250,f2244,f2245,f2246,f2251,f2248,f2286,f2287,f2305,f2306,f2334,f2336,f2347,f2349,f2367,f2364,f2386,f2370,f2371,f2372,f2374,f2375,f2400,f2403,f2404,f2405,f2407,f2408,f2429,f2444,f2427,f2432,f2433,f2434,f2436,f2456,f2460,f2461,f2462,f2464,f2489,f2488,f2486,f2485,f2497,f2506,f2507,f2508,f2509,f2511,f2512,f2513,f2514,f2515,f2516,f2517,f2518,f2519,f2520,f2521,f2522,f2523,f2524,f2525,f2526,f2528,f2529,f2530,f2531,f2532,f2533,f2831,f2830,f2827,f2824,f2821,f2820,f2819,f2818,f2817,f2816,f2815,f2814,f2813,f2812,f2811,f2810,f2809,f2806,f2803,f2802,f2801,f2797,f2796,f2790,f2787,f2786,f2780,f2779,f2772,f2746,f2745,f2744,f2739,f2738,f2733,f2732,f2731,f2726,f2718,f2714,f2710,f2709,f2704,f2703,f2702,f2701,f2700,f2696,f2694,f2688,f2687,f2681,f2680,f2673,f2670,f2669,f2663,f2660,f2659,f2658,f2656,f2654,f2653,f2651,f2650,f2649,f2648,f2646,f2645,f2644,f2643,f2642,f2641,f2635,f2634,f2627,f2626,f2620,f2619,f2612,f2611,f2609,f2607,f2605,f2602,f2601,f2595,f2594,f2589,f2574,f2573,f2572,f2571,f2569,f2568,f2567,f2566,f2565,f3267,f3266,f3263,f3260,f3257,f3256,f3255,f3254,f3253,f3252,f3251,f3250,f3249,f3248,f3247,f3246,f3245,f3242,f3239,f3238,f3237,f3233,f3232,f3226,f3223,f3222,f3216,f3215,f3208,f3182,f3181,f3180,f3175,f3174,f3169,f3168,f3167,f3162,f3154,f3150,f3146,f3145,f3140,f3139,f3138,f3137,f3136,f3132,f3130,f3124,f3123,f3117,f3116,f3109,f3106,f3105,f3099,f3096,f3095,f3093,f3090,f3088,f3087,f3084,f3083,f3082,f3081,f3078,f3077,f3076,f3075,f3074,f3073,f3067,f3066,f3059,f3058,f3052,f3051,f3044,f3043,f3041,f3039,f3037,f3034,f3033,f3027,f3026,f3021,f3006,f3005,f3004,f3003,f3001,f3000,f2999,f2998,f2997,f3620,f3619,f3616,f3613,f3610,f3609,f3608,f3607,f3606,f3605,f3604,f3603,f3602,f3601,f3600,f3599,f3598,f3595,f3592,f3591,f3590,f3586,f3585,f3579,f3576,f3575,f3569,f3568,f3561,f3535,f3534,f3533,f3528,f3527,f3522,f3521,f3520,f3515,f3507,f3503,f3499,f3498,f3493,f3492,f3491,f3490,f3489,f3485,f3483,f3477,f3476,f3470,f3469,f3462,f3459,f3458,f3452,f3449,f3448,f3446,f3443,f3441,f3440,f3437,f3436,f3435,f3434,f3431,f3430,f3429,f3428,f3427,f3426,f3420,f3419,f3412,f3411,f3405,f3404,f3397,f3396,f3395,f3394,f3393,f3391,f3390,f3384,f3383,f3378,f3363,f3362,f3361,f3360,f3358,f3357,f3356,f3354,f3708,f3726,f3770,f3771,f3773,f3788,f3799,f3800,f3802,f3851,f3852,f3853,f3854,f3905,f3907,f3908,f3909,f3910,f3911,f3912,f3913,f3914,f3915,f3916,f3917,f3918,f3919,f3920,f3921,f3922,f3923,f3924,f3925,f3926,f3927,f4039,f4569,f4573,f4576,f4579,f4583,f4586,f4589,f4592,f4595,f4598,f4601,f4604,f4607,f4610,f4613,f4616,f4619,f4622,f4625,f4072,f4073,f4074,f4075,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4646,f4647,f4648,f4649,f4650,f4124,f4125,f4126,f4127,f4651,f4652,f4653,f4654,f4655,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4667,f4668,f4669,f4670,f4671,f4177,f4179,f4180,f4181,f4182,f4183,f4184,f4185,f4186,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4677,f4679,f4681,f4683,f4686,f4688,f4690,f4692,f4694,f4696,f4698,f4700,f4702,f4704,f4706,f4708,f4710,f4712,f4714,f4571,f4941,f4744,f4751])).
fof(f4761,plain,(
  ( ! [X50,X51] : (~sP6('F_0','|dollar|exists_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X50,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X50,X51)))) ) | ($spl11 | $spl14)),
  inference(backward_demodulation,[],[f4571,f2955])).
fof(f2955,plain,(
  ( ! [X50,X51] : (~sP6('F_0','|dollar|exists_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X50,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X50,X51)))) ) | $spl11),
  inference(superposition,[],[f1313,f1572])).
fof(f4955,plain,(
  $false | ($spl11 | $spl14)),
  inference(global_subsumption,[],[f4760,f321,f320,f319,f318,f322,f573,f329,f337,f336,f341,f340,f364,f371,f370,f369,f368,f367,f379,f377,f382,f380,f387,f384,f394,f393,f574,f412,f409,f408,f407,f421,f423,f425,f424,f428,f427,f426,f431,f430,f429,f434,f433,f432,f437,f436,f435,f439,f438,f451,f450,f449,f559,f447,f446,f468,f467,f466,f465,f464,f463,f462,f472,f471,f470,f469,f476,f475,f474,f473,f480,f479,f478,f477,f484,f483,f482,f481,f558,f488,f487,f486,f485,f495,f494,f493,f492,f501,f498,f497,f496,f508,f507,f506,f557,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f525,f524,f523,f556,f521,f520,f519,f529,f528,f527,f526,f533,f532,f531,f530,f555,f575,f576,f577,f578,f580,f582,f583,f585,f587,f588,f589,f590,f591,f593,f595,f596,f597,f599,f601,f602,f604,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f617,f621,f639,f640,f643,f649,f657,f664,f681,f682,f685,f688,f695,f704,f707,f708,f711,f712,f716,f725,f729,f730,f733,f736,f740,f743,f746,f755,f756,f757,f765,f768,f772,f775,f778,f797,f800,f815,f813,f818,f836,f884,f902,f925,f923,f929,f927,f941,f938,f951,f949,f992,f1003,f1011,f1024,f1032,f1036,f1042,f1051,f1056,f1055,f1068,f1076,f1080,f1113,f1124,f1131,f1140,f1152,f1160,f1164,f1178,f1196,f1197,f1212,f1222,f1237,f1249,f1257,f1261,f1272,f1282,f1298,f1314,f1313,f1311,f1310,f1325,f1334,f1335,f1336,f1337,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1347,f1348,f1349,f1350,f1351,f1352,f1353,f1355,f1356,f1371,f1381,f1397,f1389,f1387,f1386,f1405,f1402,f1439,f1429,f1428,f1484,f1503,f1515,f1514,f1513,f1512,f1517,f1518,f1519,f1529,f1530,f1531,f1532,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1543,f1544,f1545,f1546,f1547,f1548,f1549,f1551,f1552,f1866,f1863,f1867,f1865,f1876,f1873,f1877,f1875,f1887,f1884,f1886,f1900,f1897,f1899,f1916,f1912,f1934,f1930,f1952,f1948,f1977,f1973,f2000,f2015,f2016,f2027,f2038,f2035,f2036,f2037,f2051,f2062,f2059,f2060,f2061,f2082,f2104,f2105,f2106,f2111,f2134,f2135,f2136,f2141,f2142,f2186,f2182,f2209,f2205,f2233,f2234,f2228,f2229,f2230,f2232,f2249,f2250,f2244,f2245,f2246,f2251,f2248,f2286,f2287,f2305,f2306,f2334,f2336,f2347,f2349,f2367,f2364,f2386,f2370,f2371,f2372,f2374,f2375,f2400,f2403,f2404,f2405,f2407,f2408,f2429,f2444,f2427,f2432,f2433,f2434,f2436,f2456,f2460,f2461,f2462,f2464,f2489,f2488,f2486,f2485,f2497,f2506,f2507,f2508,f2509,f2511,f2512,f2513,f2514,f2515,f2516,f2517,f2518,f2519,f2520,f2521,f2522,f2523,f2524,f2525,f2526,f2528,f2529,f2530,f2531,f2532,f2533,f2831,f2830,f2827,f2824,f2821,f2820,f2819,f2818,f2817,f2816,f2815,f2814,f2813,f2812,f2811,f2810,f2809,f2806,f2803,f2802,f2801,f2797,f2796,f2790,f2787,f2786,f2780,f2779,f2772,f2746,f2745,f2744,f2739,f2738,f2733,f2732,f2731,f2726,f2718,f2714,f2710,f2709,f2704,f2703,f2702,f2701,f2700,f2696,f2694,f2688,f2687,f2681,f2680,f2673,f2670,f2669,f2663,f2660,f2659,f2658,f2656,f2654,f2653,f2651,f2650,f2649,f2648,f2646,f2645,f2644,f2643,f2642,f2641,f2635,f2634,f2627,f2626,f2620,f2619,f2612,f2611,f2609,f2607,f2605,f2602,f2601,f2595,f2594,f2589,f2574,f2573,f2572,f2571,f2569,f2568,f2567,f2566,f2565,f3267,f3266,f3263,f3260,f3257,f3256,f3255,f3254,f3253,f3252,f3251,f3250,f3249,f3248,f3247,f3246,f3245,f3242,f3239,f3238,f3237,f3233,f3232,f3226,f3223,f3222,f3216,f3215,f3208,f3182,f3181,f3180,f3175,f3174,f3169,f3168,f3167,f3162,f3154,f3150,f3146,f3145,f3140,f3139,f3138,f3137,f3136,f3132,f3130,f3124,f3123,f3117,f3116,f3109,f3106,f3105,f3099,f3096,f3095,f3093,f3090,f3088,f3087,f3084,f3083,f3082,f3081,f3078,f3077,f3076,f3075,f3074,f3073,f3067,f3066,f3059,f3058,f3052,f3051,f3044,f3043,f3041,f3039,f3037,f3034,f3033,f3027,f3026,f3021,f3006,f3005,f3004,f3003,f3001,f3000,f2999,f2998,f2997,f3620,f3619,f3616,f3613,f3610,f3609,f3608,f3607,f3606,f3605,f3604,f3603,f3602,f3601,f3600,f3599,f3598,f3595,f3592,f3591,f3590,f3586,f3585,f3579,f3576,f3575,f3569,f3568,f3561,f3535,f3534,f3533,f3528,f3527,f3522,f3521,f3520,f3515,f3507,f3503,f3499,f3498,f3493,f3492,f3491,f3490,f3489,f3485,f3483,f3477,f3476,f3470,f3469,f3462,f3459,f3458,f3452,f3449,f3448,f3446,f3443,f3441,f3440,f3437,f3436,f3435,f3434,f3431,f3430,f3429,f3428,f3427,f3426,f3420,f3419,f3412,f3411,f3405,f3404,f3397,f3396,f3395,f3394,f3393,f3391,f3390,f3384,f3383,f3378,f3363,f3362,f3361,f3360,f3358,f3357,f3356,f3354,f3708,f3726,f3770,f3771,f3773,f3788,f3799,f3800,f3802,f3851,f3852,f3853,f3854,f3905,f3907,f3908,f3909,f3910,f3911,f3912,f3913,f3914,f3915,f3916,f3917,f3918,f3919,f3920,f3921,f3922,f3923,f3924,f3925,f3926,f3927,f4039,f4569,f4573,f4576,f4579,f4583,f4586,f4589,f4592,f4595,f4598,f4601,f4604,f4607,f4610,f4613,f4616,f4619,f4622,f4625,f4072,f4073,f4074,f4075,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4646,f4647,f4648,f4649,f4650,f4124,f4125,f4126,f4127,f4651,f4652,f4653,f4654,f4655,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4667,f4668,f4669,f4670,f4671,f4177,f4179,f4180,f4181,f4182,f4183,f4184,f4185,f4186,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4677,f4679,f4681,f4683,f4686,f4688,f4690,f4692,f4694,f4696,f4698,f4700,f4702,f4704,f4706,f4708,f4710,f4712,f4714,f4571,f4941,f4744,f4751])).
fof(f4760,plain,(
  ( ! [X48,X49] : (sP6('T_0',happ(s(fun(X48,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X48,X49)),'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) ) | ($spl11 | $spl14)),
  inference(backward_demodulation,[],[f4571,f2954])).
fof(f2954,plain,(
  ( ! [X48,X49] : (sP6('T_0',happ(s(fun(X48,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X48,X49)),'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) ) | $spl11),
  inference(superposition,[],[f1311,f1572])).
fof(f4954,plain,(
  $false | ($spl11 | $spl14)),
  inference(global_subsumption,[],[f4759,f321,f320,f319,f318,f322,f573,f329,f337,f336,f341,f340,f364,f371,f370,f369,f368,f367,f379,f377,f382,f380,f387,f384,f394,f393,f574,f412,f409,f408,f407,f421,f423,f425,f424,f428,f427,f426,f431,f430,f429,f434,f433,f432,f437,f436,f435,f439,f438,f451,f450,f449,f559,f447,f446,f468,f467,f466,f465,f464,f463,f462,f472,f471,f470,f469,f476,f475,f474,f473,f480,f479,f478,f477,f484,f483,f482,f481,f558,f488,f487,f486,f485,f495,f494,f493,f492,f501,f498,f497,f496,f508,f507,f506,f557,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f525,f524,f523,f556,f521,f520,f519,f529,f528,f527,f526,f533,f532,f531,f530,f555,f575,f576,f577,f578,f580,f582,f583,f585,f587,f588,f589,f590,f591,f593,f595,f596,f597,f599,f601,f602,f604,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f617,f621,f639,f640,f643,f649,f657,f664,f681,f682,f685,f688,f695,f704,f707,f708,f711,f712,f716,f725,f729,f730,f733,f736,f740,f743,f746,f755,f756,f757,f765,f768,f772,f775,f778,f797,f800,f815,f813,f818,f836,f884,f902,f925,f923,f929,f927,f941,f938,f951,f949,f992,f1003,f1011,f1024,f1032,f1036,f1042,f1051,f1056,f1055,f1068,f1076,f1080,f1113,f1124,f1131,f1140,f1152,f1160,f1164,f1178,f1196,f1197,f1212,f1222,f1237,f1249,f1257,f1261,f1272,f1282,f1298,f1314,f1313,f1311,f1310,f1325,f1334,f1335,f1336,f1337,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1347,f1348,f1349,f1350,f1351,f1352,f1353,f1355,f1356,f1371,f1381,f1397,f1389,f1387,f1386,f1405,f1402,f1439,f1429,f1428,f1484,f1503,f1515,f1514,f1513,f1512,f1517,f1518,f1519,f1529,f1530,f1531,f1532,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1543,f1544,f1545,f1546,f1547,f1548,f1549,f1551,f1552,f1866,f1863,f1867,f1865,f1876,f1873,f1877,f1875,f1887,f1884,f1886,f1900,f1897,f1899,f1916,f1912,f1934,f1930,f1952,f1948,f1977,f1973,f2000,f2015,f2016,f2027,f2038,f2035,f2036,f2037,f2051,f2062,f2059,f2060,f2061,f2082,f2104,f2105,f2106,f2111,f2134,f2135,f2136,f2141,f2142,f2186,f2182,f2209,f2205,f2233,f2234,f2228,f2229,f2230,f2232,f2249,f2250,f2244,f2245,f2246,f2251,f2248,f2286,f2287,f2305,f2306,f2334,f2336,f2347,f2349,f2367,f2364,f2386,f2370,f2371,f2372,f2374,f2375,f2400,f2403,f2404,f2405,f2407,f2408,f2429,f2444,f2427,f2432,f2433,f2434,f2436,f2456,f2460,f2461,f2462,f2464,f2489,f2488,f2486,f2485,f2497,f2506,f2507,f2508,f2509,f2511,f2512,f2513,f2514,f2515,f2516,f2517,f2518,f2519,f2520,f2521,f2522,f2523,f2524,f2525,f2526,f2528,f2529,f2530,f2531,f2532,f2533,f2831,f2830,f2827,f2824,f2821,f2820,f2819,f2818,f2817,f2816,f2815,f2814,f2813,f2812,f2811,f2810,f2809,f2806,f2803,f2802,f2801,f2797,f2796,f2790,f2787,f2786,f2780,f2779,f2772,f2746,f2745,f2744,f2739,f2738,f2733,f2732,f2731,f2726,f2718,f2714,f2710,f2709,f2704,f2703,f2702,f2701,f2700,f2696,f2694,f2688,f2687,f2681,f2680,f2673,f2670,f2669,f2663,f2660,f2659,f2658,f2656,f2654,f2653,f2651,f2650,f2649,f2648,f2646,f2645,f2644,f2643,f2642,f2641,f2635,f2634,f2627,f2626,f2620,f2619,f2612,f2611,f2609,f2607,f2605,f2602,f2601,f2595,f2594,f2589,f2574,f2573,f2572,f2571,f2569,f2568,f2567,f2566,f2565,f3267,f3266,f3263,f3260,f3257,f3256,f3255,f3254,f3253,f3252,f3251,f3250,f3249,f3248,f3247,f3246,f3245,f3242,f3239,f3238,f3237,f3233,f3232,f3226,f3223,f3222,f3216,f3215,f3208,f3182,f3181,f3180,f3175,f3174,f3169,f3168,f3167,f3162,f3154,f3150,f3146,f3145,f3140,f3139,f3138,f3137,f3136,f3132,f3130,f3124,f3123,f3117,f3116,f3109,f3106,f3105,f3099,f3096,f3095,f3093,f3090,f3088,f3087,f3084,f3083,f3082,f3081,f3078,f3077,f3076,f3075,f3074,f3073,f3067,f3066,f3059,f3058,f3052,f3051,f3044,f3043,f3041,f3039,f3037,f3034,f3033,f3027,f3026,f3021,f3006,f3005,f3004,f3003,f3001,f3000,f2999,f2998,f2997,f3620,f3619,f3616,f3613,f3610,f3609,f3608,f3607,f3606,f3605,f3604,f3603,f3602,f3601,f3600,f3599,f3598,f3595,f3592,f3591,f3590,f3586,f3585,f3579,f3576,f3575,f3569,f3568,f3561,f3535,f3534,f3533,f3528,f3527,f3522,f3521,f3520,f3515,f3507,f3503,f3499,f3498,f3493,f3492,f3491,f3490,f3489,f3485,f3483,f3477,f3476,f3470,f3469,f3462,f3459,f3458,f3452,f3449,f3448,f3446,f3443,f3441,f3440,f3437,f3436,f3435,f3434,f3431,f3430,f3429,f3428,f3427,f3426,f3420,f3419,f3412,f3411,f3405,f3404,f3397,f3396,f3395,f3394,f3393,f3391,f3390,f3384,f3383,f3378,f3363,f3362,f3361,f3360,f3358,f3357,f3356,f3354,f3708,f3726,f3770,f3771,f3773,f3788,f3799,f3800,f3802,f3851,f3852,f3853,f3854,f3905,f3907,f3908,f3909,f3910,f3911,f3912,f3913,f3914,f3915,f3916,f3917,f3918,f3919,f3920,f3921,f3922,f3923,f3924,f3925,f3926,f3927,f4039,f4569,f4573,f4576,f4579,f4583,f4586,f4589,f4592,f4595,f4598,f4601,f4604,f4607,f4610,f4613,f4616,f4619,f4622,f4625,f4072,f4073,f4074,f4075,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4646,f4647,f4648,f4649,f4650,f4124,f4125,f4126,f4127,f4651,f4652,f4653,f4654,f4655,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4667,f4668,f4669,f4670,f4671,f4177,f4179,f4180,f4181,f4182,f4183,f4184,f4185,f4186,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4677,f4679,f4681,f4683,f4686,f4688,f4690,f4692,f4694,f4696,f4698,f4700,f4702,f4704,f4706,f4708,f4710,f4712,f4714,f4571,f4941,f4744,f4751])).
fof(f4759,plain,(
  ( ! [X47,X46] : (~sP6('F_0',happ(s(fun(X46,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X46,X47)),'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) ) | ($spl11 | $spl14)),
  inference(backward_demodulation,[],[f4571,f2953])).
fof(f2953,plain,(
  ( ! [X47,X46] : (~sP6('F_0',happ(s(fun(X46,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X46,X47)),'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) ) | $spl11),
  inference(superposition,[],[f1310,f1572])).
fof(f4953,plain,(
  $false | ($spl1 | $spl14)),
  inference(global_subsumption,[],[f4757,f321,f320,f319,f318,f322,f573,f329,f337,f336,f341,f340,f364,f371,f370,f369,f368,f367,f379,f377,f382,f380,f387,f384,f394,f393,f574,f412,f409,f408,f407,f421,f423,f425,f424,f428,f427,f426,f431,f430,f429,f434,f433,f432,f437,f436,f435,f439,f438,f451,f450,f449,f559,f447,f446,f468,f467,f466,f465,f464,f463,f462,f472,f471,f470,f469,f476,f475,f474,f473,f480,f479,f478,f477,f484,f483,f482,f481,f558,f488,f487,f486,f485,f495,f494,f493,f492,f501,f498,f497,f496,f508,f507,f506,f557,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f525,f524,f523,f556,f521,f520,f519,f529,f528,f527,f526,f533,f532,f531,f530,f555,f575,f576,f577,f578,f580,f582,f583,f585,f587,f588,f589,f590,f591,f593,f595,f596,f597,f599,f601,f602,f604,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f617,f621,f639,f640,f643,f649,f657,f664,f681,f682,f685,f688,f695,f704,f707,f708,f711,f712,f716,f725,f729,f730,f733,f736,f740,f743,f746,f755,f756,f757,f765,f768,f772,f775,f778,f797,f800,f815,f813,f818,f836,f884,f902,f925,f923,f929,f927,f941,f938,f951,f949,f992,f1003,f1011,f1024,f1032,f1036,f1042,f1051,f1056,f1055,f1068,f1076,f1080,f1113,f1124,f1131,f1140,f1152,f1160,f1164,f1178,f1196,f1197,f1212,f1222,f1237,f1249,f1257,f1261,f1272,f1282,f1298,f1314,f1313,f1311,f1310,f1325,f1334,f1335,f1336,f1337,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1347,f1348,f1349,f1350,f1351,f1352,f1353,f1355,f1356,f1371,f1381,f1397,f1389,f1387,f1386,f1405,f1402,f1439,f1429,f1428,f1484,f1503,f1515,f1514,f1513,f1512,f1517,f1518,f1519,f1529,f1530,f1531,f1532,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1543,f1544,f1545,f1546,f1547,f1548,f1549,f1551,f1552,f1866,f1863,f1867,f1865,f1876,f1873,f1877,f1875,f1887,f1884,f1886,f1900,f1897,f1899,f1916,f1912,f1934,f1930,f1952,f1948,f1977,f1973,f2000,f2015,f2016,f2027,f2038,f2035,f2036,f2037,f2051,f2062,f2059,f2060,f2061,f2082,f2104,f2105,f2106,f2111,f2134,f2135,f2136,f2141,f2142,f2186,f2182,f2209,f2205,f2233,f2234,f2228,f2229,f2230,f2232,f2249,f2250,f2244,f2245,f2246,f2251,f2248,f2286,f2287,f2305,f2306,f2334,f2336,f2347,f2349,f2367,f2364,f2386,f2370,f2371,f2372,f2374,f2375,f2400,f2403,f2404,f2405,f2407,f2408,f2429,f2444,f2427,f2432,f2433,f2434,f2436,f2456,f2460,f2461,f2462,f2464,f2489,f2488,f2486,f2485,f2497,f2506,f2507,f2508,f2509,f2511,f2512,f2513,f2514,f2515,f2516,f2517,f2518,f2519,f2520,f2521,f2522,f2523,f2524,f2525,f2526,f2528,f2529,f2530,f2531,f2532,f2533,f2831,f2830,f2827,f2824,f2821,f2820,f2819,f2818,f2817,f2816,f2815,f2814,f2813,f2812,f2811,f2810,f2809,f2806,f2803,f2802,f2801,f2797,f2796,f2790,f2787,f2786,f2780,f2779,f2772,f2746,f2745,f2744,f2739,f2738,f2733,f2732,f2731,f2726,f2718,f2714,f2710,f2709,f2704,f2703,f2702,f2701,f2700,f2696,f2694,f2688,f2687,f2681,f2680,f2673,f2670,f2669,f2663,f2660,f2659,f2658,f2656,f2654,f2653,f2651,f2650,f2649,f2648,f2646,f2645,f2644,f2643,f2642,f2641,f2635,f2634,f2627,f2626,f2620,f2619,f2612,f2611,f2609,f2607,f2605,f2602,f2601,f2595,f2594,f2589,f2574,f2573,f2572,f2571,f2569,f2568,f2567,f2566,f2565,f3267,f3266,f3263,f3260,f3257,f3256,f3255,f3254,f3253,f3252,f3251,f3250,f3249,f3248,f3247,f3246,f3245,f3242,f3239,f3238,f3237,f3233,f3232,f3226,f3223,f3222,f3216,f3215,f3208,f3182,f3181,f3180,f3175,f3174,f3169,f3168,f3167,f3162,f3154,f3150,f3146,f3145,f3140,f3139,f3138,f3137,f3136,f3132,f3130,f3124,f3123,f3117,f3116,f3109,f3106,f3105,f3099,f3096,f3095,f3093,f3090,f3088,f3087,f3084,f3083,f3082,f3081,f3078,f3077,f3076,f3075,f3074,f3073,f3067,f3066,f3059,f3058,f3052,f3051,f3044,f3043,f3041,f3039,f3037,f3034,f3033,f3027,f3026,f3021,f3006,f3005,f3004,f3003,f3001,f3000,f2999,f2998,f2997,f3620,f3619,f3616,f3613,f3610,f3609,f3608,f3607,f3606,f3605,f3604,f3603,f3602,f3601,f3600,f3599,f3598,f3595,f3592,f3591,f3590,f3586,f3585,f3579,f3576,f3575,f3569,f3568,f3561,f3535,f3534,f3533,f3528,f3527,f3522,f3521,f3520,f3515,f3507,f3503,f3499,f3498,f3493,f3492,f3491,f3490,f3489,f3485,f3483,f3477,f3476,f3470,f3469,f3462,f3459,f3458,f3452,f3449,f3448,f3446,f3443,f3441,f3440,f3437,f3436,f3435,f3434,f3431,f3430,f3429,f3428,f3427,f3426,f3420,f3419,f3412,f3411,f3405,f3404,f3397,f3396,f3395,f3394,f3393,f3391,f3390,f3384,f3383,f3378,f3363,f3362,f3361,f3360,f3358,f3357,f3356,f3354,f3708,f3726,f3770,f3771,f3773,f3788,f3799,f3800,f3802,f3851,f3852,f3853,f3854,f3905,f3907,f3908,f3909,f3910,f3911,f3912,f3913,f3914,f3915,f3916,f3917,f3918,f3919,f3920,f3921,f3922,f3923,f3924,f3925,f3926,f3927,f4039,f4569,f4573,f4576,f4579,f4583,f4586,f4589,f4592,f4595,f4598,f4601,f4604,f4607,f4610,f4613,f4616,f4619,f4622,f4625,f4072,f4073,f4074,f4075,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4646,f4647,f4648,f4649,f4650,f4124,f4125,f4126,f4127,f4651,f4652,f4653,f4654,f4655,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4667,f4668,f4669,f4670,f4671,f4177,f4179,f4180,f4181,f4182,f4183,f4184,f4185,f4186,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4677,f4679,f4681,f4683,f4686,f4688,f4690,f4692,f4694,f4696,f4698,f4700,f4702,f4704,f4706,f4708,f4710,f4712,f4714,f4571,f4941,f4744,f4751])).
fof(f4757,plain,(
  ( ! [X8,X9] : (sP6('T_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X8,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X8,X9)))) ) | ($spl1 | $spl14)),
  inference(backward_demodulation,[],[f4571,f2463])).
fof(f2463,plain,(
  ( ! [X8,X9] : (sP6('T_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X8,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X8,X9)))) ) | $spl1),
  inference(superposition,[],[f1314,f627])).
fof(f4952,plain,(
  $false | ($spl1 | $spl14)),
  inference(global_subsumption,[],[f4755,f321,f320,f319,f318,f322,f573,f329,f337,f336,f341,f340,f364,f371,f370,f369,f368,f367,f379,f377,f382,f380,f387,f384,f394,f393,f574,f412,f409,f408,f407,f421,f423,f425,f424,f428,f427,f426,f431,f430,f429,f434,f433,f432,f437,f436,f435,f439,f438,f451,f450,f449,f559,f447,f446,f468,f467,f466,f465,f464,f463,f462,f472,f471,f470,f469,f476,f475,f474,f473,f480,f479,f478,f477,f484,f483,f482,f481,f558,f488,f487,f486,f485,f495,f494,f493,f492,f501,f498,f497,f496,f508,f507,f506,f557,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f525,f524,f523,f556,f521,f520,f519,f529,f528,f527,f526,f533,f532,f531,f530,f555,f575,f576,f577,f578,f580,f582,f583,f585,f587,f588,f589,f590,f591,f593,f595,f596,f597,f599,f601,f602,f604,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f617,f621,f639,f640,f643,f649,f657,f664,f681,f682,f685,f688,f695,f704,f707,f708,f711,f712,f716,f725,f729,f730,f733,f736,f740,f743,f746,f755,f756,f757,f765,f768,f772,f775,f778,f797,f800,f815,f813,f818,f836,f884,f902,f925,f923,f929,f927,f941,f938,f951,f949,f992,f1003,f1011,f1024,f1032,f1036,f1042,f1051,f1056,f1055,f1068,f1076,f1080,f1113,f1124,f1131,f1140,f1152,f1160,f1164,f1178,f1196,f1197,f1212,f1222,f1237,f1249,f1257,f1261,f1272,f1282,f1298,f1314,f1313,f1311,f1310,f1325,f1334,f1335,f1336,f1337,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1347,f1348,f1349,f1350,f1351,f1352,f1353,f1355,f1356,f1371,f1381,f1397,f1389,f1387,f1386,f1405,f1402,f1439,f1429,f1428,f1484,f1503,f1515,f1514,f1513,f1512,f1517,f1518,f1519,f1529,f1530,f1531,f1532,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1543,f1544,f1545,f1546,f1547,f1548,f1549,f1551,f1552,f1866,f1863,f1867,f1865,f1876,f1873,f1877,f1875,f1887,f1884,f1886,f1900,f1897,f1899,f1916,f1912,f1934,f1930,f1952,f1948,f1977,f1973,f2000,f2015,f2016,f2027,f2038,f2035,f2036,f2037,f2051,f2062,f2059,f2060,f2061,f2082,f2104,f2105,f2106,f2111,f2134,f2135,f2136,f2141,f2142,f2186,f2182,f2209,f2205,f2233,f2234,f2228,f2229,f2230,f2232,f2249,f2250,f2244,f2245,f2246,f2251,f2248,f2286,f2287,f2305,f2306,f2334,f2336,f2347,f2349,f2367,f2364,f2386,f2370,f2371,f2372,f2374,f2375,f2400,f2403,f2404,f2405,f2407,f2408,f2429,f2444,f2427,f2432,f2433,f2434,f2436,f2456,f2460,f2461,f2462,f2464,f2489,f2488,f2486,f2485,f2497,f2506,f2507,f2508,f2509,f2511,f2512,f2513,f2514,f2515,f2516,f2517,f2518,f2519,f2520,f2521,f2522,f2523,f2524,f2525,f2526,f2528,f2529,f2530,f2531,f2532,f2533,f2831,f2830,f2827,f2824,f2821,f2820,f2819,f2818,f2817,f2816,f2815,f2814,f2813,f2812,f2811,f2810,f2809,f2806,f2803,f2802,f2801,f2797,f2796,f2790,f2787,f2786,f2780,f2779,f2772,f2746,f2745,f2744,f2739,f2738,f2733,f2732,f2731,f2726,f2718,f2714,f2710,f2709,f2704,f2703,f2702,f2701,f2700,f2696,f2694,f2688,f2687,f2681,f2680,f2673,f2670,f2669,f2663,f2660,f2659,f2658,f2656,f2654,f2653,f2651,f2650,f2649,f2648,f2646,f2645,f2644,f2643,f2642,f2641,f2635,f2634,f2627,f2626,f2620,f2619,f2612,f2611,f2609,f2607,f2605,f2602,f2601,f2595,f2594,f2589,f2574,f2573,f2572,f2571,f2569,f2568,f2567,f2566,f2565,f3267,f3266,f3263,f3260,f3257,f3256,f3255,f3254,f3253,f3252,f3251,f3250,f3249,f3248,f3247,f3246,f3245,f3242,f3239,f3238,f3237,f3233,f3232,f3226,f3223,f3222,f3216,f3215,f3208,f3182,f3181,f3180,f3175,f3174,f3169,f3168,f3167,f3162,f3154,f3150,f3146,f3145,f3140,f3139,f3138,f3137,f3136,f3132,f3130,f3124,f3123,f3117,f3116,f3109,f3106,f3105,f3099,f3096,f3095,f3093,f3090,f3088,f3087,f3084,f3083,f3082,f3081,f3078,f3077,f3076,f3075,f3074,f3073,f3067,f3066,f3059,f3058,f3052,f3051,f3044,f3043,f3041,f3039,f3037,f3034,f3033,f3027,f3026,f3021,f3006,f3005,f3004,f3003,f3001,f3000,f2999,f2998,f2997,f3620,f3619,f3616,f3613,f3610,f3609,f3608,f3607,f3606,f3605,f3604,f3603,f3602,f3601,f3600,f3599,f3598,f3595,f3592,f3591,f3590,f3586,f3585,f3579,f3576,f3575,f3569,f3568,f3561,f3535,f3534,f3533,f3528,f3527,f3522,f3521,f3520,f3515,f3507,f3503,f3499,f3498,f3493,f3492,f3491,f3490,f3489,f3485,f3483,f3477,f3476,f3470,f3469,f3462,f3459,f3458,f3452,f3449,f3448,f3446,f3443,f3441,f3440,f3437,f3436,f3435,f3434,f3431,f3430,f3429,f3428,f3427,f3426,f3420,f3419,f3412,f3411,f3405,f3404,f3397,f3396,f3395,f3394,f3393,f3391,f3390,f3384,f3383,f3378,f3363,f3362,f3361,f3360,f3358,f3357,f3356,f3354,f3708,f3726,f3770,f3771,f3773,f3788,f3799,f3800,f3802,f3851,f3852,f3853,f3854,f3905,f3907,f3908,f3909,f3910,f3911,f3912,f3913,f3914,f3915,f3916,f3917,f3918,f3919,f3920,f3921,f3922,f3923,f3924,f3925,f3926,f3927,f4039,f4569,f4573,f4576,f4579,f4583,f4586,f4589,f4592,f4595,f4598,f4601,f4604,f4607,f4610,f4613,f4616,f4619,f4622,f4625,f4072,f4073,f4074,f4075,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4646,f4647,f4648,f4649,f4650,f4124,f4125,f4126,f4127,f4651,f4652,f4653,f4654,f4655,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4667,f4668,f4669,f4670,f4671,f4177,f4179,f4180,f4181,f4182,f4183,f4184,f4185,f4186,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4677,f4679,f4681,f4683,f4686,f4688,f4690,f4692,f4694,f4696,f4698,f4700,f4702,f4704,f4706,f4708,f4710,f4712,f4714,f4571,f4941,f4744,f4751])).
fof(f4755,plain,(
  ( ! [X0,X1] : (~sP3(happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X0,X1)),'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'F_0')) ) | ($spl1 | $spl14)),
  inference(backward_demodulation,[],[f4571,f2446])).
fof(f2446,plain,(
  ( ! [X0,X1] : (~sP3(happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X0,X1)),'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'F_0')) ) | $spl1),
  inference(forward_demodulation,[],[f2425,f627])).
fof(f2425,plain,(
  ( ! [X0,X1] : (~sP3(happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16))))),s(X0,X1)),'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'F_0')) ) | $spl1),
  inference(unit_resulting_resolution,[],[f573,f586,f580,f1313,f468])).
fof(f586,plain,(
  ( ! [X0,X1] : (sP4('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f567,f476])).
fof(f4951,plain,(
  $false | ($spl11 | $spl14)),
  inference(global_subsumption,[],[f4754,f321,f320,f319,f318,f322,f573,f329,f337,f336,f341,f340,f364,f371,f370,f369,f368,f367,f379,f377,f382,f380,f387,f384,f394,f393,f574,f412,f409,f408,f407,f421,f423,f425,f424,f428,f427,f426,f431,f430,f429,f434,f433,f432,f437,f436,f435,f439,f438,f451,f450,f449,f559,f447,f446,f468,f467,f466,f465,f464,f463,f462,f472,f471,f470,f469,f476,f475,f474,f473,f480,f479,f478,f477,f484,f483,f482,f481,f558,f488,f487,f486,f485,f495,f494,f493,f492,f501,f498,f497,f496,f508,f507,f506,f557,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f525,f524,f523,f556,f521,f520,f519,f529,f528,f527,f526,f533,f532,f531,f530,f555,f575,f576,f577,f578,f580,f582,f583,f585,f587,f588,f589,f590,f591,f593,f595,f596,f597,f599,f601,f602,f604,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f617,f621,f639,f640,f643,f649,f657,f664,f681,f682,f685,f688,f695,f704,f707,f708,f711,f712,f716,f725,f729,f730,f733,f736,f740,f743,f746,f755,f756,f757,f765,f768,f772,f775,f778,f797,f800,f815,f813,f818,f836,f884,f902,f925,f923,f929,f927,f941,f938,f951,f949,f992,f1003,f1011,f1024,f1032,f1036,f1042,f1051,f1056,f1055,f1068,f1076,f1080,f1113,f1124,f1131,f1140,f1152,f1160,f1164,f1178,f1196,f1197,f1212,f1222,f1237,f1249,f1257,f1261,f1272,f1282,f1298,f1314,f1313,f1311,f1310,f1325,f1334,f1335,f1336,f1337,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1347,f1348,f1349,f1350,f1351,f1352,f1353,f1355,f1356,f1371,f1381,f1397,f1389,f1387,f1386,f1405,f1402,f1439,f1429,f1428,f1484,f1503,f1515,f1514,f1513,f1512,f1517,f1518,f1519,f1529,f1530,f1531,f1532,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1543,f1544,f1545,f1546,f1547,f1548,f1549,f1551,f1552,f1866,f1863,f1867,f1865,f1876,f1873,f1877,f1875,f1887,f1884,f1886,f1900,f1897,f1899,f1916,f1912,f1934,f1930,f1952,f1948,f1977,f1973,f2000,f2015,f2016,f2027,f2038,f2035,f2036,f2037,f2051,f2062,f2059,f2060,f2061,f2082,f2104,f2105,f2106,f2111,f2134,f2135,f2136,f2141,f2142,f2186,f2182,f2209,f2205,f2233,f2234,f2228,f2229,f2230,f2232,f2249,f2250,f2244,f2245,f2246,f2251,f2248,f2286,f2287,f2305,f2306,f2334,f2336,f2347,f2349,f2367,f2364,f2386,f2370,f2371,f2372,f2374,f2375,f2400,f2403,f2404,f2405,f2407,f2408,f2429,f2444,f2427,f2432,f2433,f2434,f2436,f2456,f2460,f2461,f2462,f2464,f2489,f2488,f2486,f2485,f2497,f2506,f2507,f2508,f2509,f2511,f2512,f2513,f2514,f2515,f2516,f2517,f2518,f2519,f2520,f2521,f2522,f2523,f2524,f2525,f2526,f2528,f2529,f2530,f2531,f2532,f2533,f2831,f2830,f2827,f2824,f2821,f2820,f2819,f2818,f2817,f2816,f2815,f2814,f2813,f2812,f2811,f2810,f2809,f2806,f2803,f2802,f2801,f2797,f2796,f2790,f2787,f2786,f2780,f2779,f2772,f2746,f2745,f2744,f2739,f2738,f2733,f2732,f2731,f2726,f2718,f2714,f2710,f2709,f2704,f2703,f2702,f2701,f2700,f2696,f2694,f2688,f2687,f2681,f2680,f2673,f2670,f2669,f2663,f2660,f2659,f2658,f2656,f2654,f2653,f2651,f2650,f2649,f2648,f2646,f2645,f2644,f2643,f2642,f2641,f2635,f2634,f2627,f2626,f2620,f2619,f2612,f2611,f2609,f2607,f2605,f2602,f2601,f2595,f2594,f2589,f2574,f2573,f2572,f2571,f2569,f2568,f2567,f2566,f2565,f3267,f3266,f3263,f3260,f3257,f3256,f3255,f3254,f3253,f3252,f3251,f3250,f3249,f3248,f3247,f3246,f3245,f3242,f3239,f3238,f3237,f3233,f3232,f3226,f3223,f3222,f3216,f3215,f3208,f3182,f3181,f3180,f3175,f3174,f3169,f3168,f3167,f3162,f3154,f3150,f3146,f3145,f3140,f3139,f3138,f3137,f3136,f3132,f3130,f3124,f3123,f3117,f3116,f3109,f3106,f3105,f3099,f3096,f3095,f3093,f3090,f3088,f3087,f3084,f3083,f3082,f3081,f3078,f3077,f3076,f3075,f3074,f3073,f3067,f3066,f3059,f3058,f3052,f3051,f3044,f3043,f3041,f3039,f3037,f3034,f3033,f3027,f3026,f3021,f3006,f3005,f3004,f3003,f3001,f3000,f2999,f2998,f2997,f3620,f3619,f3616,f3613,f3610,f3609,f3608,f3607,f3606,f3605,f3604,f3603,f3602,f3601,f3600,f3599,f3598,f3595,f3592,f3591,f3590,f3586,f3585,f3579,f3576,f3575,f3569,f3568,f3561,f3535,f3534,f3533,f3528,f3527,f3522,f3521,f3520,f3515,f3507,f3503,f3499,f3498,f3493,f3492,f3491,f3490,f3489,f3485,f3483,f3477,f3476,f3470,f3469,f3462,f3459,f3458,f3452,f3449,f3448,f3446,f3443,f3441,f3440,f3437,f3436,f3435,f3434,f3431,f3430,f3429,f3428,f3427,f3426,f3420,f3419,f3412,f3411,f3405,f3404,f3397,f3396,f3395,f3394,f3393,f3391,f3390,f3384,f3383,f3378,f3363,f3362,f3361,f3360,f3358,f3357,f3356,f3354,f3708,f3726,f3770,f3771,f3773,f3788,f3799,f3800,f3802,f3851,f3852,f3853,f3854,f3905,f3907,f3908,f3909,f3910,f3911,f3912,f3913,f3914,f3915,f3916,f3917,f3918,f3919,f3920,f3921,f3922,f3923,f3924,f3925,f3926,f3927,f4039,f4569,f4573,f4576,f4579,f4583,f4586,f4589,f4592,f4595,f4598,f4601,f4604,f4607,f4610,f4613,f4616,f4619,f4622,f4625,f4072,f4073,f4074,f4075,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4646,f4647,f4648,f4649,f4650,f4124,f4125,f4126,f4127,f4651,f4652,f4653,f4654,f4655,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4667,f4668,f4669,f4670,f4671,f4177,f4179,f4180,f4181,f4182,f4183,f4184,f4185,f4186,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4677,f4679,f4681,f4683,f4686,f4688,f4690,f4692,f4694,f4696,f4698,f4700,f4702,f4704,f4706,f4708,f4710,f4712,f4714,f4571,f4941,f4744,f4751])).
fof(f4754,plain,(
  ( ! [X0,X1] : (~sP3(happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X0,X1)),'|dollar|exists_1'(s(fun(sK14,bool),sK16)),'F_0')) ) | ($spl11 | $spl14)),
  inference(backward_demodulation,[],[f4571,f2445])).
fof(f2445,plain,(
  ( ! [X0,X1] : (~sP3(happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X0,X1)),'|dollar|exists_1'(s(fun(sK14,bool),sK16)),'F_0')) ) | $spl11),
  inference(forward_demodulation,[],[f2426,f1572])).
fof(f2426,plain,(
  ( ! [X0,X1] : (~sP3(happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16))))),s(X0,X1)),'|dollar|exists_1'(s(fun(sK14,bool),sK16)),'F_0')) ) | $spl11),
  inference(unit_resulting_resolution,[],[f573,f1626,f580,f1313,f468])).
fof(f4950,plain,(
  $false | ($spl1 | $spl14)),
  inference(global_subsumption,[],[f4753,f321,f320,f319,f318,f322,f573,f329,f337,f336,f341,f340,f364,f371,f370,f369,f368,f367,f379,f377,f382,f380,f387,f384,f394,f393,f574,f412,f409,f408,f407,f421,f423,f425,f424,f428,f427,f426,f431,f430,f429,f434,f433,f432,f437,f436,f435,f439,f438,f451,f450,f449,f559,f447,f446,f468,f467,f466,f465,f464,f463,f462,f472,f471,f470,f469,f476,f475,f474,f473,f480,f479,f478,f477,f484,f483,f482,f481,f558,f488,f487,f486,f485,f495,f494,f493,f492,f501,f498,f497,f496,f508,f507,f506,f557,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f525,f524,f523,f556,f521,f520,f519,f529,f528,f527,f526,f533,f532,f531,f530,f555,f575,f576,f577,f578,f580,f582,f583,f585,f587,f588,f589,f590,f591,f593,f595,f596,f597,f599,f601,f602,f604,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f617,f621,f639,f640,f643,f649,f657,f664,f681,f682,f685,f688,f695,f704,f707,f708,f711,f712,f716,f725,f729,f730,f733,f736,f740,f743,f746,f755,f756,f757,f765,f768,f772,f775,f778,f797,f800,f815,f813,f818,f836,f884,f902,f925,f923,f929,f927,f941,f938,f951,f949,f992,f1003,f1011,f1024,f1032,f1036,f1042,f1051,f1056,f1055,f1068,f1076,f1080,f1113,f1124,f1131,f1140,f1152,f1160,f1164,f1178,f1196,f1197,f1212,f1222,f1237,f1249,f1257,f1261,f1272,f1282,f1298,f1314,f1313,f1311,f1310,f1325,f1334,f1335,f1336,f1337,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1347,f1348,f1349,f1350,f1351,f1352,f1353,f1355,f1356,f1371,f1381,f1397,f1389,f1387,f1386,f1405,f1402,f1439,f1429,f1428,f1484,f1503,f1515,f1514,f1513,f1512,f1517,f1518,f1519,f1529,f1530,f1531,f1532,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1543,f1544,f1545,f1546,f1547,f1548,f1549,f1551,f1552,f1866,f1863,f1867,f1865,f1876,f1873,f1877,f1875,f1887,f1884,f1886,f1900,f1897,f1899,f1916,f1912,f1934,f1930,f1952,f1948,f1977,f1973,f2000,f2015,f2016,f2027,f2038,f2035,f2036,f2037,f2051,f2062,f2059,f2060,f2061,f2082,f2104,f2105,f2106,f2111,f2134,f2135,f2136,f2141,f2142,f2186,f2182,f2209,f2205,f2233,f2234,f2228,f2229,f2230,f2232,f2249,f2250,f2244,f2245,f2246,f2251,f2248,f2286,f2287,f2305,f2306,f2334,f2336,f2347,f2349,f2367,f2364,f2386,f2370,f2371,f2372,f2374,f2375,f2400,f2403,f2404,f2405,f2407,f2408,f2429,f2444,f2427,f2432,f2433,f2434,f2436,f2456,f2460,f2461,f2462,f2464,f2489,f2488,f2486,f2485,f2497,f2506,f2507,f2508,f2509,f2511,f2512,f2513,f2514,f2515,f2516,f2517,f2518,f2519,f2520,f2521,f2522,f2523,f2524,f2525,f2526,f2528,f2529,f2530,f2531,f2532,f2533,f2831,f2830,f2827,f2824,f2821,f2820,f2819,f2818,f2817,f2816,f2815,f2814,f2813,f2812,f2811,f2810,f2809,f2806,f2803,f2802,f2801,f2797,f2796,f2790,f2787,f2786,f2780,f2779,f2772,f2746,f2745,f2744,f2739,f2738,f2733,f2732,f2731,f2726,f2718,f2714,f2710,f2709,f2704,f2703,f2702,f2701,f2700,f2696,f2694,f2688,f2687,f2681,f2680,f2673,f2670,f2669,f2663,f2660,f2659,f2658,f2656,f2654,f2653,f2651,f2650,f2649,f2648,f2646,f2645,f2644,f2643,f2642,f2641,f2635,f2634,f2627,f2626,f2620,f2619,f2612,f2611,f2609,f2607,f2605,f2602,f2601,f2595,f2594,f2589,f2574,f2573,f2572,f2571,f2569,f2568,f2567,f2566,f2565,f3267,f3266,f3263,f3260,f3257,f3256,f3255,f3254,f3253,f3252,f3251,f3250,f3249,f3248,f3247,f3246,f3245,f3242,f3239,f3238,f3237,f3233,f3232,f3226,f3223,f3222,f3216,f3215,f3208,f3182,f3181,f3180,f3175,f3174,f3169,f3168,f3167,f3162,f3154,f3150,f3146,f3145,f3140,f3139,f3138,f3137,f3136,f3132,f3130,f3124,f3123,f3117,f3116,f3109,f3106,f3105,f3099,f3096,f3095,f3093,f3090,f3088,f3087,f3084,f3083,f3082,f3081,f3078,f3077,f3076,f3075,f3074,f3073,f3067,f3066,f3059,f3058,f3052,f3051,f3044,f3043,f3041,f3039,f3037,f3034,f3033,f3027,f3026,f3021,f3006,f3005,f3004,f3003,f3001,f3000,f2999,f2998,f2997,f3620,f3619,f3616,f3613,f3610,f3609,f3608,f3607,f3606,f3605,f3604,f3603,f3602,f3601,f3600,f3599,f3598,f3595,f3592,f3591,f3590,f3586,f3585,f3579,f3576,f3575,f3569,f3568,f3561,f3535,f3534,f3533,f3528,f3527,f3522,f3521,f3520,f3515,f3507,f3503,f3499,f3498,f3493,f3492,f3491,f3490,f3489,f3485,f3483,f3477,f3476,f3470,f3469,f3462,f3459,f3458,f3452,f3449,f3448,f3446,f3443,f3441,f3440,f3437,f3436,f3435,f3434,f3431,f3430,f3429,f3428,f3427,f3426,f3420,f3419,f3412,f3411,f3405,f3404,f3397,f3396,f3395,f3394,f3393,f3391,f3390,f3384,f3383,f3378,f3363,f3362,f3361,f3360,f3358,f3357,f3356,f3354,f3708,f3726,f3770,f3771,f3773,f3788,f3799,f3800,f3802,f3851,f3852,f3853,f3854,f3905,f3907,f3908,f3909,f3910,f3911,f3912,f3913,f3914,f3915,f3916,f3917,f3918,f3919,f3920,f3921,f3922,f3923,f3924,f3925,f3926,f3927,f4039,f4569,f4573,f4576,f4579,f4583,f4586,f4589,f4592,f4595,f4598,f4601,f4604,f4607,f4610,f4613,f4616,f4619,f4622,f4625,f4072,f4073,f4074,f4075,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4646,f4647,f4648,f4649,f4650,f4124,f4125,f4126,f4127,f4651,f4652,f4653,f4654,f4655,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4667,f4668,f4669,f4670,f4671,f4177,f4179,f4180,f4181,f4182,f4183,f4184,f4185,f4186,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4677,f4679,f4681,f4683,f4686,f4688,f4690,f4692,f4694,f4696,f4698,f4700,f4702,f4704,f4706,f4708,f4710,f4712,f4714,f4571,f4941,f4744,f4751])).
fof(f4753,plain,(
  ( ! [X8,X9] : (~sP6('F_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X8,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X8,X9)))) ) | ($spl1 | $spl14)),
  inference(backward_demodulation,[],[f4571,f2435])).
fof(f2435,plain,(
  ( ! [X8,X9] : (~sP6('F_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X8,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X8,X9)))) ) | $spl1),
  inference(superposition,[],[f1313,f627])).
fof(f4949,plain,(
  $false | $spl14),
  inference(global_subsumption,[],[f321,f320,f319,f318,f322,f573,f329,f337,f336,f341,f340,f364,f371,f370,f369,f368,f367,f379,f377,f382,f380,f387,f384,f394,f393,f574,f412,f409,f408,f407,f421,f423,f425,f424,f428,f427,f426,f431,f430,f429,f434,f433,f432,f437,f436,f435,f439,f438,f451,f450,f449,f559,f447,f446,f468,f467,f466,f465,f464,f463,f462,f472,f471,f470,f469,f476,f475,f474,f473,f480,f479,f478,f477,f484,f483,f482,f481,f558,f488,f487,f486,f485,f495,f494,f493,f492,f501,f498,f497,f496,f508,f507,f506,f557,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f525,f524,f523,f556,f521,f520,f519,f529,f528,f527,f526,f533,f532,f531,f530,f555,f575,f576,f577,f578,f580,f582,f583,f585,f587,f588,f589,f590,f591,f593,f595,f596,f597,f599,f601,f602,f604,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f617,f621,f639,f640,f643,f649,f657,f664,f681,f682,f685,f688,f695,f704,f707,f708,f711,f712,f716,f725,f729,f730,f733,f736,f740,f743,f746,f755,f756,f757,f765,f768,f772,f775,f778,f797,f800,f815,f813,f818,f836,f884,f902,f925,f923,f929,f927,f941,f938,f951,f949,f992,f1003,f1011,f1024,f1032,f1036,f1042,f1051,f1056,f1055,f1068,f1076,f1080,f1113,f1124,f1131,f1140,f1152,f1160,f1164,f1178,f1196,f1197,f1212,f1222,f1237,f1249,f1257,f1261,f1272,f1282,f1298,f1314,f1313,f1311,f1310,f1325,f1334,f1335,f1336,f1337,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1347,f1348,f1349,f1350,f1351,f1352,f1353,f1355,f1356,f1371,f1381,f1397,f1389,f1387,f1386,f1405,f1402,f1439,f1429,f1428,f1484,f1503,f1515,f1514,f1513,f1512,f1517,f1518,f1519,f1529,f1530,f1531,f1532,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1543,f1544,f1545,f1546,f1547,f1548,f1549,f1551,f1552,f1866,f1863,f1867,f1865,f1876,f1873,f1877,f1875,f1887,f1884,f1886,f1900,f1897,f1899,f1916,f1912,f1934,f1930,f1952,f1948,f1977,f1973,f2000,f2015,f2016,f2027,f2038,f2035,f2036,f2037,f2051,f2062,f2059,f2060,f2061,f2082,f2104,f2105,f2106,f2111,f2134,f2135,f2136,f2141,f2142,f2186,f2182,f2209,f2205,f2233,f2234,f2228,f2229,f2230,f2232,f2249,f2250,f2244,f2245,f2246,f2251,f2248,f2286,f2287,f2305,f2306,f2334,f2336,f2347,f2349,f2367,f2364,f2386,f2370,f2371,f2372,f2374,f2375,f2400,f2403,f2404,f2405,f2407,f2408,f2429,f2444,f2427,f2432,f2433,f2434,f2436,f2456,f2460,f2461,f2462,f2464,f2489,f2488,f2486,f2485,f2497,f2506,f2507,f2508,f2509,f2511,f2512,f2513,f2514,f2515,f2516,f2517,f2518,f2519,f2520,f2521,f2522,f2523,f2524,f2525,f2526,f2528,f2529,f2530,f2531,f2532,f2533,f2831,f2830,f2827,f2824,f2821,f2820,f2819,f2818,f2817,f2816,f2815,f2814,f2813,f2812,f2811,f2810,f2809,f2806,f2803,f2802,f2801,f2797,f2796,f2790,f2787,f2786,f2780,f2779,f2772,f2746,f2745,f2744,f2739,f2738,f2733,f2732,f2731,f2726,f2718,f2714,f2710,f2709,f2704,f2703,f2702,f2701,f2700,f2696,f2694,f2688,f2687,f2681,f2680,f2673,f2670,f2669,f2663,f2660,f2659,f2658,f2656,f2654,f2653,f2651,f2650,f2649,f2648,f2646,f2645,f2644,f2643,f2642,f2641,f2635,f2634,f2627,f2626,f2620,f2619,f2612,f2611,f2609,f2607,f2605,f2602,f2601,f2595,f2594,f2589,f2574,f2573,f2572,f2571,f2569,f2568,f2567,f2566,f2565,f3267,f3266,f3263,f3260,f3257,f3256,f3255,f3254,f3253,f3252,f3251,f3250,f3249,f3248,f3247,f3246,f3245,f3242,f3239,f3238,f3237,f3233,f3232,f3226,f3223,f3222,f3216,f3215,f3208,f3182,f3181,f3180,f3175,f3174,f3169,f3168,f3167,f3162,f3154,f3150,f3146,f3145,f3140,f3139,f3138,f3137,f3136,f3132,f3130,f3124,f3123,f3117,f3116,f3109,f3106,f3105,f3099,f3096,f3095,f3093,f3090,f3088,f3087,f3084,f3083,f3082,f3081,f3078,f3077,f3076,f3075,f3074,f3073,f3067,f3066,f3059,f3058,f3052,f3051,f3044,f3043,f3041,f3039,f3037,f3034,f3033,f3027,f3026,f3021,f3006,f3005,f3004,f3003,f3001,f3000,f2999,f2998,f2997,f3620,f3619,f3616,f3613,f3610,f3609,f3608,f3607,f3606,f3605,f3604,f3603,f3602,f3601,f3600,f3599,f3598,f3595,f3592,f3591,f3590,f3586,f3585,f3579,f3576,f3575,f3569,f3568,f3561,f3535,f3534,f3533,f3528,f3527,f3522,f3521,f3520,f3515,f3507,f3503,f3499,f3498,f3493,f3492,f3491,f3490,f3489,f3485,f3483,f3477,f3476,f3470,f3469,f3462,f3459,f3458,f3452,f3449,f3448,f3446,f3443,f3441,f3440,f3437,f3436,f3435,f3434,f3431,f3430,f3429,f3428,f3427,f3426,f3420,f3419,f3412,f3411,f3405,f3404,f3397,f3396,f3395,f3394,f3393,f3391,f3390,f3384,f3383,f3378,f3363,f3362,f3361,f3360,f3358,f3357,f3356,f3354,f3708,f3726,f3770,f3771,f3773,f3788,f3799,f3800,f3802,f3851,f3852,f3853,f3854,f3905,f3907,f3908,f3909,f3910,f3911,f3912,f3913,f3914,f3915,f3916,f3917,f3918,f3919,f3920,f3921,f3922,f3923,f3924,f3925,f3926,f3927,f4039,f4569,f4573,f4576,f4579,f4583,f4586,f4589,f4592,f4595,f4598,f4601,f4604,f4607,f4610,f4613,f4616,f4619,f4622,f4625,f4072,f4073,f4074,f4075,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4646,f4647,f4648,f4649,f4650,f4124,f4125,f4126,f4127,f4651,f4652,f4653,f4654,f4655,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4667,f4668,f4669,f4670,f4671,f4177,f4179,f4180,f4181,f4182,f4183,f4184,f4185,f4186,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4677,f4679,f4681,f4683,f4686,f4688,f4690,f4692,f4694,f4696,f4698,f4700,f4702,f4704,f4706,f4708,f4710,f4712,f4714,f4571,f4941,f4744,f4751])).
fof(f4948,plain,(
  $false | ($spl1 | $spl14)),
  inference(global_subsumption,[],[f4750,f321,f320,f319,f318,f322,f573,f329,f337,f336,f341,f340,f364,f371,f370,f369,f368,f367,f379,f377,f382,f380,f387,f384,f394,f393,f574,f412,f409,f408,f407,f421,f423,f425,f424,f428,f427,f426,f431,f430,f429,f434,f433,f432,f437,f436,f435,f439,f438,f451,f450,f449,f559,f447,f446,f468,f467,f466,f465,f464,f463,f462,f472,f471,f470,f469,f476,f475,f474,f473,f480,f479,f478,f477,f484,f483,f482,f481,f558,f488,f487,f486,f485,f495,f494,f493,f492,f501,f498,f497,f496,f508,f507,f506,f557,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f525,f524,f523,f556,f521,f520,f519,f529,f528,f527,f526,f533,f532,f531,f530,f555,f575,f576,f577,f578,f580,f582,f583,f585,f587,f588,f589,f590,f591,f593,f595,f596,f597,f599,f601,f602,f604,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f617,f621,f639,f640,f643,f649,f657,f664,f681,f682,f685,f688,f695,f704,f707,f708,f711,f712,f716,f725,f729,f730,f733,f736,f740,f743,f746,f755,f756,f757,f765,f768,f772,f775,f778,f797,f800,f815,f813,f818,f836,f884,f902,f925,f923,f929,f927,f941,f938,f951,f949,f992,f1003,f1011,f1024,f1032,f1036,f1042,f1051,f1056,f1055,f1068,f1076,f1080,f1113,f1124,f1131,f1140,f1152,f1160,f1164,f1178,f1196,f1197,f1212,f1222,f1237,f1249,f1257,f1261,f1272,f1282,f1298,f1314,f1313,f1311,f1310,f1325,f1334,f1335,f1336,f1337,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1347,f1348,f1349,f1350,f1351,f1352,f1353,f1355,f1356,f1371,f1381,f1397,f1389,f1387,f1386,f1405,f1402,f1439,f1429,f1428,f1484,f1503,f1515,f1514,f1513,f1512,f1517,f1518,f1519,f1529,f1530,f1531,f1532,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1543,f1544,f1545,f1546,f1547,f1548,f1549,f1551,f1552,f1866,f1863,f1867,f1865,f1876,f1873,f1877,f1875,f1887,f1884,f1886,f1900,f1897,f1899,f1916,f1912,f1934,f1930,f1952,f1948,f1977,f1973,f2000,f2015,f2016,f2027,f2038,f2035,f2036,f2037,f2051,f2062,f2059,f2060,f2061,f2082,f2104,f2105,f2106,f2111,f2134,f2135,f2136,f2141,f2142,f2186,f2182,f2209,f2205,f2233,f2234,f2228,f2229,f2230,f2232,f2249,f2250,f2244,f2245,f2246,f2251,f2248,f2286,f2287,f2305,f2306,f2334,f2336,f2347,f2349,f2367,f2364,f2386,f2370,f2371,f2372,f2374,f2375,f2400,f2403,f2404,f2405,f2407,f2408,f2429,f2444,f2427,f2432,f2433,f2434,f2436,f2456,f2460,f2461,f2462,f2464,f2489,f2488,f2486,f2485,f2497,f2506,f2507,f2508,f2509,f2511,f2512,f2513,f2514,f2515,f2516,f2517,f2518,f2519,f2520,f2521,f2522,f2523,f2524,f2525,f2526,f2528,f2529,f2530,f2531,f2532,f2533,f2831,f2830,f2827,f2824,f2821,f2820,f2819,f2818,f2817,f2816,f2815,f2814,f2813,f2812,f2811,f2810,f2809,f2806,f2803,f2802,f2801,f2797,f2796,f2790,f2787,f2786,f2780,f2779,f2772,f2746,f2745,f2744,f2739,f2738,f2733,f2732,f2731,f2726,f2718,f2714,f2710,f2709,f2704,f2703,f2702,f2701,f2700,f2696,f2694,f2688,f2687,f2681,f2680,f2673,f2670,f2669,f2663,f2660,f2659,f2658,f2656,f2654,f2653,f2651,f2650,f2649,f2648,f2646,f2645,f2644,f2643,f2642,f2641,f2635,f2634,f2627,f2626,f2620,f2619,f2612,f2611,f2609,f2607,f2605,f2602,f2601,f2595,f2594,f2589,f2574,f2573,f2572,f2571,f2569,f2568,f2567,f2566,f2565,f3267,f3266,f3263,f3260,f3257,f3256,f3255,f3254,f3253,f3252,f3251,f3250,f3249,f3248,f3247,f3246,f3245,f3242,f3239,f3238,f3237,f3233,f3232,f3226,f3223,f3222,f3216,f3215,f3208,f3182,f3181,f3180,f3175,f3174,f3169,f3168,f3167,f3162,f3154,f3150,f3146,f3145,f3140,f3139,f3138,f3137,f3136,f3132,f3130,f3124,f3123,f3117,f3116,f3109,f3106,f3105,f3099,f3096,f3095,f3093,f3090,f3088,f3087,f3084,f3083,f3082,f3081,f3078,f3077,f3076,f3075,f3074,f3073,f3067,f3066,f3059,f3058,f3052,f3051,f3044,f3043,f3041,f3039,f3037,f3034,f3033,f3027,f3026,f3021,f3006,f3005,f3004,f3003,f3001,f3000,f2999,f2998,f2997,f3620,f3619,f3616,f3613,f3610,f3609,f3608,f3607,f3606,f3605,f3604,f3603,f3602,f3601,f3600,f3599,f3598,f3595,f3592,f3591,f3590,f3586,f3585,f3579,f3576,f3575,f3569,f3568,f3561,f3535,f3534,f3533,f3528,f3527,f3522,f3521,f3520,f3515,f3507,f3503,f3499,f3498,f3493,f3492,f3491,f3490,f3489,f3485,f3483,f3477,f3476,f3470,f3469,f3462,f3459,f3458,f3452,f3449,f3448,f3446,f3443,f3441,f3440,f3437,f3436,f3435,f3434,f3431,f3430,f3429,f3428,f3427,f3426,f3420,f3419,f3412,f3411,f3405,f3404,f3397,f3396,f3395,f3394,f3393,f3391,f3390,f3384,f3383,f3378,f3363,f3362,f3361,f3360,f3358,f3357,f3356,f3354,f3708,f3726,f3770,f3771,f3773,f3788,f3799,f3800,f3802,f3851,f3852,f3853,f3854,f3905,f3907,f3908,f3909,f3910,f3911,f3912,f3913,f3914,f3915,f3916,f3917,f3918,f3919,f3920,f3921,f3922,f3923,f3924,f3925,f3926,f3927,f4039,f4569,f4573,f4576,f4579,f4583,f4586,f4589,f4592,f4595,f4598,f4601,f4604,f4607,f4610,f4613,f4616,f4619,f4622,f4625,f4072,f4073,f4074,f4075,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4646,f4647,f4648,f4649,f4650,f4124,f4125,f4126,f4127,f4651,f4652,f4653,f4654,f4655,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4667,f4668,f4669,f4670,f4671,f4177,f4179,f4180,f4181,f4182,f4183,f4184,f4185,f4186,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4677,f4679,f4681,f4683,f4686,f4688,f4690,f4692,f4694,f4696,f4698,f4700,f4702,f4704,f4706,f4708,f4710,f4712,f4714,f4571,f4941,f4744])).
fof(f4750,plain,(
  ( ! [X8,X9] : (sP6('T_0',happ(s(fun(X8,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X8,X9)),'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) ) | ($spl1 | $spl14)),
  inference(backward_demodulation,[],[f4571,f2406])).
fof(f2406,plain,(
  ( ! [X8,X9] : (sP6('T_0',happ(s(fun(X8,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X8,X9)),'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) ) | $spl1),
  inference(superposition,[],[f1311,f627])).
fof(f4947,plain,(
  $false | ($spl1 | $spl14)),
  inference(global_subsumption,[],[f4748,f321,f320,f319,f318,f322,f573,f329,f337,f336,f341,f340,f364,f371,f370,f369,f368,f367,f379,f377,f382,f380,f387,f384,f394,f393,f574,f412,f409,f408,f407,f421,f423,f425,f424,f428,f427,f426,f431,f430,f429,f434,f433,f432,f437,f436,f435,f439,f438,f451,f450,f449,f559,f447,f446,f468,f467,f466,f465,f464,f463,f462,f472,f471,f470,f469,f476,f475,f474,f473,f480,f479,f478,f477,f484,f483,f482,f481,f558,f488,f487,f486,f485,f495,f494,f493,f492,f501,f498,f497,f496,f508,f507,f506,f557,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f525,f524,f523,f556,f521,f520,f519,f529,f528,f527,f526,f533,f532,f531,f530,f555,f575,f576,f577,f578,f580,f582,f583,f585,f587,f588,f589,f590,f591,f593,f595,f596,f597,f599,f601,f602,f604,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f617,f621,f639,f640,f643,f649,f657,f664,f681,f682,f685,f688,f695,f704,f707,f708,f711,f712,f716,f725,f729,f730,f733,f736,f740,f743,f746,f755,f756,f757,f765,f768,f772,f775,f778,f797,f800,f815,f813,f818,f836,f884,f902,f925,f923,f929,f927,f941,f938,f951,f949,f992,f1003,f1011,f1024,f1032,f1036,f1042,f1051,f1056,f1055,f1068,f1076,f1080,f1113,f1124,f1131,f1140,f1152,f1160,f1164,f1178,f1196,f1197,f1212,f1222,f1237,f1249,f1257,f1261,f1272,f1282,f1298,f1314,f1313,f1311,f1310,f1325,f1334,f1335,f1336,f1337,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1347,f1348,f1349,f1350,f1351,f1352,f1353,f1355,f1356,f1371,f1381,f1397,f1389,f1387,f1386,f1405,f1402,f1439,f1429,f1428,f1484,f1503,f1515,f1514,f1513,f1512,f1517,f1518,f1519,f1529,f1530,f1531,f1532,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1543,f1544,f1545,f1546,f1547,f1548,f1549,f1551,f1552,f1866,f1863,f1867,f1865,f1876,f1873,f1877,f1875,f1887,f1884,f1886,f1900,f1897,f1899,f1916,f1912,f1934,f1930,f1952,f1948,f1977,f1973,f2000,f2015,f2016,f2027,f2038,f2035,f2036,f2037,f2051,f2062,f2059,f2060,f2061,f2082,f2104,f2105,f2106,f2111,f2134,f2135,f2136,f2141,f2142,f2186,f2182,f2209,f2205,f2233,f2234,f2228,f2229,f2230,f2232,f2249,f2250,f2244,f2245,f2246,f2251,f2248,f2286,f2287,f2305,f2306,f2334,f2336,f2347,f2349,f2367,f2364,f2386,f2370,f2371,f2372,f2374,f2375,f2400,f2403,f2404,f2405,f2407,f2408,f2429,f2444,f2427,f2432,f2433,f2434,f2436,f2456,f2460,f2461,f2462,f2464,f2489,f2488,f2486,f2485,f2497,f2506,f2507,f2508,f2509,f2511,f2512,f2513,f2514,f2515,f2516,f2517,f2518,f2519,f2520,f2521,f2522,f2523,f2524,f2525,f2526,f2528,f2529,f2530,f2531,f2532,f2533,f2831,f2830,f2827,f2824,f2821,f2820,f2819,f2818,f2817,f2816,f2815,f2814,f2813,f2812,f2811,f2810,f2809,f2806,f2803,f2802,f2801,f2797,f2796,f2790,f2787,f2786,f2780,f2779,f2772,f2746,f2745,f2744,f2739,f2738,f2733,f2732,f2731,f2726,f2718,f2714,f2710,f2709,f2704,f2703,f2702,f2701,f2700,f2696,f2694,f2688,f2687,f2681,f2680,f2673,f2670,f2669,f2663,f2660,f2659,f2658,f2656,f2654,f2653,f2651,f2650,f2649,f2648,f2646,f2645,f2644,f2643,f2642,f2641,f2635,f2634,f2627,f2626,f2620,f2619,f2612,f2611,f2609,f2607,f2605,f2602,f2601,f2595,f2594,f2589,f2574,f2573,f2572,f2571,f2569,f2568,f2567,f2566,f2565,f3267,f3266,f3263,f3260,f3257,f3256,f3255,f3254,f3253,f3252,f3251,f3250,f3249,f3248,f3247,f3246,f3245,f3242,f3239,f3238,f3237,f3233,f3232,f3226,f3223,f3222,f3216,f3215,f3208,f3182,f3181,f3180,f3175,f3174,f3169,f3168,f3167,f3162,f3154,f3150,f3146,f3145,f3140,f3139,f3138,f3137,f3136,f3132,f3130,f3124,f3123,f3117,f3116,f3109,f3106,f3105,f3099,f3096,f3095,f3093,f3090,f3088,f3087,f3084,f3083,f3082,f3081,f3078,f3077,f3076,f3075,f3074,f3073,f3067,f3066,f3059,f3058,f3052,f3051,f3044,f3043,f3041,f3039,f3037,f3034,f3033,f3027,f3026,f3021,f3006,f3005,f3004,f3003,f3001,f3000,f2999,f2998,f2997,f3620,f3619,f3616,f3613,f3610,f3609,f3608,f3607,f3606,f3605,f3604,f3603,f3602,f3601,f3600,f3599,f3598,f3595,f3592,f3591,f3590,f3586,f3585,f3579,f3576,f3575,f3569,f3568,f3561,f3535,f3534,f3533,f3528,f3527,f3522,f3521,f3520,f3515,f3507,f3503,f3499,f3498,f3493,f3492,f3491,f3490,f3489,f3485,f3483,f3477,f3476,f3470,f3469,f3462,f3459,f3458,f3452,f3449,f3448,f3446,f3443,f3441,f3440,f3437,f3436,f3435,f3434,f3431,f3430,f3429,f3428,f3427,f3426,f3420,f3419,f3412,f3411,f3405,f3404,f3397,f3396,f3395,f3394,f3393,f3391,f3390,f3384,f3383,f3378,f3363,f3362,f3361,f3360,f3358,f3357,f3356,f3354,f3708,f3726,f3770,f3771,f3773,f3788,f3799,f3800,f3802,f3851,f3852,f3853,f3854,f3905,f3907,f3908,f3909,f3910,f3911,f3912,f3913,f3914,f3915,f3916,f3917,f3918,f3919,f3920,f3921,f3922,f3923,f3924,f3925,f3926,f3927,f4039,f4569,f4573,f4576,f4579,f4583,f4586,f4589,f4592,f4595,f4598,f4601,f4604,f4607,f4610,f4613,f4616,f4619,f4622,f4625,f4072,f4073,f4074,f4075,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4646,f4647,f4648,f4649,f4650,f4124,f4125,f4126,f4127,f4651,f4652,f4653,f4654,f4655,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4667,f4668,f4669,f4670,f4671,f4177,f4179,f4180,f4181,f4182,f4183,f4184,f4185,f4186,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4677,f4679,f4681,f4683,f4686,f4688,f4690,f4692,f4694,f4696,f4698,f4700,f4702,f4704,f4706,f4708,f4710,f4712,f4714,f4571,f4941,f4744])).
fof(f4748,plain,(
  ( ! [X0,X1] : (~sP3('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X0,X1)),'F_0')) ) | ($spl1 | $spl14)),
  inference(backward_demodulation,[],[f4571,f2383])).
fof(f2383,plain,(
  ( ! [X0,X1] : (~sP3('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X0,X1)),'F_0')) ) | $spl1),
  inference(forward_demodulation,[],[f2365,f627])).
fof(f2365,plain,(
  ( ! [X0,X1] : (~sP3('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16))))),s(X0,X1)),'F_0')) ) | $spl1),
  inference(unit_resulting_resolution,[],[f573,f584,f579,f1310,f468])).
fof(f579,plain,(
  ( ! [X0,X1] : (sP5(X0,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f567,f471])).
fof(f584,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f567,f475])).
fof(f4946,plain,(
  $false | ($spl11 | $spl14)),
  inference(global_subsumption,[],[f4747,f321,f320,f319,f318,f322,f573,f329,f337,f336,f341,f340,f364,f371,f370,f369,f368,f367,f379,f377,f382,f380,f387,f384,f394,f393,f574,f412,f409,f408,f407,f421,f423,f425,f424,f428,f427,f426,f431,f430,f429,f434,f433,f432,f437,f436,f435,f439,f438,f451,f450,f449,f559,f447,f446,f468,f467,f466,f465,f464,f463,f462,f472,f471,f470,f469,f476,f475,f474,f473,f480,f479,f478,f477,f484,f483,f482,f481,f558,f488,f487,f486,f485,f495,f494,f493,f492,f501,f498,f497,f496,f508,f507,f506,f557,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f525,f524,f523,f556,f521,f520,f519,f529,f528,f527,f526,f533,f532,f531,f530,f555,f575,f576,f577,f578,f580,f582,f583,f585,f587,f588,f589,f590,f591,f593,f595,f596,f597,f599,f601,f602,f604,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f617,f621,f639,f640,f643,f649,f657,f664,f681,f682,f685,f688,f695,f704,f707,f708,f711,f712,f716,f725,f729,f730,f733,f736,f740,f743,f746,f755,f756,f757,f765,f768,f772,f775,f778,f797,f800,f815,f813,f818,f836,f884,f902,f925,f923,f929,f927,f941,f938,f951,f949,f992,f1003,f1011,f1024,f1032,f1036,f1042,f1051,f1056,f1055,f1068,f1076,f1080,f1113,f1124,f1131,f1140,f1152,f1160,f1164,f1178,f1196,f1197,f1212,f1222,f1237,f1249,f1257,f1261,f1272,f1282,f1298,f1314,f1313,f1311,f1310,f1325,f1334,f1335,f1336,f1337,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1347,f1348,f1349,f1350,f1351,f1352,f1353,f1355,f1356,f1371,f1381,f1397,f1389,f1387,f1386,f1405,f1402,f1439,f1429,f1428,f1484,f1503,f1515,f1514,f1513,f1512,f1517,f1518,f1519,f1529,f1530,f1531,f1532,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1543,f1544,f1545,f1546,f1547,f1548,f1549,f1551,f1552,f1866,f1863,f1867,f1865,f1876,f1873,f1877,f1875,f1887,f1884,f1886,f1900,f1897,f1899,f1916,f1912,f1934,f1930,f1952,f1948,f1977,f1973,f2000,f2015,f2016,f2027,f2038,f2035,f2036,f2037,f2051,f2062,f2059,f2060,f2061,f2082,f2104,f2105,f2106,f2111,f2134,f2135,f2136,f2141,f2142,f2186,f2182,f2209,f2205,f2233,f2234,f2228,f2229,f2230,f2232,f2249,f2250,f2244,f2245,f2246,f2251,f2248,f2286,f2287,f2305,f2306,f2334,f2336,f2347,f2349,f2367,f2364,f2386,f2370,f2371,f2372,f2374,f2375,f2400,f2403,f2404,f2405,f2407,f2408,f2429,f2444,f2427,f2432,f2433,f2434,f2436,f2456,f2460,f2461,f2462,f2464,f2489,f2488,f2486,f2485,f2497,f2506,f2507,f2508,f2509,f2511,f2512,f2513,f2514,f2515,f2516,f2517,f2518,f2519,f2520,f2521,f2522,f2523,f2524,f2525,f2526,f2528,f2529,f2530,f2531,f2532,f2533,f2831,f2830,f2827,f2824,f2821,f2820,f2819,f2818,f2817,f2816,f2815,f2814,f2813,f2812,f2811,f2810,f2809,f2806,f2803,f2802,f2801,f2797,f2796,f2790,f2787,f2786,f2780,f2779,f2772,f2746,f2745,f2744,f2739,f2738,f2733,f2732,f2731,f2726,f2718,f2714,f2710,f2709,f2704,f2703,f2702,f2701,f2700,f2696,f2694,f2688,f2687,f2681,f2680,f2673,f2670,f2669,f2663,f2660,f2659,f2658,f2656,f2654,f2653,f2651,f2650,f2649,f2648,f2646,f2645,f2644,f2643,f2642,f2641,f2635,f2634,f2627,f2626,f2620,f2619,f2612,f2611,f2609,f2607,f2605,f2602,f2601,f2595,f2594,f2589,f2574,f2573,f2572,f2571,f2569,f2568,f2567,f2566,f2565,f3267,f3266,f3263,f3260,f3257,f3256,f3255,f3254,f3253,f3252,f3251,f3250,f3249,f3248,f3247,f3246,f3245,f3242,f3239,f3238,f3237,f3233,f3232,f3226,f3223,f3222,f3216,f3215,f3208,f3182,f3181,f3180,f3175,f3174,f3169,f3168,f3167,f3162,f3154,f3150,f3146,f3145,f3140,f3139,f3138,f3137,f3136,f3132,f3130,f3124,f3123,f3117,f3116,f3109,f3106,f3105,f3099,f3096,f3095,f3093,f3090,f3088,f3087,f3084,f3083,f3082,f3081,f3078,f3077,f3076,f3075,f3074,f3073,f3067,f3066,f3059,f3058,f3052,f3051,f3044,f3043,f3041,f3039,f3037,f3034,f3033,f3027,f3026,f3021,f3006,f3005,f3004,f3003,f3001,f3000,f2999,f2998,f2997,f3620,f3619,f3616,f3613,f3610,f3609,f3608,f3607,f3606,f3605,f3604,f3603,f3602,f3601,f3600,f3599,f3598,f3595,f3592,f3591,f3590,f3586,f3585,f3579,f3576,f3575,f3569,f3568,f3561,f3535,f3534,f3533,f3528,f3527,f3522,f3521,f3520,f3515,f3507,f3503,f3499,f3498,f3493,f3492,f3491,f3490,f3489,f3485,f3483,f3477,f3476,f3470,f3469,f3462,f3459,f3458,f3452,f3449,f3448,f3446,f3443,f3441,f3440,f3437,f3436,f3435,f3434,f3431,f3430,f3429,f3428,f3427,f3426,f3420,f3419,f3412,f3411,f3405,f3404,f3397,f3396,f3395,f3394,f3393,f3391,f3390,f3384,f3383,f3378,f3363,f3362,f3361,f3360,f3358,f3357,f3356,f3354,f3708,f3726,f3770,f3771,f3773,f3788,f3799,f3800,f3802,f3851,f3852,f3853,f3854,f3905,f3907,f3908,f3909,f3910,f3911,f3912,f3913,f3914,f3915,f3916,f3917,f3918,f3919,f3920,f3921,f3922,f3923,f3924,f3925,f3926,f3927,f4039,f4569,f4573,f4576,f4579,f4583,f4586,f4589,f4592,f4595,f4598,f4601,f4604,f4607,f4610,f4613,f4616,f4619,f4622,f4625,f4072,f4073,f4074,f4075,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4646,f4647,f4648,f4649,f4650,f4124,f4125,f4126,f4127,f4651,f4652,f4653,f4654,f4655,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4667,f4668,f4669,f4670,f4671,f4177,f4179,f4180,f4181,f4182,f4183,f4184,f4185,f4186,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4677,f4679,f4681,f4683,f4686,f4688,f4690,f4692,f4694,f4696,f4698,f4700,f4702,f4704,f4706,f4708,f4710,f4712,f4714,f4571,f4941,f4744])).
fof(f4747,plain,(
  ( ! [X0,X1] : (~sP3('|dollar|exists_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X0,X1)),'F_0')) ) | ($spl11 | $spl14)),
  inference(backward_demodulation,[],[f4571,f2382])).
fof(f2382,plain,(
  ( ! [X0,X1] : (~sP3('|dollar|exists_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X0,X1)),'F_0')) ) | $spl11),
  inference(forward_demodulation,[],[f2366,f1572])).
fof(f2366,plain,(
  ( ! [X0,X1] : (~sP3('|dollar|exists_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16))))),s(X0,X1)),'F_0')) ) | $spl11),
  inference(unit_resulting_resolution,[],[f573,f1625,f1618,f1310,f468])).
fof(f4945,plain,(
  $false | ($spl1 | $spl14)),
  inference(global_subsumption,[],[f4746,f321,f320,f319,f318,f322,f573,f329,f337,f336,f341,f340,f364,f371,f370,f369,f368,f367,f379,f377,f382,f380,f387,f384,f394,f393,f574,f412,f409,f408,f407,f421,f423,f425,f424,f428,f427,f426,f431,f430,f429,f434,f433,f432,f437,f436,f435,f439,f438,f451,f450,f449,f559,f447,f446,f468,f467,f466,f465,f464,f463,f462,f472,f471,f470,f469,f476,f475,f474,f473,f480,f479,f478,f477,f484,f483,f482,f481,f558,f488,f487,f486,f485,f495,f494,f493,f492,f501,f498,f497,f496,f508,f507,f506,f557,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f525,f524,f523,f556,f521,f520,f519,f529,f528,f527,f526,f533,f532,f531,f530,f555,f575,f576,f577,f578,f580,f582,f583,f585,f587,f588,f589,f590,f591,f593,f595,f596,f597,f599,f601,f602,f604,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f617,f621,f639,f640,f643,f649,f657,f664,f681,f682,f685,f688,f695,f704,f707,f708,f711,f712,f716,f725,f729,f730,f733,f736,f740,f743,f746,f755,f756,f757,f765,f768,f772,f775,f778,f797,f800,f815,f813,f818,f836,f884,f902,f925,f923,f929,f927,f941,f938,f951,f949,f992,f1003,f1011,f1024,f1032,f1036,f1042,f1051,f1056,f1055,f1068,f1076,f1080,f1113,f1124,f1131,f1140,f1152,f1160,f1164,f1178,f1196,f1197,f1212,f1222,f1237,f1249,f1257,f1261,f1272,f1282,f1298,f1314,f1313,f1311,f1310,f1325,f1334,f1335,f1336,f1337,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1347,f1348,f1349,f1350,f1351,f1352,f1353,f1355,f1356,f1371,f1381,f1397,f1389,f1387,f1386,f1405,f1402,f1439,f1429,f1428,f1484,f1503,f1515,f1514,f1513,f1512,f1517,f1518,f1519,f1529,f1530,f1531,f1532,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1543,f1544,f1545,f1546,f1547,f1548,f1549,f1551,f1552,f1866,f1863,f1867,f1865,f1876,f1873,f1877,f1875,f1887,f1884,f1886,f1900,f1897,f1899,f1916,f1912,f1934,f1930,f1952,f1948,f1977,f1973,f2000,f2015,f2016,f2027,f2038,f2035,f2036,f2037,f2051,f2062,f2059,f2060,f2061,f2082,f2104,f2105,f2106,f2111,f2134,f2135,f2136,f2141,f2142,f2186,f2182,f2209,f2205,f2233,f2234,f2228,f2229,f2230,f2232,f2249,f2250,f2244,f2245,f2246,f2251,f2248,f2286,f2287,f2305,f2306,f2334,f2336,f2347,f2349,f2367,f2364,f2386,f2370,f2371,f2372,f2374,f2375,f2400,f2403,f2404,f2405,f2407,f2408,f2429,f2444,f2427,f2432,f2433,f2434,f2436,f2456,f2460,f2461,f2462,f2464,f2489,f2488,f2486,f2485,f2497,f2506,f2507,f2508,f2509,f2511,f2512,f2513,f2514,f2515,f2516,f2517,f2518,f2519,f2520,f2521,f2522,f2523,f2524,f2525,f2526,f2528,f2529,f2530,f2531,f2532,f2533,f2831,f2830,f2827,f2824,f2821,f2820,f2819,f2818,f2817,f2816,f2815,f2814,f2813,f2812,f2811,f2810,f2809,f2806,f2803,f2802,f2801,f2797,f2796,f2790,f2787,f2786,f2780,f2779,f2772,f2746,f2745,f2744,f2739,f2738,f2733,f2732,f2731,f2726,f2718,f2714,f2710,f2709,f2704,f2703,f2702,f2701,f2700,f2696,f2694,f2688,f2687,f2681,f2680,f2673,f2670,f2669,f2663,f2660,f2659,f2658,f2656,f2654,f2653,f2651,f2650,f2649,f2648,f2646,f2645,f2644,f2643,f2642,f2641,f2635,f2634,f2627,f2626,f2620,f2619,f2612,f2611,f2609,f2607,f2605,f2602,f2601,f2595,f2594,f2589,f2574,f2573,f2572,f2571,f2569,f2568,f2567,f2566,f2565,f3267,f3266,f3263,f3260,f3257,f3256,f3255,f3254,f3253,f3252,f3251,f3250,f3249,f3248,f3247,f3246,f3245,f3242,f3239,f3238,f3237,f3233,f3232,f3226,f3223,f3222,f3216,f3215,f3208,f3182,f3181,f3180,f3175,f3174,f3169,f3168,f3167,f3162,f3154,f3150,f3146,f3145,f3140,f3139,f3138,f3137,f3136,f3132,f3130,f3124,f3123,f3117,f3116,f3109,f3106,f3105,f3099,f3096,f3095,f3093,f3090,f3088,f3087,f3084,f3083,f3082,f3081,f3078,f3077,f3076,f3075,f3074,f3073,f3067,f3066,f3059,f3058,f3052,f3051,f3044,f3043,f3041,f3039,f3037,f3034,f3033,f3027,f3026,f3021,f3006,f3005,f3004,f3003,f3001,f3000,f2999,f2998,f2997,f3620,f3619,f3616,f3613,f3610,f3609,f3608,f3607,f3606,f3605,f3604,f3603,f3602,f3601,f3600,f3599,f3598,f3595,f3592,f3591,f3590,f3586,f3585,f3579,f3576,f3575,f3569,f3568,f3561,f3535,f3534,f3533,f3528,f3527,f3522,f3521,f3520,f3515,f3507,f3503,f3499,f3498,f3493,f3492,f3491,f3490,f3489,f3485,f3483,f3477,f3476,f3470,f3469,f3462,f3459,f3458,f3452,f3449,f3448,f3446,f3443,f3441,f3440,f3437,f3436,f3435,f3434,f3431,f3430,f3429,f3428,f3427,f3426,f3420,f3419,f3412,f3411,f3405,f3404,f3397,f3396,f3395,f3394,f3393,f3391,f3390,f3384,f3383,f3378,f3363,f3362,f3361,f3360,f3358,f3357,f3356,f3354,f3708,f3726,f3770,f3771,f3773,f3788,f3799,f3800,f3802,f3851,f3852,f3853,f3854,f3905,f3907,f3908,f3909,f3910,f3911,f3912,f3913,f3914,f3915,f3916,f3917,f3918,f3919,f3920,f3921,f3922,f3923,f3924,f3925,f3926,f3927,f4039,f4569,f4573,f4576,f4579,f4583,f4586,f4589,f4592,f4595,f4598,f4601,f4604,f4607,f4610,f4613,f4616,f4619,f4622,f4625,f4072,f4073,f4074,f4075,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4646,f4647,f4648,f4649,f4650,f4124,f4125,f4126,f4127,f4651,f4652,f4653,f4654,f4655,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4667,f4668,f4669,f4670,f4671,f4177,f4179,f4180,f4181,f4182,f4183,f4184,f4185,f4186,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4677,f4679,f4681,f4683,f4686,f4688,f4690,f4692,f4694,f4696,f4698,f4700,f4702,f4704,f4706,f4708,f4710,f4712,f4714,f4571,f4941,f4744])).
fof(f4746,plain,(
  ( ! [X8,X9] : (~sP6('F_0',happ(s(fun(X8,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X8,X9)),'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) ) | ($spl1 | $spl14)),
  inference(backward_demodulation,[],[f4571,f2373])).
fof(f2373,plain,(
  ( ! [X8,X9] : (~sP6('F_0',happ(s(fun(X8,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X8,X9)),'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) ) | $spl1),
  inference(superposition,[],[f1310,f627])).
fof(f4944,plain,(
  $false | $spl14),
  inference(global_subsumption,[],[f321,f320,f319,f318,f322,f573,f329,f337,f336,f341,f340,f364,f371,f370,f369,f368,f367,f379,f377,f382,f380,f387,f384,f394,f393,f574,f412,f409,f408,f407,f421,f423,f425,f424,f428,f427,f426,f431,f430,f429,f434,f433,f432,f437,f436,f435,f439,f438,f451,f450,f449,f559,f447,f446,f468,f467,f466,f465,f464,f463,f462,f472,f471,f470,f469,f476,f475,f474,f473,f480,f479,f478,f477,f484,f483,f482,f481,f558,f488,f487,f486,f485,f495,f494,f493,f492,f501,f498,f497,f496,f508,f507,f506,f557,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f525,f524,f523,f556,f521,f520,f519,f529,f528,f527,f526,f533,f532,f531,f530,f555,f575,f576,f577,f578,f580,f582,f583,f585,f587,f588,f589,f590,f591,f593,f595,f596,f597,f599,f601,f602,f604,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f617,f621,f639,f640,f643,f649,f657,f664,f681,f682,f685,f688,f695,f704,f707,f708,f711,f712,f716,f725,f729,f730,f733,f736,f740,f743,f746,f755,f756,f757,f765,f768,f772,f775,f778,f797,f800,f815,f813,f818,f836,f884,f902,f925,f923,f929,f927,f941,f938,f951,f949,f992,f1003,f1011,f1024,f1032,f1036,f1042,f1051,f1056,f1055,f1068,f1076,f1080,f1113,f1124,f1131,f1140,f1152,f1160,f1164,f1178,f1196,f1197,f1212,f1222,f1237,f1249,f1257,f1261,f1272,f1282,f1298,f1314,f1313,f1311,f1310,f1325,f1334,f1335,f1336,f1337,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1347,f1348,f1349,f1350,f1351,f1352,f1353,f1355,f1356,f1371,f1381,f1397,f1389,f1387,f1386,f1405,f1402,f1439,f1429,f1428,f1484,f1503,f1515,f1514,f1513,f1512,f1517,f1518,f1519,f1529,f1530,f1531,f1532,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1543,f1544,f1545,f1546,f1547,f1548,f1549,f1551,f1552,f1866,f1863,f1867,f1865,f1876,f1873,f1877,f1875,f1887,f1884,f1886,f1900,f1897,f1899,f1916,f1912,f1934,f1930,f1952,f1948,f1977,f1973,f2000,f2015,f2016,f2027,f2038,f2035,f2036,f2037,f2051,f2062,f2059,f2060,f2061,f2082,f2104,f2105,f2106,f2111,f2134,f2135,f2136,f2141,f2142,f2186,f2182,f2209,f2205,f2233,f2234,f2228,f2229,f2230,f2232,f2249,f2250,f2244,f2245,f2246,f2251,f2248,f2286,f2287,f2305,f2306,f2334,f2336,f2347,f2349,f2367,f2364,f2386,f2370,f2371,f2372,f2374,f2375,f2400,f2403,f2404,f2405,f2407,f2408,f2429,f2444,f2427,f2432,f2433,f2434,f2436,f2456,f2460,f2461,f2462,f2464,f2489,f2488,f2486,f2485,f2497,f2506,f2507,f2508,f2509,f2511,f2512,f2513,f2514,f2515,f2516,f2517,f2518,f2519,f2520,f2521,f2522,f2523,f2524,f2525,f2526,f2528,f2529,f2530,f2531,f2532,f2533,f2831,f2830,f2827,f2824,f2821,f2820,f2819,f2818,f2817,f2816,f2815,f2814,f2813,f2812,f2811,f2810,f2809,f2806,f2803,f2802,f2801,f2797,f2796,f2790,f2787,f2786,f2780,f2779,f2772,f2746,f2745,f2744,f2739,f2738,f2733,f2732,f2731,f2726,f2718,f2714,f2710,f2709,f2704,f2703,f2702,f2701,f2700,f2696,f2694,f2688,f2687,f2681,f2680,f2673,f2670,f2669,f2663,f2660,f2659,f2658,f2656,f2654,f2653,f2651,f2650,f2649,f2648,f2646,f2645,f2644,f2643,f2642,f2641,f2635,f2634,f2627,f2626,f2620,f2619,f2612,f2611,f2609,f2607,f2605,f2602,f2601,f2595,f2594,f2589,f2574,f2573,f2572,f2571,f2569,f2568,f2567,f2566,f2565,f3267,f3266,f3263,f3260,f3257,f3256,f3255,f3254,f3253,f3252,f3251,f3250,f3249,f3248,f3247,f3246,f3245,f3242,f3239,f3238,f3237,f3233,f3232,f3226,f3223,f3222,f3216,f3215,f3208,f3182,f3181,f3180,f3175,f3174,f3169,f3168,f3167,f3162,f3154,f3150,f3146,f3145,f3140,f3139,f3138,f3137,f3136,f3132,f3130,f3124,f3123,f3117,f3116,f3109,f3106,f3105,f3099,f3096,f3095,f3093,f3090,f3088,f3087,f3084,f3083,f3082,f3081,f3078,f3077,f3076,f3075,f3074,f3073,f3067,f3066,f3059,f3058,f3052,f3051,f3044,f3043,f3041,f3039,f3037,f3034,f3033,f3027,f3026,f3021,f3006,f3005,f3004,f3003,f3001,f3000,f2999,f2998,f2997,f3620,f3619,f3616,f3613,f3610,f3609,f3608,f3607,f3606,f3605,f3604,f3603,f3602,f3601,f3600,f3599,f3598,f3595,f3592,f3591,f3590,f3586,f3585,f3579,f3576,f3575,f3569,f3568,f3561,f3535,f3534,f3533,f3528,f3527,f3522,f3521,f3520,f3515,f3507,f3503,f3499,f3498,f3493,f3492,f3491,f3490,f3489,f3485,f3483,f3477,f3476,f3470,f3469,f3462,f3459,f3458,f3452,f3449,f3448,f3446,f3443,f3441,f3440,f3437,f3436,f3435,f3434,f3431,f3430,f3429,f3428,f3427,f3426,f3420,f3419,f3412,f3411,f3405,f3404,f3397,f3396,f3395,f3394,f3393,f3391,f3390,f3384,f3383,f3378,f3363,f3362,f3361,f3360,f3358,f3357,f3356,f3354,f3708,f3726,f3770,f3771,f3773,f3788,f3799,f3800,f3802,f3851,f3852,f3853,f3854,f3905,f3907,f3908,f3909,f3910,f3911,f3912,f3913,f3914,f3915,f3916,f3917,f3918,f3919,f3920,f3921,f3922,f3923,f3924,f3925,f3926,f3927,f4039,f4569,f4573,f4576,f4579,f4583,f4586,f4589,f4592,f4595,f4598,f4601,f4604,f4607,f4610,f4613,f4616,f4619,f4622,f4625,f4072,f4073,f4074,f4075,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4646,f4647,f4648,f4649,f4650,f4124,f4125,f4126,f4127,f4651,f4652,f4653,f4654,f4655,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4667,f4668,f4669,f4670,f4671,f4177,f4179,f4180,f4181,f4182,f4183,f4184,f4185,f4186,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4677,f4679,f4681,f4683,f4686,f4688,f4690,f4692,f4694,f4696,f4698,f4700,f4702,f4704,f4706,f4708,f4710,f4712,f4714,f4571,f4941,f4744])).
fof(f4943,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f4717,f322])).
fof(f4717,plain,(
  ~p(s(bool,'T_0')) | $spl14),
  inference(backward_demodulation,[],[f4571,f573])).
fof(f4942,plain,(
  $false | $spl14),
  inference(global_subsumption,[],[f321,f320,f319,f318,f322,f573,f329,f337,f336,f341,f340,f364,f371,f370,f369,f368,f367,f379,f377,f382,f380,f387,f384,f394,f393,f574,f412,f409,f408,f407,f421,f423,f425,f424,f428,f427,f426,f431,f430,f429,f434,f433,f432,f437,f436,f435,f439,f438,f451,f450,f449,f559,f447,f446,f468,f467,f466,f465,f464,f463,f462,f472,f471,f470,f469,f476,f475,f474,f473,f480,f479,f478,f477,f484,f483,f482,f481,f558,f488,f487,f486,f485,f495,f494,f493,f492,f501,f498,f497,f496,f508,f507,f506,f557,f504,f503,f502,f512,f511,f510,f509,f516,f515,f514,f513,f525,f524,f523,f556,f521,f520,f519,f529,f528,f527,f526,f533,f532,f531,f530,f555,f575,f576,f577,f578,f580,f582,f583,f585,f587,f588,f589,f590,f591,f593,f595,f596,f597,f599,f601,f602,f604,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f617,f621,f639,f640,f643,f649,f657,f664,f681,f682,f685,f688,f695,f704,f707,f708,f711,f712,f716,f725,f729,f730,f733,f736,f740,f743,f746,f755,f756,f757,f765,f768,f772,f775,f778,f797,f800,f815,f813,f818,f836,f884,f902,f925,f923,f929,f927,f941,f938,f951,f949,f992,f1003,f1011,f1024,f1032,f1036,f1042,f1051,f1056,f1055,f1068,f1076,f1080,f1113,f1124,f1131,f1140,f1152,f1160,f1164,f1178,f1196,f1197,f1212,f1222,f1237,f1249,f1257,f1261,f1272,f1282,f1298,f1314,f1313,f1311,f1310,f1325,f1334,f1335,f1336,f1337,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1347,f1348,f1349,f1350,f1351,f1352,f1353,f1355,f1356,f1371,f1381,f1397,f1389,f1387,f1386,f1405,f1402,f1439,f1429,f1428,f1484,f1503,f1515,f1514,f1513,f1512,f1517,f1518,f1519,f1529,f1530,f1531,f1532,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1543,f1544,f1545,f1546,f1547,f1548,f1549,f1551,f1552,f1866,f1863,f1867,f1865,f1876,f1873,f1877,f1875,f1887,f1884,f1886,f1900,f1897,f1899,f1916,f1912,f1934,f1930,f1952,f1948,f1977,f1973,f2000,f2015,f2016,f2027,f2038,f2035,f2036,f2037,f2051,f2062,f2059,f2060,f2061,f2082,f2104,f2105,f2106,f2111,f2134,f2135,f2136,f2141,f2142,f2186,f2182,f2209,f2205,f2233,f2234,f2228,f2229,f2230,f2232,f2249,f2250,f2244,f2245,f2246,f2251,f2248,f2286,f2287,f2305,f2306,f2334,f2336,f2347,f2349,f2367,f2364,f2386,f2370,f2371,f2372,f2374,f2375,f2400,f2403,f2404,f2405,f2407,f2408,f2429,f2444,f2427,f2432,f2433,f2434,f2436,f2456,f2460,f2461,f2462,f2464,f2489,f2488,f2486,f2485,f2497,f2506,f2507,f2508,f2509,f2511,f2512,f2513,f2514,f2515,f2516,f2517,f2518,f2519,f2520,f2521,f2522,f2523,f2524,f2525,f2526,f2528,f2529,f2530,f2531,f2532,f2533,f2831,f2830,f2827,f2824,f2821,f2820,f2819,f2818,f2817,f2816,f2815,f2814,f2813,f2812,f2811,f2810,f2809,f2806,f2803,f2802,f2801,f2797,f2796,f2790,f2787,f2786,f2780,f2779,f2772,f2746,f2745,f2744,f2739,f2738,f2733,f2732,f2731,f2726,f2718,f2714,f2710,f2709,f2704,f2703,f2702,f2701,f2700,f2696,f2694,f2688,f2687,f2681,f2680,f2673,f2670,f2669,f2663,f2660,f2659,f2658,f2656,f2654,f2653,f2651,f2650,f2649,f2648,f2646,f2645,f2644,f2643,f2642,f2641,f2635,f2634,f2627,f2626,f2620,f2619,f2612,f2611,f2609,f2607,f2605,f2602,f2601,f2595,f2594,f2589,f2574,f2573,f2572,f2571,f2569,f2568,f2567,f2566,f2565,f3267,f3266,f3263,f3260,f3257,f3256,f3255,f3254,f3253,f3252,f3251,f3250,f3249,f3248,f3247,f3246,f3245,f3242,f3239,f3238,f3237,f3233,f3232,f3226,f3223,f3222,f3216,f3215,f3208,f3182,f3181,f3180,f3175,f3174,f3169,f3168,f3167,f3162,f3154,f3150,f3146,f3145,f3140,f3139,f3138,f3137,f3136,f3132,f3130,f3124,f3123,f3117,f3116,f3109,f3106,f3105,f3099,f3096,f3095,f3093,f3090,f3088,f3087,f3084,f3083,f3082,f3081,f3078,f3077,f3076,f3075,f3074,f3073,f3067,f3066,f3059,f3058,f3052,f3051,f3044,f3043,f3041,f3039,f3037,f3034,f3033,f3027,f3026,f3021,f3006,f3005,f3004,f3003,f3001,f3000,f2999,f2998,f2997,f3620,f3619,f3616,f3613,f3610,f3609,f3608,f3607,f3606,f3605,f3604,f3603,f3602,f3601,f3600,f3599,f3598,f3595,f3592,f3591,f3590,f3586,f3585,f3579,f3576,f3575,f3569,f3568,f3561,f3535,f3534,f3533,f3528,f3527,f3522,f3521,f3520,f3515,f3507,f3503,f3499,f3498,f3493,f3492,f3491,f3490,f3489,f3485,f3483,f3477,f3476,f3470,f3469,f3462,f3459,f3458,f3452,f3449,f3448,f3446,f3443,f3441,f3440,f3437,f3436,f3435,f3434,f3431,f3430,f3429,f3428,f3427,f3426,f3420,f3419,f3412,f3411,f3405,f3404,f3397,f3396,f3395,f3394,f3393,f3391,f3390,f3384,f3383,f3378,f3363,f3362,f3361,f3360,f3358,f3357,f3356,f3354,f3708,f3726,f3770,f3771,f3773,f3788,f3799,f3800,f3802,f3851,f3852,f3853,f3854,f3905,f3907,f3908,f3909,f3910,f3911,f3912,f3913,f3914,f3915,f3916,f3917,f3918,f3919,f3920,f3921,f3922,f3923,f3924,f3925,f3926,f3927,f4039,f4569,f4573,f4576,f4579,f4583,f4586,f4589,f4592,f4595,f4598,f4601,f4604,f4607,f4610,f4613,f4616,f4619,f4622,f4625,f4072,f4073,f4074,f4075,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4646,f4647,f4648,f4649,f4650,f4124,f4125,f4126,f4127,f4651,f4652,f4653,f4654,f4655,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4667,f4668,f4669,f4670,f4671,f4177,f4179,f4180,f4181,f4182,f4183,f4184,f4185,f4186,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4677,f4679,f4681,f4683,f4686,f4688,f4690,f4692,f4694,f4696,f4698,f4700,f4702,f4704,f4706,f4708,f4710,f4712,f4714,f4571,f4941])).
fof(f4626,plain,(
  ( ! [X2,X1] : (sP10('T_0',X1,X2) | ~sP9('T_0',X1,X2)) ) | $spl50),
  inference(cnf_transformation,[],[f4626_D])).
fof(f4626_D,plain,(
  ( ! [X2,X1] : (sP10('T_0',X1,X2) | ~sP9('T_0',X1,X2)) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f4714,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP9('T_0',X1,X2) | sP10('T_0',X1,X2)) )),
  inference(condensation,[],[f4713])).
fof(f4713,plain,(
  ( ! [X114,X112,X115,X113] : (~sP9('T_0',X113,X114) | sP10('T_0',X113,X114) | s(bool,X112) = s(bool,X115) | s(bool,'T_0') = s(bool,X115)) )),
  inference(subsumption_resolution,[],[f4232,f902])).
fof(f4232,plain,(
  ( ! [X114,X112,X115,X113] : (~p(s(bool,X112)) | ~sP9('T_0',X113,X114) | sP10('T_0',X113,X114) | s(bool,X112) = s(bool,X115) | s(bool,'T_0') = s(bool,X115)) )),
  inference(superposition,[],[f557,f836])).
fof(f4623,plain,(
  ( ! [X2,X1] : (sP12('T_0',X1,X2) | ~sP11('T_0',X1,X2)) ) | $spl48),
  inference(cnf_transformation,[],[f4623_D])).
fof(f4623_D,plain,(
  ( ! [X2,X1] : (sP12('T_0',X1,X2) | ~sP11('T_0',X1,X2)) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f4712,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP11('T_0',X1,X2) | sP12('T_0',X1,X2)) )),
  inference(condensation,[],[f4711])).
fof(f4711,plain,(
  ( ! [X111,X109,X110,X108] : (~sP11('T_0',X109,X110) | sP12('T_0',X109,X110) | s(bool,X108) = s(bool,X111) | s(bool,'T_0') = s(bool,X111)) )),
  inference(subsumption_resolution,[],[f4231,f902])).
fof(f4231,plain,(
  ( ! [X111,X109,X110,X108] : (~p(s(bool,X108)) | ~sP11('T_0',X109,X110) | sP12('T_0',X109,X110) | s(bool,X108) = s(bool,X111) | s(bool,'T_0') = s(bool,X111)) )),
  inference(superposition,[],[f556,f836])).
fof(f4620,plain,(
  ( ! [X2,X1] : (~sP12('T_0',X2,X1) | p(s(bool,X2)) | ~p(s(bool,X1))) ) | $spl46),
  inference(cnf_transformation,[],[f4620_D])).
fof(f4620_D,plain,(
  ( ! [X2,X1] : (~sP12('T_0',X2,X1) | p(s(bool,X2)) | ~p(s(bool,X1))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f4710,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP12('T_0',X2,X1)) )),
  inference(condensation,[],[f4709])).
fof(f4709,plain,(
  ( ! [X103,X101,X102,X104] : (~p(s(bool,X102)) | p(s(bool,X103)) | ~sP12('T_0',X103,X102) | s(bool,X101) = s(bool,X104) | s(bool,'T_0') = s(bool,X104)) )),
  inference(subsumption_resolution,[],[f4229,f902])).
fof(f4229,plain,(
  ( ! [X103,X101,X102,X104] : (~p(s(bool,X101)) | ~p(s(bool,X102)) | p(s(bool,X103)) | ~sP12('T_0',X103,X102) | s(bool,X101) = s(bool,X104) | s(bool,'T_0') = s(bool,X104)) )),
  inference(superposition,[],[f533,f836])).
fof(f4617,plain,(
  ( ! [X2,X1] : (sP12('T_0',X1,X2) | ~p(s(bool,X1))) ) | $spl44),
  inference(cnf_transformation,[],[f4617_D])).
fof(f4617_D,plain,(
  ( ! [X2,X1] : (sP12('T_0',X1,X2) | ~p(s(bool,X1))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f4708,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP12('T_0',X1,X2)) )),
  inference(condensation,[],[f4707])).
fof(f4707,plain,(
  ( ! [X99,X97,X100,X98] : (~p(s(bool,X98)) | sP12('T_0',X98,X99) | s(bool,X97) = s(bool,X100) | s(bool,'T_0') = s(bool,X100)) )),
  inference(subsumption_resolution,[],[f4228,f902])).
fof(f4228,plain,(
  ( ! [X99,X97,X100,X98] : (~p(s(bool,X97)) | ~p(s(bool,X98)) | sP12('T_0',X98,X99) | s(bool,X97) = s(bool,X100) | s(bool,'T_0') = s(bool,X100)) )),
  inference(superposition,[],[f532,f836])).
fof(f4614,plain,(
  ( ! [X2,X1] : (sP12('T_0',X2,X1) | p(s(bool,X1))) ) | $spl42),
  inference(cnf_transformation,[],[f4614_D])).
fof(f4614_D,plain,(
  ( ! [X2,X1] : (sP12('T_0',X2,X1) | p(s(bool,X1))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f4706,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP12('T_0',X2,X1)) )),
  inference(condensation,[],[f4705])).
fof(f4705,plain,(
  ( ! [X94,X95,X93,X96] : (p(s(bool,X94)) | sP12('T_0',X95,X94) | s(bool,X93) = s(bool,X96) | s(bool,'T_0') = s(bool,X96)) )),
  inference(subsumption_resolution,[],[f4227,f902])).
fof(f4227,plain,(
  ( ! [X94,X95,X93,X96] : (~p(s(bool,X93)) | p(s(bool,X94)) | sP12('T_0',X95,X94) | s(bool,X93) = s(bool,X96) | s(bool,'T_0') = s(bool,X96)) )),
  inference(superposition,[],[f531,f836])).
fof(f4611,plain,(
  ( ! [X2,X1] : (~sP11(X2,X1,'T_0') | ~p(s(bool,X2)) | p(s(bool,X1))) ) | $spl40),
  inference(cnf_transformation,[],[f4611_D])).
fof(f4611_D,plain,(
  ( ! [X2,X1] : (~sP11(X2,X1,'T_0') | ~p(s(bool,X2)) | p(s(bool,X1))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f4704,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP11(X2,X1,'T_0')) )),
  inference(condensation,[],[f4703])).
fof(f4703,plain,(
  ( ! [X92,X90,X91,X89] : (p(s(bool,X90)) | ~p(s(bool,X91)) | ~sP11(X91,X90,'T_0') | s(bool,X89) = s(bool,X92) | s(bool,'T_0') = s(bool,X92)) )),
  inference(subsumption_resolution,[],[f4226,f902])).
fof(f4226,plain,(
  ( ! [X92,X90,X91,X89] : (~p(s(bool,X89)) | p(s(bool,X90)) | ~p(s(bool,X91)) | ~sP11(X91,X90,'T_0') | s(bool,X89) = s(bool,X92) | s(bool,'T_0') = s(bool,X92)) )),
  inference(superposition,[],[f526,f836])).
fof(f4608,plain,(
  ( ! [X2,X1] : (~sP10('T_0',X2,X1) | p(s(bool,X2)) | p(s(bool,X1))) ) | $spl38),
  inference(cnf_transformation,[],[f4608_D])).
fof(f4608_D,plain,(
  ( ! [X2,X1] : (~sP10('T_0',X2,X1) | p(s(bool,X2)) | p(s(bool,X1))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f4702,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP10('T_0',X2,X1)) )),
  inference(condensation,[],[f4701])).
fof(f4701,plain,(
  ( ! [X88,X87,X85,X86] : (p(s(bool,X86)) | p(s(bool,X87)) | ~sP10('T_0',X87,X86) | s(bool,X85) = s(bool,X88) | s(bool,'T_0') = s(bool,X88)) )),
  inference(subsumption_resolution,[],[f4225,f902])).
fof(f4225,plain,(
  ( ! [X88,X87,X85,X86] : (~p(s(bool,X85)) | p(s(bool,X86)) | p(s(bool,X87)) | ~sP10('T_0',X87,X86) | s(bool,X85) = s(bool,X88) | s(bool,'T_0') = s(bool,X88)) )),
  inference(superposition,[],[f516,f836])).
fof(f4605,plain,(
  ( ! [X2,X1] : (sP10('T_0',X1,X2) | ~p(s(bool,X1))) ) | $spl36),
  inference(cnf_transformation,[],[f4605_D])).
fof(f4605_D,plain,(
  ( ! [X2,X1] : (sP10('T_0',X1,X2) | ~p(s(bool,X1))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f4700,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP10('T_0',X1,X2)) )),
  inference(condensation,[],[f4699])).
fof(f4699,plain,(
  ( ! [X83,X81,X84,X82] : (~p(s(bool,X82)) | sP10('T_0',X82,X83) | s(bool,X81) = s(bool,X84) | s(bool,'T_0') = s(bool,X84)) )),
  inference(subsumption_resolution,[],[f4224,f902])).
fof(f4224,plain,(
  ( ! [X83,X81,X84,X82] : (~p(s(bool,X81)) | ~p(s(bool,X82)) | sP10('T_0',X82,X83) | s(bool,X81) = s(bool,X84) | s(bool,'T_0') = s(bool,X84)) )),
  inference(superposition,[],[f515,f836])).
fof(f4602,plain,(
  ( ! [X2,X1] : (sP10('T_0',X2,X1) | ~p(s(bool,X1))) ) | $spl34),
  inference(cnf_transformation,[],[f4602_D])).
fof(f4602_D,plain,(
  ( ! [X2,X1] : (sP10('T_0',X2,X1) | ~p(s(bool,X1))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f4698,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP10('T_0',X2,X1)) )),
  inference(condensation,[],[f4697])).
fof(f4697,plain,(
  ( ! [X80,X78,X79,X77] : (~p(s(bool,X78)) | sP10('T_0',X79,X78) | s(bool,X77) = s(bool,X80) | s(bool,'T_0') = s(bool,X80)) )),
  inference(subsumption_resolution,[],[f4223,f902])).
fof(f4223,plain,(
  ( ! [X80,X78,X79,X77] : (~p(s(bool,X77)) | ~p(s(bool,X78)) | sP10('T_0',X79,X78) | s(bool,X77) = s(bool,X80) | s(bool,'T_0') = s(bool,X80)) )),
  inference(superposition,[],[f514,f836])).
fof(f4599,plain,(
  ( ! [X2,X1] : (~sP9('T_0',X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) ) | $spl32),
  inference(cnf_transformation,[],[f4599_D])).
fof(f4599_D,plain,(
  ( ! [X2,X1] : (~sP9('T_0',X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f4696,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f4695])).
fof(f4695,plain,(
  ( ! [X76,X74,X75,X73] : (p(s(bool,X74)) | p(s(bool,X75)) | ~sP9('T_0',X74,X75) | s(bool,X73) = s(bool,X76) | s(bool,'T_0') = s(bool,X76)) )),
  inference(subsumption_resolution,[],[f4222,f902])).
fof(f4222,plain,(
  ( ! [X76,X74,X75,X73] : (~p(s(bool,X73)) | p(s(bool,X74)) | p(s(bool,X75)) | ~sP9('T_0',X74,X75) | s(bool,X73) = s(bool,X76) | s(bool,'T_0') = s(bool,X76)) )),
  inference(superposition,[],[f509,f836])).
fof(f4596,plain,(
  ( ! [X2,X1] : (~sP8(X2,X1,'T_0') | p(s(bool,X2)) | ~p(s(bool,X1))) ) | $spl30),
  inference(cnf_transformation,[],[f4596_D])).
fof(f4596_D,plain,(
  ( ! [X2,X1] : (~sP8(X2,X1,'T_0') | p(s(bool,X2)) | ~p(s(bool,X1))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f4694,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP8(X2,X1,'T_0')) )),
  inference(condensation,[],[f4693])).
fof(f4693,plain,(
  ( ! [X70,X72,X71,X69] : (~p(s(bool,X70)) | p(s(bool,X71)) | ~sP8(X71,X70,'T_0') | s(bool,X69) = s(bool,X72) | s(bool,'T_0') = s(bool,X72)) )),
  inference(subsumption_resolution,[],[f4221,f902])).
fof(f4221,plain,(
  ( ! [X70,X72,X71,X69] : (~p(s(bool,X69)) | ~p(s(bool,X70)) | p(s(bool,X71)) | ~sP8(X71,X70,'T_0') | s(bool,X69) = s(bool,X72) | s(bool,'T_0') = s(bool,X72)) )),
  inference(superposition,[],[f501,f836])).
fof(f4593,plain,(
  ( ! [X2,X1] : (~sP7(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl28),
  inference(cnf_transformation,[],[f4593_D])).
fof(f4593_D,plain,(
  ( ! [X2,X1] : (~sP7(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f4692,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X1,'T_0',X2)) )),
  inference(condensation,[],[f4691])).
fof(f4691,plain,(
  ( ! [X68,X66,X67,X65] : (~p(s(bool,X66)) | p(s(bool,X67)) | ~sP7(X66,'T_0',X67) | s(bool,X65) = s(bool,X68) | s(bool,'T_0') = s(bool,X68)) )),
  inference(subsumption_resolution,[],[f4220,f902])).
fof(f4220,plain,(
  ( ! [X68,X66,X67,X65] : (~p(s(bool,X65)) | ~p(s(bool,X66)) | p(s(bool,X67)) | ~sP7(X66,'T_0',X67) | s(bool,X65) = s(bool,X68) | s(bool,'T_0') = s(bool,X68)) )),
  inference(superposition,[],[f492,f836])).
fof(f4590,plain,(
  ( ! [X2,X1] : (~sP4(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl26),
  inference(cnf_transformation,[],[f4590_D])).
fof(f4590_D,plain,(
  ( ! [X2,X1] : (~sP4(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f4690,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X1,'T_0',X2)) )),
  inference(condensation,[],[f4689])).
fof(f4689,plain,(
  ( ! [X61,X64,X62,X63] : (~p(s(bool,X62)) | p(s(bool,X63)) | ~sP4(X62,'T_0',X63) | s(bool,X61) = s(bool,X64) | s(bool,'T_0') = s(bool,X64)) )),
  inference(subsumption_resolution,[],[f4219,f902])).
fof(f4219,plain,(
  ( ! [X61,X64,X62,X63] : (~p(s(bool,X61)) | ~p(s(bool,X62)) | p(s(bool,X63)) | ~sP4(X62,'T_0',X63) | s(bool,X61) = s(bool,X64) | s(bool,'T_0') = s(bool,X64)) )),
  inference(superposition,[],[f473,f836])).
fof(f4587,plain,(
  ( ! [X2,X1] : (~sP5(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl24),
  inference(cnf_transformation,[],[f4587_D])).
fof(f4587_D,plain,(
  ( ! [X2,X1] : (~sP5(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f4688,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP5(X1,'T_0',X2)) )),
  inference(condensation,[],[f4687])).
fof(f4687,plain,(
  ( ! [X59,X57,X60,X58] : (~p(s(bool,X58)) | p(s(bool,X59)) | ~sP5(X58,'T_0',X59) | s(bool,X57) = s(bool,X60) | s(bool,'T_0') = s(bool,X60)) )),
  inference(subsumption_resolution,[],[f4218,f902])).
fof(f4218,plain,(
  ( ! [X59,X57,X60,X58] : (~p(s(bool,X57)) | ~p(s(bool,X58)) | p(s(bool,X59)) | ~sP5(X58,'T_0',X59) | s(bool,X57) = s(bool,X60) | s(bool,'T_0') = s(bool,X60)) )),
  inference(superposition,[],[f469,f836])).
fof(f4584,plain,(
  ( ! [X2,X1] : (~sP4(X2,'T_0',X1) | sP6(X1,X2,'T_0') | ~sP5(X1,'T_0',X2)) ) | $spl22),
  inference(cnf_transformation,[],[f4584_D])).
fof(f4584_D,plain,(
  ( ! [X2,X1] : (~sP4(X2,'T_0',X1) | sP6(X1,X2,'T_0') | ~sP5(X1,'T_0',X2)) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f4686,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP6(X1,X2,'T_0') | ~sP5(X1,'T_0',X2) | ~sP4(X2,'T_0',X1)) )),
  inference(subsumption_resolution,[],[f4685,f589])).
fof(f4685,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP6(X1,X2,'T_0') | ~sP5(X1,'T_0',X2) | ~sP4(X2,'T_0',X1) | ~sP3('T_0',X2,X1)) )),
  inference(condensation,[],[f4684])).
fof(f4684,plain,(
  ( ! [X54,X56,X55,X53] : (sP6(X54,X55,'T_0') | ~sP5(X54,'T_0',X55) | ~sP4(X55,'T_0',X54) | ~sP3('T_0',X55,X54) | s(bool,X53) = s(bool,X56) | s(bool,'T_0') = s(bool,X56)) )),
  inference(subsumption_resolution,[],[f4217,f902])).
fof(f4217,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,X53)) | sP6(X54,X55,'T_0') | ~sP5(X54,'T_0',X55) | ~sP4(X55,'T_0',X54) | ~sP3('T_0',X55,X54) | s(bool,X53) = s(bool,X56) | s(bool,'T_0') = s(bool,X56)) )),
  inference(superposition,[],[f466,f836])).
fof(f4580,plain,(
  ( ! [X2,X1] : (~sP6(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl20),
  inference(cnf_transformation,[],[f4580_D])).
fof(f4580_D,plain,(
  ( ! [X2,X1] : (~sP6(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f4683,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X1,'T_0',X2)) )),
  inference(condensation,[],[f4682])).
fof(f4682,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,X50)) | p(s(bool,X51)) | ~sP6(X50,'T_0',X51) | s(bool,X49) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(subsumption_resolution,[],[f4216,f902])).
fof(f4216,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,X49)) | ~p(s(bool,X50)) | p(s(bool,X51)) | ~sP6(X50,'T_0',X51) | s(bool,X49) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(superposition,[],[f465,f836])).
fof(f4577,plain,(
  ( ! [X2,X1] : (sP2(X1,X2,'T_0') | ~p(s(bool,X1))) ) | $spl18),
  inference(cnf_transformation,[],[f4577_D])).
fof(f4577_D,plain,(
  ( ! [X2,X1] : (sP2(X1,X2,'T_0') | ~p(s(bool,X1))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f4681,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP2(X1,X2,'T_0')) )),
  inference(condensation,[],[f4680])).
fof(f4680,plain,(
  ( ! [X47,X45,X48,X46] : (~p(s(bool,X46)) | sP2(X46,X47,'T_0') | s(bool,X45) = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) )),
  inference(subsumption_resolution,[],[f4215,f902])).
fof(f4215,plain,(
  ( ! [X47,X45,X48,X46] : (~p(s(bool,X45)) | ~p(s(bool,X46)) | sP2(X46,X47,'T_0') | s(bool,X45) = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) )),
  inference(superposition,[],[f451,f836])).
fof(f4574,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1(X1,'T_0')) ) | $spl16),
  inference(cnf_transformation,[],[f4574_D])).
fof(f4574_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1(X1,'T_0')) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f4679,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f4678])).
fof(f4678,plain,(
  ( ! [X35,X36,X34] : (p(s(bool,X35)) | sP1(X35,'T_0') | s(bool,X34) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(subsumption_resolution,[],[f4212,f902])).
fof(f4212,plain,(
  ( ! [X35,X36,X34] : (~p(s(bool,X34)) | p(s(bool,X35)) | sP1(X35,'T_0') | s(bool,X34) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f412,f836])).
fof(f4570,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1('T_0',X1)) ) | $spl12),
  inference(cnf_transformation,[],[f4570_D])).
fof(f4570_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1('T_0',X1)) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f4677,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f4676])).
fof(f4676,plain,(
  ( ! [X33,X31,X32] : (p(s(bool,X32)) | sP1('T_0',X32) | s(bool,X31) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f4211,f902])).
fof(f4211,plain,(
  ( ! [X33,X31,X32] : (~p(s(bool,X31)) | p(s(bool,X32)) | sP1('T_0',X32) | s(bool,X31) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f409,f836])).
fof(f4625,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP9('T_0',X1,X2) | sP10('T_0',X1,X2)) )),
  inference(condensation,[],[f4624])).
fof(f4624,plain,(
  ( ! [X114,X112,X115,X113] : (~sP9('T_0',X113,X114) | sP10('T_0',X113,X114) | s(bool,X112) = s(bool,X115) | s(bool,'T_0') = s(bool,X115)) )),
  inference(subsumption_resolution,[],[f4070,f902])).
fof(f4070,plain,(
  ( ! [X114,X112,X115,X113] : (~p(s(bool,X112)) | ~sP9('T_0',X113,X114) | sP10('T_0',X113,X114) | s(bool,X112) = s(bool,X115) | s(bool,'T_0') = s(bool,X115)) )),
  inference(superposition,[],[f557,f836])).
fof(f4622,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP11('T_0',X1,X2) | sP12('T_0',X1,X2)) )),
  inference(condensation,[],[f4621])).
fof(f4621,plain,(
  ( ! [X111,X109,X110,X108] : (~sP11('T_0',X109,X110) | sP12('T_0',X109,X110) | s(bool,X108) = s(bool,X111) | s(bool,'T_0') = s(bool,X111)) )),
  inference(subsumption_resolution,[],[f4069,f902])).
fof(f4069,plain,(
  ( ! [X111,X109,X110,X108] : (~p(s(bool,X108)) | ~sP11('T_0',X109,X110) | sP12('T_0',X109,X110) | s(bool,X108) = s(bool,X111) | s(bool,'T_0') = s(bool,X111)) )),
  inference(superposition,[],[f556,f836])).
fof(f4619,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP12('T_0',X2,X1)) )),
  inference(condensation,[],[f4618])).
fof(f4618,plain,(
  ( ! [X103,X101,X102,X104] : (~p(s(bool,X102)) | p(s(bool,X103)) | ~sP12('T_0',X103,X102) | s(bool,X101) = s(bool,X104) | s(bool,'T_0') = s(bool,X104)) )),
  inference(subsumption_resolution,[],[f4067,f902])).
fof(f4067,plain,(
  ( ! [X103,X101,X102,X104] : (~p(s(bool,X101)) | ~p(s(bool,X102)) | p(s(bool,X103)) | ~sP12('T_0',X103,X102) | s(bool,X101) = s(bool,X104) | s(bool,'T_0') = s(bool,X104)) )),
  inference(superposition,[],[f533,f836])).
fof(f4616,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP12('T_0',X1,X2)) )),
  inference(condensation,[],[f4615])).
fof(f4615,plain,(
  ( ! [X99,X97,X100,X98] : (~p(s(bool,X98)) | sP12('T_0',X98,X99) | s(bool,X97) = s(bool,X100) | s(bool,'T_0') = s(bool,X100)) )),
  inference(subsumption_resolution,[],[f4066,f902])).
fof(f4066,plain,(
  ( ! [X99,X97,X100,X98] : (~p(s(bool,X97)) | ~p(s(bool,X98)) | sP12('T_0',X98,X99) | s(bool,X97) = s(bool,X100) | s(bool,'T_0') = s(bool,X100)) )),
  inference(superposition,[],[f532,f836])).
fof(f4613,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP12('T_0',X2,X1)) )),
  inference(condensation,[],[f4612])).
fof(f4612,plain,(
  ( ! [X94,X95,X93,X96] : (p(s(bool,X94)) | sP12('T_0',X95,X94) | s(bool,X93) = s(bool,X96) | s(bool,'T_0') = s(bool,X96)) )),
  inference(subsumption_resolution,[],[f4065,f902])).
fof(f4065,plain,(
  ( ! [X94,X95,X93,X96] : (~p(s(bool,X93)) | p(s(bool,X94)) | sP12('T_0',X95,X94) | s(bool,X93) = s(bool,X96) | s(bool,'T_0') = s(bool,X96)) )),
  inference(superposition,[],[f531,f836])).
fof(f4610,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP11(X2,X1,'T_0')) )),
  inference(condensation,[],[f4609])).
fof(f4609,plain,(
  ( ! [X92,X90,X91,X89] : (p(s(bool,X90)) | ~p(s(bool,X91)) | ~sP11(X91,X90,'T_0') | s(bool,X89) = s(bool,X92) | s(bool,'T_0') = s(bool,X92)) )),
  inference(subsumption_resolution,[],[f4064,f902])).
fof(f4064,plain,(
  ( ! [X92,X90,X91,X89] : (~p(s(bool,X89)) | p(s(bool,X90)) | ~p(s(bool,X91)) | ~sP11(X91,X90,'T_0') | s(bool,X89) = s(bool,X92) | s(bool,'T_0') = s(bool,X92)) )),
  inference(superposition,[],[f526,f836])).
fof(f4607,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP10('T_0',X2,X1)) )),
  inference(condensation,[],[f4606])).
fof(f4606,plain,(
  ( ! [X88,X87,X85,X86] : (p(s(bool,X86)) | p(s(bool,X87)) | ~sP10('T_0',X87,X86) | s(bool,X85) = s(bool,X88) | s(bool,'T_0') = s(bool,X88)) )),
  inference(subsumption_resolution,[],[f4063,f902])).
fof(f4063,plain,(
  ( ! [X88,X87,X85,X86] : (~p(s(bool,X85)) | p(s(bool,X86)) | p(s(bool,X87)) | ~sP10('T_0',X87,X86) | s(bool,X85) = s(bool,X88) | s(bool,'T_0') = s(bool,X88)) )),
  inference(superposition,[],[f516,f836])).
fof(f4604,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP10('T_0',X1,X2)) )),
  inference(condensation,[],[f4603])).
fof(f4603,plain,(
  ( ! [X83,X81,X84,X82] : (~p(s(bool,X82)) | sP10('T_0',X82,X83) | s(bool,X81) = s(bool,X84) | s(bool,'T_0') = s(bool,X84)) )),
  inference(subsumption_resolution,[],[f4062,f902])).
fof(f4062,plain,(
  ( ! [X83,X81,X84,X82] : (~p(s(bool,X81)) | ~p(s(bool,X82)) | sP10('T_0',X82,X83) | s(bool,X81) = s(bool,X84) | s(bool,'T_0') = s(bool,X84)) )),
  inference(superposition,[],[f515,f836])).
fof(f4601,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP10('T_0',X2,X1)) )),
  inference(condensation,[],[f4600])).
fof(f4600,plain,(
  ( ! [X80,X78,X79,X77] : (~p(s(bool,X78)) | sP10('T_0',X79,X78) | s(bool,X77) = s(bool,X80) | s(bool,'T_0') = s(bool,X80)) )),
  inference(subsumption_resolution,[],[f4061,f902])).
fof(f4061,plain,(
  ( ! [X80,X78,X79,X77] : (~p(s(bool,X77)) | ~p(s(bool,X78)) | sP10('T_0',X79,X78) | s(bool,X77) = s(bool,X80) | s(bool,'T_0') = s(bool,X80)) )),
  inference(superposition,[],[f514,f836])).
fof(f4598,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f4597])).
fof(f4597,plain,(
  ( ! [X76,X74,X75,X73] : (p(s(bool,X74)) | p(s(bool,X75)) | ~sP9('T_0',X74,X75) | s(bool,X73) = s(bool,X76) | s(bool,'T_0') = s(bool,X76)) )),
  inference(subsumption_resolution,[],[f4060,f902])).
fof(f4060,plain,(
  ( ! [X76,X74,X75,X73] : (~p(s(bool,X73)) | p(s(bool,X74)) | p(s(bool,X75)) | ~sP9('T_0',X74,X75) | s(bool,X73) = s(bool,X76) | s(bool,'T_0') = s(bool,X76)) )),
  inference(superposition,[],[f509,f836])).
fof(f4595,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP8(X2,X1,'T_0')) )),
  inference(condensation,[],[f4594])).
fof(f4594,plain,(
  ( ! [X70,X72,X71,X69] : (~p(s(bool,X70)) | p(s(bool,X71)) | ~sP8(X71,X70,'T_0') | s(bool,X69) = s(bool,X72) | s(bool,'T_0') = s(bool,X72)) )),
  inference(subsumption_resolution,[],[f4059,f902])).
fof(f4059,plain,(
  ( ! [X70,X72,X71,X69] : (~p(s(bool,X69)) | ~p(s(bool,X70)) | p(s(bool,X71)) | ~sP8(X71,X70,'T_0') | s(bool,X69) = s(bool,X72) | s(bool,'T_0') = s(bool,X72)) )),
  inference(superposition,[],[f501,f836])).
fof(f4592,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X1,'T_0',X2)) )),
  inference(condensation,[],[f4591])).
fof(f4591,plain,(
  ( ! [X68,X66,X67,X65] : (~p(s(bool,X66)) | p(s(bool,X67)) | ~sP7(X66,'T_0',X67) | s(bool,X65) = s(bool,X68) | s(bool,'T_0') = s(bool,X68)) )),
  inference(subsumption_resolution,[],[f4058,f902])).
fof(f4058,plain,(
  ( ! [X68,X66,X67,X65] : (~p(s(bool,X65)) | ~p(s(bool,X66)) | p(s(bool,X67)) | ~sP7(X66,'T_0',X67) | s(bool,X65) = s(bool,X68) | s(bool,'T_0') = s(bool,X68)) )),
  inference(superposition,[],[f492,f836])).
fof(f4589,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X1,'T_0',X2)) )),
  inference(condensation,[],[f4588])).
fof(f4588,plain,(
  ( ! [X61,X64,X62,X63] : (~p(s(bool,X62)) | p(s(bool,X63)) | ~sP4(X62,'T_0',X63) | s(bool,X61) = s(bool,X64) | s(bool,'T_0') = s(bool,X64)) )),
  inference(subsumption_resolution,[],[f4057,f902])).
fof(f4057,plain,(
  ( ! [X61,X64,X62,X63] : (~p(s(bool,X61)) | ~p(s(bool,X62)) | p(s(bool,X63)) | ~sP4(X62,'T_0',X63) | s(bool,X61) = s(bool,X64) | s(bool,'T_0') = s(bool,X64)) )),
  inference(superposition,[],[f473,f836])).
fof(f4586,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP5(X1,'T_0',X2)) )),
  inference(condensation,[],[f4585])).
fof(f4585,plain,(
  ( ! [X59,X57,X60,X58] : (~p(s(bool,X58)) | p(s(bool,X59)) | ~sP5(X58,'T_0',X59) | s(bool,X57) = s(bool,X60) | s(bool,'T_0') = s(bool,X60)) )),
  inference(subsumption_resolution,[],[f4056,f902])).
fof(f4056,plain,(
  ( ! [X59,X57,X60,X58] : (~p(s(bool,X57)) | ~p(s(bool,X58)) | p(s(bool,X59)) | ~sP5(X58,'T_0',X59) | s(bool,X57) = s(bool,X60) | s(bool,'T_0') = s(bool,X60)) )),
  inference(superposition,[],[f469,f836])).
fof(f4583,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP6(X1,X2,'T_0') | ~sP5(X1,'T_0',X2) | ~sP4(X2,'T_0',X1)) )),
  inference(subsumption_resolution,[],[f4582,f589])).
fof(f4582,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP6(X1,X2,'T_0') | ~sP5(X1,'T_0',X2) | ~sP4(X2,'T_0',X1) | ~sP3('T_0',X2,X1)) )),
  inference(condensation,[],[f4581])).
fof(f4581,plain,(
  ( ! [X54,X56,X55,X53] : (sP6(X54,X55,'T_0') | ~sP5(X54,'T_0',X55) | ~sP4(X55,'T_0',X54) | ~sP3('T_0',X55,X54) | s(bool,X53) = s(bool,X56) | s(bool,'T_0') = s(bool,X56)) )),
  inference(subsumption_resolution,[],[f4055,f902])).
fof(f4055,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,X53)) | sP6(X54,X55,'T_0') | ~sP5(X54,'T_0',X55) | ~sP4(X55,'T_0',X54) | ~sP3('T_0',X55,X54) | s(bool,X53) = s(bool,X56) | s(bool,'T_0') = s(bool,X56)) )),
  inference(superposition,[],[f466,f836])).
fof(f4579,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X1,'T_0',X2)) )),
  inference(condensation,[],[f4578])).
fof(f4578,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,X50)) | p(s(bool,X51)) | ~sP6(X50,'T_0',X51) | s(bool,X49) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(subsumption_resolution,[],[f4054,f902])).
fof(f4054,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,X49)) | ~p(s(bool,X50)) | p(s(bool,X51)) | ~sP6(X50,'T_0',X51) | s(bool,X49) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(superposition,[],[f465,f836])).
fof(f4576,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP2(X1,X2,'T_0')) )),
  inference(condensation,[],[f4575])).
fof(f4575,plain,(
  ( ! [X47,X45,X48,X46] : (~p(s(bool,X46)) | sP2(X46,X47,'T_0') | s(bool,X45) = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) )),
  inference(subsumption_resolution,[],[f4053,f902])).
fof(f4053,plain,(
  ( ! [X47,X45,X48,X46] : (~p(s(bool,X45)) | ~p(s(bool,X46)) | sP2(X46,X47,'T_0') | s(bool,X45) = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) )),
  inference(superposition,[],[f451,f836])).
fof(f4573,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f4572])).
fof(f4572,plain,(
  ( ! [X35,X36,X34] : (p(s(bool,X35)) | sP1(X35,'T_0') | s(bool,X34) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(subsumption_resolution,[],[f4050,f902])).
fof(f4050,plain,(
  ( ! [X35,X36,X34] : (~p(s(bool,X34)) | p(s(bool,X35)) | sP1(X35,'T_0') | s(bool,X34) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f412,f836])).
fof(f4569,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f4568])).
fof(f4568,plain,(
  ( ! [X33,X31,X32] : (p(s(bool,X32)) | sP1('T_0',X32) | s(bool,X31) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f4049,f902])).
fof(f4049,plain,(
  ( ! [X33,X31,X32] : (~p(s(bool,X31)) | p(s(bool,X32)) | sP1('T_0',X32) | s(bool,X31) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f409,f836])).
fof(f1557,plain,(
  s(sK14,sK15) = s(sK14,'const.min.@_1'(s(fun(sK14,bool),sK16))) | $spl8),
  inference(cnf_transformation,[],[f1557_D])).
fof(f1557_D,plain,(
  s(sK14,sK15) = s(sK14,'const.min.@_1'(s(fun(sK14,bool),sK16))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f1559,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | $spl11),
  inference(cnf_transformation,[],[f1559_D])).
fof(f1559_D,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f1520,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | s(sK14,sK15) = s(sK14,'const.min.@_1'(s(fun(sK14,bool),sK16))) | $spl6),
  inference(superposition,[],[f572,f364])).
fof(f572,plain,(
  ( ! [X5] : (~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,X5)))) | s(sK14,sK15) = s(sK14,X5)) ) | $spl6),
  inference(cnf_transformation,[],[f572_D])).
fof(f572_D,plain,(
  ( ! [X5] : (~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,X5)))) | s(sK14,sK15) = s(sK14,X5)) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f319,plain,(
  ( ! [X5] : (s(sK14,sK15) = s(sK14,X5) | ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,X5)))) | p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16))))) )),
  inference(cnf_transformation,[],[f206])).
fof(f567,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | $spl1),
  inference(cnf_transformation,[],[f567_D])).
fof(f567_D,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f571,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | $spl4),
  inference(cnf_transformation,[],[f571_D])).
fof(f571_D,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f320,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16))))),
  inference(cnf_transformation,[],[f206])).
fof(f569,plain,(
  s(sK14,sK15) != s(sK14,sK17) | $spl3),
  inference(cnf_transformation,[],[f569_D])).
fof(f569_D,plain,(
  s(sK14,sK15) != s(sK14,sK17) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f321,plain,(
  s(sK14,sK15) != s(sK14,sK17) | ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16))))),
  inference(cnf_transformation,[],[f206])).
% SZS output end Proof for thm.quantHeuristics
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 1755
% Passive clauses: 7416
% Generated clauses: 49333
% Final active clauses: 1396
% Final passive clauses: 2973
% Input formulas: 55
% Initial clauses: 226
% 
% Duplicate literals: 424
% Trivial inequalities: 43
% Fw subsumption resolutions: 2005
% Fw demodulations: 7184
% Bw demodulations: 2421
% Condensations: 1121
% Global subsumptions: 48
% 
% Simple tautologies: 111
% Equational tautologies: 84
% Forward subsumptions: 29396
% Bw demodulations to eq. taut.: 13
% 
% Unit resulting resolution: 23833
% Factoring: 1
% Forward superposition: 8419
% Backward superposition: 3273
% Self superposition: 29
% Equality factoring: 28
% Equality resolution: 1
% 
% Split clauses: 118
% Split components: 241
% SAT solver clauses: 9085
% SAT solver unit clauses: 1657
% SAT solver binary clauses: 4204
% 
% Sat splits: 118
% Sat splitting refutations: 1383
% 
% Memory used [KB]: 28656
% Time elapsed: 2.406 s
% ------------------------------
% ------------------------------
% Success in time 3.699 s
