Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.pred_set
% (67454)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 553
% Passive clauses: 9478
% Generated clauses: 17379
% Final active clauses: 551
% Final passive clauses: 8425
% Input formulas: 47
% Initial clauses: 199
% 
% Duplicate literals: 300
% Trivial inequalities: 48
% Fw subsumption resolutions: 1452
% Fw demodulations: 70
% Bw demodulations: 1
% 
% Simple tautologies: 232
% Equational tautologies: 58
% Forward subsumptions: 5235
% 
% Binary resolution: 10621
% Factoring: 1023
% Forward superposition: 1229
% Backward superposition: 2229
% Self superposition: 34
% Equality factoring: 34
% Equality resolution: 18
% 
% Split clauses: 561
% Split components: 1125
% SAT solver clauses: 859
% SAT solver unit clauses: 18
% SAT solver binary clauses: 167
% 
% Sat splits: 561
% 
% TWLsolver clauses: 858
% TWLsolver calls for satisfiability: 327
% 
% Memory used [KB]: 8955
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.pred_set
% (67476)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 455
% Passive clauses: 3125
% Generated clauses: 13868
% Final active clauses: 356
% Final passive clauses: 1937
% Input formulas: 47
% Initial clauses: 199
% 
% Duplicate literals: 241
% Trivial inequalities: 52
% Fw subsumption resolutions: 1153
% Fw demodulations: 387
% Bw demodulations: 628
% Condensations: 494
% Global subsumptions: 17
% 
% Simple tautologies: 289
% Equational tautologies: 19
% Forward subsumptions: 4578
% Bw demodulations to eq. taut.: 148
% 
% Binary resolution: 4102
% Factoring: 145
% Forward superposition: 1485
% Backward superposition: 1482
% Self superposition: 26
% Equality factoring: 30
% Equality resolution: 13
% 
% Split clauses: 185
% Split components: 380
% SAT solver clauses: 3618
% SAT solver unit clauses: 203
% SAT solver binary clauses: 90
% SAT solver learnt clauses: 34
% SAT solver learnt literals: 10
% 
% Sat splits: 185
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 3550
% TWLsolver calls for satisfiability: 1918
% 
% Memory used [KB]: 13688
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.pred_set
% (67513)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1589
% Passive clauses: 9690
% Generated clauses: 104450
% Final active clauses: 490
% Final passive clauses: 861
% Input formulas: 47
% Initial clauses: 199
% 
% Duplicate literals: 7593
% Trivial inequalities: 275
% Fw subsumption resolutions: 4680
% Fw demodulations: 10120
% Bw demodulations: 4757
% Condensations: 9931
% Global subsumptions: 689
% 
% Simple tautologies: 396
% Equational tautologies: 1525
% Forward subsumptions: 56764
% Fw demodulations to eq. taut.: 3
% Bw demodulations to eq. taut.: 741
% 
% Unit resulting resolution: 14081
% Factoring: 3
% Forward superposition: 26951
% Backward superposition: 17208
% Self superposition: 364
% Equality factoring: 232
% Equality resolution: 14
% 
% Split clauses: 699
% Split components: 1416
% SAT solver clauses: 13136
% SAT solver unit clauses: 679
% SAT solver binary clauses: 1616
% 
% Sat splits: 699
% Sat splitting refutations: 1356
% 
% Memory used [KB]: 54370
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.pred_set
% (67894)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1976
% Passive clauses: 57743
% Generated clauses: 205080
% Final active clauses: 1634
% Final passive clauses: 52630
% Input formulas: 47
% Initial clauses: 197
% 
% Selected by SInE selection: 45
% SInE iterations: 1
% 
% Duplicate literals: 7867
% Trivial inequalities: 974
% Fw subsumption resolutions: 21010
% Fw demodulations: 2152
% Bw demodulations: 276
% 
% Simple tautologies: 4720
% Equational tautologies: 1516
% Forward subsumptions: 101398
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 153199
% Factoring: 778
% Forward superposition: 13267
% Backward superposition: 3833
% Self superposition: 9
% Equality factoring: 2
% Equality resolution: 42
% 
% Split clauses: 9154
% Split components: 19940
% SAT solver clauses: 9659
% SAT solver unit clauses: 32
% SAT solver binary clauses: 4460
% 
% Sat splits: 9154
% Sat splitting refutations: 78
% 
% Memory used [KB]: 52834
% Time elapsed: 4.0000 s
% ------------------------------
% ------------------------------
% remaining time: 142 next slice time: 6
lrs+1002_2:3_br=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:sas=minisat:stl=30:sac=on:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none:sp=reverse_arity:urr=on_4 on thm.pred_set
% (67994)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 353
% Passive clauses: 2555
% Generated clauses: 14133
% Final active clauses: 347
% Final passive clauses: 2126
% Input formulas: 47
% Initial clauses: 199
% 
% Duplicate literals: 576
% Trivial inequalities: 225
% Fw subsumption resolutions: 1956
% Fw demodulations: 48
% Bw demodulations: 29
% Global subsumptions: 26
% 
% Simple tautologies: 622
% Equational tautologies: 49
% Forward subsumptions: 7712
% 
% Unit resulting resolution: 1225
% Factoring: 18
% Forward superposition: 4567
% Backward superposition: 5094
% Self superposition: 92
% Equality factoring: 6
% Equality resolution: 22
% 
% Split clauses: 19
% Split components: 38
% SAT solver clauses: 2653
% SAT solver unit clauses: 198
% SAT solver binary clauses: 127
% 
% Sat splits: 19
% Sat splitting refutations: 51
% 
% Memory used [KB]: 7931
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 135 next slice time: 15
lrs+1011_12_bs=on:bsr=unit_only:cond=on:gs=on:gsaa=from_current:gsssp=full:nwc=1.1:sas=minisat:stl=60:sos=all:sac=on:sdd=off:sser=off:sfr=on:ssfp=100000:ssfq=1.2:ssnc=none:sp=reverse_arity:updr=off_14 on thm.pred_set
% (68037)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 986
% Passive clauses: 5372
% Generated clauses: 14959
% Final active clauses: 763
% Final passive clauses: 4225
% Input formulas: 47
% Initial clauses: 199
% 
% Duplicate literals: 760
% Trivial inequalities: 68
% Fw subsumption resolutions: 705
% Bw subsumption resolutions: 230
% Fw demodulations: 76
% Bw demodulations: 1
% Condensations: 1491
% Global subsumptions: 2
% 
% Simple tautologies: 1154
% Equational tautologies: 44
% Forward subsumptions: 5232
% Backward subsumptions: 276
% 
% Binary resolution: 7837
% Factoring: 42
% Forward superposition: 2014
% Backward superposition: 1269
% Self superposition: 40
% Equality factoring: 2
% Equality resolution: 20
% 
% Split clauses: 2
% Split components: 4
% SAT solver clauses: 5403
% SAT solver unit clauses: 91
% SAT solver binary clauses: 577
% 
% Sat splits: 2
% Sat splitting refutations: 25
% 
% Memory used [KB]: 10362
% Time elapsed: 1.600 s
% ------------------------------
% ------------------------------
% remaining time: 119 next slice time: 9
lrs+10_8:1_bd=preordered:bs=on:ccuc=first:fsr=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:nicw=on:sas=minisat:stl=120:sos=on:sscc=on:sser=off:sfr=on:ssfp=4000:ssfq=1.0:ssnc=none:sp=reverse_arity:urr=on_7 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f5927,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f261,f466,f260,f468,f474,f473,f477,f472,f479,f471,f470,f481,f469,f483,f475,f485,f500,f499,f498,f497,f506,f496,f508,f495,f494,f493,f492,f510,f491,f490,f489,f488,f502,f514,f515,f518,f516,f484,f681,f680,f677,f676,f675,f673,f672,f671,f722,f670,f720,f669,f668,f724,f660,f726,f658,f657,f656,f730,f655,f728,f654,f653,f732,f649,f734,f634,f736,f631,f630,f629,f740,f628,f738,f627,f626,f742,f623,f622,f621,f620,f619,f618,f617,f616,f615,f750,f614,f613,f612,f746,f611,f610,f609,f608,f607,f606,f754,f605,f744,f604,f603,f752,f602,f601,f600,f599,f748,f598,f597,f756,f595,f594,f593,f760,f592,f758,f591,f590,f762,f581,f580,f579,f766,f578,f764,f577,f576,f768,f574,f573,f572,f772,f571,f770,f570,f569,f774,f568,f567,f566,f778,f565,f776,f564,f563,f780,f562,f561,f560,f559,f558,f557,f556,f555,f554,f790,f553,f788,f552,f551,f550,f786,f549,f782,f548,f547,f784,f541,f540,f538,f537,f536,f535,f533,f532,f530,f529,f523,f522,f718,f521,f520,f683,f684,f793,f792,f812,f811,f820,f810,f822,f802,f801,f824,f800,f826,f839,f838,f837,f836,f844,f835,f846,f834,f833,f832,f831,f848,f830,f850,f829,f852,f828,f827,f861,f860,f859,f863,f872,f864,f870,f865,f868,f882,f881,f880,f884,f893,f885,f891,f886,f889,f958,f957,f961,f968,f967,f966,f971,f980,f972,f978,f973,f976,f1195,f1194,f1193,f1197,f1207,f1198,f1205,f1199,f1203,f1246,f1252,f1254,f1322,f1328,f1333,f1339,f1345,f1350,f1375,f1379,f1385,f1389,f1395,f1399,f1405,f1409,f1472,f1471,f1575,f1457,f1456,f1585,f1821,f1827,f1834,f1854,f1858,f1866,f1870,f1930,f1929,f1928,f1932,f1942,f1933,f1940,f1934,f1938,f1950,f1949,f1948,f1952,f1961,f1953,f1959,f1954,f1957,f1963,f1964,f1944,f1927,f1924,f1887,f1966,f842,f1968,f1973,f2029,f2134,f2028,f2027,f2136,f2012,f2148,f2011,f2010,f2150,f2065,f2066,f2164,f2160,f2068,f2163,f2069,f2165,f2070,f2071,f2072,f2167,f2073,f2166,f2074,f2168,f2075,f2169,f2078,f2170,f2081,f2171,f2082,f2172,f2083,f2173,f2089,f2174,f2096,f2175,f2097,f2176,f2098,f2177,f2103,f2178,f2104,f2179,f2106,f2180,f2108,f2181,f2109,f2182,f2110,f2183,f1997,f2253,f2326,f2254,f2327,f2256,f2328,f2257,f2329,f2261,f2330,f2267,f2331,f2268,f2332,f2269,f2333,f2270,f2334,f2271,f2335,f2272,f2336,f2276,f2337,f2277,f2338,f2278,f2339,f2279,f2340,f2282,f2341,f2284,f2347,f2373,f845,f2372,f2371,f843,f2354,f849,f2353,f2352,f847,f2412,f2502,f2413,f2414,f2503,f2415,f2504,f2416,f2417,f2418,f2506,f2419,f2505,f2420,f2507,f2421,f2508,f2424,f2509,f2427,f2510,f2428,f2511,f2429,f2512,f2435,f2513,f2442,f2514,f2443,f2515,f2444,f2516,f2449,f2517,f2450,f2518,f2452,f2519,f2454,f2520,f2455,f2521,f2456,f2522,f2527,f2528,f2529,f347,f343,f2530,f2531,f487,f2537,f2541,f2757,f2756,f2755,f2759,f2768,f2760,f2766,f2761,f2764,f2776,f2775,f2774,f2778,f2787,f2779,f2785,f2780,f2783,f2805,f2914,f2804,f2803,f2916,f2793,f2925,f2792,f2791,f2927,f2931,f2854,f2940,f2855,f2941,f2857,f2942,f2858,f2943,f2862,f2944,f2868,f2945,f2869,f2946,f2870,f2947,f2871,f2948,f2872,f2949,f2873,f2950,f2877,f2951,f2878,f2952,f2879,f2953,f2880,f2954,f2883,f2955,f2885,f2956,f2586,f2961,f2585,f504,f2968,f2970,f3295,f3301,f3312,f3313,f3319,f3318,f3317,f3322,f3316,f1941,f3358,f3365,f3376,f3375,f3383,f3390,f3389,f3397,f3414,f3413,f2135,f3412,f2133,f3409,f3408,f2149,f3407,f2147,f3485,f3493,f3484,f3495,f3483,f3497,f3482,f3499,f3502,f3503,f3519,f3532,f3548,f3574,f3573,f3578,f3586,f3579,f3584,f3594,f3598,f3602,f3609,f3608,f3615,f3623,f3622,f3626,f3636,f3635,f3639,f3634,f3633,f3641,f3658,f3653,f2784,f3673,f3672,f3671,f3670,f3680,f3692,f3691,f3699,f3708,f3707,f3715,f3726,f3725,f3724,f3723,f3733,f3739,f3738,f3756,f3752,f3758,f3772,f3775,f3786,f3803,f3799,f3804,f3818,f3821,f3757,f3823,f3822,f3826,f3829,f3837,f3841,f3849,f3854,f3853,f3859,f3933,f3949,f3989,f3986,f3995,f3996,f3984,f3998,f3999,f2162,f3982,f3981,f3980,f3979,f3978,f3977,f4001,f3976,f3975,f4003,f4004,f4005,f4007,f4009,f3784,f3780,f4011,f3287,f3282,f4013,f4051,f4057,f4055,f512,f1958,f1960,f4062,f4061,f4066,f4074,f4067,f4072,f4079,f4083,f4087,f4091,f4090,f4103,f4089,f4088,f4105,f4107,f4106,f4121,f4123,f4135,f4122,f4137,f4138,f4151,f4182,f4184,f4202,f4201,f4216,f4569,f4568,f4567,f4606,f4566,f4565,f4564,f4612,f4560,f4559,f4558,f4557,f4556,f4555,f4554,f4620,f4553,f4622,f4552,f4610,f4551,f4550,f4608,f4549,f4624,f4541,f4540,f4539,f4537,f4536,f4535,f4534,f4632,f4533,f4634,f4532,f4524,f4523,f4522,f4521,f4517,f4514,f4513,f4512,f4511,f4510,f4509,f4508,f4507,f4506,f4503,f4502,f4501,f4644,f4499,f4642,f4498,f4497,f4647,f4488,f4650,f4483,f4482,f4481,f4657,f4477,f4655,f4476,f4475,f4662,f4473,f4665,f4448,f4672,f4427,f4678,f4414,f4413,f4412,f4686,f4411,f4684,f4410,f4409,f4688,f4406,f4405,f4404,f4403,f4402,f4401,f4400,f4399,f4398,f4696,f4397,f4396,f4395,f4692,f4394,f4393,f4392,f4391,f4390,f4389,f4700,f4388,f4690,f4387,f4386,f4698,f4385,f4384,f4383,f4382,f4694,f4381,f4380,f4702,f4375,f4374,f4373,f4709,f4369,f4707,f4368,f4367,f4714,f4362,f4361,f4360,f4618,f4358,f4357,f4356,f4630,f4355,f4614,f4354,f4353,f4616,f4351,f4626,f4350,f4349,f4628,f4346,f4718,f4345,f4720,f4344,f4722,f4343,f4342,f4726,f4341,f4728,f4340,f4730,f4338,f4732,f4337,f4734,f4336,f4736,f4327,f4326,f4325,f4743,f4321,f4741,f4320,f4319,f4748,f4314,f4313,f4312,f4755,f4308,f4753,f4307,f4306,f4760,f4305,f4762,f4304,f4764,f4303,f4766,f4302,f4768,f4301,f4300,f4299,f4772,f4298,f4770,f4297,f4296,f4724,f4295,f4294,f4293,f4292,f4291,f4290,f4289,f4288,f4287,f4782,f4286,f4780,f4285,f4284,f4283,f4778,f4282,f4774,f4281,f4280,f4776,f4275,f4274,f4272,f4271,f4270,f4269,f4267,f4266,f4264,f4263,f4257,f4256,f4636,f4255,f4254,f4785,f4786,f4572,f4573,f4787,f4784,f3062,f1970,f2342,f5387,f4792,f4794,f4796,f4798,f4800,f4803,f4805,f4807,f4809,f4811,f4813,f5428,f5429,f4821,f4826,f4831,f5430,f5431,f4840,f4845,f4847,f4849,f4851,f4876,f4880,f4896,f4906,f4911,f4928,f4935,f4940,f4956,f4959,f4961,f4963,f4973,f4975,f4977,f4980,f4982,f4984,f4992,f4994,f4996,f4998,f5006,f5010,f5015,f5022,f5026,f5031,f5060,f5123,f5132,f5139,f5141,f5147,f5151,f5153,f5155,f5161,f5165,f5167,f5169,f5182,f5187,f5191,f5195,f5205,f5233,f5238,f5243,f5245,f5264,f5432,f5433,f5434,f5435,f5436,f5437,f5438,f5439,f5440,f5441,f5442,f5443,f5444,f5445,f5446,f5447,f5448,f5449,f5317,f5319,f5325,f5450,f5329,f5331,f5336,f5338,f5343,f5345,f5347,f5349,f5351,f5353,f5355,f5357,f5451,f5452,f5453,f5454,f5455,f5456,f5457,f5458,f5459,f5460,f5461,f5462,f5463,f5464,f5465,f5466,f5467,f5468,f5469,f5470,f5471,f5472,f5473,f5474,f5475,f5476,f5477,f5478,f5479,f5480,f5482,f5483,f5484,f5485,f5486,f5487,f5516,f5517,f5518,f5519,f5521,f5522,f5523,f5524,f5525,f5527,f5529,f5530,f5531,f5532,f5534,f5536,f5537,f5538,f5539,f5546,f5547,f5548,f5551,f5559,f5561,f5562,f5563,f5564,f5565,f5566,f5567,f5568,f5569,f5570,f5571,f5572,f5574,f5575,f5577,f5580,f5583,f5584,f5585,f5586,f5588,f5589,f5590,f5591,f5592,f5594,f5595,f5596,f5597,f5598,f5600,f5601,f5602,f5603,f5605,f5606,f5607,f5608,f5613,f5614,f5615,f5619,f5620,f5621,f5622,f5625,f5628,f5631,f5634,f5635,f5636,f5637,f5638,f5639,f5640,f5641,f5642,f5643,f5644,f5645,f5646,f5647,f5649,f5650,f5651,f5652,f5654,f5656,f5658,f5659,f5660,f5661,f5663,f5665,f5666,f5667,f5668,f5670,f5671,f5672,f5673,f5675,f5676,f5677,f5678,f5679,f5680,f5681,f5682,f5683,f5684,f5686,f5688,f5690,f5692,f5693,f5695,f5696,f5697,f5704,f5707,f5710,f5718,f5719,f5720,f5721,f5722,f5723,f5724,f5725,f5726,f5727,f5729,f5730,f5731,f5732,f5733,f5734,f5735,f5736,f5737,f5738,f5739,f5741,f5742,f5743,f5744,f5745,f5746,f5747,f5748,f5749,f5750,f5751,f5753,f5755,f5756,f5757,f5758,f5759,f5760,f5761,f5762,f5763,f5764,f5765,f5766,f5767,f5768,f5769,f5770,f5771,f5772,f5773,f5775,f5776,f5777,f5778,f5785,f5788,f5791,f5795,f5796,f5798,f5799,f5800,f5807,f5810,f5813,f5817,f5818,f5819,f5820,f5821,f5822,f5823,f5824,f5825,f5826,f5827,f5828,f5829,f5832,f5833,f5835,f5837,f5838,f5839,f5840,f5842,f5843,f5844,f5845,f5847,f5848,f5849,f5850,f5851,f5852,f5853,f5854,f5856,f5857,f5858,f5859,f5861,f5862,f5863,f5864,f5865,f5866,f5867,f5868,f5869,f5870,f5873,f5874,f5875,f5876,f5877,f5878,f5879,f5880,f5881,f5882,f5883,f5884,f5885,f5886,f5887,f5888,f5889,f5890,f5891,f5892,f5893,f5894,f5895,f5896,f5897,f5898,f5899,f5900,f5901,f5902,f5903,f5904,f5905,f5906,f5907,f5908,f5909,f5910,f5911,f5912,f5913,f5914,f5915,f5916,f5917,f5918,f5919,f5920,f5921,f5922,f5923,f5924,f5925,f5926])).
fof(f5926,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853,f4850,f4848,f4846,f4842,f4841,f4837,f4836,f4828,f4827,f4823,f4822,f4812,f4810,f4808,f4806,f4804,f4802,f4801,f4799,f4797,f4795,f4793,f5385,f5386,f5388,f5389,f5390,f5391,f5392,f5393,f5394,f5395,f5400,f5401,f5402,f5403,f5404,f5405,f5406,f5407,f5408,f5409,f5410,f5411,f5412,f5413,f5414,f5415,f5416,f5417,f5418,f5419,f5420,f5421,f5422,f5423,f5424,f5425,f5426])).
fof(f5426,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl190),
  inference(superposition,[],[f1970,f269])).
fof(f269,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.pred_set.inj_image_countable_IFF',aHL_BOOL_CASES)).
fof(f5425,plain,(
  ( ! [X47,X48] : (~p(s(bool,X47)) | sP11('F_0',X47,X48)) ) | $spl190),
  inference(resolution,[],[f1970,f446])).
fof(f446,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f257])).
fof(f257,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f256])).
fof(f256,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f166])).
fof(f166,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP11(X2,X0,X1))),
  inference(definition_folding,[],[f144,f165,f164])).
fof(f164,plain,(
  ! [X2,X0,X1] : (sP10(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f165,plain,(
  ! [X2,X0,X1] : (sP11(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f144,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f114])).
fof(f114,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f113])).
fof(f113,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X23,X24,X25] : ((p(s(bool,X25)) <=> (p(s(bool,X24)) => p(s(bool,X23)))) <=> ((p(s(bool,X25)) | p(s(bool,X24))) & (p(s(bool,X25)) | ~p(s(bool,X23))) & (~p(s(bool,X24)) | p(s(bool,X23)) | ~p(s(bool,X25)))))),
  file('thm.pred_set.inj_image_countable_IFF','thm.sat.sat_27_prove_ltb')).
fof(f5424,plain,(
  ( ! [X45,X46] : (p(s(bool,X45)) | sP11('F_0',X46,X45)) ) | $spl190),
  inference(resolution,[],[f1970,f445])).
fof(f445,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f257])).
fof(f5423,plain,(
  ( ! [X43,X44] : (p(s(bool,X43)) | sP11(X44,X43,'F_0') | p(s(bool,X44))) ) | $spl190),
  inference(resolution,[],[f1970,f444])).
fof(f444,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP11(X2,X0,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f257])).
fof(f5422,plain,(
  ( ! [X41,X42] : (sP10(X41,'F_0',X42)) ) | $spl190),
  inference(resolution,[],[f1970,f442])).
fof(f442,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f255])).
fof(f255,plain,(
  ! [X0,X1,X2] : ((~sP10(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP10(X0,X1,X2)))),
  inference(rectify,[],[f254])).
fof(f254,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(flattening,[],[f253])).
fof(f253,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(nnf_transformation,[],[f164])).
fof(f5421,plain,(
  ( ! [X39,X40] : (~p(s(bool,X39)) | sP9('F_0',X39,X40)) ) | $spl190),
  inference(resolution,[],[f1970,f429])).
fof(f429,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f249])).
fof(f249,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f248])).
fof(f248,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f163])).
fof(f163,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP9(X2,X0,X1))),
  inference(definition_folding,[],[f112,f162,f161])).
fof(f161,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f162,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f112,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f111])).
fof(f111,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X23,X24,X25] : ((p(s(bool,X25)) <=> (p(s(bool,X24)) | p(s(bool,X23)))) <=> ((p(s(bool,X25)) | ~p(s(bool,X24))) & (p(s(bool,X25)) | ~p(s(bool,X23))) & (p(s(bool,X24)) | p(s(bool,X23)) | ~p(s(bool,X25)))))),
  file('thm.pred_set.inj_image_countable_IFF','thm.sat.sat_25_prove_ltb')).
fof(f5420,plain,(
  ( ! [X37,X38] : (~p(s(bool,X37)) | sP9('F_0',X38,X37)) ) | $spl190),
  inference(resolution,[],[f1970,f428])).
fof(f428,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f249])).
fof(f5419,plain,(
  ( ! [X35,X36] : (sP8(X35,'F_0',X36)) ) | $spl190),
  inference(resolution,[],[f1970,f425])).
fof(f425,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f247])).
fof(f247,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f246])).
fof(f246,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f245])).
fof(f245,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f161])).
fof(f5418,plain,(
  ( ! [X33,X34] : (sP8(X33,X34,'F_0')) ) | $spl190),
  inference(resolution,[],[f1970,f424])).
fof(f424,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f247])).
fof(f5417,plain,(
  ( ! [X31,X32] : (~p(s(bool,X31)) | ~p(s(bool,X32)) | sP7('F_0',X31,X32)) ) | $spl190),
  inference(resolution,[],[f1970,f412])).
fof(f412,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | sP7(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f241])).
fof(f241,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f240])).
fof(f240,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f160])).
fof(f160,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP7(X2,X0,X1))),
  inference(definition_folding,[],[f110,f159,f158])).
fof(f158,plain,(
  ! [X0,X1,X2] : (sP6(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f159,plain,(
  ! [X2,X0,X1] : (sP7(X2,X0,X1) <=> (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f110,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f109])).
fof(f109,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X23,X24,X25] : ((p(s(bool,X25)) <=> (p(s(bool,X24)) & p(s(bool,X23)))) <=> ((p(s(bool,X25)) | ~p(s(bool,X24)) | ~p(s(bool,X23))) & (p(s(bool,X24)) | ~p(s(bool,X25))) & (p(s(bool,X23)) | ~p(s(bool,X25)))))),
  file('thm.pred_set.inj_image_countable_IFF','thm.sat.sat_23_prove_ltb')).
fof(f5416,plain,(
  ( ! [X30,X29] : (sP6(X29,X30,'F_0')) ) | $spl190),
  inference(resolution,[],[f1970,f407])).
fof(f407,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f239])).
fof(f239,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(flattening,[],[f238])).
fof(f238,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(nnf_transformation,[],[f158])).
fof(f5415,plain,(
  ( ! [X28,X27] : (sP2('F_0',X27,X28)) ) | $spl190),
  inference(resolution,[],[f1970,f394])).
fof(f394,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f233])).
fof(f233,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(flattening,[],[f232])).
fof(f232,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(nnf_transformation,[],[f153])).
fof(f153,plain,(
  ! [X0,X1,X2] : (sP2(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f5414,plain,(
  ( ! [X26,X25] : (sP2(X25,'F_0',X26)) ) | $spl190),
  inference(resolution,[],[f1970,f393])).
fof(f393,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f233])).
fof(f5413,plain,(
  ( ! [X24,X23] : (sP2(X23,X24,'F_0')) ) | $spl190),
  inference(resolution,[],[f1970,f392])).
fof(f392,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f233])).
fof(f5412,plain,(
  ( ! [X21,X22] : (sP3(X21,X22,'F_0')) ) | $spl190),
  inference(resolution,[],[f1970,f388])).
fof(f388,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f231])).
fof(f231,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f230])).
fof(f230,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(flattening,[],[f229])).
fof(f229,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(nnf_transformation,[],[f154])).
fof(f154,plain,(
  ! [X1,X0,X2] : (sP3(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f5411,plain,(
  ( ! [X19,X20] : (sP4(X19,X20,'F_0')) ) | $spl190),
  inference(resolution,[],[f1970,f384])).
fof(f384,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f228])).
fof(f228,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f227])).
fof(f227,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(flattening,[],[f226])).
fof(f226,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(nnf_transformation,[],[f155])).
fof(f155,plain,(
  ! [X2,X0,X1] : (sP4(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f5410,plain,(
  ( ! [X18] : (~p(s(bool,X18)) | ~sP1('F_0',X18)) ) | $spl190),
  inference(resolution,[],[f1970,f366])).
fof(f366,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | ~sP1(X1,X0)) )),
  inference(cnf_transformation,[],[f219])).
fof(f219,plain,(
  ! [X0,X1] : ((((p(s(bool,X1)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)))) | sP1(X1,X0)) & (~sP1(X1,X0) | ((~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X0)) | p(s(bool,X1))))))),
  inference(nnf_transformation,[],[f152])).
fof(f152,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> sP1(X1,X0))),
  inference(definition_folding,[],[f101,f151])).
fof(f151,plain,(
  ! [X1,X0] : (sP1(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f101,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  inference(flattening,[],[f100])).
fof(f100,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X24,X25] : ((p(s(bool,X25)) <=> ~p(s(bool,X24))) <=> ((p(s(bool,X25)) | p(s(bool,X24))) & (~p(s(bool,X24)) | ~p(s(bool,X25)))))),
  file('thm.pred_set.inj_image_countable_IFF','thm.sat.sat_29_prove_ltb')).
fof(f5409,plain,(
  ( ! [X17] : (p(s(bool,X17)) | sP1('F_0',X17)) ) | $spl190),
  inference(resolution,[],[f1970,f365])).
fof(f365,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP1(X1,X0)) )),
  inference(cnf_transformation,[],[f219])).
fof(f5408,plain,(
  ( ! [X16] : (sP1(X16,'F_0') | p(s(bool,X16))) ) | $spl190),
  inference(resolution,[],[f1970,f364])).
fof(f364,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP1(X1,X0) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f219])).
fof(f5407,plain,(
  ( ! [X15] : (p(s(bool,X15)) | sP1(X15,'F_0')) ) | $spl190),
  inference(resolution,[],[f1970,f363])).
fof(f363,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f218])).
fof(f218,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1)))),
  inference(rectify,[],[f217])).
fof(f217,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0)))),
  inference(flattening,[],[f216])).
fof(f216,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP1(X1,X0)))),
  inference(nnf_transformation,[],[f151])).
fof(f5406,plain,(
  ( ! [X14] : (p(s(bool,X14)) | sP1('F_0',X14)) ) | $spl190),
  inference(resolution,[],[f1970,f360])).
fof(f360,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f218])).
fof(f5405,plain,(
  ( ! [X13] : (~p(s(bool,X13)) | ~sP1(X13,'F_0')) ) | $spl190),
  inference(resolution,[],[f1970,f359])).
fof(f359,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | ~sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f218])).
fof(f5404,plain,(
  ( ! [X12] : (sP0('F_0',X12)) ) | $spl190),
  inference(resolution,[],[f1970,f350])).
fof(f350,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f213])).
fof(f213,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f212])).
fof(f212,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0')))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f150])).
fof(f150,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(definition_folding,[],[f137,f149])).
fof(f149,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f137,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f136])).
fof(f136,plain,(
  ! [X0,X1] : (((p(s(bool,X1)) | p(s(bool,X0))) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | (p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(ennf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(flattening,[],[f96])).
fof(f96,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X6,X7] : ((~(p(s(bool,X7)) | p(s(bool,X6))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X7)) => p(s(bool,'F_0'))) => (~p(s(bool,X6)) => p(s(bool,'F_0')))))),
  file('thm.pred_set.inj_image_countable_IFF','thm.sat.sat_9_prove_ltb')).
fof(f5403,plain,(
  ( ! [X11] : (sP0('F_0',X11)) ) | $spl190),
  inference(resolution,[],[f1970,f345])).
fof(f345,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f211])).
fof(f211,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f210])).
fof(f210,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f149])).
fof(f5402,plain,(
  ( ! [X10] : (sP0(X10,'F_0')) ) | $spl190),
  inference(resolution,[],[f1970,f344])).
fof(f344,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f211])).
fof(f5401,plain,(
  ( ! [X9] : (~p(s(bool,X9)) | s(bool,'F_0') = s(bool,X9)) ) | $spl190),
  inference(resolution,[],[f1970,f342])).
fof(f342,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f209])).
fof(f209,plain,(
  ! [X0,X1] : ((s(bool,X0) != s(bool,X1) | ((~p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | p(s(bool,X0))))) & ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(flattening,[],[f208])).
fof(f208,plain,(
  ! [X0,X1] : ((s(bool,X0) != s(bool,X1) | ((~p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | p(s(bool,X0))))) & (((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0)))) | s(bool,X0) = s(bool,X1)))),
  inference(nnf_transformation,[],[f135])).
fof(f135,plain,(
  ! [X0,X1] : (s(bool,X0) = s(bool,X1) <=> ((~p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | p(s(bool,X0)))))),
  inference(ennf_transformation,[],[f95])).
fof(f95,plain,(
  ! [X0,X1] : (s(bool,X0) = s(bool,X1) <=> ((p(s(bool,X0)) => p(s(bool,X1))) & (p(s(bool,X1)) => p(s(bool,X0)))))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X12,X13] : (s(bool,X12) = s(bool,X13) <=> ((p(s(bool,X12)) => p(s(bool,X13))) & (p(s(bool,X13)) => p(s(bool,X12)))))),
  file('thm.pred_set.inj_image_countable_IFF','thm.bool.EQ_IMP_THM')).
fof(f5400,plain,(
  ( ! [X8] : (~p(s(bool,X8)) | s(bool,'F_0') = s(bool,X8)) ) | $spl190),
  inference(resolution,[],[f1970,f334])).
fof(f334,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f130])).
fof(f130,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f129])).
fof(f129,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f91])).
fof(f91,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X12,X13] : ((p(s(bool,X12)) => p(s(bool,X13))) => ((p(s(bool,X13)) => p(s(bool,X12))) => s(bool,X12) = s(bool,X13)))),
  file('thm.pred_set.inj_image_countable_IFF','thm.bool.IMP_ANTISYM_AX')).
fof(f5395,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl190),
  inference(resolution,[],[f1970,f278])).
fof(f278,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f171])).
fof(f171,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X14] : (s(bool,'T_0') = s(bool,X14) <=> p(s(bool,X14)))),
  file('thm.pred_set.inj_image_countable_IFF','thm.bool.EQ_CLAUSES|split|0')).
fof(f5394,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl190),
  inference(resolution,[],[f1970,f276])).
fof(f276,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f170])).
fof(f170,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X14] : (s(bool,'T_0') = s(bool,X14) <=> p(s(bool,X14)))),
  file('thm.pred_set.inj_image_countable_IFF','thm.bool.EQ_CLAUSES|split|1')).
fof(f5393,plain,(
  ( ! [X6,X7] : (sP0(X6,X7)) ) | $spl190),
  inference(resolution,[],[f1970,f351])).
fof(f351,plain,(
  ( ! [X0,X1] : (~p(s(bool,'F_0')) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f213])).
fof(f5392,plain,(
  ( ! [X4,X5] : (sP0(X4,X5)) ) | $spl190),
  inference(resolution,[],[f1970,f346])).
fof(f346,plain,(
  ( ! [X0,X1] : (~p(s(bool,'F_0')) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f211])).
fof(f5391,plain,(
  ( ! [X3] : (~p(s(bool,X3))) ) | $spl190),
  inference(resolution,[],[f1970,f311])).
fof(f311,plain,(
  ( ! [X0] : (~p(s(bool,'F_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f193])).
fof(f193,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f192])).
fof(f192,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f127])).
fof(f127,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f73])).
fof(f73,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X14] : ((p(s(bool,X14)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X14)))),
  file('thm.pred_set.inj_image_countable_IFF','thm.bool.IMP_CLAUSES|split|4')).
fof(f5390,plain,(
  ( ! [X2] : (p(s(bool,X2)) | ~p(s(bool,'T_0'))) ) | $spl190),
  inference(resolution,[],[f1970,f303])).
fof(f303,plain,(
  ( ! [X0] : (~p(s(bool,'F_0')) | p(s(bool,X0)) | ~p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f187])).
fof(f187,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f186])).
fof(f186,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f124])).
fof(f124,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f70])).
fof(f70,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X14] : ((p(s(bool,'F_0')) => p(s(bool,X14))) <=> p(s(bool,'T_0')))),
  file('thm.pred_set.inj_image_countable_IFF','thm.bool.IMP_CLAUSES|split|2')).
fof(f5389,plain,(
  ( ! [X1] : (p(s(bool,X1))) ) | $spl190),
  inference(resolution,[],[f1970,f291])).
fof(f291,plain,(
  ( ! [X0] : (~p(s(bool,'F_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f179])).
fof(f179,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | ~p(s(bool,X0)) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,'F_0')) & p(s(bool,X0)))))),
  inference(flattening,[],[f178])).
fof(f178,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) | ~p(s(bool,X0))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,'F_0')) & p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f66])).
fof(f66,plain,(
  ! [X0] : ((p(s(bool,'F_0')) & p(s(bool,X0))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X14] : ((p(s(bool,'F_0')) & p(s(bool,X14))) <=> p(s(bool,'F_0')))),
  file('thm.pred_set.inj_image_countable_IFF','thm.bool.AND_CLAUSES|split|2')).
fof(f5388,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl190),
  inference(resolution,[],[f1970,f287])).
fof(f287,plain,(
  ( ! [X0] : (~p(s(bool,'F_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f177])).
fof(f177,plain,(
  ! [X0] : ((~p(s(bool,X0)) | ~p(s(bool,'F_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,X0)) & p(s(bool,'F_0')))))),
  inference(flattening,[],[f176])).
fof(f176,plain,(
  ! [X0] : (((~p(s(bool,X0)) | ~p(s(bool,'F_0'))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,X0)) & p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f65])).
fof(f65,plain,(
  ! [X0] : ((p(s(bool,X0)) & p(s(bool,'F_0'))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X14] : ((p(s(bool,X14)) & p(s(bool,'F_0'))) <=> p(s(bool,'F_0')))),
  file('thm.pred_set.inj_image_countable_IFF','thm.bool.AND_CLAUSES|split|3')).
fof(f5386,plain,(
  ~p(s(bool,'T_0')) | $spl190),
  inference(resolution,[],[f1970,f267])).
fof(f267,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f169])).
fof(f169,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f52])).
fof(f52,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f23])).
fof(f23,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.pred_set.inj_image_countable_IFF','thm.bool.NOT_CLAUSES|split|2')).
fof(f5385,plain,(
  ~p(s(bool,'T_0')) | $spl190),
  inference(resolution,[],[f1970,f265])).
fof(f265,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f168])).
fof(f168,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f51])).
fof(f51,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f22])).
fof(f22,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.pred_set.inj_image_countable_IFF','thm.bool.NOT_CLAUSES|split|1')).
fof(f4793,plain,(
  ~p(s(bool,'T_0')) | $spl190),
  inference(unit_resulting_resolution,[],[f1970,f265])).
fof(f4795,plain,(
  ~p(s(bool,'T_0')) | $spl190),
  inference(unit_resulting_resolution,[],[f1970,f267])).
fof(f4797,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl190),
  inference(unit_resulting_resolution,[],[f1970,f287])).
fof(f4799,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl190),
  inference(unit_resulting_resolution,[],[f1970,f291])).
fof(f4801,plain,(
  ~p(s(bool,'T_0')) | $spl190),
  inference(unit_resulting_resolution,[],[f268,f1970,f303])).
fof(f268,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f53])).
fof(f53,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.pred_set.inj_image_countable_IFF',aHL_FALSITY)).
fof(f4802,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl190),
  inference(unit_resulting_resolution,[],[f263,f1970,f303])).
fof(f263,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.inj_image_countable_IFF',aHL_TRUTH)).
fof(f4804,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl190),
  inference(unit_resulting_resolution,[],[f1970,f311])).
fof(f4806,plain,(
  ( ! [X0,X1] : (sP0(X0,X1)) ) | $spl190),
  inference(unit_resulting_resolution,[],[f1970,f346])).
fof(f4808,plain,(
  ( ! [X0,X1] : (sP0(X0,X1)) ) | $spl190),
  inference(unit_resulting_resolution,[],[f1970,f351])).
fof(f4810,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f1970,f276])).
fof(f4812,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f1970,f278])).
fof(f4822,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f263,f1970,f334])).
fof(f4823,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f262,f1970,f334])).
fof(f262,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f9])).
fof(f9,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.inj_image_countable_IFF','thm.bool.TRUTH')).
fof(f4827,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f263,f1970,f334])).
fof(f4828,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f262,f1970,f334])).
fof(f4836,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f263,f1970,f342])).
fof(f4837,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f262,f1970,f342])).
fof(f4841,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f263,f1970,f342])).
fof(f4842,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f262,f1970,f342])).
fof(f4846,plain,(
  ( ! [X0] : (sP0(X0,'F_0')) ) | $spl190),
  inference(unit_resulting_resolution,[],[f1970,f344])).
fof(f4848,plain,(
  ( ! [X0] : (sP0('F_0',X0)) ) | $spl190),
  inference(unit_resulting_resolution,[],[f1970,f345])).
fof(f4850,plain,(
  ( ! [X0] : (sP0('F_0',X0)) ) | $spl190),
  inference(unit_resulting_resolution,[],[f1970,f350])).
fof(f4853,plain,(
  ~sP1('F_0','T_0') | $spl190),
  inference(unit_resulting_resolution,[],[f263,f1970,f359])).
fof(f4854,plain,(
  ~sP1('F_0','T_0') | $spl190),
  inference(unit_resulting_resolution,[],[f262,f1970,f359])).
fof(f4855,plain,(
  ~sP1('F_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f1970,f1970,f359])).
fof(f4857,plain,(
  ~sP1('T_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f263,f1970,f359])).
fof(f4858,plain,(
  ~sP1('T_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f262,f1970,f359])).
fof(f4859,plain,(
  ~sP1('F_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f1970,f1970,f359])).
fof(f4860,plain,(
  sP1('F_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f268,f1970,f360])).
fof(f4861,plain,(
  sP1('F_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f268,f1970,f363])).
fof(f4862,plain,(
  sP1('F_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f268,f1970,f364])).
fof(f4863,plain,(
  sP1('F_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f268,f1970,f365])).
fof(f4865,plain,(
  ~sP1('T_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f263,f1970,f366])).
fof(f4866,plain,(
  ~sP1('T_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f262,f1970,f366])).
fof(f4867,plain,(
  ~sP1('F_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f1970,f1970,f366])).
fof(f4869,plain,(
  ~sP1('F_0','T_0') | $spl190),
  inference(unit_resulting_resolution,[],[f263,f1970,f366])).
fof(f4870,plain,(
  ~sP1('F_0','T_0') | $spl190),
  inference(unit_resulting_resolution,[],[f262,f1970,f366])).
fof(f4871,plain,(
  ~sP1('F_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f1970,f1970,f366])).
fof(f4875,plain,(
  ~sP5('F_0','T_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f268,f263,f1970,f379])).
fof(f379,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f225])).
fof(f225,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | (sP2(X2,X1,X0) & sP3(X1,X2,X0) & sP4(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP2(X2,X1,X0) | ~sP3(X1,X2,X0) | ~sP4(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f224])).
fof(f224,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP5(X2,X1,X0)))),
  inference(flattening,[],[f223])).
fof(f223,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP5(X2,X1,X0)))),
  inference(nnf_transformation,[],[f156])).
fof(f156,plain,(
  ! [X2,X1,X0] : (sP5(X2,X1,X0) <=> (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f4879,plain,(
  ~sP5('F_0','T_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f268,f262,f1970,f379])).
fof(f4882,plain,(
  ~sP5('F_0','F_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f268,f1970,f1970,f379])).
fof(f4886,plain,(
  ~sP5('T_0','F_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f263,f268,f1970,f379])).
fof(f4887,plain,(
  ~sP5('T_0','F_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f262,f268,f1970,f379])).
fof(f4888,plain,(
  ~sP5('F_0','F_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f1970,f268,f1970,f379])).
fof(f4905,plain,(
  ~sP4('F_0','T_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f268,f263,f1970,f383])).
fof(f383,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f228])).
fof(f4910,plain,(
  ~sP4('F_0','T_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f268,f262,f1970,f383])).
fof(f4915,plain,(
  ~sP4('F_0','F_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f268,f1970,f1970,f383])).
fof(f4923,plain,(
  ~sP4('T_0','F_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f263,f268,f1970,f383])).
fof(f4924,plain,(
  ~sP4('T_0','F_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f262,f268,f1970,f383])).
fof(f4925,plain,(
  ~sP4('F_0','F_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f1970,f268,f1970,f383])).
fof(f4927,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'F_0')) ) | $spl190),
  inference(unit_resulting_resolution,[],[f1970,f384])).
fof(f4934,plain,(
  ~sP3('F_0','T_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f268,f263,f1970,f387])).
fof(f387,plain,(
  ( ! [X2,X0,X1] : (~sP3(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f231])).
fof(f4939,plain,(
  ~sP3('F_0','T_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f268,f262,f1970,f387])).
fof(f4943,plain,(
  ~sP3('F_0','F_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f268,f1970,f1970,f387])).
fof(f4952,plain,(
  ~sP3('T_0','F_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f263,f268,f1970,f387])).
fof(f4953,plain,(
  ~sP3('T_0','F_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f262,f268,f1970,f387])).
fof(f4954,plain,(
  ~sP3('F_0','F_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f1970,f268,f1970,f387])).
fof(f4958,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'F_0')) ) | $spl190),
  inference(unit_resulting_resolution,[],[f1970,f388])).
fof(f4960,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'F_0')) ) | $spl190),
  inference(unit_resulting_resolution,[],[f1970,f392])).
fof(f4962,plain,(
  ( ! [X0,X1] : (sP2(X0,'F_0',X1)) ) | $spl190),
  inference(unit_resulting_resolution,[],[f1970,f393])).
fof(f4964,plain,(
  ( ! [X0,X1] : (sP2('F_0',X0,X1)) ) | $spl190),
  inference(unit_resulting_resolution,[],[f1970,f394])).
fof(f4971,plain,(
  ( ! [X0,X1] : (sP5('F_0','|dollar|exists_1'(s(fun(X0,bool),X1)),happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) ) | $spl190),
  inference(unit_resulting_resolution,[],[f317,f1970,f396])).
fof(f396,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | sP5(X2,X1,X0) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f234])).
fof(f234,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP5(X2,X1,X0)) & (~sP5(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f157])).
fof(f157,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP5(X2,X1,X0))),
  inference(definition_folding,[],[f108,f156,f155,f154,f153])).
fof(f108,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f107])).
fof(f107,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X23,X24,X25] : ((p(s(bool,X25)) <=> s(bool,X23) = s(bool,X24)) <=> ((p(s(bool,X25)) | p(s(bool,X24)) | p(s(bool,X23))) & (p(s(bool,X25)) | ~p(s(bool,X23)) | ~p(s(bool,X24))) & (p(s(bool,X24)) | ~p(s(bool,X23)) | ~p(s(bool,X25))) & (p(s(bool,X23)) | ~p(s(bool,X24)) | ~p(s(bool,X25)))))),
  file('thm.pred_set.inj_image_countable_IFF','thm.sat.sat_2|49|_prove_ltb')).
fof(f317,plain,(
  ( ! [X0,X1] : (s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) )),
  inference(cnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0,X1] : s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f79])).
fof(f79,plain,(
  ! [X0] : ! [X1] : s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X7] : ! [X5] : s(bool,'|dollar|exists_1'(s(fun(X7,bool),X5))) = s(bool,happ(s(fun(X7,bool),X5),s(X7,'const.min.@_1'(s(fun(X7,bool),X5)))))),
  file('thm.pred_set.inj_image_countable_IFF','thm.bool.EXISTS_DEF')).
fof(f4978,plain,(
  ( ! [X0,X1] : (sP5('F_0',happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))),'|dollar|exists_1'(s(fun(X0,bool),X1)))) ) | $spl190),
  inference(unit_resulting_resolution,[],[f317,f1970,f396])).
fof(f4999,plain,(
  ( ! [X0] : (~sP7('F_0',X0,'F_0')) ) | $spl190),
  inference(unit_resulting_resolution,[],[f268,f1970,f400])).
fof(f400,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f237])).
fof(f237,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | (sP6(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP6(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f236])).
fof(f236,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP7(X2,X0,X1)))),
  inference(flattening,[],[f235])).
fof(f235,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP7(X2,X0,X1)))),
  inference(nnf_transformation,[],[f159])).
fof(f5000,plain,(
  ( ! [X0] : (~sP7('F_0','F_0',X0)) ) | $spl190),
  inference(unit_resulting_resolution,[],[f268,f1970,f401])).
fof(f401,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f237])).
fof(f5042,plain,(
  ~sP6('F_0','T_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f268,f263,f1970,f406])).
fof(f406,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f239])).
fof(f5045,plain,(
  ~sP6('F_0','T_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f268,f262,f1970,f406])).
fof(f5048,plain,(
  ~sP6('F_0','F_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f268,f1970,f1970,f406])).
fof(f5055,plain,(
  ~sP6('T_0','F_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f263,f268,f1970,f406])).
fof(f5056,plain,(
  ~sP6('T_0','F_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f262,f268,f1970,f406])).
fof(f5057,plain,(
  ~sP6('F_0','F_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f1970,f268,f1970,f406])).
fof(f5059,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'F_0')) ) | $spl190),
  inference(unit_resulting_resolution,[],[f1970,f407])).
fof(f5065,plain,(
  sP7('T_0','F_0','T_0') | $spl190),
  inference(unit_resulting_resolution,[],[f263,f263,f1970,f412])).
fof(f5066,plain,(
  sP7('T_0','F_0','T_0') | $spl190),
  inference(unit_resulting_resolution,[],[f262,f263,f1970,f412])).
fof(f5067,plain,(
  sP7('T_0','F_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f1970,f263,f1970,f412])).
fof(f5068,plain,(
  sP7('T_0','F_0','T_0') | $spl190),
  inference(unit_resulting_resolution,[],[f263,f262,f1970,f412])).
fof(f5069,plain,(
  sP7('T_0','F_0','T_0') | $spl190),
  inference(unit_resulting_resolution,[],[f262,f262,f1970,f412])).
fof(f5070,plain,(
  sP7('T_0','F_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f1970,f262,f1970,f412])).
fof(f5071,plain,(
  sP7('F_0','F_0','T_0') | $spl190),
  inference(unit_resulting_resolution,[],[f263,f1970,f1970,f412])).
fof(f5072,plain,(
  sP7('F_0','F_0','T_0') | $spl190),
  inference(unit_resulting_resolution,[],[f262,f1970,f1970,f412])).
fof(f5073,plain,(
  sP7('F_0','F_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f1970,f1970,f1970,f412])).
fof(f5075,plain,(
  sP7('T_0','T_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f263,f263,f1970,f412])).
fof(f5076,plain,(
  sP7('T_0','T_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f262,f263,f1970,f412])).
fof(f5077,plain,(
  sP7('F_0','T_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f1970,f263,f1970,f412])).
fof(f5079,plain,(
  sP7('T_0','T_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f263,f262,f1970,f412])).
fof(f5080,plain,(
  sP7('T_0','T_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f262,f262,f1970,f412])).
fof(f5081,plain,(
  sP7('F_0','T_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f1970,f262,f1970,f412])).
fof(f5084,plain,(
  sP7('T_0','F_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f263,f1970,f1970,f412])).
fof(f5085,plain,(
  sP7('T_0','F_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f262,f1970,f1970,f412])).
fof(f5086,plain,(
  sP7('F_0','F_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f1970,f1970,f1970,f412])).
fof(f5090,plain,(
  sP7('F_0','T_0','T_0') | $spl190),
  inference(unit_resulting_resolution,[],[f263,f263,f1970,f412])).
fof(f5091,plain,(
  sP7('F_0','T_0','T_0') | $spl190),
  inference(unit_resulting_resolution,[],[f262,f263,f1970,f412])).
fof(f5092,plain,(
  sP7('F_0','T_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f1970,f263,f1970,f412])).
fof(f5094,plain,(
  sP7('F_0','T_0','T_0') | $spl190),
  inference(unit_resulting_resolution,[],[f263,f262,f1970,f412])).
fof(f5095,plain,(
  sP7('F_0','T_0','T_0') | $spl190),
  inference(unit_resulting_resolution,[],[f262,f262,f1970,f412])).
fof(f5096,plain,(
  sP7('F_0','T_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f1970,f262,f1970,f412])).
fof(f5097,plain,(
  sP7('F_0','F_0','T_0') | $spl190),
  inference(unit_resulting_resolution,[],[f263,f1970,f1970,f412])).
fof(f5098,plain,(
  sP7('F_0','F_0','T_0') | $spl190),
  inference(unit_resulting_resolution,[],[f262,f1970,f1970,f412])).
fof(f5099,plain,(
  sP7('F_0','F_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f1970,f1970,f1970,f412])).
fof(f5100,plain,(
  ( ! [X0] : (~sP7('F_0',X0,'F_0')) ) | $spl190),
  inference(unit_resulting_resolution,[],[f268,f1970,f413])).
fof(f413,plain,(
  ( ! [X2,X0,X1] : (~sP7(X2,X0,X1) | ~p(s(bool,X2)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f241])).
fof(f5101,plain,(
  ( ! [X0] : (~sP7('F_0','F_0',X0)) ) | $spl190),
  inference(unit_resulting_resolution,[],[f268,f1970,f414])).
fof(f414,plain,(
  ( ! [X2,X0,X1] : (~sP7(X2,X0,X1) | ~p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f241])).
fof(f5102,plain,(
  ~sP7('F_0','F_0','T_0') | $spl190),
  inference(unit_resulting_resolution,[],[f268,f263,f1970,f415])).
fof(f415,plain,(
  ( ! [X2,X0,X1] : (~sP7(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f241])).
fof(f5103,plain,(
  ~sP7('F_0','F_0','T_0') | $spl190),
  inference(unit_resulting_resolution,[],[f268,f262,f1970,f415])).
fof(f5104,plain,(
  ~sP7('F_0','F_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f268,f1970,f1970,f415])).
fof(f5105,plain,(
  ~sP7('F_0','T_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f268,f263,f1970,f415])).
fof(f5106,plain,(
  ~sP7('F_0','T_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f268,f262,f1970,f415])).
fof(f5107,plain,(
  ~sP7('F_0','F_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f268,f1970,f1970,f415])).
fof(f5111,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'F_0')) ) | $spl190),
  inference(unit_resulting_resolution,[],[f268,f1970,f416])).
fof(f416,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f244])).
fof(f244,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP8(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f243])).
fof(f243,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1) | sP9(X2,X0,X1)))),
  inference(flattening,[],[f242])).
fof(f242,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1)) | sP9(X2,X0,X1)))),
  inference(nnf_transformation,[],[f162])).
fof(f5115,plain,(
  ( ! [X0] : (~sP9('F_0','F_0',X0)) ) | $spl190),
  inference(unit_resulting_resolution,[],[f268,f1970,f417])).
fof(f417,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f244])).
fof(f5137,plain,(
  ~sP8('F_0','F_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f268,f268,f1970,f423])).
fof(f423,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f247])).
fof(f5138,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'F_0')) ) | $spl190),
  inference(unit_resulting_resolution,[],[f1970,f424])).
fof(f5140,plain,(
  ( ! [X0,X1] : (sP8(X0,'F_0',X1)) ) | $spl190),
  inference(unit_resulting_resolution,[],[f1970,f425])).
fof(f5144,plain,(
  ( ! [X0] : (sP9('T_0',X0,'F_0')) ) | $spl190),
  inference(unit_resulting_resolution,[],[f263,f1970,f428])).
fof(f5145,plain,(
  ( ! [X0] : (sP9('T_0',X0,'F_0')) ) | $spl190),
  inference(unit_resulting_resolution,[],[f262,f1970,f428])).
fof(f5146,plain,(
  ( ! [X0] : (sP9('F_0',X0,'F_0')) ) | $spl190),
  inference(unit_resulting_resolution,[],[f1970,f1970,f428])).
fof(f5150,plain,(
  ( ! [X0] : (sP9('F_0',X0,'T_0')) ) | $spl190),
  inference(unit_resulting_resolution,[],[f263,f1970,f428])).
fof(f5152,plain,(
  ( ! [X0] : (sP9('F_0',X0,'T_0')) ) | $spl190),
  inference(unit_resulting_resolution,[],[f262,f1970,f428])).
fof(f5154,plain,(
  ( ! [X0] : (sP9('F_0',X0,'F_0')) ) | $spl190),
  inference(unit_resulting_resolution,[],[f1970,f1970,f428])).
fof(f5158,plain,(
  ( ! [X0] : (sP9('T_0','F_0',X0)) ) | $spl190),
  inference(unit_resulting_resolution,[],[f263,f1970,f429])).
fof(f5159,plain,(
  ( ! [X0] : (sP9('T_0','F_0',X0)) ) | $spl190),
  inference(unit_resulting_resolution,[],[f262,f1970,f429])).
fof(f5160,plain,(
  ( ! [X0] : (sP9('F_0','F_0',X0)) ) | $spl190),
  inference(unit_resulting_resolution,[],[f1970,f1970,f429])).
fof(f5164,plain,(
  ( ! [X0] : (sP9('F_0','T_0',X0)) ) | $spl190),
  inference(unit_resulting_resolution,[],[f263,f1970,f429])).
fof(f5166,plain,(
  ( ! [X0] : (sP9('F_0','T_0',X0)) ) | $spl190),
  inference(unit_resulting_resolution,[],[f262,f1970,f429])).
fof(f5168,plain,(
  ( ! [X0] : (sP9('F_0','F_0',X0)) ) | $spl190),
  inference(unit_resulting_resolution,[],[f1970,f1970,f429])).
fof(f5170,plain,(
  ~sP9('F_0','F_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f268,f268,f1970,f430])).
fof(f430,plain,(
  ( ! [X2,X0,X1] : (~sP9(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f249])).
fof(f5174,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'F_0')) ) | $spl190),
  inference(unit_resulting_resolution,[],[f268,f1970,f431])).
fof(f431,plain,(
  ( ! [X2,X0,X1] : (~sP9(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f249])).
fof(f5178,plain,(
  ( ! [X0] : (~sP9('F_0','F_0',X0)) ) | $spl190),
  inference(unit_resulting_resolution,[],[f268,f1970,f432])).
fof(f432,plain,(
  ( ! [X2,X0,X1] : (~sP9(X2,X0,X1) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f249])).
fof(f5181,plain,(
  ( ! [X0] : (~sP11('F_0','F_0',X0)) ) | $spl190),
  inference(unit_resulting_resolution,[],[f268,f1970,f434])).
fof(f434,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f252])).
fof(f252,plain,(
  ! [X0,X1,X2] : ((~sP11(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP10(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP10(X0,X1,X2) | sP11(X0,X1,X2)))),
  inference(rectify,[],[f251])).
fof(f251,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1) | sP11(X2,X0,X1)))),
  inference(flattening,[],[f250])).
fof(f250,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1)) | sP11(X2,X0,X1)))),
  inference(nnf_transformation,[],[f165])).
fof(f5219,plain,(
  ~sP10('F_0','F_0','T_0') | $spl190),
  inference(unit_resulting_resolution,[],[f263,f268,f1970,f440])).
fof(f440,plain,(
  ( ! [X2,X0,X1] : (~sP10(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f255])).
fof(f5220,plain,(
  ~sP10('F_0','F_0','T_0') | $spl190),
  inference(unit_resulting_resolution,[],[f262,f268,f1970,f440])).
fof(f5221,plain,(
  ~sP10('F_0','F_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f1970,f268,f1970,f440])).
fof(f5228,plain,(
  ~sP10('T_0','F_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f263,f268,f1970,f440])).
fof(f5229,plain,(
  ~sP10('T_0','F_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f262,f268,f1970,f440])).
fof(f5230,plain,(
  ~sP10('F_0','F_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f1970,f268,f1970,f440])).
fof(f5232,plain,(
  ( ! [X0,X1] : (sP10(X0,'F_0',X1)) ) | $spl190),
  inference(unit_resulting_resolution,[],[f1970,f442])).
fof(f5235,plain,(
  sP11('F_0','F_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f268,f268,f1970,f444])).
fof(f5237,plain,(
  ( ! [X0] : (sP11('F_0',X0,'F_0')) ) | $spl190),
  inference(unit_resulting_resolution,[],[f268,f1970,f445])).
fof(f5242,plain,(
  ( ! [X0] : (sP11('T_0','F_0',X0)) ) | $spl190),
  inference(unit_resulting_resolution,[],[f263,f1970,f446])).
fof(f5244,plain,(
  ( ! [X0] : (sP11('T_0','F_0',X0)) ) | $spl190),
  inference(unit_resulting_resolution,[],[f262,f1970,f446])).
fof(f5246,plain,(
  ( ! [X0] : (sP11('F_0','F_0',X0)) ) | $spl190),
  inference(unit_resulting_resolution,[],[f1970,f1970,f446])).
fof(f5248,plain,(
  ( ! [X0] : (sP11('F_0','T_0',X0)) ) | $spl190),
  inference(unit_resulting_resolution,[],[f263,f1970,f446])).
fof(f5249,plain,(
  ( ! [X0] : (sP11('F_0','T_0',X0)) ) | $spl190),
  inference(unit_resulting_resolution,[],[f262,f1970,f446])).
fof(f5250,plain,(
  ( ! [X0] : (sP11('F_0','F_0',X0)) ) | $spl190),
  inference(unit_resulting_resolution,[],[f1970,f1970,f446])).
fof(f5252,plain,(
  ~sP11('T_0','F_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f268,f263,f1970,f447])).
fof(f447,plain,(
  ( ! [X2,X0,X1] : (~sP11(X2,X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f257])).
fof(f5254,plain,(
  ~sP11('T_0','F_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f268,f262,f1970,f447])).
fof(f5256,plain,(
  ~sP11('F_0','F_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f268,f1970,f1970,f447])).
fof(f5258,plain,(
  ~sP11('F_0','F_0','T_0') | $spl190),
  inference(unit_resulting_resolution,[],[f268,f263,f1970,f447])).
fof(f5259,plain,(
  ~sP11('F_0','F_0','T_0') | $spl190),
  inference(unit_resulting_resolution,[],[f268,f262,f1970,f447])).
fof(f5260,plain,(
  ~sP11('F_0','F_0','F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f268,f1970,f1970,f447])).
fof(f5263,plain,(
  ( ! [X0] : (~sP11('F_0','F_0',X0)) ) | $spl190),
  inference(unit_resulting_resolution,[],[f268,f1970,f449])).
fof(f449,plain,(
  ( ! [X2,X0,X1] : (~sP11(X2,X0,X1) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f257])).
fof(f5925,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853,f4850,f4848,f4846,f4842,f4841,f4837,f4836,f4828,f4827,f4823,f4822,f4812,f4810,f4808,f4806,f4804,f4802,f4801,f4799,f4797,f4795,f4793,f5385,f5386,f5388,f5389,f5390,f5391,f5392,f5393,f5394,f5395,f5400,f5401,f5402,f5403,f5404,f5405,f5406,f5407,f5408,f5409,f5410,f5411,f5412,f5413,f5414,f5415,f5416,f5417,f5418,f5419,f5420,f5421,f5422,f5423,f5424,f5425])).
fof(f5924,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853,f4850,f4848,f4846,f4842,f4841,f4837,f4836,f4828,f4827,f4823,f4822,f4812,f4810,f4808,f4806,f4804,f4802,f4801,f4799,f4797,f4795,f4793,f5385,f5386,f5388,f5389,f5390,f5391,f5392,f5393,f5394,f5395,f5400,f5401,f5402,f5403,f5404,f5405,f5406,f5407,f5408,f5409,f5410,f5411,f5412,f5413,f5414,f5415,f5416,f5417,f5418,f5419,f5420,f5421,f5422,f5423,f5424])).
fof(f5923,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853,f4850,f4848,f4846,f4842,f4841,f4837,f4836,f4828,f4827,f4823,f4822,f4812,f4810,f4808,f4806,f4804,f4802,f4801,f4799,f4797,f4795,f4793,f5385,f5386,f5388,f5389,f5390,f5391,f5392,f5393,f5394,f5395,f5400,f5401,f5402,f5403,f5404,f5405,f5406,f5407,f5408,f5409,f5410,f5411,f5412,f5413,f5414,f5415,f5416,f5417,f5418,f5419,f5420,f5421,f5422,f5423])).
fof(f5922,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853,f4850,f4848,f4846,f4842,f4841,f4837,f4836,f4828,f4827,f4823,f4822,f4812,f4810,f4808,f4806,f4804,f4802,f4801,f4799,f4797,f4795,f4793,f5385,f5386,f5388,f5389,f5390,f5391,f5392,f5393,f5394,f5395,f5400,f5401,f5402,f5403,f5404,f5405,f5406,f5407,f5408,f5409,f5410,f5411,f5412,f5413,f5414,f5415,f5416,f5417,f5418,f5419,f5420,f5421,f5422])).
fof(f5921,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853,f4850,f4848,f4846,f4842,f4841,f4837,f4836,f4828,f4827,f4823,f4822,f4812,f4810,f4808,f4806,f4804,f4802,f4801,f4799,f4797,f4795,f4793,f5385,f5386,f5388,f5389,f5390,f5391,f5392,f5393,f5394,f5395,f5400,f5401,f5402,f5403,f5404,f5405,f5406,f5407,f5408,f5409,f5410,f5411,f5412,f5413,f5414,f5415,f5416,f5417,f5418,f5419,f5420,f5421])).
fof(f5920,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853,f4850,f4848,f4846,f4842,f4841,f4837,f4836,f4828,f4827,f4823,f4822,f4812,f4810,f4808,f4806,f4804,f4802,f4801,f4799,f4797,f4795,f4793,f5385,f5386,f5388,f5389,f5390,f5391,f5392,f5393,f5394,f5395,f5400,f5401,f5402,f5403,f5404,f5405,f5406,f5407,f5408,f5409,f5410,f5411,f5412,f5413,f5414,f5415,f5416,f5417,f5418,f5419,f5420])).
fof(f5919,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853,f4850,f4848,f4846,f4842,f4841,f4837,f4836,f4828,f4827,f4823,f4822,f4812,f4810,f4808,f4806,f4804,f4802,f4801,f4799,f4797,f4795,f4793,f5385,f5386,f5388,f5389,f5390,f5391,f5392,f5393,f5394,f5395,f5400,f5401,f5402,f5403,f5404,f5405,f5406,f5407,f5408,f5409,f5410,f5411,f5412,f5413,f5414,f5415,f5416,f5417,f5418,f5419])).
fof(f5918,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853,f4850,f4848,f4846,f4842,f4841,f4837,f4836,f4828,f4827,f4823,f4822,f4812,f4810,f4808,f4806,f4804,f4802,f4801,f4799,f4797,f4795,f4793,f5385,f5386,f5388,f5389,f5390,f5391,f5392,f5393,f5394,f5395,f5400,f5401,f5402,f5403,f5404,f5405,f5406,f5407,f5408,f5409,f5410,f5411,f5412,f5413,f5414,f5415,f5416,f5417,f5418])).
fof(f5917,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853,f4850,f4848,f4846,f4842,f4841,f4837,f4836,f4828,f4827,f4823,f4822,f4812,f4810,f4808,f4806,f4804,f4802,f4801,f4799,f4797,f4795,f4793,f5385,f5386,f5388,f5389,f5390,f5391,f5392,f5393,f5394,f5395,f5400,f5401,f5402,f5403,f5404,f5405,f5406,f5407,f5408,f5409,f5410,f5411,f5412,f5413,f5414,f5415,f5416,f5417])).
fof(f5916,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853,f4850,f4848,f4846,f4842,f4841,f4837,f4836,f4828,f4827,f4823,f4822,f4812,f4810,f4808,f4806,f4804,f4802,f4801,f4799,f4797,f4795,f4793,f5385,f5386,f5388,f5389,f5390,f5391,f5392,f5393,f5394,f5395,f5400,f5401,f5402,f5403,f5404,f5405,f5406,f5407,f5408,f5409,f5410,f5411,f5412,f5413,f5414,f5415,f5416])).
fof(f5915,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853,f4850,f4848,f4846,f4842,f4841,f4837,f4836,f4828,f4827,f4823,f4822,f4812,f4810,f4808,f4806,f4804,f4802,f4801,f4799,f4797,f4795,f4793,f5385,f5386,f5388,f5389,f5390,f5391,f5392,f5393,f5394,f5395,f5400,f5401,f5402,f5403,f5404,f5405,f5406,f5407,f5408,f5409,f5410,f5411,f5412,f5413,f5414,f5415])).
fof(f5914,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853,f4850,f4848,f4846,f4842,f4841,f4837,f4836,f4828,f4827,f4823,f4822,f4812,f4810,f4808,f4806,f4804,f4802,f4801,f4799,f4797,f4795,f4793,f5385,f5386,f5388,f5389,f5390,f5391,f5392,f5393,f5394,f5395,f5400,f5401,f5402,f5403,f5404,f5405,f5406,f5407,f5408,f5409,f5410,f5411,f5412,f5413,f5414])).
fof(f5913,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853,f4850,f4848,f4846,f4842,f4841,f4837,f4836,f4828,f4827,f4823,f4822,f4812,f4810,f4808,f4806,f4804,f4802,f4801,f4799,f4797,f4795,f4793,f5385,f5386,f5388,f5389,f5390,f5391,f5392,f5393,f5394,f5395,f5400,f5401,f5402,f5403,f5404,f5405,f5406,f5407,f5408,f5409,f5410,f5411,f5412,f5413])).
fof(f5912,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853,f4850,f4848,f4846,f4842,f4841,f4837,f4836,f4828,f4827,f4823,f4822,f4812,f4810,f4808,f4806,f4804,f4802,f4801,f4799,f4797,f4795,f4793,f5385,f5386,f5388,f5389,f5390,f5391,f5392,f5393,f5394,f5395,f5400,f5401,f5402,f5403,f5404,f5405,f5406,f5407,f5408,f5409,f5410,f5411,f5412])).
fof(f5911,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853,f4850,f4848,f4846,f4842,f4841,f4837,f4836,f4828,f4827,f4823,f4822,f4812,f4810,f4808,f4806,f4804,f4802,f4801,f4799,f4797,f4795,f4793,f5385,f5386,f5388,f5389,f5390,f5391,f5392,f5393,f5394,f5395,f5400,f5401,f5402,f5403,f5404,f5405,f5406,f5407,f5408,f5409,f5410,f5411])).
fof(f5910,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853,f4850,f4848,f4846,f4842,f4841,f4837,f4836,f4828,f4827,f4823,f4822,f4812,f4810,f4808,f4806,f4804,f4802,f4801,f4799,f4797,f4795,f4793,f5385,f5386,f5388,f5389,f5390,f5391,f5392,f5393,f5394,f5395,f5400,f5401,f5402,f5403,f5404,f5405,f5406,f5407,f5408,f5409,f5410])).
fof(f5909,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853,f4850,f4848,f4846,f4842,f4841,f4837,f4836,f4828,f4827,f4823,f4822,f4812,f4810,f4808,f4806,f4804,f4802,f4801,f4799,f4797,f4795,f4793,f5385,f5386,f5388,f5389,f5390,f5391,f5392,f5393,f5394,f5395,f5400,f5401,f5402,f5403,f5404,f5405,f5406,f5407,f5408,f5409])).
fof(f5908,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853,f4850,f4848,f4846,f4842,f4841,f4837,f4836,f4828,f4827,f4823,f4822,f4812,f4810,f4808,f4806,f4804,f4802,f4801,f4799,f4797,f4795,f4793,f5385,f5386,f5388,f5389,f5390,f5391,f5392,f5393,f5394,f5395,f5400,f5401,f5402,f5403,f5404,f5405,f5406,f5407,f5408])).
fof(f5907,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853,f4850,f4848,f4846,f4842,f4841,f4837,f4836,f4828,f4827,f4823,f4822,f4812,f4810,f4808,f4806,f4804,f4802,f4801,f4799,f4797,f4795,f4793,f5385,f5386,f5388,f5389,f5390,f5391,f5392,f5393,f5394,f5395,f5400,f5401,f5402,f5403,f5404,f5405,f5406,f5407])).
fof(f5906,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853,f4850,f4848,f4846,f4842,f4841,f4837,f4836,f4828,f4827,f4823,f4822,f4812,f4810,f4808,f4806,f4804,f4802,f4801,f4799,f4797,f4795,f4793,f5385,f5386,f5388,f5389,f5390,f5391,f5392,f5393,f5394,f5395,f5400,f5401,f5402,f5403,f5404,f5405,f5406])).
fof(f5905,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853,f4850,f4848,f4846,f4842,f4841,f4837,f4836,f4828,f4827,f4823,f4822,f4812,f4810,f4808,f4806,f4804,f4802,f4801,f4799,f4797,f4795,f4793,f5385,f5386,f5388,f5389,f5390,f5391,f5392,f5393,f5394,f5395,f5400,f5401,f5402,f5403,f5404,f5405])).
fof(f5904,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853,f4850,f4848,f4846,f4842,f4841,f4837,f4836,f4828,f4827,f4823,f4822,f4812,f4810,f4808,f4806,f4804,f4802,f4801,f4799,f4797,f4795,f4793,f5385,f5386,f5388,f5389,f5390,f5391,f5392,f5393,f5394,f5395,f5400,f5401,f5402,f5403,f5404])).
fof(f5903,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853,f4850,f4848,f4846,f4842,f4841,f4837,f4836,f4828,f4827,f4823,f4822,f4812,f4810,f4808,f4806,f4804,f4802,f4801,f4799,f4797,f4795,f4793,f5385,f5386,f5388,f5389,f5390,f5391,f5392,f5393,f5394,f5395,f5400,f5401,f5402,f5403])).
fof(f5902,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853,f4850,f4848,f4846,f4842,f4841,f4837,f4836,f4828,f4827,f4823,f4822,f4812,f4810,f4808,f4806,f4804,f4802,f4801,f4799,f4797,f4795,f4793,f5385,f5386,f5388,f5389,f5390,f5391,f5392,f5393,f5394,f5395,f5400,f5401,f5402])).
fof(f5901,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853,f4850,f4848,f4846,f4842,f4841,f4837,f4836,f4828,f4827,f4823,f4822,f4812,f4810,f4808,f4806,f4804,f4802,f4801,f4799,f4797,f4795,f4793,f5385,f5386,f5388,f5389,f5390,f5391,f5392,f5393,f5394,f5395,f5400,f5401])).
fof(f5900,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853,f4850,f4848,f4846,f4842,f4841,f4837,f4836,f4828,f4827,f4823,f4822,f4812,f4810,f4808,f4806,f4804,f4802,f4801,f4799,f4797,f4795,f4793,f5385,f5386,f5388,f5389,f5390,f5391,f5392,f5393,f5394,f5395,f5400])).
fof(f5899,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853,f4850,f4848,f4846,f4842,f4841,f4837,f4836,f4828,f4827,f4823,f4822,f4812,f4810,f4808,f4806,f4804,f4802,f4801,f4799,f4797,f4795,f4793,f5385,f5386,f5388,f5389,f5390,f5391,f5392,f5393,f5394,f5395])).
fof(f5898,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853,f4850,f4848,f4846,f4842,f4841,f4837,f4836,f4828,f4827,f4823,f4822,f4812,f4810,f4808,f4806,f4804,f4802,f4801,f4799,f4797,f4795,f4793,f5385,f5386,f5388,f5389,f5390,f5391,f5392,f5393,f5394])).
fof(f5897,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853,f4850,f4848,f4846,f4842,f4841,f4837,f4836,f4828,f4827,f4823,f4822,f4812,f4810,f4808,f4806,f4804,f4802,f4801,f4799,f4797,f4795,f4793,f5385,f5386,f5388,f5389,f5390,f5391,f5392,f5393])).
fof(f5896,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853,f4850,f4848,f4846,f4842,f4841,f4837,f4836,f4828,f4827,f4823,f4822,f4812,f4810,f4808,f4806,f4804,f4802,f4801,f4799,f4797,f4795,f4793,f5385,f5386,f5388,f5389,f5390,f5391,f5392])).
fof(f5895,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853,f4850,f4848,f4846,f4842,f4841,f4837,f4836,f4828,f4827,f4823,f4822,f4812,f4810,f4808,f4806,f4804,f4802,f4801,f4799,f4797,f4795,f4793,f5385,f5386,f5388,f5389,f5390,f5391])).
fof(f5894,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853,f4850,f4848,f4846,f4842,f4841,f4837,f4836,f4828,f4827,f4823,f4822,f4812,f4810,f4808,f4806,f4804,f4802,f4801,f4799,f4797,f4795,f4793,f5385,f5386,f5388,f5389,f5390])).
fof(f5893,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853,f4850,f4848,f4846,f4842,f4841,f4837,f4836,f4828,f4827,f4823,f4822,f4812,f4810,f4808,f4806,f4804,f4802,f4801,f4799,f4797,f4795,f4793,f5385,f5386,f5388,f5389])).
fof(f5892,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853,f4850,f4848,f4846,f4842,f4841,f4837,f4836,f4828,f4827,f4823,f4822,f4812,f4810,f4808,f4806,f4804,f4802,f4801,f4799,f4797,f4795,f4793,f5385,f5386,f5388])).
fof(f5891,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853,f4850,f4848,f4846,f4842,f4841,f4837,f4836,f4828,f4827,f4823,f4822,f4812,f4810,f4808,f4806,f4804,f4802,f4801,f4799,f4797,f4795,f4793,f5385,f5386])).
fof(f5890,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853,f4850,f4848,f4846,f4842,f4841,f4837,f4836,f4828,f4827,f4823,f4822,f4812,f4810,f4808,f4806,f4804,f4802,f4801,f4799,f4797,f4795,f4793,f5385])).
fof(f5889,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853,f4850,f4848,f4846,f4842,f4841,f4837,f4836,f4828,f4827,f4823,f4822,f4812,f4810,f4808,f4806,f4804,f4802,f4801,f4799,f4797,f4795,f4793])).
fof(f5888,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853,f4850,f4848,f4846,f4842,f4841,f4837,f4836,f4828,f4827,f4823,f4822,f4812,f4810,f4808,f4806,f4804,f4802,f4801,f4799,f4797,f4795])).
fof(f5887,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853,f4850,f4848,f4846,f4842,f4841,f4837,f4836,f4828,f4827,f4823,f4822,f4812,f4810,f4808,f4806,f4804,f4802,f4801,f4799,f4797])).
fof(f5886,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853,f4850,f4848,f4846,f4842,f4841,f4837,f4836,f4828,f4827,f4823,f4822,f4812,f4810,f4808,f4806,f4804,f4802,f4801,f4799])).
fof(f5885,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853,f4850,f4848,f4846,f4842,f4841,f4837,f4836,f4828,f4827,f4823,f4822,f4812,f4810,f4808,f4806,f4804,f4802,f4801])).
fof(f5884,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853,f4850,f4848,f4846,f4842,f4841,f4837,f4836,f4828,f4827,f4823,f4822,f4812,f4810,f4808,f4806,f4804,f4802])).
fof(f5883,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853,f4850,f4848,f4846,f4842,f4841,f4837,f4836,f4828,f4827,f4823,f4822,f4812,f4810,f4808,f4806,f4804])).
fof(f5882,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853,f4850,f4848,f4846,f4842,f4841,f4837,f4836,f4828,f4827,f4823,f4822,f4812,f4810,f4808,f4806])).
fof(f5881,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853,f4850,f4848,f4846,f4842,f4841,f4837,f4836,f4828,f4827,f4823,f4822,f4812,f4810,f4808])).
fof(f5880,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853,f4850,f4848,f4846,f4842,f4841,f4837,f4836,f4828,f4827,f4823,f4822,f4812,f4810])).
fof(f5879,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853,f4850,f4848,f4846,f4842,f4841,f4837,f4836,f4828,f4827,f4823,f4822,f4812])).
fof(f5878,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853,f4850,f4848,f4846,f4842,f4841,f4837,f4836,f4828,f4827,f4823,f4822])).
fof(f5877,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853,f4850,f4848,f4846,f4842,f4841,f4837,f4836,f4828,f4827,f4823])).
fof(f5876,plain,(
  $false | ($spl29 | $spl190)),
  inference(global_subsumption,[],[f4825,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853,f4850,f4848,f4846,f4842,f4841,f4837,f4836,f4828,f4827])).
fof(f4825,plain,(
  ~p(s(bool,'T_0')) | ($spl29 | $spl190)),
  inference(unit_resulting_resolution,[],[f518,f1970,f334])).
fof(f5875,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853,f4850,f4848,f4846,f4842,f4841,f4837,f4836,f4828,f4827])).
fof(f5874,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853,f4850,f4848,f4846,f4842,f4841,f4837,f4836,f4828])).
fof(f5873,plain,(
  $false | ($spl29 | $spl190)),
  inference(global_subsumption,[],[f4830,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853,f4850,f4848,f4846,f4842,f4841,f4837,f4836])).
fof(f4830,plain,(
  ~p(s(bool,'T_0')) | ($spl29 | $spl190)),
  inference(unit_resulting_resolution,[],[f518,f1970,f334])).
fof(f5870,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853,f4850,f4848,f4846,f4842,f4841,f4837,f4836])).
fof(f5869,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853,f4850,f4848,f4846,f4842,f4841,f4837])).
fof(f5868,plain,(
  $false | ($spl29 | $spl190)),
  inference(global_subsumption,[],[f4839,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853,f4850,f4848,f4846,f4842,f4841])).
fof(f4839,plain,(
  ~p(s(bool,'T_0')) | ($spl29 | $spl190)),
  inference(unit_resulting_resolution,[],[f518,f1970,f342])).
fof(f5867,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853,f4850,f4848,f4846,f4842,f4841])).
fof(f5866,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853,f4850,f4848,f4846,f4842])).
fof(f5865,plain,(
  $false | ($spl29 | $spl190)),
  inference(global_subsumption,[],[f4844,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853,f4850,f4848,f4846])).
fof(f4844,plain,(
  ~p(s(bool,'T_0')) | ($spl29 | $spl190)),
  inference(unit_resulting_resolution,[],[f518,f1970,f342])).
fof(f5864,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853,f4850,f4848,f4846])).
fof(f5863,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853,f4850,f4848])).
fof(f5862,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853,f4850])).
fof(f5861,plain,(
  $false | ($spl102 | $spl104 | $spl190)),
  inference(global_subsumption,[],[f5860,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853])).
fof(f5860,plain,(
  ~p(s(bool,'T_0')) | ($spl102 | $spl104 | $spl190)),
  inference(forward_demodulation,[],[f4852,f792])).
fof(f4852,plain,(
  ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) | ($spl102 | $spl190)),
  inference(unit_resulting_resolution,[],[f790,f1970,f359])).
fof(f5859,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854,f4853])).
fof(f5858,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855,f4854])).
fof(f5857,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857,f4855])).
fof(f5856,plain,(
  $false | ($spl100 | $spl104 | $spl190)),
  inference(global_subsumption,[],[f5855,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857])).
fof(f5855,plain,(
  ~p(s(bool,'T_0')) | ($spl100 | $spl104 | $spl190)),
  inference(forward_demodulation,[],[f4856,f792])).
fof(f4856,plain,(
  ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) | ($spl100 | $spl190)),
  inference(unit_resulting_resolution,[],[f788,f1970,f359])).
fof(f5854,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858,f4857])).
fof(f5853,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859,f4858])).
fof(f5852,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860,f4859])).
fof(f5851,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861,f4860])).
fof(f5850,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862,f4861])).
fof(f5849,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863,f4862])).
fof(f5848,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865,f4863])).
fof(f5847,plain,(
  $false | ($spl100 | $spl104 | $spl190)),
  inference(global_subsumption,[],[f5846,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865])).
fof(f5846,plain,(
  ~p(s(bool,'T_0')) | ($spl100 | $spl104 | $spl190)),
  inference(forward_demodulation,[],[f4864,f792])).
fof(f4864,plain,(
  ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) | ($spl100 | $spl190)),
  inference(unit_resulting_resolution,[],[f788,f1970,f366])).
fof(f5845,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866,f4865])).
fof(f5844,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867,f4866])).
fof(f5843,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869,f4867])).
fof(f5842,plain,(
  $false | ($spl102 | $spl104 | $spl190)),
  inference(global_subsumption,[],[f5841,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869])).
fof(f5841,plain,(
  ~p(s(bool,'T_0')) | ($spl102 | $spl104 | $spl190)),
  inference(forward_demodulation,[],[f4868,f792])).
fof(f4868,plain,(
  ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) | ($spl102 | $spl190)),
  inference(unit_resulting_resolution,[],[f790,f1970,f366])).
fof(f5840,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870,f4869])).
fof(f5839,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871,f4870])).
fof(f5838,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875,f4871])).
fof(f5837,plain,(
  $false | ($spl108 | $spl190)),
  inference(global_subsumption,[],[f4872,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875])).
fof(f4872,plain,(
  ~p(s(bool,'T_0')) | ($spl108 | $spl190)),
  inference(unit_resulting_resolution,[],[f822,f268,f1970,f379])).
fof(f5835,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879,f4875])).
fof(f5833,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882,f4879])).
fof(f5832,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886,f4882])).
fof(f5829,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887,f4886])).
fof(f5828,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888,f4887])).
fof(f5827,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905,f4888])).
fof(f5826,plain,(
  $false | ($spl2 | $spl93 | $spl190)),
  inference(global_subsumption,[],[f4889,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905])).
fof(f4889,plain,(
  ~sP4('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0','T_0') | ($spl2 | $spl93 | $spl190)),
  inference(unit_resulting_resolution,[],[f583,f582,f780,f1970,f380])).
fof(f380,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X2,X1) | ~p(s(bool,X2)) | sP5(X0,X1,X2) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f225])).
fof(f582,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f468,f388])).
fof(f583,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f468,f392])).
fof(f5825,plain,(
  $false | ($spl2 | $spl89 | $spl190)),
  inference(global_subsumption,[],[f4890,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905])).
fof(f4890,plain,(
  ~sP4('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | ($spl2 | $spl89 | $spl190)),
  inference(unit_resulting_resolution,[],[f584,f582,f776,f1970,f380])).
fof(f584,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f468,f393])).
fof(f5824,plain,(
  $false | ($spl2 | $spl190)),
  inference(global_subsumption,[],[f4891,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905])).
fof(f4891,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f584,f582,f575,f1970,f380])).
fof(f575,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f468,f384])).
fof(f5823,plain,(
  $false | ($spl2 | $spl190 | $spl292)),
  inference(global_subsumption,[],[f4892,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905])).
fof(f4892,plain,(
  sP5('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') | ($spl2 | $spl190 | $spl292)),
  inference(unit_resulting_resolution,[],[f584,f2509,f575,f1970,f380])).
fof(f5822,plain,(
  $false | ($spl2 | $spl91 | $spl190)),
  inference(global_subsumption,[],[f4893,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905])).
fof(f4893,plain,(
  ~sP3('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0','T_0') | ($spl2 | $spl91 | $spl190)),
  inference(unit_resulting_resolution,[],[f584,f778,f575,f1970,f380])).
fof(f5821,plain,(
  $false | ($spl2 | $spl89 | $spl190)),
  inference(global_subsumption,[],[f4894,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905])).
fof(f4894,plain,(
  ~sP3('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | ($spl2 | $spl89 | $spl190)),
  inference(unit_resulting_resolution,[],[f584,f776,f575,f1970,f380])).
fof(f5820,plain,(
  $false | ($spl2 | $spl89 | $spl190)),
  inference(global_subsumption,[],[f4895,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905])).
fof(f4895,plain,(
  ~sP2('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | ($spl2 | $spl89 | $spl190)),
  inference(unit_resulting_resolution,[],[f582,f776,f575,f1970,f380])).
fof(f5819,plain,(
  $false | ($spl2 | $spl190 | $spl290 | $spl296)),
  inference(global_subsumption,[],[f4897,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905])).
fof(f4897,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') | ($spl2 | $spl190 | $spl290 | $spl296)),
  inference(unit_resulting_resolution,[],[f2511,f582,f2508,f1970,f380])).
fof(f5818,plain,(
  $false | ($spl190 | $spl290 | $spl292 | $spl296)),
  inference(global_subsumption,[],[f4898,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905])).
fof(f4898,plain,(
  sP5('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') | ($spl190 | $spl290 | $spl292 | $spl296)),
  inference(unit_resulting_resolution,[],[f2511,f2509,f2508,f1970,f380])).
fof(f5817,plain,(
  $false | ($spl146 | $spl190)),
  inference(global_subsumption,[],[f4900,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905])).
fof(f4900,plain,(
  ~p(s(bool,'T_0')) | ($spl146 | $spl190)),
  inference(unit_resulting_resolution,[],[f1203,f268,f1970,f383])).
fof(f5813,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910,f4905])).
fof(f5810,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915,f4910])).
fof(f5807,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923,f4915])).
fof(f5800,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924,f4923])).
fof(f5799,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925,f4924])).
fof(f5798,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927,f4925])).
fof(f5796,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934,f4927])).
fof(f5795,plain,(
  $false | ($spl148 | $spl190)),
  inference(global_subsumption,[],[f4929,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934])).
fof(f4929,plain,(
  ~p(s(bool,'T_0')) | ($spl148 | $spl190)),
  inference(unit_resulting_resolution,[],[f1205,f268,f1970,f387])).
fof(f5791,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939,f4934])).
fof(f5788,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943,f4939])).
fof(f5785,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952,f4943])).
fof(f5778,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953,f4952])).
fof(f5777,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954,f4953])).
fof(f5776,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958,f4954])).
fof(f5775,plain,(
  $false | ($spl142 | $spl190)),
  inference(global_subsumption,[],[f4955,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958])).
fof(f4955,plain,(
  ~p(s(bool,'T_0')) | ($spl142 | $spl190)),
  inference(unit_resulting_resolution,[],[f978,f268,f1970,f387])).
fof(f5773,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960,f4958])).
fof(f5772,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962,f4960])).
fof(f5771,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964,f4962])).
fof(f5770,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971,f4964])).
fof(f5769,plain,(
  $false | ($spl119 | $spl190)),
  inference(global_subsumption,[],[f4965,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971])).
fof(f4965,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | ($spl119 | $spl190)),
  inference(unit_resulting_resolution,[],[f845,f1970,f396])).
fof(f5768,plain,(
  $false | ($spl159 | $spl190)),
  inference(global_subsumption,[],[f4966,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971])).
fof(f4966,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | ($spl159 | $spl190)),
  inference(unit_resulting_resolution,[],[f1379,f1970,f396])).
fof(f5767,plain,(
  $false | ($spl21 | $spl190)),
  inference(global_subsumption,[],[f4967,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971])).
fof(f4967,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | ($spl21 | $spl190)),
  inference(unit_resulting_resolution,[],[f507,f1970,f396])).
fof(f507,plain,(
  ~sP5('F_0','F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | $spl21),
  inference(cnf_transformation,[],[f507_D])).
fof(f507_D,plain,(
  ~sP5('F_0','F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f5766,plain,(
  $false | ($spl161 | $spl190)),
  inference(global_subsumption,[],[f4968,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971])).
fof(f4968,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | ($spl161 | $spl190)),
  inference(unit_resulting_resolution,[],[f1389,f1970,f396])).
fof(f5765,plain,(
  $false | ($spl25 | $spl190)),
  inference(global_subsumption,[],[f4969,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971])).
fof(f4969,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | ($spl25 | $spl190)),
  inference(unit_resulting_resolution,[],[f511,f1970,f396])).
fof(f511,plain,(
  ~sP5('F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0') | $spl25),
  inference(cnf_transformation,[],[f511_D])).
fof(f511_D,plain,(
  ~sP5('F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0') <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f5764,plain,(
  $false | ($spl123 | $spl190)),
  inference(global_subsumption,[],[f4970,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971])).
fof(f4970,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | ($spl123 | $spl190)),
  inference(unit_resulting_resolution,[],[f849,f1970,f396])).
fof(f5763,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978,f4971])).
fof(f5762,plain,(
  $false | ($spl104 | $spl190)),
  inference(global_subsumption,[],[f4972,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978])).
fof(f4972,plain,(
  sP5('F_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | ($spl104 | $spl190)),
  inference(unit_resulting_resolution,[],[f792,f1970,f396])).
fof(f5761,plain,(
  $false | ($spl12 | $spl190)),
  inference(global_subsumption,[],[f4974,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978])).
fof(f4974,plain,(
  sP5('F_0','T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | ($spl12 | $spl190)),
  inference(unit_resulting_resolution,[],[f485,f1970,f396])).
fof(f5760,plain,(
  $false | ($spl14 | $spl190)),
  inference(global_subsumption,[],[f4976,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978])).
fof(f4976,plain,(
  sP5('F_0','F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | ($spl14 | $spl190)),
  inference(unit_resulting_resolution,[],[f486,f1970,f396])).
fof(f486,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | $spl14),
  inference(cnf_transformation,[],[f486_D])).
fof(f486_D,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f5759,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999,f4978])).
fof(f5758,plain,(
  $false | ($spl104 | $spl190)),
  inference(global_subsumption,[],[f4979,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999])).
fof(f4979,plain,(
  sP5('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | ($spl104 | $spl190)),
  inference(unit_resulting_resolution,[],[f792,f1970,f396])).
fof(f5757,plain,(
  $false | ($spl12 | $spl190)),
  inference(global_subsumption,[],[f4981,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999])).
fof(f4981,plain,(
  sP5('F_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | ($spl12 | $spl190)),
  inference(unit_resulting_resolution,[],[f485,f1970,f396])).
fof(f5756,plain,(
  $false | ($spl14 | $spl190)),
  inference(global_subsumption,[],[f4983,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999])).
fof(f4983,plain,(
  sP5('F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0') | ($spl14 | $spl190)),
  inference(unit_resulting_resolution,[],[f486,f1970,f396])).
fof(f5755,plain,(
  $false | ($spl6 | $spl12 | $spl190)),
  inference(global_subsumption,[],[f5754,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999])).
fof(f5754,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | ($spl6 | $spl12 | $spl190)),
  inference(forward_demodulation,[],[f4985,f485])).
fof(f4985,plain,(
  s(bool,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | ($spl6 | $spl190)),
  inference(unit_resulting_resolution,[],[f479,f1970,f397])).
fof(f397,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | ~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f234])).
fof(f5753,plain,(
  $false | ($spl10 | $spl12 | $spl190)),
  inference(global_subsumption,[],[f5752,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999])).
fof(f5752,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | ($spl10 | $spl12 | $spl190)),
  inference(forward_demodulation,[],[f4986,f485])).
fof(f4986,plain,(
  s(bool,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) = s(bool,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | ($spl10 | $spl190)),
  inference(unit_resulting_resolution,[],[f483,f1970,f397])).
fof(f5751,plain,(
  $false | ($spl190 | $spl196)),
  inference(global_subsumption,[],[f4987,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999])).
fof(f4987,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | ($spl190 | $spl196)),
  inference(unit_resulting_resolution,[],[f2147,f1970,f397])).
fof(f5750,plain,(
  $false | ($spl108 | $spl190)),
  inference(global_subsumption,[],[f4988,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999])).
fof(f4988,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl108 | $spl190)),
  inference(unit_resulting_resolution,[],[f822,f1970,f397])).
fof(f5749,plain,(
  $false | ($spl190 | $spl192)),
  inference(global_subsumption,[],[f4989,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999])).
fof(f4989,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | ($spl190 | $spl192)),
  inference(unit_resulting_resolution,[],[f2133,f1970,f397])).
fof(f5748,plain,(
  $false | ($spl112 | $spl190)),
  inference(global_subsumption,[],[f4990,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999])).
fof(f4990,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl112 | $spl190)),
  inference(unit_resulting_resolution,[],[f826,f1970,f397])).
fof(f5747,plain,(
  $false | ($spl29 | $spl190)),
  inference(global_subsumption,[],[f4991,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999])).
fof(f4991,plain,(
  ~sP5('F_0','T_0','F_0') | ($spl29 | $spl190)),
  inference(unit_resulting_resolution,[],[f518,f1970,f397])).
fof(f5746,plain,(
  $false | ($spl27 | $spl190)),
  inference(global_subsumption,[],[f4993,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999])).
fof(f4993,plain,(
  ~sP5('F_0','T_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | ($spl27 | $spl190)),
  inference(unit_resulting_resolution,[],[f513,f1970,f397])).
fof(f513,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | $spl27),
  inference(cnf_transformation,[],[f513_D])).
fof(f513_D,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) <=> ~$spl27),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl27])])).
fof(f5745,plain,(
  $false | ($spl29 | $spl190)),
  inference(global_subsumption,[],[f4995,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999])).
fof(f4995,plain,(
  ~sP5('F_0','F_0','T_0') | ($spl29 | $spl190)),
  inference(unit_resulting_resolution,[],[f518,f1970,f397])).
fof(f5744,plain,(
  $false | ($spl27 | $spl190)),
  inference(global_subsumption,[],[f4997,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999])).
fof(f4997,plain,(
  ~sP5('F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'T_0') | ($spl27 | $spl190)),
  inference(unit_resulting_resolution,[],[f513,f1970,f397])).
fof(f5743,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000,f4999])).
fof(f5742,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042,f5000])).
fof(f5741,plain,(
  $false | ($spl2 | $spl190)),
  inference(global_subsumption,[],[f5001,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042])).
fof(f5001,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f625,f596,f1970,f402])).
fof(f402,plain,(
  ( ! [X2,X0,X1] : (~sP6(X1,X2,X0) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | sP7(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f237])).
fof(f596,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f468,f407])).
fof(f625,plain,(
  ( ! [X0] : (~sP7('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0',X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f268,f468,f414])).
fof(f5739,plain,(
  $false | ($spl2 | $spl190)),
  inference(global_subsumption,[],[f5004,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042])).
fof(f5004,plain,(
  ~sP6('F_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f625,f263,f1970,f402])).
fof(f5738,plain,(
  $false | ($spl2 | $spl190)),
  inference(global_subsumption,[],[f5005,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042])).
fof(f5005,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0','T_0') | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f596,f263,f1970,f402])).
fof(f5737,plain,(
  $false | ($spl190 | $spl300)),
  inference(global_subsumption,[],[f5007,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042])).
fof(f5007,plain,(
  sP7('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0','T_0') | ($spl190 | $spl300)),
  inference(unit_resulting_resolution,[],[f2513,f263,f1970,f402])).
fof(f5736,plain,(
  $false | ($spl2 | $spl190)),
  inference(global_subsumption,[],[f5008,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042])).
fof(f5008,plain,(
  ~sP6('F_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f625,f262,f1970,f402])).
fof(f5735,plain,(
  $false | ($spl2 | $spl190)),
  inference(global_subsumption,[],[f5009,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042])).
fof(f5009,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0','T_0') | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f596,f262,f1970,f402])).
fof(f5734,plain,(
  $false | ($spl190 | $spl300)),
  inference(global_subsumption,[],[f5011,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042])).
fof(f5011,plain,(
  sP7('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0','T_0') | ($spl190 | $spl300)),
  inference(unit_resulting_resolution,[],[f2513,f262,f1970,f402])).
fof(f5733,plain,(
  $false | ($spl2 | $spl190)),
  inference(global_subsumption,[],[f5012,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042])).
fof(f5012,plain,(
  ~sP6('F_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f625,f1970,f1970,f402])).
fof(f5732,plain,(
  $false | ($spl2 | $spl190)),
  inference(global_subsumption,[],[f5013,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042])).
fof(f5013,plain,(
  ~sP6('F_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f624,f1970,f1970,f402])).
fof(f624,plain,(
  ( ! [X0] : (~sP7('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X0,'F_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f268,f468,f413])).
fof(f5731,plain,(
  $false | ($spl2 | $spl190)),
  inference(global_subsumption,[],[f5014,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042])).
fof(f5014,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0','F_0') | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f596,f1970,f1970,f402])).
fof(f5730,plain,(
  $false | ($spl190 | $spl300)),
  inference(global_subsumption,[],[f5016,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042])).
fof(f5016,plain,(
  sP7('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0','F_0') | ($spl190 | $spl300)),
  inference(unit_resulting_resolution,[],[f2513,f1970,f1970,f402])).
fof(f5729,plain,(
  $false | ($spl2 | $spl190)),
  inference(global_subsumption,[],[f5017,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042])).
fof(f5017,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f624,f596,f1970,f402])).
fof(f5727,plain,(
  $false | ($spl2 | $spl190)),
  inference(global_subsumption,[],[f5020,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042])).
fof(f5020,plain,(
  ~sP6('T_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f624,f263,f1970,f402])).
fof(f5726,plain,(
  $false | ($spl2 | $spl190)),
  inference(global_subsumption,[],[f5021,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042])).
fof(f5021,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','F_0') | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f596,f263,f1970,f402])).
fof(f5725,plain,(
  $false | ($spl190 | $spl300)),
  inference(global_subsumption,[],[f5023,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042])).
fof(f5023,plain,(
  sP7('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','F_0') | ($spl190 | $spl300)),
  inference(unit_resulting_resolution,[],[f2513,f263,f1970,f402])).
fof(f5724,plain,(
  $false | ($spl2 | $spl190)),
  inference(global_subsumption,[],[f5024,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042])).
fof(f5024,plain,(
  ~sP6('T_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f624,f262,f1970,f402])).
fof(f5723,plain,(
  $false | ($spl2 | $spl190)),
  inference(global_subsumption,[],[f5025,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042])).
fof(f5025,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','F_0') | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f596,f262,f1970,f402])).
fof(f5722,plain,(
  $false | ($spl190 | $spl300)),
  inference(global_subsumption,[],[f5027,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042])).
fof(f5027,plain,(
  sP7('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','F_0') | ($spl190 | $spl300)),
  inference(unit_resulting_resolution,[],[f2513,f262,f1970,f402])).
fof(f5721,plain,(
  $false | ($spl2 | $spl190)),
  inference(global_subsumption,[],[f5028,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042])).
fof(f5028,plain,(
  ~sP6('F_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f625,f1970,f1970,f402])).
fof(f5720,plain,(
  $false | ($spl2 | $spl190)),
  inference(global_subsumption,[],[f5029,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042])).
fof(f5029,plain,(
  ~sP6('F_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f624,f1970,f1970,f402])).
fof(f5719,plain,(
  $false | ($spl2 | $spl190)),
  inference(global_subsumption,[],[f5030,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042])).
fof(f5030,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0','F_0') | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f596,f1970,f1970,f402])).
fof(f5718,plain,(
  $false | ($spl190 | $spl300)),
  inference(global_subsumption,[],[f5032,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042])).
fof(f5032,plain,(
  sP7('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0','F_0') | ($spl190 | $spl300)),
  inference(unit_resulting_resolution,[],[f2513,f1970,f1970,f402])).
fof(f5710,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045,f5042])).
fof(f5707,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048,f5045])).
fof(f5704,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055,f5048])).
fof(f5697,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056,f5055])).
fof(f5696,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057,f5056])).
fof(f5695,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059,f5057])).
fof(f5693,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065,f5059])).
fof(f5692,plain,(
  $false | ($spl2 | $spl104 | $spl190)),
  inference(global_subsumption,[],[f5691,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065])).
fof(f5691,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl104 | $spl190)),
  inference(forward_demodulation,[],[f5061,f792])).
fof(f5061,plain,(
  ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f625,f263,f1970,f412])).
fof(f5690,plain,(
  $false | ($spl2 | $spl104 | $spl190)),
  inference(global_subsumption,[],[f5689,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065])).
fof(f5689,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl104 | $spl190)),
  inference(forward_demodulation,[],[f5062,f792])).
fof(f5062,plain,(
  ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f625,f262,f1970,f412])).
fof(f5688,plain,(
  $false | ($spl2 | $spl104 | $spl190)),
  inference(global_subsumption,[],[f5687,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065])).
fof(f5687,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl104 | $spl190)),
  inference(forward_demodulation,[],[f5063,f792])).
fof(f5063,plain,(
  ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f625,f1970,f1970,f412])).
fof(f5686,plain,(
  $false | ($spl2 | $spl104 | $spl190)),
  inference(global_subsumption,[],[f5685,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065])).
fof(f5685,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl104 | $spl190)),
  inference(forward_demodulation,[],[f5064,f792])).
fof(f5064,plain,(
  ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f624,f1970,f1970,f412])).
fof(f5684,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066,f5065])).
fof(f5683,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067,f5066])).
fof(f5682,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068,f5067])).
fof(f5681,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069,f5068])).
fof(f5680,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070,f5069])).
fof(f5679,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071,f5070])).
fof(f5678,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072,f5071])).
fof(f5677,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073,f5072])).
fof(f5676,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075,f5073])).
fof(f5675,plain,(
  $false | ($spl2 | $spl104 | $spl190)),
  inference(global_subsumption,[],[f5674,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075])).
fof(f5674,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl104 | $spl190)),
  inference(forward_demodulation,[],[f5074,f792])).
fof(f5074,plain,(
  ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f624,f263,f1970,f412])).
fof(f5673,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076,f5075])).
fof(f5672,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077,f5076])).
fof(f5671,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079,f5077])).
fof(f5670,plain,(
  $false | ($spl2 | $spl104 | $spl190)),
  inference(global_subsumption,[],[f5669,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079])).
fof(f5669,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl104 | $spl190)),
  inference(forward_demodulation,[],[f5078,f792])).
fof(f5078,plain,(
  ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f624,f262,f1970,f412])).
fof(f5668,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080,f5079])).
fof(f5667,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081,f5080])).
fof(f5666,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084,f5081])).
fof(f5665,plain,(
  $false | ($spl2 | $spl104 | $spl190)),
  inference(global_subsumption,[],[f5664,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084])).
fof(f5664,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl104 | $spl190)),
  inference(forward_demodulation,[],[f5082,f792])).
fof(f5082,plain,(
  ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f625,f1970,f1970,f412])).
fof(f5663,plain,(
  $false | ($spl2 | $spl104 | $spl190)),
  inference(global_subsumption,[],[f5662,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084])).
fof(f5662,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl104 | $spl190)),
  inference(forward_demodulation,[],[f5083,f792])).
fof(f5083,plain,(
  ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f624,f1970,f1970,f412])).
fof(f5661,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085,f5084])).
fof(f5660,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086,f5085])).
fof(f5659,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090,f5086])).
fof(f5658,plain,(
  $false | ($spl55 | $spl104 | $spl190)),
  inference(global_subsumption,[],[f5657,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090])).
fof(f5657,plain,(
  ~p(s(bool,'T_0')) | ($spl55 | $spl104 | $spl190)),
  inference(forward_demodulation,[],[f5087,f792])).
fof(f5087,plain,(
  ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) | ($spl55 | $spl190)),
  inference(unit_resulting_resolution,[],[f742,f263,f1970,f412])).
fof(f5656,plain,(
  $false | ($spl55 | $spl104 | $spl190)),
  inference(global_subsumption,[],[f5655,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090])).
fof(f5655,plain,(
  ~p(s(bool,'T_0')) | ($spl55 | $spl104 | $spl190)),
  inference(forward_demodulation,[],[f5088,f792])).
fof(f5088,plain,(
  ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) | ($spl55 | $spl190)),
  inference(unit_resulting_resolution,[],[f742,f262,f1970,f412])).
fof(f5654,plain,(
  $false | ($spl53 | $spl104 | $spl190)),
  inference(global_subsumption,[],[f5653,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090])).
fof(f5653,plain,(
  ~p(s(bool,'T_0')) | ($spl53 | $spl104 | $spl190)),
  inference(forward_demodulation,[],[f5089,f792])).
fof(f5089,plain,(
  ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) | ($spl53 | $spl190)),
  inference(unit_resulting_resolution,[],[f740,f263,f1970,f412])).
fof(f5652,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091,f5090])).
fof(f5651,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092,f5091])).
fof(f5650,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094,f5092])).
fof(f5649,plain,(
  $false | ($spl53 | $spl104 | $spl190)),
  inference(global_subsumption,[],[f5648,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094])).
fof(f5648,plain,(
  ~p(s(bool,'T_0')) | ($spl53 | $spl104 | $spl190)),
  inference(forward_demodulation,[],[f5093,f792])).
fof(f5093,plain,(
  ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) | ($spl53 | $spl190)),
  inference(unit_resulting_resolution,[],[f740,f262,f1970,f412])).
fof(f5647,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095,f5094])).
fof(f5646,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096,f5095])).
fof(f5645,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097,f5096])).
fof(f5644,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098,f5097])).
fof(f5643,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099,f5098])).
fof(f5642,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100,f5099])).
fof(f5641,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101,f5100])).
fof(f5640,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102,f5101])).
fof(f5639,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103,f5102])).
fof(f5638,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104,f5103])).
fof(f5637,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105,f5104])).
fof(f5636,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106,f5105])).
fof(f5635,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107,f5106])).
fof(f5634,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111,f5107])).
fof(f5631,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115,f5111])).
fof(f5628,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137,f5115])).
fof(f5625,plain,(
  $false | ($spl2 | $spl190)),
  inference(global_subsumption,[],[f5118,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137])).
fof(f5118,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f651,f636,f1970,f420])).
fof(f420,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X1,X2) | sP9(X0,X1,X2) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f244])).
fof(f636,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f468,f425])).
fof(f651,plain,(
  ( ! [X0] : (~sP9('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f268,f468,f432])).
fof(f5622,plain,(
  $false | ($spl2 | $spl190)),
  inference(global_subsumption,[],[f5121,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137])).
fof(f5121,plain,(
  ~sP8('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f651,f268,f1970,f420])).
fof(f5621,plain,(
  $false | ($spl2 | $spl190)),
  inference(global_subsumption,[],[f5122,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137])).
fof(f5122,plain,(
  sP9('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f636,f268,f1970,f420])).
fof(f5620,plain,(
  $false | ($spl190 | $spl306)),
  inference(global_subsumption,[],[f5124,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137])).
fof(f5124,plain,(
  sP9('F_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') | ($spl190 | $spl306)),
  inference(unit_resulting_resolution,[],[f2516,f268,f1970,f420])).
fof(f5619,plain,(
  $false | ($spl2 | $spl190)),
  inference(global_subsumption,[],[f5125,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137])).
fof(f5125,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f635,f650,f1970,f421])).
fof(f421,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X1,X2) | sP9(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f244])).
fof(f650,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f268,f468,f431])).
fof(f635,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f468,f424])).
fof(f5615,plain,(
  $false | ($spl2 | $spl190)),
  inference(global_subsumption,[],[f5129,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137])).
fof(f5129,plain,(
  sP9('F_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f635,f268,f1970,f421])).
fof(f5614,plain,(
  $false | ($spl190 | $spl304)),
  inference(global_subsumption,[],[f5130,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137])).
fof(f5130,plain,(
  sP9('F_0','F_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | ($spl190 | $spl304)),
  inference(unit_resulting_resolution,[],[f2515,f268,f1970,f421])).
fof(f5613,plain,(
  $false | ($spl2 | $spl190)),
  inference(global_subsumption,[],[f5131,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137])).
fof(f5131,plain,(
  ~sP8('F_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f650,f268,f1970,f421])).
fof(f5608,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138,f5137])).
fof(f5607,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140,f5138])).
fof(f5606,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144,f5140])).
fof(f5605,plain,(
  $false | ($spl47 | $spl104 | $spl190)),
  inference(global_subsumption,[],[f5604,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144])).
fof(f5604,plain,(
  ~p(s(bool,'T_0')) | ($spl47 | $spl104 | $spl190)),
  inference(forward_demodulation,[],[f5143,f792])).
fof(f5143,plain,(
  ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) | ($spl47 | $spl190)),
  inference(unit_resulting_resolution,[],[f734,f1970,f428])).
fof(f5603,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145,f5144])).
fof(f5602,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146,f5145])).
fof(f5601,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150,f5146])).
fof(f5600,plain,(
  $false | ($spl2 | $spl104 | $spl190)),
  inference(global_subsumption,[],[f5599,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150])).
fof(f5599,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl104 | $spl190)),
  inference(forward_demodulation,[],[f5148,f792])).
fof(f5148,plain,(
  ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f650,f1970,f428])).
fof(f5598,plain,(
  $false | ($spl2 | $spl190)),
  inference(global_subsumption,[],[f5149,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150])).
fof(f5149,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f651,f1970,f428])).
fof(f5597,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152,f5150])).
fof(f5596,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154,f5152])).
fof(f5595,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158,f5154])).
fof(f5594,plain,(
  $false | ($spl47 | $spl104 | $spl190)),
  inference(global_subsumption,[],[f5593,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158])).
fof(f5593,plain,(
  ~p(s(bool,'T_0')) | ($spl47 | $spl104 | $spl190)),
  inference(forward_demodulation,[],[f5157,f792])).
fof(f5157,plain,(
  ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) | ($spl47 | $spl190)),
  inference(unit_resulting_resolution,[],[f734,f1970,f429])).
fof(f5592,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159,f5158])).
fof(f5591,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160,f5159])).
fof(f5590,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164,f5160])).
fof(f5589,plain,(
  $false | ($spl2 | $spl190)),
  inference(global_subsumption,[],[f5162,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164])).
fof(f5162,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f650,f1970,f429])).
fof(f5588,plain,(
  $false | ($spl2 | $spl104 | $spl190)),
  inference(global_subsumption,[],[f5587,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164])).
fof(f5587,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl104 | $spl190)),
  inference(forward_demodulation,[],[f5163,f792])).
fof(f5163,plain,(
  ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f651,f1970,f429])).
fof(f5586,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166,f5164])).
fof(f5585,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168,f5166])).
fof(f5584,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170,f5168])).
fof(f5583,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174,f5170])).
fof(f5580,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178,f5174])).
fof(f5577,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181,f5178])).
fof(f5575,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219,f5181])).
fof(f5574,plain,(
  $false | ($spl2 | $spl190)),
  inference(global_subsumption,[],[f5183,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219])).
fof(f5183,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f674,f659,f1970,f437])).
fof(f437,plain,(
  ( ! [X2,X0,X1] : (~sP10(X0,X1,X2) | sP11(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f252])).
fof(f659,plain,(
  ( ! [X0,X1] : (sP10(X0,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f468,f442])).
fof(f674,plain,(
  ( ! [X0] : (~sP11('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f268,f468,f449])).
fof(f5572,plain,(
  $false | ($spl2 | $spl190)),
  inference(global_subsumption,[],[f5185,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219])).
fof(f5185,plain,(
  ~sP10('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f674,f263,f1970,f437])).
fof(f5571,plain,(
  $false | ($spl2 | $spl190)),
  inference(global_subsumption,[],[f5186,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219])).
fof(f5186,plain,(
  sP11('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f659,f263,f1970,f437])).
fof(f5570,plain,(
  $false | ($spl190 | $spl312)),
  inference(global_subsumption,[],[f5188,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219])).
fof(f5188,plain,(
  sP11('F_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | ($spl190 | $spl312)),
  inference(unit_resulting_resolution,[],[f2519,f263,f1970,f437])).
fof(f5569,plain,(
  $false | ($spl2 | $spl190)),
  inference(global_subsumption,[],[f5189,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219])).
fof(f5189,plain,(
  ~sP10('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f674,f262,f1970,f437])).
fof(f5568,plain,(
  $false | ($spl2 | $spl190)),
  inference(global_subsumption,[],[f5190,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219])).
fof(f5190,plain,(
  sP11('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f659,f262,f1970,f437])).
fof(f5567,plain,(
  $false | ($spl190 | $spl312)),
  inference(global_subsumption,[],[f5192,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219])).
fof(f5192,plain,(
  sP11('F_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | ($spl190 | $spl312)),
  inference(unit_resulting_resolution,[],[f2519,f262,f1970,f437])).
fof(f5566,plain,(
  $false | ($spl2 | $spl190)),
  inference(global_subsumption,[],[f5193,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219])).
fof(f5193,plain,(
  ~sP10('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f674,f1970,f1970,f437])).
fof(f5565,plain,(
  $false | ($spl2 | $spl190)),
  inference(global_subsumption,[],[f5194,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219])).
fof(f5194,plain,(
  sP11('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f659,f1970,f1970,f437])).
fof(f5564,plain,(
  $false | ($spl190 | $spl312)),
  inference(global_subsumption,[],[f5196,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219])).
fof(f5196,plain,(
  sP11('F_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') | ($spl190 | $spl312)),
  inference(unit_resulting_resolution,[],[f2519,f1970,f1970,f437])).
fof(f5563,plain,(
  $false | ($spl2 | $spl190)),
  inference(global_subsumption,[],[f5202,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219])).
fof(f5202,plain,(
  sP11('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f659,f1970,f1970,f437])).
fof(f5562,plain,(
  $false | ($spl190 | $spl312)),
  inference(global_subsumption,[],[f5203,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219])).
fof(f5203,plain,(
  sP11('F_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') | ($spl190 | $spl312)),
  inference(unit_resulting_resolution,[],[f2519,f1970,f1970,f437])).
fof(f5561,plain,(
  $false | ($spl2 | $spl190)),
  inference(global_subsumption,[],[f5204,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219])).
fof(f5204,plain,(
  ~sP10('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f674,f1970,f1970,f437])).
fof(f5559,plain,(
  $false | ($spl190 | $spl274)),
  inference(global_subsumption,[],[f5208,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219])).
fof(f5208,plain,(
  sP11('F_0','F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | ($spl190 | $spl274)),
  inference(unit_resulting_resolution,[],[f2341,f268,f1970,f438])).
fof(f438,plain,(
  ( ! [X2,X0,X1] : (~sP10(X0,X1,X2) | sP11(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f252])).
fof(f5551,plain,(
  $false | ($spl104 | $spl152 | $spl190)),
  inference(global_subsumption,[],[f5550,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219])).
fof(f5550,plain,(
  ~p(s(bool,'T_0')) | ($spl104 | $spl152 | $spl190)),
  inference(forward_demodulation,[],[f5217,f792])).
fof(f5217,plain,(
  ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) | ($spl152 | $spl190)),
  inference(unit_resulting_resolution,[],[f1254,f268,f1970,f440])).
fof(f5548,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220,f5219])).
fof(f5547,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221,f5220])).
fof(f5546,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228,f5221])).
fof(f5539,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229,f5228])).
fof(f5538,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230,f5229])).
fof(f5537,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232,f5230])).
fof(f5536,plain,(
  $false | ($spl2 | $spl104 | $spl190)),
  inference(global_subsumption,[],[f5535,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232])).
fof(f5535,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl104 | $spl190)),
  inference(forward_demodulation,[],[f5231,f792])).
fof(f5231,plain,(
  ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f1164,f268,f1970,f440])).
fof(f1164,plain,(
  ( ! [X0] : (sP10('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X0,'F_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f661,f435])).
fof(f435,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f252])).
fof(f661,plain,(
  ( ! [X0] : (sP11('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X0,'F_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f268,f468,f445])).
fof(f5534,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235,f5232])).
fof(f5532,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237,f5235])).
fof(f5531,plain,(
  $false | ($spl2 | $spl190)),
  inference(global_subsumption,[],[f5236,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237])).
fof(f5236,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f674,f1970,f445])).
fof(f5530,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242,f5237])).
fof(f5529,plain,(
  $false | ($spl35 | $spl104 | $spl190)),
  inference(global_subsumption,[],[f5528,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242])).
fof(f5528,plain,(
  ~p(s(bool,'T_0')) | ($spl35 | $spl104 | $spl190)),
  inference(forward_demodulation,[],[f5239,f792])).
fof(f5239,plain,(
  ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) | ($spl35 | $spl190)),
  inference(unit_resulting_resolution,[],[f722,f1970,f446])).
fof(f5527,plain,(
  $false | ($spl33 | $spl104 | $spl190)),
  inference(global_subsumption,[],[f5526,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242])).
fof(f5526,plain,(
  ~p(s(bool,'T_0')) | ($spl33 | $spl104 | $spl190)),
  inference(forward_demodulation,[],[f5240,f792])).
fof(f5240,plain,(
  ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) | ($spl33 | $spl190)),
  inference(unit_resulting_resolution,[],[f720,f1970,f446])).
fof(f5525,plain,(
  $false | ($spl37 | $spl190)),
  inference(global_subsumption,[],[f5241,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242])).
fof(f5241,plain,(
  ~p(s(bool,'T_0')) | ($spl37 | $spl190)),
  inference(unit_resulting_resolution,[],[f724,f1970,f446])).
fof(f5524,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244,f5242])).
fof(f5523,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246,f5244])).
fof(f5522,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248,f5246])).
fof(f5521,plain,(
  $false | ($spl2 | $spl104 | $spl190)),
  inference(global_subsumption,[],[f5520,f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248])).
fof(f5520,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl104 | $spl190)),
  inference(forward_demodulation,[],[f5247,f792])).
fof(f5247,plain,(
  ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f674,f1970,f446])).
fof(f5519,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249,f5248])).
fof(f5518,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250,f5249])).
fof(f5517,plain,(
  $false | $spl190),
  inference(global_subsumption,[],[f261,f260,f347,f343,f1970,f5263,f5260,f5259,f5258,f5256,f5254,f5252,f5250])).
fof(f5516,plain,(
  $false | ($spl2 | $spl104 | $spl190)),
  inference(global_subsumption,[],[f5515,f261,f260,f468,f681,f680,f677,f676,f675,f674,f673,f672,f671,f670,f669,f668,f667,f666,f663,f661,f660,f659,f658,f657,f656,f655,f654,f653,f651,f650,f649,f648,f647,f644,f642,f641,f638,f636,f635,f634,f631,f630,f629,f628,f627,f626,f625,f624,f623,f622,f621,f620,f619,f618,f617,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f590,f587,f586,f585,f584,f583,f582,f581,f580,f579,f578,f577,f576,f575,f574,f573,f572,f571,f570,f569,f568,f567,f566,f565,f564,f563,f562,f561,f560,f559,f558,f557,f556,f555,f554,f553,f552,f551,f550,f549,f548,f547,f546,f545,f541,f540,f538,f537,f536,f535,f533,f532,f530,f529,f523,f522,f521,f520,f682,f683,f684,f689,f695,f696,f697,f700,f707,f710,f711,f713,f714,f715,f716,f793,f942,f948,f944,f958,f957,f1040,f1053,f1079,f1110,f1152,f1164,f1212,f1706,f1712,f1718,f1724,f1888,f1893,f1897,f1898,f1899,f1900,f1901,f1902,f1903,f1904,f1905,f1906,f1907,f1908,f1909,f1910,f1911,f1912,f1913,f1887,f347,f343,f1970,f5356,f5354,f5352,f5350,f5348,f5346,f5342,f5341,f5340,f5339,f5335,f5334,f5333,f5332,f5330,f5328,f5324,f5321,f5318,f5316,f5314,f5312,f5263,f5260,f5259,f5258,f5256,f5254,f5252])).
fof(f5312,plain,(
  ( ! [X0] : (sP9('F_0',X0,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) ) | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f635,f1970,f462])).
fof(f462,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X1,X2) | sP9(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(duplicate_literal_removal,[],[f419])).
fof(f419,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f244])).
fof(f5314,plain,(
  ( ! [X0] : (sP9('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X0)) ) | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f636,f1970,f462])).
fof(f5316,plain,(
  ( ! [X0] : (~sP8('F_0',X0,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) ) | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f650,f1970,f462])).
fof(f5318,plain,(
  ( ! [X0] : (~sP8('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X0)) ) | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f651,f1970,f462])).
fof(f5321,plain,(
  ( ! [X0] : (sP11('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X0)) ) | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f659,f1970,f463])).
fof(f463,plain,(
  ( ! [X2,X0,X1] : (~sP10(X0,X1,X2) | sP11(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(duplicate_literal_removal,[],[f436])).
fof(f436,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f252])).
fof(f5324,plain,(
  ( ! [X0] : (~sP10('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X0)) ) | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f674,f1970,f463])).
fof(f5328,plain,(
  ~sP1('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f1970,f695])).
fof(f5330,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f1970,f700])).
fof(f5332,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','F_0') | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f263,f1970,f707])).
fof(f5333,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','F_0') | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f262,f1970,f707])).
fof(f5334,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0','F_0') | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f1970,f1970,f707])).
fof(f5335,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f624,f1970,f707])).
fof(f5339,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0','T_0') | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f263,f1970,f707])).
fof(f5340,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0','T_0') | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f262,f1970,f707])).
fof(f5341,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0','F_0') | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f1970,f1970,f707])).
fof(f5342,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f625,f1970,f707])).
fof(f5346,plain,(
  ( ! [X0] : (sP9('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X0,'F_0')) ) | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f1970,f710])).
fof(f5348,plain,(
  ( ! [X0] : (sP9('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0',X0)) ) | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f1970,f711])).
fof(f5350,plain,(
  ( ! [X0] : (sP11('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0',X0)) ) | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f1970,f715])).
fof(f5352,plain,(
  ( ! [X0] : (sP9('F_0',X0,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) ) | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f1970,f1040])).
fof(f5354,plain,(
  ( ! [X0] : (sP9('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X0)) ) | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f1970,f1053])).
fof(f5356,plain,(
  ( ! [X0] : (sP11('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X0)) ) | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f1970,f1152])).
fof(f1913,plain,(
  ( ! [X111,X114,X112,X115,X113] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,X111),'const.combin.o_2'(s(fun(sK12,X111),X112),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun(X111,bool),X113)))) | sP10('const.pred_set.INJ_3'(s(fun(sK12,X111),X112),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))),s(fun(X111,bool),X113)),X114,X115)) ) | $spl2),
  inference(resolution,[],[f682,f443])).
fof(f443,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f255])).
fof(f1912,plain,(
  ( ! [X109,X107,X110,X108,X106] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,X106),'const.combin.o_2'(s(fun(sK12,X106),X107),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun(X106,bool),X108)))) | sP10(X109,X110,'const.pred_set.INJ_3'(s(fun(sK12,X106),X107),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))),s(fun(X106,bool),X108)))) ) | $spl2),
  inference(resolution,[],[f682,f441])).
fof(f441,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f255])).
fof(f1911,plain,(
  ( ! [X103,X101,X105,X102,X104] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,X101),'const.combin.o_2'(s(fun(sK12,X101),X102),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun(X101,bool),X103)))) | sP9(X104,X105,'const.pred_set.INJ_3'(s(fun(sK12,X101),X102),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))),s(fun(X101,bool),X103))) | p(s(bool,X105)) | p(s(bool,X104))) ) | $spl2),
  inference(resolution,[],[f682,f427])).
fof(f427,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f249])).
fof(f1910,plain,(
  ( ! [X99,X97,X100,X98,X96] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,X96),'const.combin.o_2'(s(fun(sK12,X96),X97),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun(X96,bool),X98)))) | sP9('const.pred_set.INJ_3'(s(fun(sK12,X96),X97),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))),s(fun(X96,bool),X98)),X99,X100) | p(s(bool,X99)) | p(s(bool,X100))) ) | $spl2),
  inference(resolution,[],[f682,f427])).
fof(f1909,plain,(
  ( ! [X94,X92,X95,X93,X91] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,X91),'const.combin.o_2'(s(fun(sK12,X91),X92),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun(X91,bool),X93)))) | sP9(X94,'const.pred_set.INJ_3'(s(fun(sK12,X91),X92),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))),s(fun(X91,bool),X93)),X95) | p(s(bool,X95)) | p(s(bool,X94))) ) | $spl2),
  inference(resolution,[],[f682,f427])).
fof(f1908,plain,(
  ( ! [X90,X88,X87,X89,X86] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,X86),'const.combin.o_2'(s(fun(sK12,X86),X87),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun(X86,bool),X88)))) | sP8('const.pred_set.INJ_3'(s(fun(sK12,X86),X87),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))),s(fun(X86,bool),X88)),X89,X90)) ) | $spl2),
  inference(resolution,[],[f682,f426])).
fof(f426,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f247])).
fof(f1907,plain,(
  ( ! [X85,X83,X81,X84,X82] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,X81),'const.combin.o_2'(s(fun(sK12,X81),X82),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun(X81,bool),X83)))) | sP7('const.pred_set.INJ_3'(s(fun(sK12,X81),X82),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))),s(fun(X81,bool),X83)),X84,X85) | p(s(bool,X84))) ) | $spl2),
  inference(resolution,[],[f682,f411])).
fof(f411,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f241])).
fof(f1906,plain,(
  ( ! [X80,X78,X76,X79,X77] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,X76),'const.combin.o_2'(s(fun(sK12,X76),X77),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun(X76,bool),X78)))) | sP7(X79,'const.pred_set.INJ_3'(s(fun(sK12,X76),X77),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))),s(fun(X76,bool),X78)),X80) | p(s(bool,X79))) ) | $spl2),
  inference(resolution,[],[f682,f411])).
fof(f1905,plain,(
  ( ! [X74,X72,X71,X75,X73] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,X71),'const.combin.o_2'(s(fun(sK12,X71),X72),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun(X71,bool),X73)))) | sP7(X74,X75,'const.pred_set.INJ_3'(s(fun(sK12,X71),X72),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))),s(fun(X71,bool),X73))) | p(s(bool,X74))) ) | $spl2),
  inference(resolution,[],[f682,f410])).
fof(f410,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f241])).
fof(f1904,plain,(
  ( ! [X70,X68,X66,X69,X67] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,X66),'const.combin.o_2'(s(fun(sK12,X66),X67),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun(X66,bool),X68)))) | sP7('const.pred_set.INJ_3'(s(fun(sK12,X66),X67),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))),s(fun(X66,bool),X68)),X69,X70) | p(s(bool,X70))) ) | $spl2),
  inference(resolution,[],[f682,f410])).
fof(f1903,plain,(
  ( ! [X61,X64,X62,X65,X63] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,X61),'const.combin.o_2'(s(fun(sK12,X61),X62),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun(X61,bool),X63)))) | sP6('const.pred_set.INJ_3'(s(fun(sK12,X61),X62),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))),s(fun(X61,bool),X63)),X64,X65)) ) | $spl2),
  inference(resolution,[],[f682,f409])).
fof(f409,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f239])).
fof(f1902,plain,(
  ( ! [X59,X57,X60,X58,X56] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,X56),'const.combin.o_2'(s(fun(sK12,X56),X57),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun(X56,bool),X58)))) | sP6(X59,'const.pred_set.INJ_3'(s(fun(sK12,X56),X57),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))),s(fun(X56,bool),X58)),X60)) ) | $spl2),
  inference(resolution,[],[f682,f408])).
fof(f408,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f239])).
fof(f1901,plain,(
  ( ! [X54,X52,X55,X53,X51] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,X51),'const.combin.o_2'(s(fun(sK12,X51),X52),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun(X51,bool),X53)))) | sP5('const.pred_set.INJ_3'(s(fun(sK12,X51),X52),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))),s(fun(X51,bool),X53)),X54,X55) | s(bool,X54) = s(bool,X55)) ) | $spl2),
  inference(resolution,[],[f682,f395])).
fof(f395,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f234])).
fof(f1900,plain,(
  ( ! [X47,X50,X48,X46,X49] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,X46),'const.combin.o_2'(s(fun(sK12,X46),X47),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun(X46,bool),X48)))) | sP3('const.pred_set.INJ_3'(s(fun(sK12,X46),X47),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))),s(fun(X46,bool),X48)),X49,X50)) ) | $spl2),
  inference(resolution,[],[f682,f390])).
fof(f390,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f231])).
fof(f1899,plain,(
  ( ! [X45,X43,X41,X44,X42] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,X41),'const.combin.o_2'(s(fun(sK12,X41),X42),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun(X41,bool),X43)))) | sP3(X44,'const.pred_set.INJ_3'(s(fun(sK12,X41),X42),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))),s(fun(X41,bool),X43)),X45)) ) | $spl2),
  inference(resolution,[],[f682,f389])).
fof(f389,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f231])).
fof(f1898,plain,(
  ( ! [X39,X37,X38,X36,X40] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,X36),'const.combin.o_2'(s(fun(sK12,X36),X37),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun(X36,bool),X38)))) | sP4('const.pred_set.INJ_3'(s(fun(sK12,X36),X37),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))),s(fun(X36,bool),X38)),X39,X40)) ) | $spl2),
  inference(resolution,[],[f682,f386])).
fof(f386,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f228])).
fof(f1897,plain,(
  ( ! [X35,X33,X31,X34,X32] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,X31),'const.combin.o_2'(s(fun(sK12,X31),X32),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun(X31,bool),X33)))) | sP4(X34,'const.pred_set.INJ_3'(s(fun(sK12,X31),X32),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))),s(fun(X31,bool),X33)),X35)) ) | $spl2),
  inference(resolution,[],[f682,f385])).
fof(f385,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f228])).
fof(f1893,plain,(
  ( ! [X17,X15,X18,X16] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,X15),'const.combin.o_2'(s(fun(sK12,X15),X16),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun(X15,bool),X17)))) | p(s(bool,X18)) | s(bool,X18) = s(bool,'const.pred_set.INJ_3'(s(fun(sK12,X15),X16),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))),s(fun(X15,bool),X17)))) ) | $spl2),
  inference(resolution,[],[f682,f331])).
fof(f331,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f130])).
fof(f1888,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,X0),'const.combin.o_2'(s(fun(sK12,X0),X1),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun(X0,bool),X2)))) | s(bool,'F_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK12,X0),X1),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))),s(fun(X0,bool),X2)))) ) | $spl2),
  inference(resolution,[],[f682,f280])).
fof(f280,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f172])).
fof(f172,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f61])).
fof(f61,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f60])).
fof(f60,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X14] : (s(bool,'F_0') = s(bool,X14) <=> ~p(s(bool,X14)))),
  file('thm.pred_set.inj_image_countable_IFF','thm.bool.EQ_CLAUSES|split|2')).
fof(f1724,plain,(
  ( ! [X0] : (sP5('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),happ(s(fun(bool,bool),X0),s(bool,'F_0')),'|dollar|exists_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X0)))) ) | $spl2),
  inference(superposition,[],[f587,f269])).
fof(f1718,plain,(
  ( ! [X0,X1] : (sP4('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'|dollar|exists_1'(s(fun(X0,bool),X1)),happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f587,f378])).
fof(f378,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP4(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f225])).
fof(f1712,plain,(
  ( ! [X0] : (sP5('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'|dollar|exists_1'(s(fun(bool,bool),X0)),happ(s(fun(bool,bool),X0),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X0)))) ) | $spl2),
  inference(superposition,[],[f586,f269])).
fof(f1706,plain,(
  ( ! [X0,X1] : (sP4('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))),'|dollar|exists_1'(s(fun(X0,bool),X1)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f586,f378])).
fof(f1212,plain,(
  ( ! [X0] : (sP10('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0',X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f666,f435])).
fof(f1152,plain,(
  ( ! [X2,X3] : (sP11(X2,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X3) | ~p(s(bool,X2))) ) | $spl2),
  inference(resolution,[],[f659,f463])).
fof(f1110,plain,(
  ( ! [X0] : (sP8('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0',X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f647,f418])).
fof(f418,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f244])).
fof(f1079,plain,(
  ( ! [X0] : (sP8('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X0,'T_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f641,f418])).
fof(f1053,plain,(
  ( ! [X2,X3] : (sP9(X2,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X3) | ~p(s(bool,X2))) ) | $spl2),
  inference(resolution,[],[f636,f462])).
fof(f1040,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | ~p(s(bool,X2))) ) | $spl2),
  inference(resolution,[],[f635,f462])).
fof(f944,plain,(
  ( ! [X6,X7] : (p(s(bool,X6)) | sP5(X6,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X7) | ~sP3('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X7,X6) | ~sP2(X7,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X6)) ) | $spl2),
  inference(resolution,[],[f575,f382])).
fof(f382,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X2,X1) | p(s(bool,X0)) | sP5(X0,X1,X2) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f225])).
fof(f948,plain,(
  ( ! [X2,X3] : (~p(s(bool,X2)) | sP5(X3,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X2) | ~sP3('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X2,X3)) ) | $spl2),
  inference(global_subsumption,[],[f261,f260,f468,f681,f680,f677,f676,f675,f674,f673,f672,f671,f670,f669,f668,f667,f666,f663,f661,f660,f659,f658,f657,f656,f655,f654,f653,f651,f650,f649,f648,f647,f644,f642,f641,f638,f636,f635,f634,f631,f630,f629,f628,f627,f626,f625,f624,f623,f622,f621,f620,f619,f618,f617,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f590,f587,f586,f585,f584,f583,f582,f581,f580,f579,f578,f577,f576,f575,f574,f573,f572,f571,f570,f569,f568,f567,f566,f565,f564,f563,f562,f561,f560,f559,f558,f557,f556,f555,f554,f553,f552,f551,f550,f549,f548,f547,f546,f545,f541,f540,f538,f537,f536,f535,f533,f532,f530,f529,f523,f522,f521,f520,f682,f683,f684,f689,f695,f696,f697,f700,f707,f710,f711,f713,f714,f715,f716,f793,f942])).
fof(f942,plain,(
  ( ! [X2,X3] : (~p(s(bool,X2)) | sP5(X3,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X2) | ~sP3('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X2,X3) | ~sP2(X2,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X3)) ) | $spl2),
  inference(resolution,[],[f575,f380])).
fof(f716,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl2),
  inference(superposition,[],[f468,f269])).
fof(f715,plain,(
  ( ! [X43,X44] : (sP11('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X43,X44) | ~p(s(bool,X43))) ) | $spl2),
  inference(resolution,[],[f468,f446])).
fof(f714,plain,(
  ( ! [X41,X42] : (sP11('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X42,X41) | p(s(bool,X41))) ) | $spl2),
  inference(resolution,[],[f468,f445])).
fof(f713,plain,(
  ( ! [X39,X40] : (sP11(X40,X39,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | p(s(bool,X39)) | p(s(bool,X40))) ) | $spl2),
  inference(resolution,[],[f468,f444])).
fof(f711,plain,(
  ( ! [X35,X36] : (sP9('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X35,X36) | ~p(s(bool,X35))) ) | $spl2),
  inference(resolution,[],[f468,f429])).
fof(f710,plain,(
  ( ! [X33,X34] : (sP9('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X34,X33) | ~p(s(bool,X33))) ) | $spl2),
  inference(resolution,[],[f468,f428])).
fof(f707,plain,(
  ( ! [X28,X27] : (sP7('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X27,X28) | ~p(s(bool,X28)) | ~p(s(bool,X27))) ) | $spl2),
  inference(resolution,[],[f468,f412])).
fof(f700,plain,(
  ( ! [X14] : (~sP1('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X14) | ~p(s(bool,X14))) ) | $spl2),
  inference(resolution,[],[f468,f366])).
fof(f697,plain,(
  ( ! [X11] : (sP1(X11,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | p(s(bool,X11))) ) | $spl2),
  inference(resolution,[],[f468,f363])).
fof(f696,plain,(
  ( ! [X10] : (sP1('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X10) | p(s(bool,X10))) ) | $spl2),
  inference(resolution,[],[f468,f360])).
fof(f695,plain,(
  ( ! [X9] : (~sP1(X9,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | ~p(s(bool,X9))) ) | $spl2),
  inference(resolution,[],[f468,f359])).
fof(f689,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | s(bool,X3) = s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) ) | $spl2),
  inference(resolution,[],[f468,f334])).
fof(f682,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK12,X0),X1),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))),s(fun(X0,bool),X2)))) | p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,X0),'const.combin.o_2'(s(fun(sK12,X0),X1),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun(X0,bool),X2))))) ) | $spl2),
  inference(resolution,[],[f468,f450])).
fof(f450,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X3),s(fun(X1,bool),X5),s(fun(X0,bool),X6)))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X2),X4),s(fun(X0,bool),X6),s(fun(X2,bool),X7)))) | p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X2),'const.combin.o_2'(s(fun(X0,X2),X4),s(fun(X1,X0),X3))),s(fun(X1,bool),X5),s(fun(X2,bool),X7))))) )),
  inference(cnf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (~p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X3),s(fun(X1,bool),X5),s(fun(X0,bool),X6)))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X2),X4),s(fun(X0,bool),X6),s(fun(X2,bool),X7)))) | p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X2),'const.combin.o_2'(s(fun(X0,X2),X4),s(fun(X1,X0),X3))),s(fun(X1,bool),X5),s(fun(X2,bool),X7)))))),
  inference(flattening,[],[f145])).
fof(f145,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((~p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X3),s(fun(X1,bool),X5),s(fun(X0,bool),X6)))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X2),X4),s(fun(X0,bool),X6),s(fun(X2,bool),X7))))) | p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X2),'const.combin.o_2'(s(fun(X0,X2),X4),s(fun(X1,X0),X3))),s(fun(X1,bool),X5),s(fun(X2,bool),X7)))))),
  inference(ennf_transformation,[],[f116])).
fof(f116,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X3),s(fun(X1,bool),X5),s(fun(X0,bool),X6)))) & p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X2),X4),s(fun(X0,bool),X6),s(fun(X2,bool),X7))))) => p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X2),'const.combin.o_2'(s(fun(X0,X2),X4),s(fun(X1,X0),X3))),s(fun(X1,bool),X5),s(fun(X2,bool),X7)))))),
  inference(flattening,[],[f115])).
fof(f115,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5,X6,X7] : ((p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X3),s(fun(X1,bool),X5),s(fun(X0,bool),X6)))) & p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X2),X4),s(fun(X0,bool),X6),s(fun(X2,bool),X7))))) => p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X2),'const.combin.o_2'(s(fun(X0,X2),X4),s(fun(X1,X0),X3))),s(fun(X1,bool),X5),s(fun(X2,bool),X7)))))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X6,X7,X20] : ! [X9,X21,X8,X14,X22] : ((p(s(bool,'const.pred_set.INJ_3'(s(fun(X7,X6),X9),s(fun(X7,bool),X8),s(fun(X6,bool),X14)))) & p(s(bool,'const.pred_set.INJ_3'(s(fun(X6,X20),X21),s(fun(X6,bool),X14),s(fun(X20,bool),X22))))) => p(s(bool,'const.pred_set.INJ_3'(s(fun(X7,X20),'const.combin.o_2'(s(fun(X6,X20),X21),s(fun(X7,X6),X9))),s(fun(X7,bool),X8),s(fun(X20,bool),X22)))))),
  file('thm.pred_set.inj_image_countable_IFF','thm.pred_set.INJ_COMPOSE')).
fof(f545,plain,(
  ( ! [X0] : (sP0(X0,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f268,f468,f349])).
fof(f349,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,'F_0')) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f213])).
fof(f546,plain,(
  ( ! [X0] : (sP0('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f468,f350])).
fof(f585,plain,(
  ( ! [X0,X1] : (sP2('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X0,X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f468,f394])).
fof(f586,plain,(
  ( ! [X0,X1] : (sP5('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'|dollar|exists_1'(s(fun(X0,bool),X1)),happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f317,f468,f396])).
fof(f587,plain,(
  ( ! [X0,X1] : (sP5('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))),'|dollar|exists_1'(s(fun(X0,bool),X1)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f317,f468,f396])).
fof(f638,plain,(
  ( ! [X0] : (sP9('T_0',X0,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f262,f468,f428])).
fof(f641,plain,(
  ( ! [X0] : (sP9('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X0,'T_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f262,f468,f428])).
fof(f642,plain,(
  ( ! [X0] : (sP9('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X0,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f468,f468,f428])).
fof(f644,plain,(
  ( ! [X0] : (sP9('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f262,f468,f429])).
fof(f647,plain,(
  ( ! [X0] : (sP9('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0',X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f262,f468,f429])).
fof(f648,plain,(
  ( ! [X0] : (sP9('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f468,f468,f429])).
fof(f663,plain,(
  ( ! [X0] : (sP11('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f262,f468,f446])).
fof(f666,plain,(
  ( ! [X0] : (sP11('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0',X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f262,f468,f446])).
fof(f667,plain,(
  ( ! [X0] : (sP11('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f468,f468,f446])).
fof(f5515,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl104 | $spl190)),
  inference(forward_demodulation,[],[f5251,f792])).
fof(f5251,plain,(
  ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f661,f268,f1970,f447])).
fof(f5487,plain,(
  $false | ($spl2 | $spl190)),
  inference(global_subsumption,[],[f261,f260,f468,f681,f680,f677,f676,f675,f674,f673,f672,f671,f670,f669,f668,f667,f666,f663,f661,f660,f659,f658,f657,f656,f655,f654,f653,f651,f650,f649,f648,f647,f644,f642,f641,f638,f636,f635,f634,f631,f630,f629,f628,f627,f626,f625,f624,f623,f622,f621,f620,f619,f618,f617,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f590,f587,f586,f585,f584,f583,f582,f581,f580,f579,f578,f577,f576,f575,f574,f573,f572,f571,f570,f569,f568,f567,f566,f565,f564,f563,f562,f561,f560,f559,f558,f557,f556,f555,f554,f553,f552,f551,f550,f549,f548,f547,f546,f545,f541,f540,f538,f537,f536,f535,f533,f532,f530,f529,f523,f522,f521,f520,f682,f683,f684,f689,f695,f696,f697,f700,f707,f710,f711,f713,f714,f715,f716,f793,f942,f948,f944,f958,f957,f1040,f1053,f1079,f1110,f1152,f1164,f1212,f1706,f1712,f1718,f1724,f1888,f1893,f1897,f1898,f1899,f1900,f1901,f1902,f1903,f1904,f1905,f1906,f1907,f1908,f1909,f1910,f1911,f1912,f1913,f1887,f347,f343,f1970,f5356,f5354,f5352,f5350,f5348,f5346,f5342,f5341,f5340,f5339,f5335,f5334,f5333,f5332,f5330,f5328,f5324,f5321,f5318,f5316,f5314,f5312])).
fof(f5486,plain,(
  $false | ($spl2 | $spl190)),
  inference(global_subsumption,[],[f261,f260,f468,f681,f680,f677,f676,f675,f674,f673,f672,f671,f670,f669,f668,f667,f666,f663,f661,f660,f659,f658,f657,f656,f655,f654,f653,f651,f650,f649,f648,f647,f644,f642,f641,f638,f636,f635,f634,f631,f630,f629,f628,f627,f626,f625,f624,f623,f622,f621,f620,f619,f618,f617,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f590,f587,f586,f585,f584,f583,f582,f581,f580,f579,f578,f577,f576,f575,f574,f573,f572,f571,f570,f569,f568,f567,f566,f565,f564,f563,f562,f561,f560,f559,f558,f557,f556,f555,f554,f553,f552,f551,f550,f549,f548,f547,f546,f545,f541,f540,f538,f537,f536,f535,f533,f532,f530,f529,f523,f522,f521,f520,f682,f683,f684,f689,f695,f696,f697,f700,f707,f710,f711,f713,f714,f715,f716,f793,f942,f948,f944,f958,f957,f1040,f1053,f1079,f1110,f1152,f1164,f1212,f1706,f1712,f1718,f1724,f1888,f1893,f1897,f1898,f1899,f1900,f1901,f1902,f1903,f1904,f1905,f1906,f1907,f1908,f1909,f1910,f1911,f1912,f1913,f1887,f347,f343,f1970,f5356,f5354,f5352,f5350,f5348,f5346,f5342,f5341,f5340,f5339,f5335,f5334,f5333,f5332,f5330,f5328,f5324,f5321,f5318,f5316,f5314])).
fof(f5485,plain,(
  $false | ($spl2 | $spl190)),
  inference(global_subsumption,[],[f261,f260,f468,f681,f680,f677,f676,f675,f674,f673,f672,f671,f670,f669,f668,f667,f666,f663,f661,f660,f659,f658,f657,f656,f655,f654,f653,f651,f650,f649,f648,f647,f644,f642,f641,f638,f636,f635,f634,f631,f630,f629,f628,f627,f626,f625,f624,f623,f622,f621,f620,f619,f618,f617,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f590,f587,f586,f585,f584,f583,f582,f581,f580,f579,f578,f577,f576,f575,f574,f573,f572,f571,f570,f569,f568,f567,f566,f565,f564,f563,f562,f561,f560,f559,f558,f557,f556,f555,f554,f553,f552,f551,f550,f549,f548,f547,f546,f545,f541,f540,f538,f537,f536,f535,f533,f532,f530,f529,f523,f522,f521,f520,f682,f683,f684,f689,f695,f696,f697,f700,f707,f710,f711,f713,f714,f715,f716,f793,f942,f948,f944,f958,f957,f1040,f1053,f1079,f1110,f1152,f1164,f1212,f1706,f1712,f1718,f1724,f1888,f1893,f1897,f1898,f1899,f1900,f1901,f1902,f1903,f1904,f1905,f1906,f1907,f1908,f1909,f1910,f1911,f1912,f1913,f1887,f347,f343,f1970,f5356,f5354,f5352,f5350,f5348,f5346,f5342,f5341,f5340,f5339,f5335,f5334,f5333,f5332,f5330,f5328,f5324,f5321,f5318,f5316])).
fof(f5484,plain,(
  $false | ($spl2 | $spl190)),
  inference(global_subsumption,[],[f261,f260,f468,f681,f680,f677,f676,f675,f674,f673,f672,f671,f670,f669,f668,f667,f666,f663,f661,f660,f659,f658,f657,f656,f655,f654,f653,f651,f650,f649,f648,f647,f644,f642,f641,f638,f636,f635,f634,f631,f630,f629,f628,f627,f626,f625,f624,f623,f622,f621,f620,f619,f618,f617,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f590,f587,f586,f585,f584,f583,f582,f581,f580,f579,f578,f577,f576,f575,f574,f573,f572,f571,f570,f569,f568,f567,f566,f565,f564,f563,f562,f561,f560,f559,f558,f557,f556,f555,f554,f553,f552,f551,f550,f549,f548,f547,f546,f545,f541,f540,f538,f537,f536,f535,f533,f532,f530,f529,f523,f522,f521,f520,f682,f683,f684,f689,f695,f696,f697,f700,f707,f710,f711,f713,f714,f715,f716,f793,f942,f948,f944,f958,f957,f1040,f1053,f1079,f1110,f1152,f1164,f1212,f1706,f1712,f1718,f1724,f1888,f1893,f1897,f1898,f1899,f1900,f1901,f1902,f1903,f1904,f1905,f1906,f1907,f1908,f1909,f1910,f1911,f1912,f1913,f1887,f347,f343,f1970,f5356,f5354,f5352,f5350,f5348,f5346,f5342,f5341,f5340,f5339,f5335,f5334,f5333,f5332,f5330,f5328,f5324,f5321,f5318])).
fof(f5483,plain,(
  $false | ($spl2 | $spl190)),
  inference(global_subsumption,[],[f261,f260,f468,f681,f680,f677,f676,f675,f674,f673,f672,f671,f670,f669,f668,f667,f666,f663,f661,f660,f659,f658,f657,f656,f655,f654,f653,f651,f650,f649,f648,f647,f644,f642,f641,f638,f636,f635,f634,f631,f630,f629,f628,f627,f626,f625,f624,f623,f622,f621,f620,f619,f618,f617,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f590,f587,f586,f585,f584,f583,f582,f581,f580,f579,f578,f577,f576,f575,f574,f573,f572,f571,f570,f569,f568,f567,f566,f565,f564,f563,f562,f561,f560,f559,f558,f557,f556,f555,f554,f553,f552,f551,f550,f549,f548,f547,f546,f545,f541,f540,f538,f537,f536,f535,f533,f532,f530,f529,f523,f522,f521,f520,f682,f683,f684,f689,f695,f696,f697,f700,f707,f710,f711,f713,f714,f715,f716,f793,f942,f948,f944,f958,f957,f1040,f1053,f1079,f1110,f1152,f1164,f1212,f1706,f1712,f1718,f1724,f1888,f1893,f1897,f1898,f1899,f1900,f1901,f1902,f1903,f1904,f1905,f1906,f1907,f1908,f1909,f1910,f1911,f1912,f1913,f1887,f347,f343,f1970,f5356,f5354,f5352,f5350,f5348,f5346,f5342,f5341,f5340,f5339,f5335,f5334,f5333,f5332,f5330,f5328,f5324,f5321])).
fof(f5482,plain,(
  $false | ($spl2 | $spl190)),
  inference(global_subsumption,[],[f261,f260,f468,f681,f680,f677,f676,f675,f674,f673,f672,f671,f670,f669,f668,f667,f666,f663,f661,f660,f659,f658,f657,f656,f655,f654,f653,f651,f650,f649,f648,f647,f644,f642,f641,f638,f636,f635,f634,f631,f630,f629,f628,f627,f626,f625,f624,f623,f622,f621,f620,f619,f618,f617,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f590,f587,f586,f585,f584,f583,f582,f581,f580,f579,f578,f577,f576,f575,f574,f573,f572,f571,f570,f569,f568,f567,f566,f565,f564,f563,f562,f561,f560,f559,f558,f557,f556,f555,f554,f553,f552,f551,f550,f549,f548,f547,f546,f545,f541,f540,f538,f537,f536,f535,f533,f532,f530,f529,f523,f522,f521,f520,f682,f683,f684,f689,f695,f696,f697,f700,f707,f710,f711,f713,f714,f715,f716,f793,f942,f948,f944,f958,f957,f1040,f1053,f1079,f1110,f1152,f1164,f1212,f1706,f1712,f1718,f1724,f1888,f1893,f1897,f1898,f1899,f1900,f1901,f1902,f1903,f1904,f1905,f1906,f1907,f1908,f1909,f1910,f1911,f1912,f1913,f1887,f347,f343,f1970,f5356,f5354,f5352,f5350,f5348,f5346,f5342,f5341,f5340,f5339,f5335,f5334,f5333,f5332,f5330,f5328,f5324])).
fof(f5480,plain,(
  $false | ($spl2 | $spl190)),
  inference(global_subsumption,[],[f261,f260,f468,f681,f680,f677,f676,f675,f674,f673,f672,f671,f670,f669,f668,f667,f666,f663,f661,f660,f659,f658,f657,f656,f655,f654,f653,f651,f650,f649,f648,f647,f644,f642,f641,f638,f636,f635,f634,f631,f630,f629,f628,f627,f626,f625,f624,f623,f622,f621,f620,f619,f618,f617,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f590,f587,f586,f585,f584,f583,f582,f581,f580,f579,f578,f577,f576,f575,f574,f573,f572,f571,f570,f569,f568,f567,f566,f565,f564,f563,f562,f561,f560,f559,f558,f557,f556,f555,f554,f553,f552,f551,f550,f549,f548,f547,f546,f545,f541,f540,f538,f537,f536,f535,f533,f532,f530,f529,f523,f522,f521,f520,f682,f683,f684,f689,f695,f696,f697,f700,f707,f710,f711,f713,f714,f715,f716,f793,f942,f948,f944,f958,f957,f1040,f1053,f1079,f1110,f1152,f1164,f1212,f1706,f1712,f1718,f1724,f1888,f1893,f1897,f1898,f1899,f1900,f1901,f1902,f1903,f1904,f1905,f1906,f1907,f1908,f1909,f1910,f1911,f1912,f1913,f1887,f347,f343,f1970,f5356,f5354,f5352,f5350,f5348,f5346,f5342,f5341,f5340,f5339,f5335,f5334,f5333,f5332,f5330,f5328])).
fof(f5479,plain,(
  $false | ($spl2 | $spl190)),
  inference(global_subsumption,[],[f261,f260,f468,f681,f680,f677,f676,f675,f674,f673,f672,f671,f670,f669,f668,f667,f666,f663,f661,f660,f659,f658,f657,f656,f655,f654,f653,f651,f650,f649,f648,f647,f644,f642,f641,f638,f636,f635,f634,f631,f630,f629,f628,f627,f626,f625,f624,f623,f622,f621,f620,f619,f618,f617,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f590,f587,f586,f585,f584,f583,f582,f581,f580,f579,f578,f577,f576,f575,f574,f573,f572,f571,f570,f569,f568,f567,f566,f565,f564,f563,f562,f561,f560,f559,f558,f557,f556,f555,f554,f553,f552,f551,f550,f549,f548,f547,f546,f545,f541,f540,f538,f537,f536,f535,f533,f532,f530,f529,f523,f522,f521,f520,f682,f683,f684,f689,f695,f696,f697,f700,f707,f710,f711,f713,f714,f715,f716,f793,f942,f948,f944,f958,f957,f1040,f1053,f1079,f1110,f1152,f1164,f1212,f1706,f1712,f1718,f1724,f1888,f1893,f1897,f1898,f1899,f1900,f1901,f1902,f1903,f1904,f1905,f1906,f1907,f1908,f1909,f1910,f1911,f1912,f1913,f1887,f347,f343,f1970,f5356,f5354,f5352,f5350,f5348,f5346,f5342,f5341,f5340,f5339,f5335,f5334,f5333,f5332,f5330])).
fof(f5478,plain,(
  $false | ($spl2 | $spl190)),
  inference(global_subsumption,[],[f261,f260,f468,f681,f680,f677,f676,f675,f674,f673,f672,f671,f670,f669,f668,f667,f666,f663,f661,f660,f659,f658,f657,f656,f655,f654,f653,f651,f650,f649,f648,f647,f644,f642,f641,f638,f636,f635,f634,f631,f630,f629,f628,f627,f626,f625,f624,f623,f622,f621,f620,f619,f618,f617,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f590,f587,f586,f585,f584,f583,f582,f581,f580,f579,f578,f577,f576,f575,f574,f573,f572,f571,f570,f569,f568,f567,f566,f565,f564,f563,f562,f561,f560,f559,f558,f557,f556,f555,f554,f553,f552,f551,f550,f549,f548,f547,f546,f545,f541,f540,f538,f537,f536,f535,f533,f532,f530,f529,f523,f522,f521,f520,f682,f683,f684,f689,f695,f696,f697,f700,f707,f710,f711,f713,f714,f715,f716,f793,f942,f948,f944,f958,f957,f1040,f1053,f1079,f1110,f1152,f1164,f1212,f1706,f1712,f1718,f1724,f1888,f1893,f1897,f1898,f1899,f1900,f1901,f1902,f1903,f1904,f1905,f1906,f1907,f1908,f1909,f1910,f1911,f1912,f1913,f1887,f347,f343,f1970,f5356,f5354,f5352,f5350,f5348,f5346,f5342,f5341,f5340,f5339,f5335,f5334,f5333,f5332])).
fof(f5477,plain,(
  $false | ($spl2 | $spl190)),
  inference(global_subsumption,[],[f261,f260,f468,f681,f680,f677,f676,f675,f674,f673,f672,f671,f670,f669,f668,f667,f666,f663,f661,f660,f659,f658,f657,f656,f655,f654,f653,f651,f650,f649,f648,f647,f644,f642,f641,f638,f636,f635,f634,f631,f630,f629,f628,f627,f626,f625,f624,f623,f622,f621,f620,f619,f618,f617,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f590,f587,f586,f585,f584,f583,f582,f581,f580,f579,f578,f577,f576,f575,f574,f573,f572,f571,f570,f569,f568,f567,f566,f565,f564,f563,f562,f561,f560,f559,f558,f557,f556,f555,f554,f553,f552,f551,f550,f549,f548,f547,f546,f545,f541,f540,f538,f537,f536,f535,f533,f532,f530,f529,f523,f522,f521,f520,f682,f683,f684,f689,f695,f696,f697,f700,f707,f710,f711,f713,f714,f715,f716,f793,f942,f948,f944,f958,f957,f1040,f1053,f1079,f1110,f1152,f1164,f1212,f1706,f1712,f1718,f1724,f1888,f1893,f1897,f1898,f1899,f1900,f1901,f1902,f1903,f1904,f1905,f1906,f1907,f1908,f1909,f1910,f1911,f1912,f1913,f1887,f347,f343,f1970,f5356,f5354,f5352,f5350,f5348,f5346,f5342,f5341,f5340,f5339,f5335,f5334,f5333])).
fof(f5476,plain,(
  $false | ($spl2 | $spl190)),
  inference(global_subsumption,[],[f261,f260,f468,f681,f680,f677,f676,f675,f674,f673,f672,f671,f670,f669,f668,f667,f666,f663,f661,f660,f659,f658,f657,f656,f655,f654,f653,f651,f650,f649,f648,f647,f644,f642,f641,f638,f636,f635,f634,f631,f630,f629,f628,f627,f626,f625,f624,f623,f622,f621,f620,f619,f618,f617,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f590,f587,f586,f585,f584,f583,f582,f581,f580,f579,f578,f577,f576,f575,f574,f573,f572,f571,f570,f569,f568,f567,f566,f565,f564,f563,f562,f561,f560,f559,f558,f557,f556,f555,f554,f553,f552,f551,f550,f549,f548,f547,f546,f545,f541,f540,f538,f537,f536,f535,f533,f532,f530,f529,f523,f522,f521,f520,f682,f683,f684,f689,f695,f696,f697,f700,f707,f710,f711,f713,f714,f715,f716,f793,f942,f948,f944,f958,f957,f1040,f1053,f1079,f1110,f1152,f1164,f1212,f1706,f1712,f1718,f1724,f1888,f1893,f1897,f1898,f1899,f1900,f1901,f1902,f1903,f1904,f1905,f1906,f1907,f1908,f1909,f1910,f1911,f1912,f1913,f1887,f347,f343,f1970,f5356,f5354,f5352,f5350,f5348,f5346,f5342,f5341,f5340,f5339,f5335,f5334])).
fof(f5475,plain,(
  $false | ($spl2 | $spl190)),
  inference(global_subsumption,[],[f261,f260,f468,f681,f680,f677,f676,f675,f674,f673,f672,f671,f670,f669,f668,f667,f666,f663,f661,f660,f659,f658,f657,f656,f655,f654,f653,f651,f650,f649,f648,f647,f644,f642,f641,f638,f636,f635,f634,f631,f630,f629,f628,f627,f626,f625,f624,f623,f622,f621,f620,f619,f618,f617,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f590,f587,f586,f585,f584,f583,f582,f581,f580,f579,f578,f577,f576,f575,f574,f573,f572,f571,f570,f569,f568,f567,f566,f565,f564,f563,f562,f561,f560,f559,f558,f557,f556,f555,f554,f553,f552,f551,f550,f549,f548,f547,f546,f545,f541,f540,f538,f537,f536,f535,f533,f532,f530,f529,f523,f522,f521,f520,f682,f683,f684,f689,f695,f696,f697,f700,f707,f710,f711,f713,f714,f715,f716,f793,f942,f948,f944,f958,f957,f1040,f1053,f1079,f1110,f1152,f1164,f1212,f1706,f1712,f1718,f1724,f1888,f1893,f1897,f1898,f1899,f1900,f1901,f1902,f1903,f1904,f1905,f1906,f1907,f1908,f1909,f1910,f1911,f1912,f1913,f1887,f347,f343,f1970,f5356,f5354,f5352,f5350,f5348,f5346,f5342,f5341,f5340,f5339,f5335])).
fof(f5474,plain,(
  $false | ($spl2 | $spl190)),
  inference(global_subsumption,[],[f261,f260,f468,f681,f680,f677,f676,f675,f674,f673,f672,f671,f670,f669,f668,f667,f666,f663,f661,f660,f659,f658,f657,f656,f655,f654,f653,f651,f650,f649,f648,f647,f644,f642,f641,f638,f636,f635,f634,f631,f630,f629,f628,f627,f626,f625,f624,f623,f622,f621,f620,f619,f618,f617,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f590,f587,f586,f585,f584,f583,f582,f581,f580,f579,f578,f577,f576,f575,f574,f573,f572,f571,f570,f569,f568,f567,f566,f565,f564,f563,f562,f561,f560,f559,f558,f557,f556,f555,f554,f553,f552,f551,f550,f549,f548,f547,f546,f545,f541,f540,f538,f537,f536,f535,f533,f532,f530,f529,f523,f522,f521,f520,f682,f683,f684,f689,f695,f696,f697,f700,f707,f710,f711,f713,f714,f715,f716,f793,f942,f948,f944,f958,f957,f1040,f1053,f1079,f1110,f1152,f1164,f1212,f1706,f1712,f1718,f1724,f1888,f1893,f1897,f1898,f1899,f1900,f1901,f1902,f1903,f1904,f1905,f1906,f1907,f1908,f1909,f1910,f1911,f1912,f1913,f1887,f347,f343,f1970,f5356,f5354,f5352,f5350,f5348,f5346,f5342,f5341,f5340,f5339])).
fof(f5473,plain,(
  $false | ($spl2 | $spl190)),
  inference(global_subsumption,[],[f261,f260,f468,f681,f680,f677,f676,f675,f674,f673,f672,f671,f670,f669,f668,f667,f666,f663,f661,f660,f659,f658,f657,f656,f655,f654,f653,f651,f650,f649,f648,f647,f644,f642,f641,f638,f636,f635,f634,f631,f630,f629,f628,f627,f626,f625,f624,f623,f622,f621,f620,f619,f618,f617,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f590,f587,f586,f585,f584,f583,f582,f581,f580,f579,f578,f577,f576,f575,f574,f573,f572,f571,f570,f569,f568,f567,f566,f565,f564,f563,f562,f561,f560,f559,f558,f557,f556,f555,f554,f553,f552,f551,f550,f549,f548,f547,f546,f545,f541,f540,f538,f537,f536,f535,f533,f532,f530,f529,f523,f522,f521,f520,f682,f683,f684,f689,f695,f696,f697,f700,f707,f710,f711,f713,f714,f715,f716,f793,f942,f948,f944,f958,f957,f1040,f1053,f1079,f1110,f1152,f1164,f1212,f1706,f1712,f1718,f1724,f1888,f1893,f1897,f1898,f1899,f1900,f1901,f1902,f1903,f1904,f1905,f1906,f1907,f1908,f1909,f1910,f1911,f1912,f1913,f1887,f347,f343,f1970,f5356,f5354,f5352,f5350,f5348,f5346,f5342,f5341,f5340])).
fof(f5472,plain,(
  $false | ($spl2 | $spl190)),
  inference(global_subsumption,[],[f261,f260,f468,f681,f680,f677,f676,f675,f674,f673,f672,f671,f670,f669,f668,f667,f666,f663,f661,f660,f659,f658,f657,f656,f655,f654,f653,f651,f650,f649,f648,f647,f644,f642,f641,f638,f636,f635,f634,f631,f630,f629,f628,f627,f626,f625,f624,f623,f622,f621,f620,f619,f618,f617,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f590,f587,f586,f585,f584,f583,f582,f581,f580,f579,f578,f577,f576,f575,f574,f573,f572,f571,f570,f569,f568,f567,f566,f565,f564,f563,f562,f561,f560,f559,f558,f557,f556,f555,f554,f553,f552,f551,f550,f549,f548,f547,f546,f545,f541,f540,f538,f537,f536,f535,f533,f532,f530,f529,f523,f522,f521,f520,f682,f683,f684,f689,f695,f696,f697,f700,f707,f710,f711,f713,f714,f715,f716,f793,f942,f948,f944,f958,f957,f1040,f1053,f1079,f1110,f1152,f1164,f1212,f1706,f1712,f1718,f1724,f1888,f1893,f1897,f1898,f1899,f1900,f1901,f1902,f1903,f1904,f1905,f1906,f1907,f1908,f1909,f1910,f1911,f1912,f1913,f1887,f347,f343,f1970,f5356,f5354,f5352,f5350,f5348,f5346,f5342,f5341])).
fof(f5471,plain,(
  $false | ($spl2 | $spl190)),
  inference(global_subsumption,[],[f261,f260,f468,f681,f680,f677,f676,f675,f674,f673,f672,f671,f670,f669,f668,f667,f666,f663,f661,f660,f659,f658,f657,f656,f655,f654,f653,f651,f650,f649,f648,f647,f644,f642,f641,f638,f636,f635,f634,f631,f630,f629,f628,f627,f626,f625,f624,f623,f622,f621,f620,f619,f618,f617,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f590,f587,f586,f585,f584,f583,f582,f581,f580,f579,f578,f577,f576,f575,f574,f573,f572,f571,f570,f569,f568,f567,f566,f565,f564,f563,f562,f561,f560,f559,f558,f557,f556,f555,f554,f553,f552,f551,f550,f549,f548,f547,f546,f545,f541,f540,f538,f537,f536,f535,f533,f532,f530,f529,f523,f522,f521,f520,f682,f683,f684,f689,f695,f696,f697,f700,f707,f710,f711,f713,f714,f715,f716,f793,f942,f948,f944,f958,f957,f1040,f1053,f1079,f1110,f1152,f1164,f1212,f1706,f1712,f1718,f1724,f1888,f1893,f1897,f1898,f1899,f1900,f1901,f1902,f1903,f1904,f1905,f1906,f1907,f1908,f1909,f1910,f1911,f1912,f1913,f1887,f347,f343,f1970,f5356,f5354,f5352,f5350,f5348,f5346,f5342])).
fof(f5470,plain,(
  $false | ($spl2 | $spl190)),
  inference(global_subsumption,[],[f261,f260,f468,f681,f680,f677,f676,f675,f674,f673,f672,f671,f670,f669,f668,f667,f666,f663,f661,f660,f659,f658,f657,f656,f655,f654,f653,f651,f650,f649,f648,f647,f644,f642,f641,f638,f636,f635,f634,f631,f630,f629,f628,f627,f626,f625,f624,f623,f622,f621,f620,f619,f618,f617,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f590,f587,f586,f585,f584,f583,f582,f581,f580,f579,f578,f577,f576,f575,f574,f573,f572,f571,f570,f569,f568,f567,f566,f565,f564,f563,f562,f561,f560,f559,f558,f557,f556,f555,f554,f553,f552,f551,f550,f549,f548,f547,f546,f545,f541,f540,f538,f537,f536,f535,f533,f532,f530,f529,f523,f522,f521,f520,f682,f683,f684,f689,f695,f696,f697,f700,f707,f710,f711,f713,f714,f715,f716,f793,f942,f948,f944,f958,f957,f1040,f1053,f1079,f1110,f1152,f1164,f1212,f1706,f1712,f1718,f1724,f1888,f1893,f1897,f1898,f1899,f1900,f1901,f1902,f1903,f1904,f1905,f1906,f1907,f1908,f1909,f1910,f1911,f1912,f1913,f1887,f347,f343,f1970,f5356,f5354,f5352,f5350,f5348,f5346])).
fof(f5469,plain,(
  $false | ($spl2 | $spl190)),
  inference(global_subsumption,[],[f261,f260,f468,f681,f680,f677,f676,f675,f674,f673,f672,f671,f670,f669,f668,f667,f666,f663,f661,f660,f659,f658,f657,f656,f655,f654,f653,f651,f650,f649,f648,f647,f644,f642,f641,f638,f636,f635,f634,f631,f630,f629,f628,f627,f626,f625,f624,f623,f622,f621,f620,f619,f618,f617,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f590,f587,f586,f585,f584,f583,f582,f581,f580,f579,f578,f577,f576,f575,f574,f573,f572,f571,f570,f569,f568,f567,f566,f565,f564,f563,f562,f561,f560,f559,f558,f557,f556,f555,f554,f553,f552,f551,f550,f549,f548,f547,f546,f545,f541,f540,f538,f537,f536,f535,f533,f532,f530,f529,f523,f522,f521,f520,f682,f683,f684,f689,f695,f696,f697,f700,f707,f710,f711,f713,f714,f715,f716,f793,f942,f948,f944,f958,f957,f1040,f1053,f1079,f1110,f1152,f1164,f1212,f1706,f1712,f1718,f1724,f1888,f1893,f1897,f1898,f1899,f1900,f1901,f1902,f1903,f1904,f1905,f1906,f1907,f1908,f1909,f1910,f1911,f1912,f1913,f1887,f347,f343,f1970,f5356,f5354,f5352,f5350,f5348])).
fof(f5468,plain,(
  $false | ($spl2 | $spl190)),
  inference(global_subsumption,[],[f261,f260,f468,f681,f680,f677,f676,f675,f674,f673,f672,f671,f670,f669,f668,f667,f666,f663,f661,f660,f659,f658,f657,f656,f655,f654,f653,f651,f650,f649,f648,f647,f644,f642,f641,f638,f636,f635,f634,f631,f630,f629,f628,f627,f626,f625,f624,f623,f622,f621,f620,f619,f618,f617,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f590,f587,f586,f585,f584,f583,f582,f581,f580,f579,f578,f577,f576,f575,f574,f573,f572,f571,f570,f569,f568,f567,f566,f565,f564,f563,f562,f561,f560,f559,f558,f557,f556,f555,f554,f553,f552,f551,f550,f549,f548,f547,f546,f545,f541,f540,f538,f537,f536,f535,f533,f532,f530,f529,f523,f522,f521,f520,f682,f683,f684,f689,f695,f696,f697,f700,f707,f710,f711,f713,f714,f715,f716,f793,f942,f948,f944,f958,f957,f1040,f1053,f1079,f1110,f1152,f1164,f1212,f1706,f1712,f1718,f1724,f1888,f1893,f1897,f1898,f1899,f1900,f1901,f1902,f1903,f1904,f1905,f1906,f1907,f1908,f1909,f1910,f1911,f1912,f1913,f1887,f347,f343,f1970,f5356,f5354,f5352,f5350])).
fof(f5467,plain,(
  $false | ($spl2 | $spl190)),
  inference(global_subsumption,[],[f261,f260,f468,f681,f680,f677,f676,f675,f674,f673,f672,f671,f670,f669,f668,f667,f666,f663,f661,f660,f659,f658,f657,f656,f655,f654,f653,f651,f650,f649,f648,f647,f644,f642,f641,f638,f636,f635,f634,f631,f630,f629,f628,f627,f626,f625,f624,f623,f622,f621,f620,f619,f618,f617,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f590,f587,f586,f585,f584,f583,f582,f581,f580,f579,f578,f577,f576,f575,f574,f573,f572,f571,f570,f569,f568,f567,f566,f565,f564,f563,f562,f561,f560,f559,f558,f557,f556,f555,f554,f553,f552,f551,f550,f549,f548,f547,f546,f545,f541,f540,f538,f537,f536,f535,f533,f532,f530,f529,f523,f522,f521,f520,f682,f683,f684,f689,f695,f696,f697,f700,f707,f710,f711,f713,f714,f715,f716,f793,f942,f948,f944,f958,f957,f1040,f1053,f1079,f1110,f1152,f1164,f1212,f1706,f1712,f1718,f1724,f1888,f1893,f1897,f1898,f1899,f1900,f1901,f1902,f1903,f1904,f1905,f1906,f1907,f1908,f1909,f1910,f1911,f1912,f1913,f1887,f347,f343,f1970,f5356,f5354,f5352])).
fof(f5466,plain,(
  $false | ($spl2 | $spl190)),
  inference(global_subsumption,[],[f261,f260,f468,f681,f680,f677,f676,f675,f674,f673,f672,f671,f670,f669,f668,f667,f666,f663,f661,f660,f659,f658,f657,f656,f655,f654,f653,f651,f650,f649,f648,f647,f644,f642,f641,f638,f636,f635,f634,f631,f630,f629,f628,f627,f626,f625,f624,f623,f622,f621,f620,f619,f618,f617,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f590,f587,f586,f585,f584,f583,f582,f581,f580,f579,f578,f577,f576,f575,f574,f573,f572,f571,f570,f569,f568,f567,f566,f565,f564,f563,f562,f561,f560,f559,f558,f557,f556,f555,f554,f553,f552,f551,f550,f549,f548,f547,f546,f545,f541,f540,f538,f537,f536,f535,f533,f532,f530,f529,f523,f522,f521,f520,f682,f683,f684,f689,f695,f696,f697,f700,f707,f710,f711,f713,f714,f715,f716,f793,f942,f948,f944,f958,f957,f1040,f1053,f1079,f1110,f1152,f1164,f1212,f1706,f1712,f1718,f1724,f1888,f1893,f1897,f1898,f1899,f1900,f1901,f1902,f1903,f1904,f1905,f1906,f1907,f1908,f1909,f1910,f1911,f1912,f1913,f1887,f347,f343,f1970,f5356,f5354])).
fof(f5465,plain,(
  $false | ($spl2 | $spl190)),
  inference(global_subsumption,[],[f261,f260,f468,f681,f680,f677,f676,f675,f674,f673,f672,f671,f670,f669,f668,f667,f666,f663,f661,f660,f659,f658,f657,f656,f655,f654,f653,f651,f650,f649,f648,f647,f644,f642,f641,f638,f636,f635,f634,f631,f630,f629,f628,f627,f626,f625,f624,f623,f622,f621,f620,f619,f618,f617,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f590,f587,f586,f585,f584,f583,f582,f581,f580,f579,f578,f577,f576,f575,f574,f573,f572,f571,f570,f569,f568,f567,f566,f565,f564,f563,f562,f561,f560,f559,f558,f557,f556,f555,f554,f553,f552,f551,f550,f549,f548,f547,f546,f545,f541,f540,f538,f537,f536,f535,f533,f532,f530,f529,f523,f522,f521,f520,f682,f683,f684,f689,f695,f696,f697,f700,f707,f710,f711,f713,f714,f715,f716,f793,f942,f948,f944,f958,f957,f1040,f1053,f1079,f1110,f1152,f1164,f1212,f1706,f1712,f1718,f1724,f1888,f1893,f1897,f1898,f1899,f1900,f1901,f1902,f1903,f1904,f1905,f1906,f1907,f1908,f1909,f1910,f1911,f1912,f1913,f1887,f347,f343,f1970,f5356])).
fof(f5464,plain,(
  $false | ($spl12 | $spl104 | $spl190)),
  inference(global_subsumption,[],[f5359,f261,f260,f2373,f2372,f2371,f2354,f2353,f2352,f2412,f2413,f2414,f2415,f2416,f2417,f2418,f2419,f2420,f2421,f2424,f2427,f2428,f2429,f2431,f2432,f2433,f2434,f2435,f2442,f2443,f2444,f2449,f2450,f2452,f2454,f2455,f2456,f347,f343,f485,f3190,f3191,f3192,f3193,f4091,f4090,f4100,f4098,f4099,f4116,f4117,f4123,f4132,f4130,f4131,f4146,f4147,f1970,f5383,f5380])).
fof(f5380,plain,(
  sP5('F_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | ($spl12 | $spl190)),
  inference(unit_resulting_resolution,[],[f485,f1970,f3190])).
fof(f5383,plain,(
  sP5('F_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | ($spl12 | $spl190)),
  inference(unit_resulting_resolution,[],[f485,f1970,f3191])).
fof(f4147,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~sP5(X0,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0')) ) | $spl12),
  inference(equality_resolution,[],[f3193])).
fof(f4146,plain,(
  ( ! [X8,X7,X9] : (s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(X7,bool),X8))) | p(s(bool,X9)) | ~sP5(X9,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),happ(s(fun(X7,bool),X8),s(X7,'const.min.@_1'(s(fun(X7,bool),X8)))))) ) | $spl12),
  inference(superposition,[],[f3193,f317])).
fof(f4131,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~sP5(X0,'T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) ) | $spl12),
  inference(equality_resolution,[],[f3192])).
fof(f4130,plain,(
  ( ! [X8,X7,X9] : (s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(X7,bool),X8))) | p(s(bool,X9)) | ~sP5(X9,happ(s(fun(X7,bool),X8),s(X7,'const.min.@_1'(s(fun(X7,bool),X8)))),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) ) | $spl12),
  inference(superposition,[],[f3192,f317])).
fof(f4132,plain,(
  ( ! [X6] : (p(s(bool,X6)) | ~sP5(X6,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) ) | $spl12),
  inference(trivial_inequality_removal,[],[f4129])).
fof(f4129,plain,(
  ( ! [X6] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,X6)) | ~sP5(X6,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) ) | $spl12),
  inference(superposition,[],[f3192,f485])).
fof(f4117,plain,(
  ( ! [X0] : (sP5(X0,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | ~p(s(bool,X0))) ) | $spl12),
  inference(equality_resolution,[],[f3191])).
fof(f4116,plain,(
  ( ! [X8,X7,X9] : (s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(X7,bool),X8))) | sP5(X9,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),happ(s(fun(X7,bool),X8),s(X7,'const.min.@_1'(s(fun(X7,bool),X8))))) | ~p(s(bool,X9))) ) | $spl12),
  inference(superposition,[],[f3191,f317])).
fof(f4099,plain,(
  ( ! [X0] : (sP5(X0,'T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | ~p(s(bool,X0))) ) | $spl12),
  inference(equality_resolution,[],[f3190])).
fof(f4098,plain,(
  ( ! [X8,X7,X9] : (s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(X7,bool),X8))) | sP5(X9,happ(s(fun(X7,bool),X8),s(X7,'const.min.@_1'(s(fun(X7,bool),X8)))),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | ~p(s(bool,X9))) ) | $spl12),
  inference(superposition,[],[f3190,f317])).
fof(f4100,plain,(
  ( ! [X6] : (sP5(X6,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | ~p(s(bool,X6))) ) | $spl12),
  inference(trivial_inequality_removal,[],[f4097])).
fof(f4097,plain,(
  ( ! [X6] : (s(bool,'T_0') != s(bool,'T_0') | sP5(X6,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | ~p(s(bool,X6))) ) | $spl12),
  inference(superposition,[],[f3190,f485])).
fof(f3193,plain,(
  ( ! [X45,X46] : (s(bool,'T_0') != s(bool,X45) | p(s(bool,X46)) | ~sP5(X46,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X45)) ) | $spl12),
  inference(superposition,[],[f398,f485])).
fof(f398,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | p(s(bool,X2)) | ~sP5(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f234])).
fof(f3192,plain,(
  ( ! [X43,X44] : (s(bool,'T_0') != s(bool,X43) | p(s(bool,X44)) | ~sP5(X44,X43,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) ) | $spl12),
  inference(superposition,[],[f398,f485])).
fof(f3191,plain,(
  ( ! [X41,X42] : (s(bool,'T_0') != s(bool,X41) | sP5(X42,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X41) | ~p(s(bool,X42))) ) | $spl12),
  inference(superposition,[],[f396,f485])).
fof(f3190,plain,(
  ( ! [X39,X40] : (s(bool,'T_0') != s(bool,X39) | sP5(X40,X39,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | ~p(s(bool,X40))) ) | $spl12),
  inference(superposition,[],[f396,f485])).
fof(f2434,plain,(
  ( ! [X45,X46] : (s(bool,'T_0') != s(bool,X45) | p(s(bool,X46)) | ~sP5(X46,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X45)) ) | $spl12),
  inference(superposition,[],[f398,f485])).
fof(f2433,plain,(
  ( ! [X43,X44] : (s(bool,'T_0') != s(bool,X43) | p(s(bool,X44)) | ~sP5(X44,X43,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) ) | $spl12),
  inference(superposition,[],[f398,f485])).
fof(f2432,plain,(
  ( ! [X41,X42] : (s(bool,'T_0') != s(bool,X41) | sP5(X42,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X41) | ~p(s(bool,X42))) ) | $spl12),
  inference(superposition,[],[f396,f485])).
fof(f2431,plain,(
  ( ! [X39,X40] : (s(bool,'T_0') != s(bool,X39) | sP5(X40,X39,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | ~p(s(bool,X40))) ) | $spl12),
  inference(superposition,[],[f396,f485])).
fof(f5359,plain,(
  sP5('F_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | ($spl12 | $spl104 | $spl190)),
  inference(unit_resulting_resolution,[],[f485,f1970,f1528])).
fof(f1528,plain,(
  ( ! [X43,X44] : (s(bool,'T_0') != s(bool,X43) | sP5(X44,X43,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | ~p(s(bool,X44))) ) | $spl104),
  inference(superposition,[],[f396,f792])).
fof(f5463,plain,(
  $false | ($spl12 | $spl104 | $spl190)),
  inference(global_subsumption,[],[f5362,f261,f260,f2373,f2372,f2371,f2354,f2353,f2352,f2412,f2413,f2414,f2415,f2416,f2417,f2418,f2419,f2420,f2421,f2424,f2427,f2428,f2429,f2431,f2432,f2433,f2434,f2435,f2442,f2443,f2444,f2449,f2450,f2452,f2454,f2455,f2456,f347,f343,f485,f3190,f3191,f3192,f3193,f4091,f4090,f4100,f4098,f4099,f4116,f4117,f4123,f4132,f4130,f4131,f4146,f4147,f1970,f5383,f5380])).
fof(f5362,plain,(
  sP5('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | ($spl12 | $spl104 | $spl190)),
  inference(unit_resulting_resolution,[],[f485,f1970,f1529])).
fof(f1529,plain,(
  ( ! [X45,X46] : (s(bool,'T_0') != s(bool,X45) | sP5(X46,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X45) | ~p(s(bool,X46))) ) | $spl104),
  inference(superposition,[],[f396,f792])).
fof(f5462,plain,(
  $false | ($spl14 | $spl190)),
  inference(global_subsumption,[],[f261,f260,f2253,f2254,f2256,f2257,f2261,f2262,f2263,f2264,f2265,f2267,f2268,f2269,f2270,f2271,f2272,f2276,f2277,f2278,f2279,f2282,f2284,f347,f343,f486,f3040,f3041,f3042,f3043,f3739,f3738,f3742,f3753,f3752,f3751,f3749,f3750,f3762,f3773,f3772,f3769,f3770,f3786,f3789,f3800,f3799,f3798,f3796,f3797,f3808,f3819,f3818,f3815,f3816,f3062,f1970,f5377,f5375])).
fof(f5375,plain,(
  sP5('F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | ($spl14 | $spl190)),
  inference(unit_resulting_resolution,[],[f486,f1970,f3040])).
fof(f5377,plain,(
  sP5('F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | ($spl14 | $spl190)),
  inference(unit_resulting_resolution,[],[f486,f1970,f3041])).
fof(f3816,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~sP5(X0,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0')) ) | $spl14),
  inference(equality_resolution,[],[f3043])).
fof(f3815,plain,(
  ( ! [X8,X7,X9] : (s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(X7,bool),X8))) | p(s(bool,X9)) | ~sP5(X9,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),happ(s(fun(X7,bool),X8),s(X7,'const.min.@_1'(s(fun(X7,bool),X8)))))) ) | $spl14),
  inference(superposition,[],[f3043,f317])).
fof(f3819,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | ~sP5(X1,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl14),
  inference(trivial_inequality_removal,[],[f3810])).
fof(f3810,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,'F_0') | p(s(bool,X1)) | ~sP5(X1,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl14),
  inference(superposition,[],[f3043,f269])).
fof(f3808,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | p(s(bool,X2)) | ~sP5(X2,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X1) | s(bool,'T_0') = s(bool,X0)) ) | $spl14),
  inference(superposition,[],[f3043,f269])).
fof(f3797,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~sP5(X0,'F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) ) | $spl14),
  inference(equality_resolution,[],[f3042])).
fof(f3796,plain,(
  ( ! [X8,X7,X9] : (s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(X7,bool),X8))) | p(s(bool,X9)) | ~sP5(X9,happ(s(fun(X7,bool),X8),s(X7,'const.min.@_1'(s(fun(X7,bool),X8)))),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) ) | $spl14),
  inference(superposition,[],[f3042,f317])).
fof(f3798,plain,(
  ( ! [X5] : (p(s(bool,X5)) | ~sP5(X5,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) ) | $spl14),
  inference(trivial_inequality_removal,[],[f3794])).
fof(f3794,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,'F_0') | p(s(bool,X5)) | ~sP5(X5,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) ) | $spl14),
  inference(superposition,[],[f3042,f486])).
fof(f3800,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | ~sP5(X1,X0,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | s(bool,'T_0') = s(bool,X0)) ) | $spl14),
  inference(trivial_inequality_removal,[],[f3791])).
fof(f3791,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,'F_0') | p(s(bool,X1)) | ~sP5(X1,X0,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | s(bool,'T_0') = s(bool,X0)) ) | $spl14),
  inference(superposition,[],[f3042,f269])).
fof(f3789,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | p(s(bool,X2)) | ~sP5(X2,X1,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | s(bool,'T_0') = s(bool,X0)) ) | $spl14),
  inference(superposition,[],[f3042,f269])).
fof(f3770,plain,(
  ( ! [X0] : (sP5(X0,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0') | ~p(s(bool,X0))) ) | $spl14),
  inference(equality_resolution,[],[f3041])).
fof(f3769,plain,(
  ( ! [X8,X7,X9] : (s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(X7,bool),X8))) | sP5(X9,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),happ(s(fun(X7,bool),X8),s(X7,'const.min.@_1'(s(fun(X7,bool),X8))))) | ~p(s(bool,X9))) ) | $spl14),
  inference(superposition,[],[f3041,f317])).
fof(f3773,plain,(
  ( ! [X0,X1] : (sP5(X1,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X0) | ~p(s(bool,X1)) | s(bool,'T_0') = s(bool,X0)) ) | $spl14),
  inference(trivial_inequality_removal,[],[f3764])).
fof(f3764,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,'F_0') | sP5(X1,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X0) | ~p(s(bool,X1)) | s(bool,'T_0') = s(bool,X0)) ) | $spl14),
  inference(superposition,[],[f3041,f269])).
fof(f3762,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | sP5(X2,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X1) | ~p(s(bool,X2)) | s(bool,'T_0') = s(bool,X0)) ) | $spl14),
  inference(superposition,[],[f3041,f269])).
fof(f3750,plain,(
  ( ! [X0] : (sP5(X0,'F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | ~p(s(bool,X0))) ) | $spl14),
  inference(equality_resolution,[],[f3040])).
fof(f3749,plain,(
  ( ! [X8,X7,X9] : (s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(X7,bool),X8))) | sP5(X9,happ(s(fun(X7,bool),X8),s(X7,'const.min.@_1'(s(fun(X7,bool),X8)))),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | ~p(s(bool,X9))) ) | $spl14),
  inference(superposition,[],[f3040,f317])).
fof(f3751,plain,(
  ( ! [X5] : (sP5(X5,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | ~p(s(bool,X5))) ) | $spl14),
  inference(trivial_inequality_removal,[],[f3747])).
fof(f3747,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,'F_0') | sP5(X5,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | ~p(s(bool,X5))) ) | $spl14),
  inference(superposition,[],[f3040,f486])).
fof(f3753,plain,(
  ( ! [X0,X1] : (sP5(X1,X0,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | ~p(s(bool,X1)) | s(bool,'T_0') = s(bool,X0)) ) | $spl14),
  inference(trivial_inequality_removal,[],[f3744])).
fof(f3744,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,'F_0') | sP5(X1,X0,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | ~p(s(bool,X1)) | s(bool,'T_0') = s(bool,X0)) ) | $spl14),
  inference(superposition,[],[f3040,f269])).
fof(f3742,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | sP5(X2,X1,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | ~p(s(bool,X2)) | s(bool,'T_0') = s(bool,X0)) ) | $spl14),
  inference(superposition,[],[f3040,f269])).
fof(f3043,plain,(
  ( ! [X45,X46] : (s(bool,'F_0') != s(bool,X45) | p(s(bool,X46)) | ~sP5(X46,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X45)) ) | $spl14),
  inference(superposition,[],[f398,f486])).
fof(f3042,plain,(
  ( ! [X43,X44] : (s(bool,'F_0') != s(bool,X43) | p(s(bool,X44)) | ~sP5(X44,X43,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) ) | $spl14),
  inference(superposition,[],[f398,f486])).
fof(f3041,plain,(
  ( ! [X41,X42] : (s(bool,'F_0') != s(bool,X41) | sP5(X42,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X41) | ~p(s(bool,X42))) ) | $spl14),
  inference(superposition,[],[f396,f486])).
fof(f3040,plain,(
  ( ! [X39,X40] : (s(bool,'F_0') != s(bool,X39) | sP5(X40,X39,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | ~p(s(bool,X40))) ) | $spl14),
  inference(superposition,[],[f396,f486])).
fof(f2265,plain,(
  ( ! [X47,X46] : (s(bool,'F_0') != s(bool,X46) | p(s(bool,X47)) | ~sP5(X47,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X46)) ) | $spl14),
  inference(superposition,[],[f398,f486])).
fof(f2264,plain,(
  ( ! [X45,X44] : (s(bool,'F_0') != s(bool,X44) | p(s(bool,X45)) | ~sP5(X45,X44,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) ) | $spl14),
  inference(superposition,[],[f398,f486])).
fof(f2263,plain,(
  ( ! [X43,X42] : (s(bool,'F_0') != s(bool,X42) | sP5(X43,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X42) | ~p(s(bool,X43))) ) | $spl14),
  inference(superposition,[],[f396,f486])).
fof(f2262,plain,(
  ( ! [X41,X40] : (s(bool,'F_0') != s(bool,X40) | sP5(X41,X40,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | ~p(s(bool,X41))) ) | $spl14),
  inference(superposition,[],[f396,f486])).
fof(f5461,plain,(
  $false | ($spl14 | $spl190)),
  inference(global_subsumption,[],[f261,f260,f2253,f2254,f2256,f2257,f2261,f2262,f2263,f2264,f2265,f2267,f2268,f2269,f2270,f2271,f2272,f2276,f2277,f2278,f2279,f2282,f2284,f347,f343,f486,f3040,f3041,f3042,f3043,f3739,f3738,f3742,f3753,f3752,f3751,f3749,f3750,f3762,f3773,f3772,f3769,f3770,f3786,f3789,f3800,f3799,f3798,f3796,f3797,f3808,f3819,f3818,f3815,f3816,f3062,f1970,f5377])).
fof(f5460,plain,(
  $false | ($spl12 | $spl104 | $spl190)),
  inference(global_subsumption,[],[f5379,f261,f260,f2373,f2372,f2371,f2354,f2353,f2352,f2412,f2413,f2414,f2415,f2416,f2417,f2418,f2419,f2420,f2421,f2424,f2427,f2428,f2429,f2431,f2432,f2433,f2434,f2435,f2442,f2443,f2444,f2449,f2450,f2452,f2454,f2455,f2456,f347,f343,f485,f3190,f3191,f3192,f3193,f4091,f4090,f4100,f4098,f4099,f4116,f4117,f4123,f4132,f4130,f4131,f4146,f4147,f1970,f5383,f5380])).
fof(f5379,plain,(
  sP5('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | ($spl12 | $spl104 | $spl190)),
  inference(unit_resulting_resolution,[],[f792,f1970,f3190])).
fof(f5459,plain,(
  $false | ($spl12 | $spl190)),
  inference(global_subsumption,[],[f261,f260,f2373,f2372,f2371,f2354,f2353,f2352,f2412,f2413,f2414,f2415,f2416,f2417,f2418,f2419,f2420,f2421,f2424,f2427,f2428,f2429,f2431,f2432,f2433,f2434,f2435,f2442,f2443,f2444,f2449,f2450,f2452,f2454,f2455,f2456,f347,f343,f485,f3190,f3191,f3192,f3193,f4091,f4090,f4100,f4098,f4099,f4116,f4117,f4123,f4132,f4130,f4131,f4146,f4147,f1970,f5383,f5380])).
fof(f5458,plain,(
  $false | ($spl12 | $spl104 | $spl190)),
  inference(global_subsumption,[],[f5382,f261,f260,f2373,f2372,f2371,f2354,f2353,f2352,f2412,f2413,f2414,f2415,f2416,f2417,f2418,f2419,f2420,f2421,f2424,f2427,f2428,f2429,f2431,f2432,f2433,f2434,f2435,f2442,f2443,f2444,f2449,f2450,f2452,f2454,f2455,f2456,f347,f343,f485,f3190,f3191,f3192,f3193,f4091,f4090,f4100,f4098,f4099,f4116,f4117,f4123,f4132,f4130,f4131,f4146,f4147,f1970,f5383])).
fof(f5382,plain,(
  sP5('F_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | ($spl12 | $spl104 | $spl190)),
  inference(unit_resulting_resolution,[],[f792,f1970,f3191])).
fof(f5457,plain,(
  $false | ($spl12 | $spl190)),
  inference(global_subsumption,[],[f261,f260,f2373,f2372,f2371,f2354,f2353,f2352,f2412,f2413,f2414,f2415,f2416,f2417,f2418,f2419,f2420,f2421,f2424,f2427,f2428,f2429,f2431,f2432,f2433,f2434,f2435,f2442,f2443,f2444,f2449,f2450,f2452,f2454,f2455,f2456,f347,f343,f485,f3190,f3191,f3192,f3193,f4091,f4090,f4100,f4098,f4099,f4116,f4117,f4123,f4132,f4130,f4131,f4146,f4147,f1970,f5383])).
fof(f5456,plain,(
  $false | ($spl12 | $spl123 | $spl190)),
  inference(trivial_inequality_removal,[],[f5384])).
fof(f5384,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl12 | $spl123 | $spl190)),
  inference(unit_resulting_resolution,[],[f849,f1970,f3191])).
fof(f5455,plain,(
  $false | ($spl12 | $spl119 | $spl190)),
  inference(trivial_inequality_removal,[],[f5381])).
fof(f5381,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl12 | $spl119 | $spl190)),
  inference(unit_resulting_resolution,[],[f845,f1970,f3190])).
fof(f5454,plain,(
  $false | ($spl14 | $spl25 | $spl190)),
  inference(trivial_inequality_removal,[],[f5378])).
fof(f5378,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl14 | $spl25 | $spl190)),
  inference(unit_resulting_resolution,[],[f511,f1970,f3041])).
fof(f5453,plain,(
  $false | ($spl14 | $spl21 | $spl190)),
  inference(trivial_inequality_removal,[],[f5376])).
fof(f5376,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl14 | $spl21 | $spl190)),
  inference(unit_resulting_resolution,[],[f507,f1970,f3040])).
fof(f5452,plain,(
  $false | ($spl104 | $spl161 | $spl190)),
  inference(trivial_inequality_removal,[],[f5363])).
fof(f5363,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl104 | $spl161 | $spl190)),
  inference(unit_resulting_resolution,[],[f1389,f1970,f1529])).
fof(f5451,plain,(
  $false | ($spl104 | $spl159 | $spl190)),
  inference(trivial_inequality_removal,[],[f5360])).
fof(f5360,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl104 | $spl159 | $spl190)),
  inference(unit_resulting_resolution,[],[f1379,f1970,f1528])).
fof(f5357,plain,(
  $false | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f674,f1970,f1152])).
fof(f5355,plain,(
  $false | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f651,f1970,f1053])).
fof(f5353,plain,(
  $false | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f650,f1970,f1040])).
fof(f5351,plain,(
  $false | ($spl2 | $spl35 | $spl190)),
  inference(unit_resulting_resolution,[],[f722,f1970,f715])).
fof(f5349,plain,(
  $false | ($spl2 | $spl47 | $spl190)),
  inference(unit_resulting_resolution,[],[f734,f1970,f711])).
fof(f5347,plain,(
  $false | ($spl2 | $spl47 | $spl190)),
  inference(unit_resulting_resolution,[],[f734,f1970,f710])).
fof(f5345,plain,(
  $false | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f1970,f624,f1970,f707])).
fof(f5343,plain,(
  $false | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f263,f625,f1970,f707])).
fof(f5338,plain,(
  $false | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f1970,f625,f1970,f707])).
fof(f5336,plain,(
  $false | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f263,f624,f1970,f707])).
fof(f5331,plain,(
  $false | ($spl2 | $spl100 | $spl190)),
  inference(unit_resulting_resolution,[],[f788,f1970,f700])).
fof(f5329,plain,(
  $false | ($spl2 | $spl102 | $spl190)),
  inference(unit_resulting_resolution,[],[f790,f1970,f695])).
fof(f5450,plain,(
  $false | $spl190),
  inference(trivial_inequality_removal,[],[f5326])).
fof(f5326,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f268,f268,f1970,f464])).
fof(f464,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(duplicate_literal_removal,[],[f455])).
fof(f455,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f259])).
fof(f259,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f258])).
fof(f258,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f148])).
fof(f148,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f147])).
fof(f147,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f117])).
fof(f117,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X11,X18,X15,X19] : ((s(bool,X11) = s(bool,X18) & (p(s(bool,X18)) => s(bool,X15) = s(bool,X19))) => ((p(s(bool,X11)) => p(s(bool,X15))) <=> (p(s(bool,X18)) => p(s(bool,X19)))))),
  file('thm.pred_set.inj_image_countable_IFF','thm.bool.IMP_CONG')).
fof(f5325,plain,(
  $false | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f659,f674,f1970,f463])).
fof(f5319,plain,(
  $false | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f636,f651,f1970,f462])).
fof(f5317,plain,(
  $false | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f635,f650,f1970,f462])).
fof(f5449,plain,(
  $false | ($spl14 | $spl190)),
  inference(trivial_inequality_removal,[],[f5308])).
fof(f5308,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl14 | $spl190)),
  inference(unit_resulting_resolution,[],[f268,f1970,f486,f1970,f458])).
fof(f458,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X2) != s(bool,X3) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X2)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f259])).
fof(f5448,plain,(
  $false | ($spl12 | $spl190)),
  inference(trivial_inequality_removal,[],[f5307])).
fof(f5307,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl12 | $spl190)),
  inference(unit_resulting_resolution,[],[f268,f262,f485,f1970,f458])).
fof(f5447,plain,(
  $false | ($spl12 | $spl190)),
  inference(trivial_inequality_removal,[],[f5306])).
fof(f5306,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl12 | $spl190)),
  inference(unit_resulting_resolution,[],[f268,f263,f485,f1970,f458])).
fof(f5446,plain,(
  $false | ($spl104 | $spl190)),
  inference(trivial_inequality_removal,[],[f5305])).
fof(f5305,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl104 | $spl190)),
  inference(unit_resulting_resolution,[],[f268,f262,f792,f1970,f458])).
fof(f5445,plain,(
  $false | ($spl104 | $spl190)),
  inference(trivial_inequality_removal,[],[f5304])).
fof(f5304,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl104 | $spl190)),
  inference(unit_resulting_resolution,[],[f268,f263,f792,f1970,f458])).
fof(f5444,plain,(
  $false | ($spl14 | $spl190)),
  inference(trivial_inequality_removal,[],[f5299])).
fof(f5299,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl14 | $spl190)),
  inference(unit_resulting_resolution,[],[f1970,f268,f486,f1970,f458])).
fof(f5443,plain,(
  $false | $spl190),
  inference(trivial_inequality_removal,[],[f5292])).
fof(f5292,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f268,f268,f1970,f1970,f457])).
fof(f457,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f259])).
fof(f5442,plain,(
  $false | $spl190),
  inference(trivial_inequality_removal,[],[f5286])).
fof(f5286,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f268,f268,f1970,f1970,f457])).
fof(f5441,plain,(
  $false | $spl190),
  inference(trivial_inequality_removal,[],[f5285])).
fof(f5285,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f268,f268,f262,f1970,f457])).
fof(f5440,plain,(
  $false | $spl190),
  inference(trivial_inequality_removal,[],[f5284])).
fof(f5284,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f268,f268,f263,f1970,f457])).
fof(f5439,plain,(
  $false | ($spl14 | $spl190)),
  inference(trivial_inequality_removal,[],[f5282])).
fof(f5282,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl14 | $spl190)),
  inference(unit_resulting_resolution,[],[f268,f268,f486,f1970,f456])).
fof(f456,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X2) != s(bool,X3) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f259])).
fof(f5438,plain,(
  $false | ($spl14 | $spl190)),
  inference(trivial_inequality_removal,[],[f5278])).
fof(f5278,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl14 | $spl190)),
  inference(unit_resulting_resolution,[],[f268,f1970,f486,f1970,f454])).
fof(f454,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X2) != s(bool,X3) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f259])).
fof(f5437,plain,(
  $false | ($spl12 | $spl190)),
  inference(trivial_inequality_removal,[],[f5277])).
fof(f5277,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl12 | $spl190)),
  inference(unit_resulting_resolution,[],[f268,f262,f485,f1970,f454])).
fof(f5436,plain,(
  $false | ($spl12 | $spl190)),
  inference(trivial_inequality_removal,[],[f5276])).
fof(f5276,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl12 | $spl190)),
  inference(unit_resulting_resolution,[],[f268,f263,f485,f1970,f454])).
fof(f5435,plain,(
  $false | ($spl104 | $spl190)),
  inference(trivial_inequality_removal,[],[f5275])).
fof(f5275,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl104 | $spl190)),
  inference(unit_resulting_resolution,[],[f268,f262,f792,f1970,f454])).
fof(f5434,plain,(
  $false | ($spl104 | $spl190)),
  inference(trivial_inequality_removal,[],[f5274])).
fof(f5274,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl104 | $spl190)),
  inference(unit_resulting_resolution,[],[f268,f263,f792,f1970,f454])).
fof(f5433,plain,(
  $false | ($spl14 | $spl190)),
  inference(trivial_inequality_removal,[],[f5269])).
fof(f5269,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl14 | $spl190)),
  inference(unit_resulting_resolution,[],[f1970,f268,f486,f1970,f454])).
fof(f5432,plain,(
  $false | ($spl14 | $spl190)),
  inference(trivial_inequality_removal,[],[f5265])).
fof(f5265,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl14 | $spl190)),
  inference(unit_resulting_resolution,[],[f268,f268,f486,f1970,f452])).
fof(f452,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X2) != s(bool,X3) | ~p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,X3)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f259])).
fof(f5264,plain,(
  $false | ($spl38 | $spl190)),
  inference(unit_resulting_resolution,[],[f726,f268,f1970,f449])).
fof(f5245,plain,(
  $false | ($spl37 | $spl190)),
  inference(unit_resulting_resolution,[],[f724,f262,f1970,f446])).
fof(f5243,plain,(
  $false | ($spl37 | $spl190)),
  inference(unit_resulting_resolution,[],[f724,f263,f1970,f446])).
fof(f5238,plain,(
  $false | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f674,f268,f1970,f445])).
fof(f5233,plain,(
  $false | ($spl43 | $spl190)),
  inference(unit_resulting_resolution,[],[f730,f1970,f442])).
fof(f5205,plain,(
  $false | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f659,f674,f1970,f1970,f437])).
fof(f5195,plain,(
  $false | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f674,f659,f1970,f1970,f437])).
fof(f5191,plain,(
  $false | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f674,f659,f262,f1970,f437])).
fof(f5187,plain,(
  $false | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f674,f659,f263,f1970,f437])).
fof(f5182,plain,(
  $false | ($spl38 | $spl190)),
  inference(unit_resulting_resolution,[],[f726,f268,f1970,f434])).
fof(f5169,plain,(
  $false | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f650,f1970,f1970,f429])).
fof(f5167,plain,(
  $false | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f650,f262,f1970,f429])).
fof(f5165,plain,(
  $false | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f650,f263,f1970,f429])).
fof(f5161,plain,(
  $false | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f650,f1970,f1970,f429])).
fof(f5155,plain,(
  $false | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f651,f1970,f1970,f428])).
fof(f5153,plain,(
  $false | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f651,f262,f1970,f428])).
fof(f5151,plain,(
  $false | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f651,f263,f1970,f428])).
fof(f5147,plain,(
  $false | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f651,f1970,f1970,f428])).
fof(f5141,plain,(
  $false | ($spl49 | $spl190)),
  inference(unit_resulting_resolution,[],[f736,f1970,f425])).
fof(f5139,plain,(
  $false | ($spl49 | $spl190)),
  inference(unit_resulting_resolution,[],[f736,f1970,f424])).
fof(f5132,plain,(
  $false | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f635,f650,f268,f1970,f421])).
fof(f5123,plain,(
  $false | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f651,f636,f268,f1970,f420])).
fof(f5060,plain,(
  $false | ($spl73 | $spl190)),
  inference(unit_resulting_resolution,[],[f760,f1970,f407])).
fof(f5031,plain,(
  $false | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f625,f596,f1970,f1970,f402])).
fof(f5026,plain,(
  $false | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f624,f596,f262,f1970,f402])).
fof(f5022,plain,(
  $false | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f624,f596,f263,f1970,f402])).
fof(f5015,plain,(
  $false | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f625,f596,f1970,f1970,f402])).
fof(f5010,plain,(
  $false | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f625,f596,f262,f1970,f402])).
fof(f5006,plain,(
  $false | ($spl2 | $spl190)),
  inference(unit_resulting_resolution,[],[f625,f596,f263,f1970,f402])).
fof(f4998,plain,(
  $false | ($spl27 | $spl190 | $spl196)),
  inference(unit_resulting_resolution,[],[f2147,f513,f1970,f397])).
fof(f4996,plain,(
  $false | ($spl29 | $spl112 | $spl190)),
  inference(unit_resulting_resolution,[],[f826,f518,f1970,f397])).
fof(f4994,plain,(
  $false | ($spl27 | $spl190 | $spl192)),
  inference(unit_resulting_resolution,[],[f2133,f513,f1970,f397])).
fof(f4992,plain,(
  $false | ($spl29 | $spl108 | $spl190)),
  inference(unit_resulting_resolution,[],[f822,f518,f1970,f397])).
fof(f4984,plain,(
  $false | ($spl14 | $spl25 | $spl190)),
  inference(unit_resulting_resolution,[],[f511,f486,f1970,f396])).
fof(f4982,plain,(
  $false | ($spl12 | $spl123 | $spl190)),
  inference(unit_resulting_resolution,[],[f849,f485,f1970,f396])).
fof(f4980,plain,(
  $false | ($spl104 | $spl161 | $spl190)),
  inference(unit_resulting_resolution,[],[f1389,f792,f1970,f396])).
fof(f4977,plain,(
  $false | ($spl14 | $spl21 | $spl190)),
  inference(unit_resulting_resolution,[],[f507,f486,f1970,f396])).
fof(f4975,plain,(
  $false | ($spl12 | $spl119 | $spl190)),
  inference(unit_resulting_resolution,[],[f845,f485,f1970,f396])).
fof(f4973,plain,(
  $false | ($spl104 | $spl159 | $spl190)),
  inference(unit_resulting_resolution,[],[f1379,f792,f1970,f396])).
fof(f4963,plain,(
  $false | ($spl181 | $spl190)),
  inference(unit_resulting_resolution,[],[f1941,f1970,f393])).
fof(f4961,plain,(
  $false | ($spl181 | $spl190)),
  inference(unit_resulting_resolution,[],[f1941,f1970,f392])).
fof(f4959,plain,(
  $false | ($spl79 | $spl190)),
  inference(unit_resulting_resolution,[],[f766,f1970,f388])).
fof(f4956,plain,(
  $false | ($spl142 | $spl190)),
  inference(unit_resulting_resolution,[],[f263,f978,f268,f1970,f387])).
fof(f4940,plain,(
  $false | ($spl148 | $spl190)),
  inference(unit_resulting_resolution,[],[f1205,f268,f262,f1970,f387])).
fof(f4935,plain,(
  $false | ($spl148 | $spl190)),
  inference(unit_resulting_resolution,[],[f1205,f268,f263,f1970,f387])).
fof(f4928,plain,(
  $false | ($spl85 | $spl190)),
  inference(unit_resulting_resolution,[],[f772,f1970,f384])).
fof(f4911,plain,(
  $false | ($spl146 | $spl190)),
  inference(unit_resulting_resolution,[],[f1203,f268,f262,f1970,f383])).
fof(f4906,plain,(
  $false | ($spl146 | $spl190)),
  inference(unit_resulting_resolution,[],[f1203,f268,f263,f1970,f383])).
fof(f4896,plain,(
  $false | ($spl2 | $spl89 | $spl190)),
  inference(unit_resulting_resolution,[],[f584,f582,f776,f575,f1970,f380])).
fof(f4880,plain,(
  $false | ($spl108 | $spl190)),
  inference(unit_resulting_resolution,[],[f822,f268,f262,f1970,f379])).
fof(f4876,plain,(
  $false | ($spl108 | $spl190)),
  inference(unit_resulting_resolution,[],[f822,f268,f263,f1970,f379])).
fof(f4851,plain,(
  $false | ($spl190 | $spl323)),
  inference(unit_resulting_resolution,[],[f2541,f1970,f350])).
fof(f4849,plain,(
  $false | ($spl190 | $spl323)),
  inference(unit_resulting_resolution,[],[f2541,f1970,f345])).
fof(f4847,plain,(
  $false | ($spl190 | $spl323)),
  inference(unit_resulting_resolution,[],[f2541,f1970,f344])).
fof(f4845,plain,(
  $false | ($spl29 | $spl190)),
  inference(unit_resulting_resolution,[],[f263,f518,f1970,f342])).
fof(f4840,plain,(
  $false | ($spl29 | $spl190)),
  inference(unit_resulting_resolution,[],[f263,f518,f1970,f342])).
fof(f5431,plain,(
  $false | $spl190),
  inference(trivial_inequality_removal,[],[f4834])).
fof(f4834,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f268,f1970,f338])).
fof(f338,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | ~p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f209])).
fof(f5430,plain,(
  $false | $spl190),
  inference(trivial_inequality_removal,[],[f4832])).
fof(f4832,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f268,f1970,f337])).
fof(f337,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | ~p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f209])).
fof(f4831,plain,(
  $false | ($spl29 | $spl190)),
  inference(unit_resulting_resolution,[],[f263,f518,f1970,f334])).
fof(f4826,plain,(
  $false | ($spl29 | $spl190)),
  inference(unit_resulting_resolution,[],[f263,f518,f1970,f334])).
fof(f4821,plain,(
  $false | $spl190),
  inference(unit_resulting_resolution,[],[f1970,f1970,f311])).
fof(f5429,plain,(
  $false | $spl190),
  inference(trivial_inequality_removal,[],[f4815])).
fof(f4815,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f1970,f281])).
fof(f281,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f173])).
fof(f173,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f63])).
fof(f63,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f62])).
fof(f62,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X14] : (s(bool,'F_0') = s(bool,X14) <=> ~p(s(bool,X14)))),
  file('thm.pred_set.inj_image_countable_IFF','thm.bool.EQ_CLAUSES|split|3')).
fof(f5428,plain,(
  $false | $spl190),
  inference(trivial_inequality_removal,[],[f4814])).
fof(f4814,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl190),
  inference(unit_resulting_resolution,[],[f1970,f279])).
fof(f279,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f172])).
fof(f4813,plain,(
  $false | ($spl29 | $spl190)),
  inference(unit_resulting_resolution,[],[f518,f1970,f278])).
fof(f4811,plain,(
  $false | ($spl29 | $spl190)),
  inference(unit_resulting_resolution,[],[f518,f1970,f276])).
fof(f4809,plain,(
  $false | ($spl190 | $spl323)),
  inference(unit_resulting_resolution,[],[f2541,f1970,f351])).
fof(f4807,plain,(
  $false | ($spl190 | $spl323)),
  inference(unit_resulting_resolution,[],[f2541,f1970,f346])).
fof(f4805,plain,(
  $false | $spl190),
  inference(unit_resulting_resolution,[],[f263,f1970,f311])).
fof(f4803,plain,(
  $false | $spl190),
  inference(unit_resulting_resolution,[],[f268,f263,f1970,f303])).
fof(f4800,plain,(
  $false | $spl190),
  inference(unit_resulting_resolution,[],[f268,f1970,f291])).
fof(f4798,plain,(
  $false | $spl190),
  inference(unit_resulting_resolution,[],[f268,f1970,f287])).
fof(f4796,plain,(
  $false | $spl190),
  inference(unit_resulting_resolution,[],[f263,f1970,f267])).
fof(f4794,plain,(
  $false | $spl190),
  inference(unit_resulting_resolution,[],[f263,f1970,f265])).
fof(f4792,plain,(
  $false | $spl190),
  inference(unit_resulting_resolution,[],[f268,f1970])).
fof(f5387,plain,(
  $false | $spl190),
  inference(resolution,[],[f1970,f268])).
fof(f2342,plain,(
  ( ! [X85,X84] : (sP10('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X84,X85)) ) | $spl276),
  inference(cnf_transformation,[],[f2342_D])).
fof(f2342_D,plain,(
  ( ! [X85,X84] : (sP10('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X84,X85)) ) <=> ~$spl276),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl276])])).
fof(f1970,plain,(
  p(s(bool,'F_0')) | $spl190),
  inference(cnf_transformation,[],[f1970_D])).
fof(f1970_D,plain,(
  p(s(bool,'F_0')) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f3062,plain,(
  ( ! [X83,X84] : (p(s(bool,'F_0')) | sP10('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X83,X84)) ) | $spl14),
  inference(superposition,[],[f443,f486])).
fof(f4784,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl614),
  inference(cnf_transformation,[],[f4784_D])).
fof(f4784_D,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) <=> ~$spl614),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl614])])).
fof(f4787,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl188),
  inference(global_subsumption,[],[f261,f260,f1966,f347,f343,f4532,f4517,f4514,f4513,f4512,f4504,f4503,f4502,f4501,f4499,f4498,f4497,f4496,f4495,f4492,f4490,f4488,f4487,f4483,f4482,f4481,f4477,f4476,f4475,f4453,f4449,f4448,f4447,f4446,f4442,f4440,f4439,f4435,f4433,f4432,f4427,f4414,f4413,f4412,f4411,f4410,f4409,f4408,f4407,f4406,f4405,f4404,f4403,f4402,f4401,f4400,f4399,f4398,f4397,f4396,f4395,f4394,f4393,f4392,f4391,f4390,f4389,f4388,f4387,f4386,f4385,f4384,f4383,f4382,f4381,f4380,f4379,f4375,f4374,f4373,f4369,f4368,f4367,f4339,f4335,f4334,f4333,f4332,f4331,f4327,f4326,f4325,f4321,f4320,f4319,f4318,f4314,f4313,f4312,f4308,f4307,f4306,f4301,f4300,f4299,f4298,f4297,f4296,f4295,f4294,f4293,f4292,f4291,f4290,f4289,f4288,f4287,f4286,f4285,f4284,f4283,f4282,f4281,f4280,f4279,f4277,f4275,f4274,f4272,f4271,f4270,f4269,f4267,f4266,f4264,f4263,f4257,f4256,f4255,f4254,f4571,f4572,f4573,f4578,f4583,f4584,f4585,f4588,f4595,f4598,f4599,f4601,f4602,f4603,f4604])).
fof(f4604,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl188),
  inference(superposition,[],[f1966,f269])).
fof(f4603,plain,(
  ( ! [X43,X42] : (~p(s(bool,X42)) | sP11('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),X42,X43)) ) | $spl188),
  inference(resolution,[],[f1966,f446])).
fof(f4602,plain,(
  ( ! [X41,X40] : (p(s(bool,X40)) | sP11('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),X41,X40)) ) | $spl188),
  inference(resolution,[],[f1966,f445])).
fof(f4601,plain,(
  ( ! [X39,X38] : (p(s(bool,X38)) | sP11(X39,X38,'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | p(s(bool,X39))) ) | $spl188),
  inference(resolution,[],[f1966,f444])).
fof(f4599,plain,(
  ( ! [X35,X34] : (~p(s(bool,X34)) | sP9('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),X34,X35)) ) | $spl188),
  inference(resolution,[],[f1966,f429])).
fof(f4598,plain,(
  ( ! [X33,X32] : (~p(s(bool,X32)) | sP9('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),X33,X32)) ) | $spl188),
  inference(resolution,[],[f1966,f428])).
fof(f4595,plain,(
  ( ! [X26,X27] : (~p(s(bool,X26)) | ~p(s(bool,X27)) | sP7('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),X26,X27)) ) | $spl188),
  inference(resolution,[],[f1966,f412])).
fof(f4588,plain,(
  ( ! [X13] : (~p(s(bool,X13)) | ~sP1('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),X13)) ) | $spl188),
  inference(resolution,[],[f1966,f366])).
fof(f4585,plain,(
  ( ! [X10] : (p(s(bool,X10)) | sP1(X10,'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')))) ) | $spl188),
  inference(resolution,[],[f1966,f363])).
fof(f4584,plain,(
  ( ! [X9] : (p(s(bool,X9)) | sP1('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),X9)) ) | $spl188),
  inference(resolution,[],[f1966,f360])).
fof(f4583,plain,(
  ( ! [X8] : (~p(s(bool,X8)) | ~sP1(X8,'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')))) ) | $spl188),
  inference(resolution,[],[f1966,f359])).
fof(f4578,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | s(bool,X3) = s(bool,'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')))) ) | $spl188),
  inference(resolution,[],[f1966,f334])).
fof(f4571,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.pred_set.INJ_3'(s(fun('type.num.num',X0),X1),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(X0,bool),X2)))) | p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,X0),'const.combin.o_2'(s(fun('type.num.num',X0),X1),s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))))),s(fun(sK13,bool),sK14),s(fun(X0,bool),X2))))) ) | $spl188),
  inference(resolution,[],[f1966,f450])).
fof(f4277,plain,(
  ( ! [X0] : (sP0(X0,'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')))) ) | $spl188),
  inference(unit_resulting_resolution,[],[f1966,f344])).
fof(f4279,plain,(
  ( ! [X0] : (sP0('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),X0)) ) | $spl188),
  inference(unit_resulting_resolution,[],[f1966,f350])).
fof(f4318,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')))) ) | $spl188),
  inference(unit_resulting_resolution,[],[f1966,f384])).
fof(f4331,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')))) ) | $spl188),
  inference(unit_resulting_resolution,[],[f1966,f388])).
fof(f4332,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')))) ) | $spl188),
  inference(unit_resulting_resolution,[],[f1966,f392])).
fof(f4333,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),X1)) ) | $spl188),
  inference(unit_resulting_resolution,[],[f1966,f393])).
fof(f4334,plain,(
  ( ! [X0,X1] : (sP2('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),X0,X1)) ) | $spl188),
  inference(unit_resulting_resolution,[],[f1966,f394])).
fof(f4335,plain,(
  ( ! [X0,X1] : (sP5('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'|dollar|exists_1'(s(fun(X0,bool),X1)),happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) ) | $spl188),
  inference(unit_resulting_resolution,[],[f317,f1966,f396])).
fof(f4339,plain,(
  ( ! [X0,X1] : (sP5('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))),'|dollar|exists_1'(s(fun(X0,bool),X1)))) ) | $spl188),
  inference(unit_resulting_resolution,[],[f317,f1966,f396])).
fof(f4379,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')))) ) | $spl188),
  inference(unit_resulting_resolution,[],[f1966,f407])).
fof(f4407,plain,(
  ( ! [X0] : (~sP7('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),X0,'F_0')) ) | $spl188),
  inference(unit_resulting_resolution,[],[f268,f1966,f413])).
fof(f4408,plain,(
  ( ! [X0] : (~sP7('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0',X0)) ) | $spl188),
  inference(unit_resulting_resolution,[],[f268,f1966,f414])).
fof(f4432,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')))) ) | $spl188),
  inference(unit_resulting_resolution,[],[f1966,f424])).
fof(f4433,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),X1)) ) | $spl188),
  inference(unit_resulting_resolution,[],[f1966,f425])).
fof(f4435,plain,(
  ( ! [X0] : (sP9('T_0',X0,'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')))) ) | $spl188),
  inference(unit_resulting_resolution,[],[f262,f1966,f428])).
fof(f4439,plain,(
  ( ! [X0] : (sP9('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),X0,'T_0')) ) | $spl188),
  inference(unit_resulting_resolution,[],[f262,f1966,f428])).
fof(f4440,plain,(
  ( ! [X0] : (sP9('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),X0,'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')))) ) | $spl188),
  inference(unit_resulting_resolution,[],[f1966,f1966,f428])).
fof(f4442,plain,(
  ( ! [X0] : (sP9('T_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),X0)) ) | $spl188),
  inference(unit_resulting_resolution,[],[f262,f1966,f429])).
fof(f4446,plain,(
  ( ! [X0] : (sP9('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0',X0)) ) | $spl188),
  inference(unit_resulting_resolution,[],[f262,f1966,f429])).
fof(f4447,plain,(
  ( ! [X0] : (sP9('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),X0)) ) | $spl188),
  inference(unit_resulting_resolution,[],[f1966,f1966,f429])).
fof(f4449,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')))) ) | $spl188),
  inference(unit_resulting_resolution,[],[f268,f1966,f431])).
fof(f4453,plain,(
  ( ! [X0] : (~sP9('F_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),X0)) ) | $spl188),
  inference(unit_resulting_resolution,[],[f268,f1966,f432])).
fof(f4487,plain,(
  ( ! [X0,X1] : (sP10(X0,'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),X1)) ) | $spl188),
  inference(unit_resulting_resolution,[],[f1966,f442])).
fof(f4490,plain,(
  ( ! [X0] : (sP11('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),X0,'F_0')) ) | $spl188),
  inference(unit_resulting_resolution,[],[f268,f1966,f445])).
fof(f4492,plain,(
  ( ! [X0] : (sP11('T_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),X0)) ) | $spl188),
  inference(unit_resulting_resolution,[],[f262,f1966,f446])).
fof(f4495,plain,(
  ( ! [X0] : (sP11('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0',X0)) ) | $spl188),
  inference(unit_resulting_resolution,[],[f262,f1966,f446])).
fof(f4496,plain,(
  ( ! [X0] : (sP11('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),X0)) ) | $spl188),
  inference(unit_resulting_resolution,[],[f1966,f1966,f446])).
fof(f4504,plain,(
  ( ! [X0] : (~sP11('F_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),X0)) ) | $spl188),
  inference(unit_resulting_resolution,[],[f268,f1966,f449])).
fof(f4573,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl188),
  inference(resolution,[],[f1966,f278])).
fof(f4572,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl188),
  inference(resolution,[],[f1966,f276])).
fof(f4786,plain,(
  p(s(bool,'F_0')) | ($spl14 | $spl188)),
  inference(forward_demodulation,[],[f4570,f486])).
fof(f4570,plain,(
  p(s(bool,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) | $spl188),
  inference(resolution,[],[f1966,f319])).
fof(f319,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,'type.num.num'),X2),s(fun(X0,bool),X1),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')))) | p(s(bool,'const.pred_set.countable_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f198])).
fof(f198,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.pred_set.countable_1'(s(fun(X0,bool),X1)))) | p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,'type.num.num'),sK16(X1,X0)),s(fun(X0,bool),X1),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) & (! [X2] : ~p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,'type.num.num'),X2),s(fun(X0,bool),X1),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')))) | p(s(bool,'const.pred_set.countable_1'(s(fun(X0,bool),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK16])],[f197])).
fof(f197,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.pred_set.countable_1'(s(fun(X0,bool),X1)))) | ? [X3] : p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,'type.num.num'),X3),s(fun(X0,bool),X1),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) & (! [X2] : ~p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,'type.num.num'),X2),s(fun(X0,bool),X1),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')))) | p(s(bool,'const.pred_set.countable_1'(s(fun(X0,bool),X1))))))),
  inference(rectify,[],[f196])).
fof(f196,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.pred_set.countable_1'(s(fun(X0,bool),X1)))) | ? [X2] : p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,'type.num.num'),X2),s(fun(X0,bool),X1),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))) & (! [X2] : ~p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,'type.num.num'),X2),s(fun(X0,bool),X1),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')))) | p(s(bool,'const.pred_set.countable_1'(s(fun(X0,bool),X1))))))),
  inference(nnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0,X1] : (p(s(bool,'const.pred_set.countable_1'(s(fun(X0,bool),X1)))) <=> ? [X2] : p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,'type.num.num'),X2),s(fun(X0,bool),X1),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')))))),
  inference(flattening,[],[f81])).
fof(f81,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.pred_set.countable_1'(s(fun(X0,bool),X1)))) <=> ? [X2] : p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,'type.num.num'),X2),s(fun(X0,bool),X1),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')))))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X7] : ! [X8] : (p(s(bool,'const.pred_set.countable_1'(s(fun(X7,bool),X8)))) <=> ? [X9] : p(s(bool,'const.pred_set.INJ_3'(s(fun(X7,'type.num.num'),X9),s(fun(X7,bool),X8),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')))))),
  file('thm.pred_set.inj_image_countable_IFF','thm.pred_set.countable_def')).
fof(f4785,plain,(
  p(s(bool,'F_0')) | ($spl14 | $spl188)),
  inference(forward_demodulation,[],[f4253,f486])).
fof(f4253,plain,(
  p(s(bool,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) | $spl188),
  inference(unit_resulting_resolution,[],[f1966,f319])).
fof(f4254,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl188),
  inference(unit_resulting_resolution,[],[f1966,f276])).
fof(f4255,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl188),
  inference(unit_resulting_resolution,[],[f1966,f278])).
fof(f4636,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl513),
  inference(cnf_transformation,[],[f4636_D])).
fof(f4636_D,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) <=> ~$spl513),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl513])])).
fof(f4256,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl188),
  inference(unit_resulting_resolution,[],[f1966,f279])).
fof(f4257,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl188),
  inference(unit_resulting_resolution,[],[f1966,f281])).
fof(f4263,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl188),
  inference(unit_resulting_resolution,[],[f263,f1966,f334])).
fof(f4264,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl188),
  inference(unit_resulting_resolution,[],[f262,f1966,f334])).
fof(f4266,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl188),
  inference(unit_resulting_resolution,[],[f263,f1966,f334])).
fof(f4267,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl188),
  inference(unit_resulting_resolution,[],[f262,f1966,f334])).
fof(f4269,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl188),
  inference(unit_resulting_resolution,[],[f268,f1966,f337])).
fof(f4270,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl188),
  inference(unit_resulting_resolution,[],[f268,f1966,f338])).
fof(f4271,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl188),
  inference(unit_resulting_resolution,[],[f263,f1966,f342])).
fof(f4272,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl188),
  inference(unit_resulting_resolution,[],[f262,f1966,f342])).
fof(f4274,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl188),
  inference(unit_resulting_resolution,[],[f263,f1966,f342])).
fof(f4275,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl188),
  inference(unit_resulting_resolution,[],[f262,f1966,f342])).
fof(f4776,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0') | $spl607),
  inference(cnf_transformation,[],[f4776_D])).
fof(f4776_D,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0') <=> ~$spl607),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl607])])).
fof(f4280,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0') | $spl188),
  inference(unit_resulting_resolution,[],[f263,f1966,f359])).
fof(f4281,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0') | $spl188),
  inference(unit_resulting_resolution,[],[f262,f1966,f359])).
fof(f4774,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl605),
  inference(cnf_transformation,[],[f4774_D])).
fof(f4774_D,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) <=> ~$spl605),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl605])])).
fof(f4282,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl188),
  inference(unit_resulting_resolution,[],[f1966,f1966,f359])).
fof(f4778,plain,(
  ~sP1('T_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl609),
  inference(cnf_transformation,[],[f4778_D])).
fof(f4778_D,plain,(
  ~sP1('T_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) <=> ~$spl609),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl609])])).
fof(f4283,plain,(
  ~sP1('T_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl188),
  inference(unit_resulting_resolution,[],[f263,f1966,f359])).
fof(f4284,plain,(
  ~sP1('T_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl188),
  inference(unit_resulting_resolution,[],[f262,f1966,f359])).
fof(f4285,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl188),
  inference(unit_resulting_resolution,[],[f1966,f1966,f359])).
fof(f4780,plain,(
  sP1('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0') | $spl610),
  inference(cnf_transformation,[],[f4780_D])).
fof(f4780_D,plain,(
  sP1('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0') <=> ~$spl610),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl610])])).
fof(f4286,plain,(
  sP1('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0') | $spl188),
  inference(unit_resulting_resolution,[],[f268,f1966,f360])).
fof(f4782,plain,(
  sP1('F_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl612),
  inference(cnf_transformation,[],[f4782_D])).
fof(f4782_D,plain,(
  sP1('F_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) <=> ~$spl612),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl612])])).
fof(f4287,plain,(
  sP1('F_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl188),
  inference(unit_resulting_resolution,[],[f268,f1966,f363])).
fof(f4288,plain,(
  sP1('F_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl188),
  inference(unit_resulting_resolution,[],[f268,f1966,f364])).
fof(f4289,plain,(
  sP1('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0') | $spl188),
  inference(unit_resulting_resolution,[],[f268,f1966,f365])).
fof(f4290,plain,(
  ~sP1('T_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl188),
  inference(unit_resulting_resolution,[],[f263,f1966,f366])).
fof(f4291,plain,(
  ~sP1('T_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl188),
  inference(unit_resulting_resolution,[],[f262,f1966,f366])).
fof(f4292,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl188),
  inference(unit_resulting_resolution,[],[f1966,f1966,f366])).
fof(f4293,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0') | $spl188),
  inference(unit_resulting_resolution,[],[f263,f1966,f366])).
fof(f4294,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0') | $spl188),
  inference(unit_resulting_resolution,[],[f262,f1966,f366])).
fof(f4295,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl188),
  inference(unit_resulting_resolution,[],[f1966,f1966,f366])).
fof(f4724,plain,(
  ~sP5('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0','F_0') | $spl567),
  inference(cnf_transformation,[],[f4724_D])).
fof(f4724_D,plain,(
  ~sP5('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0','F_0') <=> ~$spl567),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl567])])).
fof(f4296,plain,(
  ~sP5('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0','F_0') | $spl188),
  inference(unit_resulting_resolution,[],[f263,f268,f1966,f379])).
fof(f4297,plain,(
  ~sP5('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0','F_0') | $spl188),
  inference(unit_resulting_resolution,[],[f262,f268,f1966,f379])).
fof(f4770,plain,(
  ~sP5('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0') | $spl601),
  inference(cnf_transformation,[],[f4770_D])).
fof(f4770_D,plain,(
  ~sP5('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0') <=> ~$spl601),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl601])])).
fof(f4298,plain,(
  ~sP5('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0') | $spl188),
  inference(unit_resulting_resolution,[],[f1966,f268,f1966,f379])).
fof(f4772,plain,(
  ~sP5('T_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0') | $spl603),
  inference(cnf_transformation,[],[f4772_D])).
fof(f4772_D,plain,(
  ~sP5('T_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0') <=> ~$spl603),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl603])])).
fof(f4299,plain,(
  ~sP5('T_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0') | $spl188),
  inference(unit_resulting_resolution,[],[f263,f268,f1966,f379])).
fof(f4300,plain,(
  ~sP5('T_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0') | $spl188),
  inference(unit_resulting_resolution,[],[f262,f268,f1966,f379])).
fof(f4301,plain,(
  ~sP5('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0') | $spl188),
  inference(unit_resulting_resolution,[],[f1966,f268,f1966,f379])).
fof(f4768,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl598),
  inference(cnf_transformation,[],[f4768_D])).
fof(f4768_D,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) <=> ~$spl598),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl598])])).
fof(f4302,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | ($spl2 | $spl188)),
  inference(unit_resulting_resolution,[],[f584,f582,f575,f1966,f380])).
fof(f4766,plain,(
  sP5('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl596),
  inference(cnf_transformation,[],[f4766_D])).
fof(f4766_D,plain,(
  sP5('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) <=> ~$spl596),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl596])])).
fof(f4303,plain,(
  sP5('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | ($spl2 | $spl188 | $spl292)),
  inference(unit_resulting_resolution,[],[f584,f2509,f575,f1966,f380])).
fof(f4764,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl594),
  inference(cnf_transformation,[],[f4764_D])).
fof(f4764_D,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) <=> ~$spl594),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl594])])).
fof(f4304,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | ($spl2 | $spl188 | $spl290 | $spl296)),
  inference(unit_resulting_resolution,[],[f2511,f582,f2508,f1966,f380])).
fof(f4762,plain,(
  sP5('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl592),
  inference(cnf_transformation,[],[f4762_D])).
fof(f4762_D,plain,(
  sP5('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) <=> ~$spl592),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl592])])).
fof(f4305,plain,(
  sP5('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | ($spl188 | $spl290 | $spl292 | $spl296)),
  inference(unit_resulting_resolution,[],[f2511,f2509,f2508,f1966,f380])).
fof(f4760,plain,(
  ~sP4('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0','F_0') | $spl591),
  inference(cnf_transformation,[],[f4760_D])).
fof(f4760_D,plain,(
  ~sP4('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0','F_0') <=> ~$spl591),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl591])])).
fof(f4306,plain,(
  ~sP4('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0','F_0') | $spl188),
  inference(unit_resulting_resolution,[],[f263,f268,f1966,f383])).
fof(f4307,plain,(
  ~sP4('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0','F_0') | $spl188),
  inference(unit_resulting_resolution,[],[f262,f268,f1966,f383])).
fof(f4753,plain,(
  ~sP4('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0') | $spl587),
  inference(cnf_transformation,[],[f4753_D])).
fof(f4753_D,plain,(
  ~sP4('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0') <=> ~$spl587),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl587])])).
fof(f4308,plain,(
  ~sP4('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0') | $spl188),
  inference(unit_resulting_resolution,[],[f1966,f268,f1966,f383])).
fof(f4755,plain,(
  ~sP4('T_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0') | $spl589),
  inference(cnf_transformation,[],[f4755_D])).
fof(f4755_D,plain,(
  ~sP4('T_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0') <=> ~$spl589),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl589])])).
fof(f4312,plain,(
  ~sP4('T_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0') | $spl188),
  inference(unit_resulting_resolution,[],[f263,f268,f1966,f383])).
fof(f4313,plain,(
  ~sP4('T_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0') | $spl188),
  inference(unit_resulting_resolution,[],[f262,f268,f1966,f383])).
fof(f4314,plain,(
  ~sP4('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0') | $spl188),
  inference(unit_resulting_resolution,[],[f1966,f268,f1966,f383])).
fof(f4748,plain,(
  ~sP3('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0','F_0') | $spl585),
  inference(cnf_transformation,[],[f4748_D])).
fof(f4748_D,plain,(
  ~sP3('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0','F_0') <=> ~$spl585),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl585])])).
fof(f4319,plain,(
  ~sP3('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0','F_0') | $spl188),
  inference(unit_resulting_resolution,[],[f263,f268,f1966,f387])).
fof(f4320,plain,(
  ~sP3('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0','F_0') | $spl188),
  inference(unit_resulting_resolution,[],[f262,f268,f1966,f387])).
fof(f4741,plain,(
  ~sP3('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0') | $spl581),
  inference(cnf_transformation,[],[f4741_D])).
fof(f4741_D,plain,(
  ~sP3('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0') <=> ~$spl581),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl581])])).
fof(f4321,plain,(
  ~sP3('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0') | $spl188),
  inference(unit_resulting_resolution,[],[f1966,f268,f1966,f387])).
fof(f4743,plain,(
  ~sP3('T_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0') | $spl583),
  inference(cnf_transformation,[],[f4743_D])).
fof(f4743_D,plain,(
  ~sP3('T_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0') <=> ~$spl583),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl583])])).
fof(f4325,plain,(
  ~sP3('T_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0') | $spl188),
  inference(unit_resulting_resolution,[],[f263,f268,f1966,f387])).
fof(f4326,plain,(
  ~sP3('T_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0') | $spl188),
  inference(unit_resulting_resolution,[],[f262,f268,f1966,f387])).
fof(f4327,plain,(
  ~sP3('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0') | $spl188),
  inference(unit_resulting_resolution,[],[f1966,f268,f1966,f387])).
fof(f4736,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl578),
  inference(cnf_transformation,[],[f4736_D])).
fof(f4736_D,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) <=> ~$spl578),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl578])])).
fof(f4336,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | ($spl104 | $spl188)),
  inference(unit_resulting_resolution,[],[f792,f1966,f396])).
fof(f4734,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl576),
  inference(cnf_transformation,[],[f4734_D])).
fof(f4734_D,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) <=> ~$spl576),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl576])])).
fof(f4337,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | ($spl12 | $spl188)),
  inference(unit_resulting_resolution,[],[f485,f1966,f396])).
fof(f4732,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | $spl574),
  inference(cnf_transformation,[],[f4732_D])).
fof(f4732_D,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) <=> ~$spl574),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl574])])).
fof(f4338,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | ($spl14 | $spl188)),
  inference(unit_resulting_resolution,[],[f486,f1966,f396])).
fof(f4730,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl572),
  inference(cnf_transformation,[],[f4730_D])).
fof(f4730_D,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') <=> ~$spl572),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl572])])).
fof(f4340,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | ($spl104 | $spl188)),
  inference(unit_resulting_resolution,[],[f792,f1966,f396])).
fof(f4728,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl570),
  inference(cnf_transformation,[],[f4728_D])).
fof(f4728_D,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') <=> ~$spl570),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl570])])).
fof(f4341,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | ($spl12 | $spl188)),
  inference(unit_resulting_resolution,[],[f485,f1966,f396])).
fof(f4726,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0') | $spl568),
  inference(cnf_transformation,[],[f4726_D])).
fof(f4726_D,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0') <=> ~$spl568),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl568])])).
fof(f4342,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0') | ($spl14 | $spl188)),
  inference(unit_resulting_resolution,[],[f486,f1966,f396])).
fof(f4343,plain,(
  ~sP5('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0','F_0') | ($spl29 | $spl188)),
  inference(unit_resulting_resolution,[],[f518,f1966,f397])).
fof(f4722,plain,(
  ~sP5('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | $spl565),
  inference(cnf_transformation,[],[f4722_D])).
fof(f4722_D,plain,(
  ~sP5('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) <=> ~$spl565),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl565])])).
fof(f4344,plain,(
  ~sP5('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | ($spl27 | $spl188)),
  inference(unit_resulting_resolution,[],[f513,f1966,f397])).
fof(f4720,plain,(
  ~sP5('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0','T_0') | $spl563),
  inference(cnf_transformation,[],[f4720_D])).
fof(f4720_D,plain,(
  ~sP5('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0','T_0') <=> ~$spl563),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl563])])).
fof(f4345,plain,(
  ~sP5('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0','T_0') | ($spl29 | $spl188)),
  inference(unit_resulting_resolution,[],[f518,f1966,f397])).
fof(f4718,plain,(
  ~sP5('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl561),
  inference(cnf_transformation,[],[f4718_D])).
fof(f4718_D,plain,(
  ~sP5('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl561),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl561])])).
fof(f4346,plain,(
  ~sP5('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'T_0') | ($spl27 | $spl188)),
  inference(unit_resulting_resolution,[],[f513,f1966,f397])).
fof(f4628,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0') | $spl504),
  inference(cnf_transformation,[],[f4628_D])).
fof(f4628_D,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0') <=> ~$spl504),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl504])])).
fof(f4349,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0') | ($spl2 | $spl188)),
  inference(unit_resulting_resolution,[],[f263,f596,f1966,f402])).
fof(f4350,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0') | ($spl2 | $spl188)),
  inference(unit_resulting_resolution,[],[f262,f596,f1966,f402])).
fof(f4626,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl502),
  inference(cnf_transformation,[],[f4626_D])).
fof(f4626_D,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) <=> ~$spl502),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl502])])).
fof(f4351,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | ($spl2 | $spl188)),
  inference(unit_resulting_resolution,[],[f1966,f596,f1966,f402])).
fof(f4616,plain,(
  sP7('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0') | $spl492),
  inference(cnf_transformation,[],[f4616_D])).
fof(f4616_D,plain,(
  sP7('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0') <=> ~$spl492),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl492])])).
fof(f4353,plain,(
  sP7('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0') | ($spl188 | $spl300)),
  inference(unit_resulting_resolution,[],[f263,f2513,f1966,f402])).
fof(f4354,plain,(
  sP7('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0') | ($spl188 | $spl300)),
  inference(unit_resulting_resolution,[],[f262,f2513,f1966,f402])).
fof(f4614,plain,(
  sP7('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl490),
  inference(cnf_transformation,[],[f4614_D])).
fof(f4614_D,plain,(
  sP7('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) <=> ~$spl490),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl490])])).
fof(f4355,plain,(
  sP7('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | ($spl188 | $spl300)),
  inference(unit_resulting_resolution,[],[f1966,f2513,f1966,f402])).
fof(f4630,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl506),
  inference(cnf_transformation,[],[f4630_D])).
fof(f4630_D,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) <=> ~$spl506),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl506])])).
fof(f4356,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | ($spl2 | $spl188)),
  inference(unit_resulting_resolution,[],[f263,f596,f1966,f402])).
fof(f4357,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | ($spl2 | $spl188)),
  inference(unit_resulting_resolution,[],[f262,f596,f1966,f402])).
fof(f4358,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | ($spl2 | $spl188)),
  inference(unit_resulting_resolution,[],[f1966,f596,f1966,f402])).
fof(f4618,plain,(
  sP7('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl494),
  inference(cnf_transformation,[],[f4618_D])).
fof(f4618_D,plain,(
  sP7('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) <=> ~$spl494),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl494])])).
fof(f4360,plain,(
  sP7('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | ($spl188 | $spl300)),
  inference(unit_resulting_resolution,[],[f263,f2513,f1966,f402])).
fof(f4361,plain,(
  sP7('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | ($spl188 | $spl300)),
  inference(unit_resulting_resolution,[],[f262,f2513,f1966,f402])).
fof(f4362,plain,(
  sP7('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | ($spl188 | $spl300)),
  inference(unit_resulting_resolution,[],[f1966,f2513,f1966,f402])).
fof(f4714,plain,(
  ~sP6('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0','F_0') | $spl559),
  inference(cnf_transformation,[],[f4714_D])).
fof(f4714_D,plain,(
  ~sP6('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0','F_0') <=> ~$spl559),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl559])])).
fof(f4367,plain,(
  ~sP6('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0','F_0') | $spl188),
  inference(unit_resulting_resolution,[],[f263,f268,f1966,f406])).
fof(f4368,plain,(
  ~sP6('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0','F_0') | $spl188),
  inference(unit_resulting_resolution,[],[f262,f268,f1966,f406])).
fof(f4707,plain,(
  ~sP6('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0') | $spl555),
  inference(cnf_transformation,[],[f4707_D])).
fof(f4707_D,plain,(
  ~sP6('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0') <=> ~$spl555),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl555])])).
fof(f4369,plain,(
  ~sP6('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0') | $spl188),
  inference(unit_resulting_resolution,[],[f1966,f268,f1966,f406])).
fof(f4709,plain,(
  ~sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0') | $spl557),
  inference(cnf_transformation,[],[f4709_D])).
fof(f4709_D,plain,(
  ~sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0') <=> ~$spl557),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl557])])).
fof(f4373,plain,(
  ~sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0') | $spl188),
  inference(unit_resulting_resolution,[],[f263,f268,f1966,f406])).
fof(f4374,plain,(
  ~sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0') | $spl188),
  inference(unit_resulting_resolution,[],[f262,f268,f1966,f406])).
fof(f4375,plain,(
  ~sP6('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0') | $spl188),
  inference(unit_resulting_resolution,[],[f1966,f268,f1966,f406])).
fof(f4702,plain,(
  sP7('T_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0') | $spl552),
  inference(cnf_transformation,[],[f4702_D])).
fof(f4702_D,plain,(
  sP7('T_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0') <=> ~$spl552),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl552])])).
fof(f4380,plain,(
  sP7('T_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0') | $spl188),
  inference(unit_resulting_resolution,[],[f263,f263,f1966,f412])).
fof(f4381,plain,(
  sP7('T_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0') | $spl188),
  inference(unit_resulting_resolution,[],[f262,f263,f1966,f412])).
fof(f4694,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0') | $spl544),
  inference(cnf_transformation,[],[f4694_D])).
fof(f4694_D,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0') <=> ~$spl544),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl544])])).
fof(f4382,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0') | $spl188),
  inference(unit_resulting_resolution,[],[f1966,f263,f1966,f412])).
fof(f4383,plain,(
  sP7('T_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0') | $spl188),
  inference(unit_resulting_resolution,[],[f263,f262,f1966,f412])).
fof(f4384,plain,(
  sP7('T_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0') | $spl188),
  inference(unit_resulting_resolution,[],[f262,f262,f1966,f412])).
fof(f4385,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0') | $spl188),
  inference(unit_resulting_resolution,[],[f1966,f262,f1966,f412])).
fof(f4698,plain,(
  sP7('T_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl548),
  inference(cnf_transformation,[],[f4698_D])).
fof(f4698_D,plain,(
  sP7('T_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) <=> ~$spl548),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl548])])).
fof(f4386,plain,(
  sP7('T_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl188),
  inference(unit_resulting_resolution,[],[f263,f1966,f1966,f412])).
fof(f4387,plain,(
  sP7('T_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl188),
  inference(unit_resulting_resolution,[],[f262,f1966,f1966,f412])).
fof(f4690,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl540),
  inference(cnf_transformation,[],[f4690_D])).
fof(f4690_D,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) <=> ~$spl540),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl540])])).
fof(f4388,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl188),
  inference(unit_resulting_resolution,[],[f1966,f1966,f1966,f412])).
fof(f4700,plain,(
  sP7('T_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl550),
  inference(cnf_transformation,[],[f4700_D])).
fof(f4700_D,plain,(
  sP7('T_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) <=> ~$spl550),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl550])])).
fof(f4389,plain,(
  sP7('T_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl188),
  inference(unit_resulting_resolution,[],[f263,f263,f1966,f412])).
fof(f4390,plain,(
  sP7('T_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl188),
  inference(unit_resulting_resolution,[],[f262,f263,f1966,f412])).
fof(f4391,plain,(
  sP7('T_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl188),
  inference(unit_resulting_resolution,[],[f1966,f263,f1966,f412])).
fof(f4392,plain,(
  sP7('T_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl188),
  inference(unit_resulting_resolution,[],[f263,f262,f1966,f412])).
fof(f4393,plain,(
  sP7('T_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl188),
  inference(unit_resulting_resolution,[],[f262,f262,f1966,f412])).
fof(f4394,plain,(
  sP7('T_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl188),
  inference(unit_resulting_resolution,[],[f1966,f262,f1966,f412])).
fof(f4692,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl542),
  inference(cnf_transformation,[],[f4692_D])).
fof(f4692_D,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) <=> ~$spl542),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl542])])).
fof(f4395,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl188),
  inference(unit_resulting_resolution,[],[f263,f1966,f1966,f412])).
fof(f4396,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl188),
  inference(unit_resulting_resolution,[],[f262,f1966,f1966,f412])).
fof(f4397,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl188),
  inference(unit_resulting_resolution,[],[f1966,f1966,f1966,f412])).
fof(f4696,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0','T_0') | $spl546),
  inference(cnf_transformation,[],[f4696_D])).
fof(f4696_D,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0','T_0') <=> ~$spl546),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl546])])).
fof(f4398,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0','T_0') | $spl188),
  inference(unit_resulting_resolution,[],[f263,f263,f1966,f412])).
fof(f4399,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0','T_0') | $spl188),
  inference(unit_resulting_resolution,[],[f262,f263,f1966,f412])).
fof(f4400,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0') | $spl188),
  inference(unit_resulting_resolution,[],[f1966,f263,f1966,f412])).
fof(f4401,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0','T_0') | $spl188),
  inference(unit_resulting_resolution,[],[f263,f262,f1966,f412])).
fof(f4402,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0','T_0') | $spl188),
  inference(unit_resulting_resolution,[],[f262,f262,f1966,f412])).
fof(f4403,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0') | $spl188),
  inference(unit_resulting_resolution,[],[f1966,f262,f1966,f412])).
fof(f4404,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl188),
  inference(unit_resulting_resolution,[],[f263,f1966,f1966,f412])).
fof(f4405,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl188),
  inference(unit_resulting_resolution,[],[f262,f1966,f1966,f412])).
fof(f4406,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl188),
  inference(unit_resulting_resolution,[],[f1966,f1966,f1966,f412])).
fof(f4688,plain,(
  ~sP7('F_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0') | $spl539),
  inference(cnf_transformation,[],[f4688_D])).
fof(f4688_D,plain,(
  ~sP7('F_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0') <=> ~$spl539),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl539])])).
fof(f4409,plain,(
  ~sP7('F_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0') | $spl188),
  inference(unit_resulting_resolution,[],[f268,f263,f1966,f415])).
fof(f4410,plain,(
  ~sP7('F_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0') | $spl188),
  inference(unit_resulting_resolution,[],[f268,f262,f1966,f415])).
fof(f4684,plain,(
  ~sP7('F_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl535),
  inference(cnf_transformation,[],[f4684_D])).
fof(f4684_D,plain,(
  ~sP7('F_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) <=> ~$spl535),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl535])])).
fof(f4411,plain,(
  ~sP7('F_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl188),
  inference(unit_resulting_resolution,[],[f268,f1966,f1966,f415])).
fof(f4686,plain,(
  ~sP7('F_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl537),
  inference(cnf_transformation,[],[f4686_D])).
fof(f4686_D,plain,(
  ~sP7('F_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) <=> ~$spl537),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl537])])).
fof(f4412,plain,(
  ~sP7('F_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl188),
  inference(unit_resulting_resolution,[],[f268,f263,f1966,f415])).
fof(f4413,plain,(
  ~sP7('F_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl188),
  inference(unit_resulting_resolution,[],[f268,f262,f1966,f415])).
fof(f4414,plain,(
  ~sP7('F_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl188),
  inference(unit_resulting_resolution,[],[f268,f1966,f1966,f415])).
fof(f4678,plain,(
  ~sP8('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0','F_0') | $spl533),
  inference(cnf_transformation,[],[f4678_D])).
fof(f4678_D,plain,(
  ~sP8('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0','F_0') <=> ~$spl533),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl533])])).
fof(f4427,plain,(
  ~sP8('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0','F_0') | $spl188),
  inference(unit_resulting_resolution,[],[f268,f268,f1966,f423])).
fof(f4672,plain,(
  ~sP9('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0','F_0') | $spl531),
  inference(cnf_transformation,[],[f4672_D])).
fof(f4672_D,plain,(
  ~sP9('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0','F_0') <=> ~$spl531),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl531])])).
fof(f4448,plain,(
  ~sP9('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0','F_0') | $spl188),
  inference(unit_resulting_resolution,[],[f268,f268,f1966,f430])).
fof(f4665,plain,(
  sP11('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl528),
  inference(cnf_transformation,[],[f4665_D])).
fof(f4665_D,plain,(
  sP11('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) <=> ~$spl528),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl528])])).
fof(f4473,plain,(
  sP11('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | ($spl188 | $spl276)),
  inference(unit_resulting_resolution,[],[f268,f2342,f1966,f439])).
fof(f439,plain,(
  ( ! [X2,X0,X1] : (~sP10(X0,X1,X2) | sP11(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f252])).
fof(f4662,plain,(
  ~sP10('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0','T_0') | $spl527),
  inference(cnf_transformation,[],[f4662_D])).
fof(f4662_D,plain,(
  ~sP10('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0','T_0') <=> ~$spl527),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl527])])).
fof(f4475,plain,(
  ~sP10('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0','T_0') | $spl188),
  inference(unit_resulting_resolution,[],[f268,f263,f1966,f440])).
fof(f4476,plain,(
  ~sP10('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0','T_0') | $spl188),
  inference(unit_resulting_resolution,[],[f268,f262,f1966,f440])).
fof(f4655,plain,(
  ~sP10('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl523),
  inference(cnf_transformation,[],[f4655_D])).
fof(f4655_D,plain,(
  ~sP10('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) <=> ~$spl523),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl523])])).
fof(f4477,plain,(
  ~sP10('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl188),
  inference(unit_resulting_resolution,[],[f268,f1966,f1966,f440])).
fof(f4657,plain,(
  ~sP10('T_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl525),
  inference(cnf_transformation,[],[f4657_D])).
fof(f4657_D,plain,(
  ~sP10('T_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) <=> ~$spl525),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl525])])).
fof(f4481,plain,(
  ~sP10('T_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl188),
  inference(unit_resulting_resolution,[],[f263,f268,f1966,f440])).
fof(f4482,plain,(
  ~sP10('T_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl188),
  inference(unit_resulting_resolution,[],[f262,f268,f1966,f440])).
fof(f4483,plain,(
  ~sP10('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl188),
  inference(unit_resulting_resolution,[],[f1966,f268,f1966,f440])).
fof(f4650,plain,(
  sP11('F_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl520),
  inference(cnf_transformation,[],[f4650_D])).
fof(f4650_D,plain,(
  sP11('F_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) <=> ~$spl520),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl520])])).
fof(f4488,plain,(
  sP11('F_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl188),
  inference(unit_resulting_resolution,[],[f268,f268,f1966,f444])).
fof(f4647,plain,(
  ~sP11('T_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl519),
  inference(cnf_transformation,[],[f4647_D])).
fof(f4647_D,plain,(
  ~sP11('T_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) <=> ~$spl519),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl519])])).
fof(f4497,plain,(
  ~sP11('T_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl188),
  inference(unit_resulting_resolution,[],[f268,f263,f1966,f447])).
fof(f4498,plain,(
  ~sP11('T_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl188),
  inference(unit_resulting_resolution,[],[f268,f262,f1966,f447])).
fof(f4642,plain,(
  ~sP11('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl515),
  inference(cnf_transformation,[],[f4642_D])).
fof(f4642_D,plain,(
  ~sP11('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) <=> ~$spl515),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl515])])).
fof(f4499,plain,(
  ~sP11('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl188),
  inference(unit_resulting_resolution,[],[f268,f1966,f1966,f447])).
fof(f4644,plain,(
  ~sP11('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0','T_0') | $spl517),
  inference(cnf_transformation,[],[f4644_D])).
fof(f4644_D,plain,(
  ~sP11('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0','T_0') <=> ~$spl517),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl517])])).
fof(f4501,plain,(
  ~sP11('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0','T_0') | $spl188),
  inference(unit_resulting_resolution,[],[f268,f263,f1966,f447])).
fof(f4502,plain,(
  ~sP11('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0','T_0') | $spl188),
  inference(unit_resulting_resolution,[],[f268,f262,f1966,f447])).
fof(f4503,plain,(
  ~sP11('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'F_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl188),
  inference(unit_resulting_resolution,[],[f268,f1966,f1966,f447])).
fof(f4506,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | ($spl14 | $spl188)),
  inference(unit_resulting_resolution,[],[f268,f268,f486,f1966,f452])).
fof(f4507,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | ($spl104 | $spl188)),
  inference(unit_resulting_resolution,[],[f268,f263,f792,f1966,f454])).
fof(f4508,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | ($spl104 | $spl188)),
  inference(unit_resulting_resolution,[],[f268,f262,f792,f1966,f454])).
fof(f4509,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | ($spl12 | $spl188)),
  inference(unit_resulting_resolution,[],[f268,f263,f485,f1966,f454])).
fof(f4510,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | ($spl12 | $spl188)),
  inference(unit_resulting_resolution,[],[f268,f262,f485,f1966,f454])).
fof(f4511,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | ($spl14 | $spl188)),
  inference(unit_resulting_resolution,[],[f268,f268,f486,f1966,f456])).
fof(f4512,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl188),
  inference(unit_resulting_resolution,[],[f268,f268,f263,f1966,f457])).
fof(f4513,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl188),
  inference(unit_resulting_resolution,[],[f268,f268,f262,f1966,f457])).
fof(f4514,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl188),
  inference(unit_resulting_resolution,[],[f268,f268,f1966,f1966,f457])).
fof(f4517,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl188),
  inference(unit_resulting_resolution,[],[f268,f268,f1966,f1966,f457])).
fof(f4521,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | ($spl104 | $spl188)),
  inference(unit_resulting_resolution,[],[f268,f263,f792,f1966,f458])).
fof(f4522,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | ($spl104 | $spl188)),
  inference(unit_resulting_resolution,[],[f268,f262,f792,f1966,f458])).
fof(f4523,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | ($spl12 | $spl188)),
  inference(unit_resulting_resolution,[],[f268,f263,f485,f1966,f458])).
fof(f4524,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | ($spl12 | $spl188)),
  inference(unit_resulting_resolution,[],[f268,f262,f485,f1966,f458])).
fof(f4532,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl188),
  inference(unit_resulting_resolution,[],[f268,f268,f1966,f464])).
fof(f4634,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl511),
  inference(cnf_transformation,[],[f4634_D])).
fof(f4634_D,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) <=> ~$spl511),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl511])])).
fof(f4533,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | ($spl2 | $spl188)),
  inference(unit_resulting_resolution,[],[f1966,f695])).
fof(f4632,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl509),
  inference(cnf_transformation,[],[f4632_D])).
fof(f4632_D,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) <=> ~$spl509),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl509])])).
fof(f4534,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | ($spl2 | $spl188)),
  inference(unit_resulting_resolution,[],[f1966,f700])).
fof(f4535,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | ($spl2 | $spl188)),
  inference(unit_resulting_resolution,[],[f263,f1966,f707])).
fof(f4536,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | ($spl2 | $spl188)),
  inference(unit_resulting_resolution,[],[f262,f1966,f707])).
fof(f4537,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | ($spl2 | $spl188)),
  inference(unit_resulting_resolution,[],[f1966,f1966,f707])).
fof(f4539,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0') | ($spl2 | $spl188)),
  inference(unit_resulting_resolution,[],[f263,f1966,f707])).
fof(f4540,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0') | ($spl2 | $spl188)),
  inference(unit_resulting_resolution,[],[f262,f1966,f707])).
fof(f4541,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | ($spl2 | $spl188)),
  inference(unit_resulting_resolution,[],[f1966,f1966,f707])).
fof(f4624,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl500),
  inference(cnf_transformation,[],[f4624_D])).
fof(f4624_D,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) <=> ~$spl500),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl500])])).
fof(f4549,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | ($spl104 | $spl188)),
  inference(unit_resulting_resolution,[],[f792,f1966,f1528])).
fof(f4608,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl484),
  inference(cnf_transformation,[],[f4608_D])).
fof(f4608_D,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) <=> ~$spl484),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl484])])).
fof(f4550,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | ($spl12 | $spl104 | $spl188)),
  inference(unit_resulting_resolution,[],[f485,f1966,f1528])).
fof(f4551,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | ($spl104 | $spl188)),
  inference(unit_resulting_resolution,[],[f792,f1966,f1529])).
fof(f4610,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl486),
  inference(cnf_transformation,[],[f4610_D])).
fof(f4610_D,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) <=> ~$spl486),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl486])])).
fof(f4552,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | ($spl12 | $spl104 | $spl188)),
  inference(unit_resulting_resolution,[],[f485,f1966,f1529])).
fof(f4622,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl499),
  inference(cnf_transformation,[],[f4622_D])).
fof(f4622_D,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) <=> ~$spl499),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl499])])).
fof(f4553,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | ($spl188 | $spl282)),
  inference(unit_resulting_resolution,[],[f1966,f2504])).
fof(f4620,plain,(
  ~sP1('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | $spl497),
  inference(cnf_transformation,[],[f4620_D])).
fof(f4620_D,plain,(
  ~sP1('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) <=> ~$spl497),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl497])])).
fof(f4554,plain,(
  ~sP1('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | ($spl188 | $spl288)),
  inference(unit_resulting_resolution,[],[f1966,f2507])).
fof(f4555,plain,(
  sP7('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | ($spl188 | $spl302)),
  inference(unit_resulting_resolution,[],[f263,f1966,f2514])).
fof(f4556,plain,(
  sP7('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | ($spl188 | $spl302)),
  inference(unit_resulting_resolution,[],[f262,f1966,f2514])).
fof(f4557,plain,(
  sP7('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | ($spl188 | $spl302)),
  inference(unit_resulting_resolution,[],[f1966,f1966,f2514])).
fof(f4558,plain,(
  sP7('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0') | ($spl188 | $spl302)),
  inference(unit_resulting_resolution,[],[f263,f1966,f2514])).
fof(f4559,plain,(
  sP7('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'T_0') | ($spl188 | $spl302)),
  inference(unit_resulting_resolution,[],[f262,f1966,f2514])).
fof(f4560,plain,(
  sP7('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))) | ($spl188 | $spl302)),
  inference(unit_resulting_resolution,[],[f1966,f1966,f2514])).
fof(f4612,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | $spl488),
  inference(cnf_transformation,[],[f4612_D])).
fof(f4612_D,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) <=> ~$spl488),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl488])])).
fof(f4564,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | ($spl14 | $spl188)),
  inference(unit_resulting_resolution,[],[f486,f1966,f3040])).
fof(f4565,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | ($spl14 | $spl188)),
  inference(unit_resulting_resolution,[],[f486,f1966,f3041])).
fof(f4566,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | ($spl12 | $spl104 | $spl188)),
  inference(unit_resulting_resolution,[],[f792,f1966,f3190])).
fof(f4606,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl482),
  inference(cnf_transformation,[],[f4606_D])).
fof(f4606_D,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) <=> ~$spl482),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl482])])).
fof(f4567,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | ($spl12 | $spl188)),
  inference(unit_resulting_resolution,[],[f485,f1966,f3190])).
fof(f4568,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | ($spl12 | $spl104 | $spl188)),
  inference(unit_resulting_resolution,[],[f792,f1966,f3191])).
fof(f4569,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | ($spl12 | $spl188)),
  inference(unit_resulting_resolution,[],[f485,f1966,f3191])).
fof(f4216,plain,(
  sP5('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl480),
  inference(cnf_transformation,[],[f4216_D])).
fof(f4216_D,plain,(
  sP5('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) <=> ~$spl480),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl480])])).
fof(f4201,plain,(
  sP5('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl104),
  inference(unit_resulting_resolution,[],[f263,f792,f1528])).
fof(f4202,plain,(
  sP5('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl104),
  inference(unit_resulting_resolution,[],[f262,f792,f1528])).
fof(f4184,plain,(
  $false | $spl183),
  inference(unit_resulting_resolution,[],[f268,f1956,f386])).
fof(f1956,plain,(
  ~sP4('F_0','F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | $spl183),
  inference(cnf_transformation,[],[f1956_D])).
fof(f1956_D,plain,(
  ~sP4('F_0','F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) <=> ~$spl183),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl183])])).
fof(f4182,plain,(
  $false | $spl183),
  inference(unit_resulting_resolution,[],[f268,f1956,f385])).
fof(f4151,plain,(
  ~sP5('F_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl479),
  inference(cnf_transformation,[],[f4151_D])).
fof(f4151_D,plain,(
  ~sP5('F_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) <=> ~$spl479),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl479])])).
fof(f4138,plain,(
  ~sP5('F_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | ($spl12 | $spl104)),
  inference(unit_resulting_resolution,[],[f268,f792,f3193])).
fof(f4137,plain,(
  ~sP5('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl477),
  inference(cnf_transformation,[],[f4137_D])).
fof(f4137_D,plain,(
  ~sP5('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) <=> ~$spl477),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl477])])).
fof(f4122,plain,(
  ~sP5('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | ($spl12 | $spl104)),
  inference(unit_resulting_resolution,[],[f268,f792,f3192])).
fof(f4135,plain,(
  ~sP5('F_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl475),
  inference(cnf_transformation,[],[f4135_D])).
fof(f4135_D,plain,(
  ~sP5('F_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) <=> ~$spl475),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl475])])).
fof(f4123,plain,(
  ~sP5('F_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl12),
  inference(unit_resulting_resolution,[],[f268,f485,f3192])).
fof(f4121,plain,(
  sP5('T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl472),
  inference(cnf_transformation,[],[f4121_D])).
fof(f4121_D,plain,(
  sP5('T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) <=> ~$spl472),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl472])])).
fof(f4106,plain,(
  sP5('T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | ($spl12 | $spl104)),
  inference(unit_resulting_resolution,[],[f263,f792,f3191])).
fof(f4107,plain,(
  sP5('T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | ($spl12 | $spl104)),
  inference(unit_resulting_resolution,[],[f262,f792,f3191])).
fof(f4105,plain,(
  sP5('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl470),
  inference(cnf_transformation,[],[f4105_D])).
fof(f4105_D,plain,(
  sP5('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) <=> ~$spl470),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl470])])).
fof(f4088,plain,(
  sP5('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | ($spl12 | $spl104)),
  inference(unit_resulting_resolution,[],[f263,f792,f3190])).
fof(f4089,plain,(
  sP5('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | ($spl12 | $spl104)),
  inference(unit_resulting_resolution,[],[f262,f792,f3190])).
fof(f4103,plain,(
  sP5('T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl468),
  inference(cnf_transformation,[],[f4103_D])).
fof(f4103_D,plain,(
  sP5('T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) <=> ~$spl468),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl468])])).
fof(f4090,plain,(
  sP5('T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl12),
  inference(unit_resulting_resolution,[],[f263,f485,f3190])).
fof(f4091,plain,(
  sP5('T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl12),
  inference(unit_resulting_resolution,[],[f262,f485,f3190])).
fof(f4087,plain,(
  sP3('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','T_0') | $spl466),
  inference(cnf_transformation,[],[f4087_D])).
fof(f4087_D,plain,(
  sP3('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','T_0') <=> ~$spl466),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl466])])).
fof(f4083,plain,(
  sP3('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','T_0') | $spl168),
  inference(resolution,[],[f1585,f377])).
fof(f377,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP3(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f225])).
fof(f4079,plain,(
  sP3('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','T_0') | $spl168),
  inference(unit_resulting_resolution,[],[f1585,f377])).
fof(f4072,plain,(
  sP4('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl462),
  inference(cnf_transformation,[],[f4072_D])).
fof(f4072_D,plain,(
  sP4('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') <=> ~$spl462),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl462])])).
fof(f4067,plain,(
  sP4('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl166),
  inference(resolution,[],[f1575,f378])).
fof(f4074,plain,(
  sP3('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl464),
  inference(cnf_transformation,[],[f4074_D])).
fof(f4074_D,plain,(
  sP3('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') <=> ~$spl464),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl464])])).
fof(f4066,plain,(
  sP3('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl166),
  inference(resolution,[],[f1575,f377])).
fof(f4061,plain,(
  sP3('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl166),
  inference(unit_resulting_resolution,[],[f1575,f377])).
fof(f4062,plain,(
  sP4('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl166),
  inference(unit_resulting_resolution,[],[f1575,f378])).
fof(f1960,plain,(
  ~sP2('F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0') | $spl187),
  inference(cnf_transformation,[],[f1960_D])).
fof(f1960_D,plain,(
  ~sP2('F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0') <=> ~$spl187),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl187])])).
fof(f1958,plain,(
  ~sP3('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0','F_0') | $spl185),
  inference(cnf_transformation,[],[f1958_D])).
fof(f1958_D,plain,(
  ~sP3('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0','F_0') <=> ~$spl185),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl185])])).
fof(f512,plain,(
  sP5('F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0') | $spl24),
  inference(cnf_transformation,[],[f512_D])).
fof(f512_D,plain,(
  sP5('F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0') <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f4055,plain,(
  p(s(bool,'F_0')) | sP5('F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0') | ~sP3('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0','F_0') | ~sP2('F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0') | $spl182),
  inference(resolution,[],[f1957,f382])).
fof(f4057,plain,(
  p(s(bool,'F_0')) | sP5('F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0') | ~sP3('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0','F_0') | ~sP2('F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0') | ($spl14 | $spl182)),
  inference(forward_demodulation,[],[f4054,f486])).
fof(f4054,plain,(
  p(s(bool,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) | sP5('F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0') | ~sP3('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0','F_0') | ~sP2('F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0') | $spl182),
  inference(resolution,[],[f1957,f381])).
fof(f381,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X2,X1) | p(s(bool,X1)) | sP5(X0,X1,X2) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f225])).
fof(f4051,plain,(
  ~sP2('F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0') | ($spl25 | $spl182 | $spl250)),
  inference(unit_resulting_resolution,[],[f268,f2329,f511,f1957,f382])).
fof(f4013,plain,(
  sP4('T_0','F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | $spl460),
  inference(cnf_transformation,[],[f4013_D])).
fof(f4013_D,plain,(
  sP4('T_0','F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) <=> ~$spl460),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl460])])).
fof(f3282,plain,(
  sP4('T_0','F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | $spl22),
  inference(unit_resulting_resolution,[],[f509,f378])).
fof(f509,plain,(
  sP5('T_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0') | $spl22),
  inference(cnf_transformation,[],[f509_D])).
fof(f509_D,plain,(
  sP5('T_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0') <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f3287,plain,(
  sP4('T_0','F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | $spl22),
  inference(resolution,[],[f509,f378])).
fof(f4011,plain,(
  sP4('F_0','T_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | $spl458),
  inference(cnf_transformation,[],[f4011_D])).
fof(f4011_D,plain,(
  sP4('F_0','T_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) <=> ~$spl458),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl458])])).
fof(f3780,plain,(
  sP4('F_0','T_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | $spl196),
  inference(unit_resulting_resolution,[],[f2147,f378])).
fof(f3784,plain,(
  sP4('F_0','T_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | $spl196),
  inference(resolution,[],[f2147,f378])).
fof(f4009,plain,(
  p(s(bool,'F_0')) | ($spl14 | $spl214)),
  inference(global_subsumption,[],[f261,f260,f2253,f2254,f2256,f2257,f2261,f2262,f2263,f2264,f2265,f2267,f2268,f2269,f2270,f2271,f2272,f2276,f2277,f2278,f2279,f2282,f2284,f347,f343,f2169,f486,f3040,f3041,f3042,f3043,f3739,f3738,f3742,f3753,f3752,f3751,f3749,f3750,f3762,f3773,f3772,f3769,f3770,f3786,f3789,f3800,f3799,f3798,f3796,f3797,f3808,f3819,f3818,f3815,f3816,f4004,f4005,f4006,f4007,f3991,f4008])).
fof(f4008,plain,(
  ( ! [X4,X5] : (p(s(bool,'F_0')) | sP5(X4,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X5) | ~sP3('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X5,X4) | ~sP2(X5,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X4)) ) | ($spl14 | $spl214)),
  inference(forward_demodulation,[],[f3992,f486])).
fof(f3992,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) | sP5(X4,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X5) | ~sP3('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X5,X4) | ~sP2(X5,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X4)) ) | $spl214),
  inference(resolution,[],[f2169,f381])).
fof(f3991,plain,(
  ( ! [X2,X3] : (~p(s(bool,X2)) | sP5(X3,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X2) | ~sP3('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X2,X3) | ~sP2(X2,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X3)) ) | $spl214),
  inference(resolution,[],[f2169,f380])).
fof(f4006,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | ~p(s(bool,X0)) | ~p(s(bool,X1))) ) | ($spl14 | $spl214)),
  inference(forward_demodulation,[],[f3990,f486])).
fof(f3990,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,X1))) ) | $spl214),
  inference(resolution,[],[f2169,f383])).
fof(f4007,plain,(
  p(s(bool,'F_0')) | ($spl14 | $spl214)),
  inference(global_subsumption,[],[f261,f260,f2253,f2254,f2256,f2257,f2261,f2262,f2263,f2264,f2265,f2267,f2268,f2269,f2270,f2271,f2272,f2276,f2277,f2278,f2279,f2282,f2284,f347,f343,f2169,f486,f3040,f3041,f3042,f3043,f3739,f3738,f3742,f3753,f3752,f3751,f3749,f3750,f3762,f3773,f3772,f3769,f3770,f3786,f3789,f3800,f3799,f3798,f3796,f3797,f3808,f3819,f3818,f3815,f3816,f4004,f4005,f4006])).
fof(f4005,plain,(
  p(s(bool,'F_0')) | ($spl14 | $spl214)),
  inference(forward_demodulation,[],[f3973,f486])).
fof(f3973,plain,(
  p(s(bool,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) | $spl214),
  inference(unit_resulting_resolution,[],[f263,f263,f2169,f383])).
fof(f4004,plain,(
  p(s(bool,'F_0')) | ($spl14 | $spl214)),
  inference(forward_demodulation,[],[f3974,f486])).
fof(f3974,plain,(
  p(s(bool,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) | $spl214),
  inference(unit_resulting_resolution,[],[f263,f262,f2169,f383])).
fof(f4003,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl456),
  inference(cnf_transformation,[],[f4003_D])).
fof(f4003_D,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl456),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl456])])).
fof(f3975,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'T_0') | ($spl2 | $spl214 | $spl250)),
  inference(unit_resulting_resolution,[],[f263,f583,f2329,f2169,f380])).
fof(f3976,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'T_0') | ($spl2 | $spl214 | $spl250)),
  inference(unit_resulting_resolution,[],[f262,f583,f2329,f2169,f380])).
fof(f4001,plain,(
  sP5('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl454),
  inference(cnf_transformation,[],[f4001_D])).
fof(f4001_D,plain,(
  sP5('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl454),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl454])])).
fof(f3977,plain,(
  sP5('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'T_0') | ($spl214 | $spl250 | $spl294)),
  inference(unit_resulting_resolution,[],[f263,f2510,f2329,f2169,f380])).
fof(f3978,plain,(
  sP5('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'T_0') | ($spl214 | $spl250 | $spl294)),
  inference(unit_resulting_resolution,[],[f262,f2510,f2329,f2169,f380])).
fof(f3979,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'T_0') | ($spl2 | $spl214)),
  inference(unit_resulting_resolution,[],[f263,f583,f582,f2169,f380])).
fof(f3980,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'T_0') | ($spl2 | $spl214)),
  inference(unit_resulting_resolution,[],[f262,f583,f582,f2169,f380])).
fof(f3981,plain,(
  sP5('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'T_0') | ($spl214 | $spl292 | $spl294)),
  inference(unit_resulting_resolution,[],[f263,f2510,f2509,f2169,f380])).
fof(f3982,plain,(
  sP5('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'T_0') | ($spl214 | $spl292 | $spl294)),
  inference(unit_resulting_resolution,[],[f262,f2510,f2509,f2169,f380])).
fof(f2162,plain,(
  ~p(s(bool,'T_0')) | $spl203),
  inference(cnf_transformation,[],[f2162_D])).
fof(f2162_D,plain,(
  ~p(s(bool,'T_0')) <=> ~$spl203),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl203])])).
fof(f3999,plain,(
  ~p(s(bool,'T_0')) | ($spl5 | $spl12 | $spl214 | $spl250 | $spl298)),
  inference(forward_demodulation,[],[f3983,f485])).
fof(f3983,plain,(
  ~p(s(bool,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) | ($spl5 | $spl214 | $spl250 | $spl298)),
  inference(unit_resulting_resolution,[],[f2512,f2329,f477,f2169,f380])).
fof(f3998,plain,(
  ~sP2('T_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl453),
  inference(cnf_transformation,[],[f3998_D])).
fof(f3998_D,plain,(
  ~sP2('T_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl453),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl453])])).
fof(f3984,plain,(
  ~sP2('T_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'T_0') | ($spl199 | $spl214 | $spl250)),
  inference(unit_resulting_resolution,[],[f263,f2329,f2149,f2169,f380])).
fof(f3996,plain,(
  p(s(bool,'F_0')) | ($spl5 | $spl14 | $spl214 | $spl250 | $spl298)),
  inference(forward_demodulation,[],[f3985,f486])).
fof(f3985,plain,(
  p(s(bool,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) | ($spl5 | $spl214 | $spl250 | $spl298)),
  inference(unit_resulting_resolution,[],[f2512,f2329,f477,f2169,f381])).
fof(f3995,plain,(
  sP5('F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl450),
  inference(cnf_transformation,[],[f3995_D])).
fof(f3995_D,plain,(
  sP5('F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) <=> ~$spl450),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl450])])).
fof(f3986,plain,(
  sP5('F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | ($spl2 | $spl214 | $spl250)),
  inference(unit_resulting_resolution,[],[f268,f585,f2329,f2169,f382])).
fof(f3989,plain,(
  ~sP2('F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0') | ($spl25 | $spl214 | $spl250)),
  inference(unit_resulting_resolution,[],[f268,f2329,f511,f2169,f382])).
fof(f3949,plain,(
  sP11('F_0','F_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl448),
  inference(cnf_transformation,[],[f3949_D])).
fof(f3949_D,plain,(
  sP11('F_0','F_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) <=> ~$spl448),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl448])])).
fof(f3933,plain,(
  sP11('F_0','F_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl314),
  inference(unit_resulting_resolution,[],[f268,f268,f2520])).
fof(f3859,plain,(
  ~sP1('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl447),
  inference(cnf_transformation,[],[f3859_D])).
fof(f3859_D,plain,(
  ~sP1('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') <=> ~$spl447),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl447])])).
fof(f3853,plain,(
  ~sP1('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl288),
  inference(unit_resulting_resolution,[],[f263,f2507])).
fof(f3854,plain,(
  ~sP1('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl288),
  inference(unit_resulting_resolution,[],[f262,f2507])).
fof(f3849,plain,(
  sP1('F_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl444),
  inference(cnf_transformation,[],[f3849_D])).
fof(f3849_D,plain,(
  sP1('F_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) <=> ~$spl444),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl444])])).
fof(f3841,plain,(
  sP1('F_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl286),
  inference(unit_resulting_resolution,[],[f268,f2506])).
fof(f3837,plain,(
  sP1('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') | $spl442),
  inference(cnf_transformation,[],[f3837_D])).
fof(f3837_D,plain,(
  sP1('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') <=> ~$spl442),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl442])])).
fof(f3829,plain,(
  sP1('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') | $spl284),
  inference(unit_resulting_resolution,[],[f268,f2505])).
fof(f3826,plain,(
  ~sP1('T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl441),
  inference(cnf_transformation,[],[f3826_D])).
fof(f3826_D,plain,(
  ~sP1('T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) <=> ~$spl441),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl441])])).
fof(f3822,plain,(
  ~sP1('T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl282),
  inference(unit_resulting_resolution,[],[f263,f2504])).
fof(f3823,plain,(
  ~sP1('T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl282),
  inference(unit_resulting_resolution,[],[f262,f2504])).
fof(f3757,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,X2)) ) | $spl428),
  inference(cnf_transformation,[],[f3757_D])).
fof(f3757_D,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl428),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl428])])).
fof(f3821,plain,(
  ( ! [X3] : (p(s(bool,X3)) | ~sP5(X3,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0')) ) | $spl438),
  inference(cnf_transformation,[],[f3821_D])).
fof(f3821_D,plain,(
  ( ! [X3] : (p(s(bool,X3)) | ~sP5(X3,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0')) ) <=> ~$spl438),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl438])])).
fof(f3818,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | ~sP5(X3,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0') | s(bool,'T_0') = s(bool,X2)) ) | $spl14),
  inference(trivial_inequality_removal,[],[f3811])).
fof(f3811,plain,(
  ( ! [X2,X3] : (s(bool,X2) != s(bool,X2) | p(s(bool,X3)) | ~sP5(X3,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0') | s(bool,'T_0') = s(bool,X2)) ) | $spl14),
  inference(superposition,[],[f3043,f269])).
fof(f3804,plain,(
  ( ! [X3] : (p(s(bool,X3)) | ~sP5(X3,'F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) ) | $spl436),
  inference(cnf_transformation,[],[f3804_D])).
fof(f3804_D,plain,(
  ( ! [X3] : (p(s(bool,X3)) | ~sP5(X3,'F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) ) <=> ~$spl436),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl436])])).
fof(f3799,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | ~sP5(X3,'F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | s(bool,'T_0') = s(bool,X2)) ) | $spl14),
  inference(trivial_inequality_removal,[],[f3792])).
fof(f3792,plain,(
  ( ! [X2,X3] : (s(bool,X2) != s(bool,X2) | p(s(bool,X3)) | ~sP5(X3,'F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | s(bool,'T_0') = s(bool,X2)) ) | $spl14),
  inference(superposition,[],[f3042,f269])).
fof(f3803,plain,(
  ~sP5('F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | $spl435),
  inference(cnf_transformation,[],[f3803_D])).
fof(f3803_D,plain,(
  ~sP5('F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) <=> ~$spl435),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl435])])).
fof(f3786,plain,(
  ~sP5('F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | $spl14),
  inference(unit_resulting_resolution,[],[f268,f486,f3042])).
fof(f3775,plain,(
  ( ! [X3] : (sP5(X3,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0') | ~p(s(bool,X3))) ) | $spl432),
  inference(cnf_transformation,[],[f3775_D])).
fof(f3775_D,plain,(
  ( ! [X3] : (sP5(X3,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0') | ~p(s(bool,X3))) ) <=> ~$spl432),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl432])])).
fof(f3772,plain,(
  ( ! [X2,X3] : (sP5(X3,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0') | ~p(s(bool,X3)) | s(bool,'T_0') = s(bool,X2)) ) | $spl14),
  inference(trivial_inequality_removal,[],[f3765])).
fof(f3765,plain,(
  ( ! [X2,X3] : (s(bool,X2) != s(bool,X2) | sP5(X3,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0') | ~p(s(bool,X3)) | s(bool,'T_0') = s(bool,X2)) ) | $spl14),
  inference(superposition,[],[f3041,f269])).
fof(f3758,plain,(
  ( ! [X3] : (sP5(X3,'F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | ~p(s(bool,X3))) ) | $spl430),
  inference(cnf_transformation,[],[f3758_D])).
fof(f3758_D,plain,(
  ( ! [X3] : (sP5(X3,'F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | ~p(s(bool,X3))) ) <=> ~$spl430),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl430])])).
fof(f3752,plain,(
  ( ! [X2,X3] : (sP5(X3,'F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | ~p(s(bool,X3)) | s(bool,'T_0') = s(bool,X2)) ) | $spl14),
  inference(trivial_inequality_removal,[],[f3745])).
fof(f3745,plain,(
  ( ! [X2,X3] : (s(bool,X2) != s(bool,X2) | sP5(X3,'F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | ~p(s(bool,X3)) | s(bool,'T_0') = s(bool,X2)) ) | $spl14),
  inference(superposition,[],[f3040,f269])).
fof(f3756,plain,(
  sP5('T_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | $spl426),
  inference(cnf_transformation,[],[f3756_D])).
fof(f3756_D,plain,(
  sP5('T_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) <=> ~$spl426),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl426])])).
fof(f3738,plain,(
  sP5('T_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | $spl14),
  inference(unit_resulting_resolution,[],[f263,f486,f3040])).
fof(f3739,plain,(
  sP5('T_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | $spl14),
  inference(unit_resulting_resolution,[],[f262,f486,f3040])).
fof(f3733,plain,(
  sP11('T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl424),
  inference(cnf_transformation,[],[f3733_D])).
fof(f3733_D,plain,(
  sP11('T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') <=> ~$spl424),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl424])])).
fof(f3723,plain,(
  sP11('T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl312),
  inference(unit_resulting_resolution,[],[f263,f263,f2519,f437])).
fof(f3724,plain,(
  sP11('T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl312),
  inference(unit_resulting_resolution,[],[f262,f263,f2519,f437])).
fof(f3725,plain,(
  sP11('T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl312),
  inference(unit_resulting_resolution,[],[f263,f262,f2519,f437])).
fof(f3726,plain,(
  sP11('T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl312),
  inference(unit_resulting_resolution,[],[f262,f262,f2519,f437])).
fof(f3715,plain,(
  sP9('T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') | $spl422),
  inference(cnf_transformation,[],[f3715_D])).
fof(f3715_D,plain,(
  sP9('T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') <=> ~$spl422),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl422])])).
fof(f3707,plain,(
  sP9('T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') | $spl306),
  inference(unit_resulting_resolution,[],[f268,f263,f2516,f420])).
fof(f3708,plain,(
  sP9('T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') | $spl306),
  inference(unit_resulting_resolution,[],[f268,f262,f2516,f420])).
fof(f3699,plain,(
  sP9('T_0','F_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl420),
  inference(cnf_transformation,[],[f3699_D])).
fof(f3699_D,plain,(
  sP9('T_0','F_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) <=> ~$spl420),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl420])])).
fof(f3691,plain,(
  sP9('T_0','F_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl304),
  inference(unit_resulting_resolution,[],[f263,f268,f2515,f421])).
fof(f3692,plain,(
  sP9('T_0','F_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl304),
  inference(unit_resulting_resolution,[],[f262,f268,f2515,f421])).
fof(f3680,plain,(
  sP7('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','T_0') | $spl418),
  inference(cnf_transformation,[],[f3680_D])).
fof(f3680_D,plain,(
  sP7('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','T_0') <=> ~$spl418),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl418])])).
fof(f3670,plain,(
  sP7('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','T_0') | $spl300),
  inference(unit_resulting_resolution,[],[f263,f263,f2513,f402])).
fof(f3671,plain,(
  sP7('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','T_0') | $spl300),
  inference(unit_resulting_resolution,[],[f262,f263,f2513,f402])).
fof(f3672,plain,(
  sP7('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','T_0') | $spl300),
  inference(unit_resulting_resolution,[],[f263,f262,f2513,f402])).
fof(f3673,plain,(
  sP7('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','T_0') | $spl300),
  inference(unit_resulting_resolution,[],[f262,f262,f2513,f402])).
fof(f2784,plain,(
  ~sP3('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','F_0') | $spl333),
  inference(cnf_transformation,[],[f2784_D])).
fof(f2784_D,plain,(
  ~sP3('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','F_0') <=> ~$spl333),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl333])])).
fof(f3653,plain,(
  ~sP3('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','F_0') | ($spl123 | $spl290 | $spl296)),
  inference(unit_resulting_resolution,[],[f263,f2508,f849,f2511,f380])).
fof(f3658,plain,(
  ~sP3('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','F_0') | ($spl123 | $spl290 | $spl296)),
  inference(unit_resulting_resolution,[],[f268,f2508,f849,f2511,f382])).
fof(f3641,plain,(
  sP5('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl416),
  inference(cnf_transformation,[],[f3641_D])).
fof(f3641_D,plain,(
  sP5('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') <=> ~$spl416),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl416])])).
fof(f3633,plain,(
  sP5('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | ($spl290 | $spl292 | $spl294)),
  inference(unit_resulting_resolution,[],[f263,f2508,f2509,f2510,f380])).
fof(f3634,plain,(
  sP5('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | ($spl290 | $spl292 | $spl294)),
  inference(unit_resulting_resolution,[],[f262,f2508,f2509,f2510,f380])).
fof(f3639,plain,(
  sP5('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | $spl414),
  inference(cnf_transformation,[],[f3639_D])).
fof(f3639_D,plain,(
  sP5('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) <=> ~$spl414),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl414])])).
fof(f3635,plain,(
  sP5('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | ($spl244 | $spl292 | $spl294)),
  inference(unit_resulting_resolution,[],[f268,f2326,f2509,f2510,f381])).
fof(f3636,plain,(
  sP5('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | ($spl244 | $spl248 | $spl294)),
  inference(unit_resulting_resolution,[],[f268,f2326,f2328,f2510,f381])).
fof(f3626,plain,(
  sP5('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl412),
  inference(cnf_transformation,[],[f3626_D])).
fof(f3626_D,plain,(
  sP5('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') <=> ~$spl412),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl412])])).
fof(f3622,plain,(
  sP5('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | ($spl2 | $spl292)),
  inference(unit_resulting_resolution,[],[f263,f584,f575,f2509,f380])).
fof(f3623,plain,(
  sP5('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | ($spl2 | $spl292)),
  inference(unit_resulting_resolution,[],[f262,f584,f575,f2509,f380])).
fof(f3615,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl410),
  inference(cnf_transformation,[],[f3615_D])).
fof(f3615_D,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') <=> ~$spl410),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl410])])).
fof(f3608,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | ($spl2 | $spl290)),
  inference(unit_resulting_resolution,[],[f263,f583,f582,f2508,f380])).
fof(f3609,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | ($spl2 | $spl290)),
  inference(unit_resulting_resolution,[],[f262,f583,f582,f2508,f380])).
fof(f3602,plain,(
  sP3('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','T_0') | $spl408),
  inference(cnf_transformation,[],[f3602_D])).
fof(f3602_D,plain,(
  sP3('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','T_0') <=> ~$spl408),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl408])])).
fof(f3598,plain,(
  sP3('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','T_0') | $spl120),
  inference(resolution,[],[f847,f377])).
fof(f3594,plain,(
  sP3('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','T_0') | $spl120),
  inference(unit_resulting_resolution,[],[f847,f377])).
fof(f3584,plain,(
  sP4('T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl404),
  inference(cnf_transformation,[],[f3584_D])).
fof(f3584_D,plain,(
  sP4('T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') <=> ~$spl404),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl404])])).
fof(f3579,plain,(
  sP4('T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl116),
  inference(resolution,[],[f843,f378])).
fof(f3586,plain,(
  sP3('T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl406),
  inference(cnf_transformation,[],[f3586_D])).
fof(f3586_D,plain,(
  sP3('T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') <=> ~$spl406),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl406])])).
fof(f3578,plain,(
  sP3('T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl116),
  inference(resolution,[],[f843,f377])).
fof(f3573,plain,(
  sP3('T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl116),
  inference(unit_resulting_resolution,[],[f843,f377])).
fof(f3574,plain,(
  sP4('T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl116),
  inference(unit_resulting_resolution,[],[f843,f378])).
fof(f3548,plain,(
  sP9('F_0','F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | $spl402),
  inference(cnf_transformation,[],[f3548_D])).
fof(f3548_D,plain,(
  sP9('F_0','F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) <=> ~$spl402),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl402])])).
fof(f3532,plain,(
  sP9('F_0','F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | $spl272),
  inference(unit_resulting_resolution,[],[f268,f268,f2340])).
fof(f3519,plain,(
  sP9('F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0') | $spl400),
  inference(cnf_transformation,[],[f3519_D])).
fof(f3519_D,plain,(
  sP9('F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0') <=> ~$spl400),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl400])])).
fof(f3503,plain,(
  sP9('F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0') | $spl268),
  inference(unit_resulting_resolution,[],[f268,f268,f2338])).
fof(f3502,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') != s(bool,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) | p(s(bool,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))))),
  introduced(theory_axiom,[])).
fof(f3499,plain,(
  sP5('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'T_0','F_0') | $spl398),
  inference(cnf_transformation,[],[f3499_D])).
fof(f3499_D,plain,(
  sP5('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'T_0','F_0') <=> ~$spl398),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl398])])).
fof(f3482,plain,(
  sP5('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'T_0','F_0') | ($spl29 | $spl252)),
  inference(unit_resulting_resolution,[],[f518,f2330])).
fof(f3497,plain,(
  sP5('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'T_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | $spl396),
  inference(cnf_transformation,[],[f3497_D])).
fof(f3497_D,plain,(
  sP5('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'T_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) <=> ~$spl396),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl396])])).
fof(f3483,plain,(
  sP5('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'T_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | ($spl27 | $spl252)),
  inference(unit_resulting_resolution,[],[f513,f2330])).
fof(f3495,plain,(
  sP5('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0','T_0') | $spl394),
  inference(cnf_transformation,[],[f3495_D])).
fof(f3495_D,plain,(
  sP5('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0','T_0') <=> ~$spl394),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl394])])).
fof(f3484,plain,(
  sP5('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0','T_0') | ($spl29 | $spl252)),
  inference(unit_resulting_resolution,[],[f518,f2330])).
fof(f3493,plain,(
  sP5('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl392),
  inference(cnf_transformation,[],[f3493_D])).
fof(f3493_D,plain,(
  sP5('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl392),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl392])])).
fof(f3485,plain,(
  sP5('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'T_0') | ($spl27 | $spl252)),
  inference(unit_resulting_resolution,[],[f513,f2330])).
fof(f2147,plain,(
  sP5('F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl196),
  inference(cnf_transformation,[],[f2147_D])).
fof(f2147_D,plain,(
  sP5('F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f3407,plain,(
  sP5('F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl27),
  inference(unit_resulting_resolution,[],[f268,f513,f395])).
fof(f2149,plain,(
  ~sP5('T_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl199),
  inference(cnf_transformation,[],[f2149_D])).
fof(f2149_D,plain,(
  ~sP5('T_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl199),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl199])])).
fof(f3408,plain,(
  ~sP5('T_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl27),
  inference(unit_resulting_resolution,[],[f263,f513,f397])).
fof(f3409,plain,(
  ~sP5('T_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl27),
  inference(unit_resulting_resolution,[],[f262,f513,f397])).
fof(f2133,plain,(
  sP5('F_0','T_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | $spl192),
  inference(cnf_transformation,[],[f2133_D])).
fof(f2133_D,plain,(
  sP5('F_0','T_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f3412,plain,(
  sP5('F_0','T_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | $spl27),
  inference(unit_resulting_resolution,[],[f268,f513,f395])).
fof(f2135,plain,(
  ~sP5('T_0','T_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | $spl195),
  inference(cnf_transformation,[],[f2135_D])).
fof(f2135_D,plain,(
  ~sP5('T_0','T_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) <=> ~$spl195),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl195])])).
fof(f3413,plain,(
  ~sP5('T_0','T_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | $spl27),
  inference(unit_resulting_resolution,[],[f263,f513,f397])).
fof(f3414,plain,(
  ~sP5('T_0','T_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | $spl27),
  inference(unit_resulting_resolution,[],[f262,f513,f397])).
fof(f3397,plain,(
  sP11('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0','T_0') | $spl390),
  inference(cnf_transformation,[],[f3397_D])).
fof(f3397_D,plain,(
  sP11('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0','T_0') <=> ~$spl390),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl390])])).
fof(f3389,plain,(
  sP11('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0','T_0') | $spl276),
  inference(unit_resulting_resolution,[],[f263,f268,f2342,f439])).
fof(f3390,plain,(
  sP11('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0','T_0') | $spl276),
  inference(unit_resulting_resolution,[],[f262,f268,f2342,f439])).
fof(f3383,plain,(
  sP11('T_0','F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | $spl388),
  inference(cnf_transformation,[],[f3383_D])).
fof(f3383_D,plain,(
  sP11('T_0','F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) <=> ~$spl388),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl388])])).
fof(f3375,plain,(
  sP11('T_0','F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | $spl274),
  inference(unit_resulting_resolution,[],[f263,f268,f2341,f438])).
fof(f3376,plain,(
  sP11('T_0','F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | $spl274),
  inference(unit_resulting_resolution,[],[f262,f268,f2341,f438])).
fof(f3365,plain,(
  sP9('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0','F_0') | $spl386),
  inference(cnf_transformation,[],[f3365_D])).
fof(f3365_D,plain,(
  sP9('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0','F_0') <=> ~$spl386),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl386])])).
fof(f3358,plain,(
  sP9('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0','F_0') | $spl266),
  inference(unit_resulting_resolution,[],[f268,f268,f2337,f422])).
fof(f422,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X1,X2) | sP9(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f244])).
fof(f1941,plain,(
  ~sP2('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0','F_0') | $spl181),
  inference(cnf_transformation,[],[f1941_D])).
fof(f1941_D,plain,(
  ~sP2('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0','F_0') <=> ~$spl181),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl181])])).
fof(f3316,plain,(
  ~sP2('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0','F_0') | ($spl21 | $spl244 | $spl248)),
  inference(unit_resulting_resolution,[],[f268,f2326,f507,f2328,f381])).
fof(f3322,plain,(
  sP5('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | $spl384),
  inference(cnf_transformation,[],[f3322_D])).
fof(f3322_D,plain,(
  sP5('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) <=> ~$spl384),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl384])])).
fof(f3317,plain,(
  sP5('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | ($spl2 | $spl244 | $spl248)),
  inference(unit_resulting_resolution,[],[f268,f584,f2326,f2328,f382])).
fof(f3318,plain,(
  sP5('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | ($spl2 | $spl248)),
  inference(unit_resulting_resolution,[],[f268,f584,f575,f2328,f382])).
fof(f3319,plain,(
  ~sP2('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0','F_0') | ($spl21 | $spl244 | $spl248)),
  inference(unit_resulting_resolution,[],[f268,f2326,f507,f2328,f382])).
fof(f3313,plain,(
  ( ! [X6,X7] : (sP5('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X6,X7) | ~sP2(X7,X6,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | ~sP3(X6,X7,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) ) | $spl382),
  inference(cnf_transformation,[],[f3313_D])).
fof(f3313_D,plain,(
  ( ! [X6,X7] : (sP5('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X6,X7) | ~sP2(X7,X6,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | ~sP3(X6,X7,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) ) <=> ~$spl382),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl382])])).
fof(f3312,plain,(
  ( ! [X6,X7] : (p(s(bool,'F_0')) | sP5('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X6,X7) | ~sP3(X6,X7,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | ~sP2(X7,X6,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) ) | ($spl14 | $spl246)),
  inference(forward_demodulation,[],[f3309,f486])).
fof(f3309,plain,(
  ( ! [X6,X7] : (p(s(bool,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) | sP5('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X6,X7) | ~sP3(X6,X7,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | ~sP2(X7,X6,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) ) | $spl246),
  inference(resolution,[],[f2327,f382])).
fof(f3301,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | $spl380),
  inference(cnf_transformation,[],[f3301_D])).
fof(f3301_D,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) <=> ~$spl380),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl380])])).
fof(f3295,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | ($spl2 | $spl244)),
  inference(unit_resulting_resolution,[],[f268,f583,f582,f2326,f381])).
fof(f2970,plain,(
  p(s(bool,'F_0')) | ($spl14 | $spl204 | $spl320)),
  inference(global_subsumption,[],[f261,f260,f2253,f2254,f2256,f2257,f2261,f2262,f2263,f2264,f2265,f2267,f2268,f2269,f2270,f2271,f2272,f2276,f2277,f2278,f2279,f2282,f2284,f347,f343,f2163,f2530,f2537,f486,f2968,f2969])).
fof(f2969,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0))) ) | ($spl14 | $spl204 | $spl320)),
  inference(forward_demodulation,[],[f2967,f486])).
fof(f2967,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))))) ) | ($spl204 | $spl320)),
  inference(resolution,[],[f2163,f2530])).
fof(f2968,plain,(
  p(s(bool,'F_0')) | ($spl14 | $spl204 | $spl320)),
  inference(forward_demodulation,[],[f2966,f486])).
fof(f2966,plain,(
  p(s(bool,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) | ($spl204 | $spl320)),
  inference(unit_resulting_resolution,[],[f268,f2163,f2530])).
fof(f504,plain,(
  p(s(bool,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) | $spl16),
  inference(cnf_transformation,[],[f504_D])).
fof(f504_D,plain,(
  p(s(bool,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f2585,plain,(
  p(s(bool,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) | ($spl200 | $spl320)),
  inference(unit_resulting_resolution,[],[f268,f2160,f2530])).
fof(f2961,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl378),
  inference(cnf_transformation,[],[f2961_D])).
fof(f2961_D,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) <=> ~$spl378),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl378])])).
fof(f2586,plain,(
  ( ! [X0] : (p(s(bool,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) | p(s(bool,X0))) ) | ($spl200 | $spl320)),
  inference(resolution,[],[f2160,f2530])).
fof(f2956,plain,(
  ( ! [X83,X84] : (sP10('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X83,X84)) ) | $spl376),
  inference(cnf_transformation,[],[f2956_D])).
fof(f2956_D,plain,(
  ( ! [X83,X84] : (sP10('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X83,X84)) ) <=> ~$spl376),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl376])])).
fof(f2885,plain,(
  ( ! [X83,X84] : (p(s(bool,'F_0')) | sP10('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X83,X84)) ) | $spl124),
  inference(superposition,[],[f443,f852])).
fof(f2955,plain,(
  ( ! [X80,X79] : (sP10(X79,X80,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) ) | $spl374),
  inference(cnf_transformation,[],[f2955_D])).
fof(f2955_D,plain,(
  ( ! [X80,X79] : (sP10(X79,X80,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) ) <=> ~$spl374),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl374])])).
fof(f2883,plain,(
  ( ! [X80,X79] : (p(s(bool,'F_0')) | sP10(X79,X80,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) ) | $spl124),
  inference(superposition,[],[f441,f852])).
fof(f2954,plain,(
  ( ! [X74,X73] : (sP9(X73,X74,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | p(s(bool,X73)) | p(s(bool,X74))) ) | $spl372),
  inference(cnf_transformation,[],[f2954_D])).
fof(f2954_D,plain,(
  ( ! [X74,X73] : (sP9(X73,X74,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | p(s(bool,X73)) | p(s(bool,X74))) ) <=> ~$spl372),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl372])])).
fof(f2880,plain,(
  ( ! [X74,X73] : (p(s(bool,'F_0')) | sP9(X73,X74,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | p(s(bool,X74)) | p(s(bool,X73))) ) | $spl124),
  inference(superposition,[],[f427,f852])).
fof(f2953,plain,(
  ( ! [X72,X71] : (sP9('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X71,X72) | p(s(bool,X72)) | p(s(bool,X71))) ) | $spl370),
  inference(cnf_transformation,[],[f2953_D])).
fof(f2953_D,plain,(
  ( ! [X72,X71] : (sP9('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X71,X72) | p(s(bool,X72)) | p(s(bool,X71))) ) <=> ~$spl370),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl370])])).
fof(f2879,plain,(
  ( ! [X72,X71] : (p(s(bool,'F_0')) | sP9('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X71,X72) | p(s(bool,X71)) | p(s(bool,X72))) ) | $spl124),
  inference(superposition,[],[f427,f852])).
fof(f2952,plain,(
  ( ! [X70,X69] : (sP9(X69,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X70) | p(s(bool,X69)) | p(s(bool,X70))) ) | $spl368),
  inference(cnf_transformation,[],[f2952_D])).
fof(f2952_D,plain,(
  ( ! [X70,X69] : (sP9(X69,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X70) | p(s(bool,X69)) | p(s(bool,X70))) ) <=> ~$spl368),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl368])])).
fof(f2878,plain,(
  ( ! [X70,X69] : (p(s(bool,'F_0')) | sP9(X69,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X70) | p(s(bool,X70)) | p(s(bool,X69))) ) | $spl124),
  inference(superposition,[],[f427,f852])).
fof(f2951,plain,(
  ( ! [X68,X67] : (sP8('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X67,X68)) ) | $spl366),
  inference(cnf_transformation,[],[f2951_D])).
fof(f2951_D,plain,(
  ( ! [X68,X67] : (sP8('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X67,X68)) ) <=> ~$spl366),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl366])])).
fof(f2877,plain,(
  ( ! [X68,X67] : (p(s(bool,'F_0')) | sP8('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X67,X68)) ) | $spl124),
  inference(superposition,[],[f426,f852])).
fof(f2950,plain,(
  ( ! [X59,X60] : (sP7('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X59,X60) | p(s(bool,X59))) ) | $spl364),
  inference(cnf_transformation,[],[f2950_D])).
fof(f2950_D,plain,(
  ( ! [X59,X60] : (sP7('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X59,X60) | p(s(bool,X59))) ) <=> ~$spl364),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl364])])).
fof(f2873,plain,(
  ( ! [X59,X60] : (p(s(bool,'F_0')) | sP7('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X59,X60) | p(s(bool,X59))) ) | $spl124),
  inference(superposition,[],[f411,f852])).
fof(f2949,plain,(
  ( ! [X57,X58] : (sP7(X57,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X58) | p(s(bool,X57))) ) | $spl362),
  inference(cnf_transformation,[],[f2949_D])).
fof(f2949_D,plain,(
  ( ! [X57,X58] : (sP7(X57,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X58) | p(s(bool,X57))) ) <=> ~$spl362),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl362])])).
fof(f2872,plain,(
  ( ! [X57,X58] : (p(s(bool,'F_0')) | sP7(X57,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X58) | p(s(bool,X57))) ) | $spl124),
  inference(superposition,[],[f411,f852])).
fof(f2948,plain,(
  ( ! [X56,X55] : (sP7(X55,X56,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | p(s(bool,X55))) ) | $spl360),
  inference(cnf_transformation,[],[f2948_D])).
fof(f2948_D,plain,(
  ( ! [X56,X55] : (sP7(X55,X56,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | p(s(bool,X55))) ) <=> ~$spl360),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl360])])).
fof(f2871,plain,(
  ( ! [X56,X55] : (p(s(bool,'F_0')) | sP7(X55,X56,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | p(s(bool,X55))) ) | $spl124),
  inference(superposition,[],[f410,f852])).
fof(f2947,plain,(
  ( ! [X54,X53] : (sP7('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X53,X54) | p(s(bool,X54))) ) | $spl358),
  inference(cnf_transformation,[],[f2947_D])).
fof(f2947_D,plain,(
  ( ! [X54,X53] : (sP7('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X53,X54) | p(s(bool,X54))) ) <=> ~$spl358),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl358])])).
fof(f2870,plain,(
  ( ! [X54,X53] : (p(s(bool,'F_0')) | sP7('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X53,X54) | p(s(bool,X54))) ) | $spl124),
  inference(superposition,[],[f410,f852])).
fof(f2946,plain,(
  ( ! [X52,X51] : (sP6('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X51,X52)) ) | $spl356),
  inference(cnf_transformation,[],[f2946_D])).
fof(f2946_D,plain,(
  ( ! [X52,X51] : (sP6('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X51,X52)) ) <=> ~$spl356),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl356])])).
fof(f2869,plain,(
  ( ! [X52,X51] : (p(s(bool,'F_0')) | sP6('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X51,X52)) ) | $spl124),
  inference(superposition,[],[f409,f852])).
fof(f2945,plain,(
  ( ! [X50,X49] : (sP6(X49,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X50)) ) | $spl354),
  inference(cnf_transformation,[],[f2945_D])).
fof(f2945_D,plain,(
  ( ! [X50,X49] : (sP6(X49,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X50)) ) <=> ~$spl354),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl354])])).
fof(f2868,plain,(
  ( ! [X50,X49] : (p(s(bool,'F_0')) | sP6(X49,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X50)) ) | $spl124),
  inference(superposition,[],[f408,f852])).
fof(f2944,plain,(
  ( ! [X37,X38] : (sP5('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X37,X38) | s(bool,X37) = s(bool,X38)) ) | $spl352),
  inference(cnf_transformation,[],[f2944_D])).
fof(f2944_D,plain,(
  ( ! [X37,X38] : (sP5('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X37,X38) | s(bool,X37) = s(bool,X38)) ) <=> ~$spl352),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl352])])).
fof(f2862,plain,(
  ( ! [X37,X38] : (p(s(bool,'F_0')) | sP5('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X37,X38) | s(bool,X37) = s(bool,X38)) ) | $spl124),
  inference(superposition,[],[f395,f852])).
fof(f2943,plain,(
  ( ! [X30,X29] : (sP3('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X29,X30)) ) | $spl350),
  inference(cnf_transformation,[],[f2943_D])).
fof(f2943_D,plain,(
  ( ! [X30,X29] : (sP3('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X29,X30)) ) <=> ~$spl350),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl350])])).
fof(f2858,plain,(
  ( ! [X30,X29] : (p(s(bool,'F_0')) | sP3('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X29,X30)) ) | $spl124),
  inference(superposition,[],[f390,f852])).
fof(f2942,plain,(
  ( ! [X28,X27] : (sP3(X27,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X28)) ) | $spl348),
  inference(cnf_transformation,[],[f2942_D])).
fof(f2942_D,plain,(
  ( ! [X28,X27] : (sP3(X27,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X28)) ) <=> ~$spl348),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl348])])).
fof(f2857,plain,(
  ( ! [X28,X27] : (p(s(bool,'F_0')) | sP3(X27,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X28)) ) | $spl124),
  inference(superposition,[],[f389,f852])).
fof(f2941,plain,(
  ( ! [X24,X23] : (sP4('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X23,X24)) ) | $spl346),
  inference(cnf_transformation,[],[f2941_D])).
fof(f2941_D,plain,(
  ( ! [X24,X23] : (sP4('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X23,X24)) ) <=> ~$spl346),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl346])])).
fof(f2855,plain,(
  ( ! [X24,X23] : (p(s(bool,'F_0')) | sP4('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X23,X24)) ) | $spl124),
  inference(superposition,[],[f386,f852])).
fof(f2940,plain,(
  ( ! [X21,X22] : (sP4(X21,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X22)) ) | $spl344),
  inference(cnf_transformation,[],[f2940_D])).
fof(f2940_D,plain,(
  ( ! [X21,X22] : (sP4(X21,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X22)) ) <=> ~$spl344),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl344])])).
fof(f2854,plain,(
  ( ! [X21,X22] : (p(s(bool,'F_0')) | sP4(X21,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X22)) ) | $spl124),
  inference(superposition,[],[f385,f852])).
fof(f2931,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'F_0')) | ($spl26 | $spl124)),
  inference(forward_demodulation,[],[f2818,f514])).
fof(f2818,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) | $spl124),
  inference(superposition,[],[f368,f852])).
fof(f368,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.pred_set.countable_1'(s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X1,X0),X2),s(fun(X1,bool),X3)))))) | ~p(s(bool,'const.pred_set.countable_1'(s(fun(X1,bool),X3))))) )),
  inference(cnf_transformation,[],[f140])).
fof(f140,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.pred_set.countable_1'(s(fun(X1,bool),X3)))) | p(s(bool,'const.pred_set.countable_1'(s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X1,X0),X2),s(fun(X1,bool),X3)))))))),
  inference(ennf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.pred_set.countable_1'(s(fun(X1,bool),X3)))) => p(s(bool,'const.pred_set.countable_1'(s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X1,X0),X2),s(fun(X1,bool),X3)))))))),
  inference(flattening,[],[f102])).
fof(f102,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.pred_set.countable_1'(s(fun(X1,bool),X3)))) => p(s(bool,'const.pred_set.countable_1'(s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X1,X0),X2),s(fun(X1,bool),X3)))))))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X6,X7] : ! [X9,X8] : (p(s(bool,'const.pred_set.countable_1'(s(fun(X7,bool),X8)))) => p(s(bool,'const.pred_set.countable_1'(s(fun(X6,bool),'const.pred_set.IMAGE_2'(s(fun(X7,X6),X9),s(fun(X7,bool),X8)))))))),
  file('thm.pred_set.inj_image_countable_IFF','thm.pred_set.image_countable')).
fof(f2927,plain,(
  sP5('T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') | $spl342),
  inference(cnf_transformation,[],[f2927_D])).
fof(f2927_D,plain,(
  sP5('T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') <=> ~$spl342),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl342])])).
fof(f2791,plain,(
  sP5('T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') | $spl124),
  inference(unit_resulting_resolution,[],[f263,f852,f396])).
fof(f2792,plain,(
  sP5('T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') | $spl124),
  inference(unit_resulting_resolution,[],[f262,f852,f396])).
fof(f2925,plain,(
  ~sP5('F_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') | $spl341),
  inference(cnf_transformation,[],[f2925_D])).
fof(f2925_D,plain,(
  ~sP5('F_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') <=> ~$spl341),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl341])])).
fof(f2793,plain,(
  ~sP5('F_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') | $spl124),
  inference(unit_resulting_resolution,[],[f268,f852,f398])).
fof(f2916,plain,(
  sP5('T_0','F_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl338),
  inference(cnf_transformation,[],[f2916_D])).
fof(f2916_D,plain,(
  sP5('T_0','F_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) <=> ~$spl338),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl338])])).
fof(f2803,plain,(
  sP5('T_0','F_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl124),
  inference(unit_resulting_resolution,[],[f263,f852,f396])).
fof(f2804,plain,(
  sP5('T_0','F_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl124),
  inference(unit_resulting_resolution,[],[f262,f852,f396])).
fof(f2914,plain,(
  ~sP5('F_0','F_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl337),
  inference(cnf_transformation,[],[f2914_D])).
fof(f2914_D,plain,(
  ~sP5('F_0','F_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) <=> ~$spl337),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl337])])).
fof(f2805,plain,(
  ~sP5('F_0','F_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl124),
  inference(unit_resulting_resolution,[],[f268,f852,f398])).
fof(f2783,plain,(
  sP4('F_0','T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl330),
  inference(cnf_transformation,[],[f2783_D])).
fof(f2783_D,plain,(
  sP4('F_0','T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) <=> ~$spl330),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl330])])).
fof(f2780,plain,(
  sP4('F_0','T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl122),
  inference(resolution,[],[f850,f378])).
fof(f2785,plain,(
  sP3('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','F_0') | $spl332),
  inference(cnf_transformation,[],[f2785_D])).
fof(f2785_D,plain,(
  sP3('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','F_0') <=> ~$spl332),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl332])])).
fof(f2779,plain,(
  sP3('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','F_0') | $spl122),
  inference(resolution,[],[f850,f377])).
fof(f2787,plain,(
  sP2('T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') | $spl334),
  inference(cnf_transformation,[],[f2787_D])).
fof(f2787_D,plain,(
  sP2('T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') <=> ~$spl334),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl334])])).
fof(f2778,plain,(
  sP2('T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') | $spl122),
  inference(resolution,[],[f850,f376])).
fof(f376,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f225])).
fof(f2774,plain,(
  sP2('T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') | $spl122),
  inference(unit_resulting_resolution,[],[f850,f376])).
fof(f2775,plain,(
  sP3('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','F_0') | $spl122),
  inference(unit_resulting_resolution,[],[f850,f377])).
fof(f2776,plain,(
  sP4('F_0','T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl122),
  inference(unit_resulting_resolution,[],[f850,f378])).
fof(f2764,plain,(
  sP4('F_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl324),
  inference(cnf_transformation,[],[f2764_D])).
fof(f2764_D,plain,(
  sP4('F_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') <=> ~$spl324),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl324])])).
fof(f2761,plain,(
  sP4('F_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl118),
  inference(resolution,[],[f846,f378])).
fof(f2766,plain,(
  sP3('T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') | $spl326),
  inference(cnf_transformation,[],[f2766_D])).
fof(f2766_D,plain,(
  sP3('T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') <=> ~$spl326),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl326])])).
fof(f2760,plain,(
  sP3('T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') | $spl118),
  inference(resolution,[],[f846,f377])).
fof(f2768,plain,(
  sP2('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','F_0') | $spl328),
  inference(cnf_transformation,[],[f2768_D])).
fof(f2768_D,plain,(
  sP2('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','F_0') <=> ~$spl328),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl328])])).
fof(f2759,plain,(
  sP2('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','F_0') | $spl118),
  inference(resolution,[],[f846,f376])).
fof(f2755,plain,(
  sP2('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','F_0') | $spl118),
  inference(unit_resulting_resolution,[],[f846,f376])).
fof(f2756,plain,(
  sP3('T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') | $spl118),
  inference(unit_resulting_resolution,[],[f846,f377])).
fof(f2757,plain,(
  sP4('F_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl118),
  inference(unit_resulting_resolution,[],[f846,f378])).
fof(f2541,plain,(
  ~sP0('F_0','F_0') | $spl323),
  inference(cnf_transformation,[],[f2541_D])).
fof(f2541_D,plain,(
  ~sP0('F_0','F_0') <=> ~$spl323),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl323])])).
fof(f2537,plain,(
  ~sP0('F_0','F_0') | $spl320),
  inference(unit_resulting_resolution,[],[f268,f268,f2530])).
fof(f487,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | $spl15),
  inference(cnf_transformation,[],[f487_D])).
fof(f487_D,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f2531,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | ($spl1 | $spl124)),
  inference(backward_demodulation,[],[f852,f466])).
fof(f2530,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) ) | $spl320),
  inference(cnf_transformation,[],[f2530_D])).
fof(f2530_D,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) ) <=> ~$spl320),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl320])])).
fof(f343,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f211])).
fof(f347,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f213])).
fof(f2529,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | p(s(bool,'T_0')) | ~p(s(bool,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))))),
  introduced(theory_axiom,[])).
fof(f2528,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | p(s(bool,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) | ~p(s(bool,'F_0'))),
  introduced(theory_axiom,[])).
fof(f2527,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | s(bool,'T_0') != s(bool,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | p(s(bool,'T_0')) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))))),
  introduced(theory_axiom,[])).
fof(f2522,plain,(
  ( ! [X90,X89] : (sP11('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X89,X90) | ~p(s(bool,X89))) ) | $spl318),
  inference(cnf_transformation,[],[f2522_D])).
fof(f2522_D,plain,(
  ( ! [X90,X89] : (sP11('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X89,X90) | ~p(s(bool,X89))) ) <=> ~$spl318),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl318])])).
fof(f2456,plain,(
  ( ! [X90,X89] : (~p(s(bool,'T_0')) | ~p(s(bool,X89)) | sP11('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X89,X90)) ) | $spl12),
  inference(superposition,[],[f446,f485])).
fof(f2521,plain,(
  ( ! [X88,X87] : (sP11('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X88,X87) | p(s(bool,X87))) ) | $spl316),
  inference(cnf_transformation,[],[f2521_D])).
fof(f2521_D,plain,(
  ( ! [X88,X87] : (sP11('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X88,X87) | p(s(bool,X87))) ) <=> ~$spl316),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl316])])).
fof(f2455,plain,(
  ( ! [X88,X87] : (~p(s(bool,'T_0')) | p(s(bool,X87)) | sP11('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X88,X87)) ) | $spl12),
  inference(superposition,[],[f445,f485])).
fof(f2520,plain,(
  ( ! [X85,X86] : (sP11(X86,X85,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | p(s(bool,X86)) | p(s(bool,X85))) ) | $spl314),
  inference(cnf_transformation,[],[f2520_D])).
fof(f2520_D,plain,(
  ( ! [X85,X86] : (sP11(X86,X85,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | p(s(bool,X86)) | p(s(bool,X85))) ) <=> ~$spl314),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl314])])).
fof(f2454,plain,(
  ( ! [X85,X86] : (~p(s(bool,'T_0')) | p(s(bool,X85)) | sP11(X86,X85,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | p(s(bool,X86))) ) | $spl12),
  inference(superposition,[],[f444,f485])).
fof(f2519,plain,(
  ( ! [X81,X82] : (sP10(X81,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X82)) ) | $spl312),
  inference(cnf_transformation,[],[f2519_D])).
fof(f2519_D,plain,(
  ( ! [X81,X82] : (sP10(X81,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X82)) ) <=> ~$spl312),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl312])])).
fof(f2452,plain,(
  ( ! [X81,X82] : (~p(s(bool,'T_0')) | sP10(X81,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X82)) ) | $spl12),
  inference(superposition,[],[f442,f485])).
fof(f2518,plain,(
  ( ! [X78,X77] : (sP9('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X77,X78) | ~p(s(bool,X77))) ) | $spl310),
  inference(cnf_transformation,[],[f2518_D])).
fof(f2518_D,plain,(
  ( ! [X78,X77] : (sP9('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X77,X78) | ~p(s(bool,X77))) ) <=> ~$spl310),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl310])])).
fof(f2450,plain,(
  ( ! [X78,X77] : (~p(s(bool,'T_0')) | ~p(s(bool,X77)) | sP9('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X77,X78)) ) | $spl12),
  inference(superposition,[],[f429,f485])).
fof(f2517,plain,(
  ( ! [X76,X75] : (sP9('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X76,X75) | ~p(s(bool,X75))) ) | $spl308),
  inference(cnf_transformation,[],[f2517_D])).
fof(f2517_D,plain,(
  ( ! [X76,X75] : (sP9('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X76,X75) | ~p(s(bool,X75))) ) <=> ~$spl308),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl308])])).
fof(f2449,plain,(
  ( ! [X76,X75] : (~p(s(bool,'T_0')) | ~p(s(bool,X75)) | sP9('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X76,X75)) ) | $spl12),
  inference(superposition,[],[f428,f485])).
fof(f2516,plain,(
  ( ! [X66,X65] : (sP8(X65,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X66)) ) | $spl306),
  inference(cnf_transformation,[],[f2516_D])).
fof(f2516_D,plain,(
  ( ! [X66,X65] : (sP8(X65,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X66)) ) <=> ~$spl306),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl306])])).
fof(f2444,plain,(
  ( ! [X66,X65] : (~p(s(bool,'T_0')) | sP8(X65,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X66)) ) | $spl12),
  inference(superposition,[],[f425,f485])).
fof(f2515,plain,(
  ( ! [X64,X63] : (sP8(X63,X64,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) ) | $spl304),
  inference(cnf_transformation,[],[f2515_D])).
fof(f2515_D,plain,(
  ( ! [X64,X63] : (sP8(X63,X64,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) ) <=> ~$spl304),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl304])])).
fof(f2443,plain,(
  ( ! [X64,X63] : (~p(s(bool,'T_0')) | sP8(X63,X64,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) ) | $spl12),
  inference(superposition,[],[f424,f485])).
fof(f2514,plain,(
  ( ! [X61,X62] : (sP7('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X61,X62) | ~p(s(bool,X61)) | ~p(s(bool,X62))) ) | $spl302),
  inference(cnf_transformation,[],[f2514_D])).
fof(f2514_D,plain,(
  ( ! [X61,X62] : (sP7('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X61,X62) | ~p(s(bool,X61)) | ~p(s(bool,X62))) ) <=> ~$spl302),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl302])])).
fof(f2442,plain,(
  ( ! [X61,X62] : (~p(s(bool,'T_0')) | ~p(s(bool,X61)) | ~p(s(bool,X62)) | sP7('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X61,X62)) ) | $spl12),
  inference(superposition,[],[f412,f485])).
fof(f2513,plain,(
  ( ! [X47,X48] : (sP6(X47,X48,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) ) | $spl300),
  inference(cnf_transformation,[],[f2513_D])).
fof(f2513_D,plain,(
  ( ! [X47,X48] : (sP6(X47,X48,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) ) <=> ~$spl300),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl300])])).
fof(f2435,plain,(
  ( ! [X47,X48] : (~p(s(bool,'T_0')) | sP6(X47,X48,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) ) | $spl12),
  inference(superposition,[],[f407,f485])).
fof(f2512,plain,(
  ( ! [X35,X36] : (sP2('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X35,X36)) ) | $spl298),
  inference(cnf_transformation,[],[f2512_D])).
fof(f2512_D,plain,(
  ( ! [X35,X36] : (sP2('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X35,X36)) ) <=> ~$spl298),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl298])])).
fof(f2429,plain,(
  ( ! [X35,X36] : (~p(s(bool,'T_0')) | sP2('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X35,X36)) ) | $spl12),
  inference(superposition,[],[f394,f485])).
fof(f2511,plain,(
  ( ! [X33,X34] : (sP2(X33,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X34)) ) | $spl296),
  inference(cnf_transformation,[],[f2511_D])).
fof(f2511_D,plain,(
  ( ! [X33,X34] : (sP2(X33,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X34)) ) <=> ~$spl296),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl296])])).
fof(f2428,plain,(
  ( ! [X33,X34] : (~p(s(bool,'T_0')) | sP2(X33,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X34)) ) | $spl12),
  inference(superposition,[],[f393,f485])).
fof(f2510,plain,(
  ( ! [X31,X32] : (sP2(X31,X32,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) ) | $spl294),
  inference(cnf_transformation,[],[f2510_D])).
fof(f2510_D,plain,(
  ( ! [X31,X32] : (sP2(X31,X32,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) ) <=> ~$spl294),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl294])])).
fof(f2427,plain,(
  ( ! [X31,X32] : (~p(s(bool,'T_0')) | sP2(X31,X32,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) ) | $spl12),
  inference(superposition,[],[f392,f485])).
fof(f2509,plain,(
  ( ! [X26,X25] : (sP3(X25,X26,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) ) | $spl292),
  inference(cnf_transformation,[],[f2509_D])).
fof(f2509_D,plain,(
  ( ! [X26,X25] : (sP3(X25,X26,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) ) <=> ~$spl292),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl292])])).
fof(f2424,plain,(
  ( ! [X26,X25] : (~p(s(bool,'T_0')) | sP3(X25,X26,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) ) | $spl12),
  inference(superposition,[],[f388,f485])).
fof(f2508,plain,(
  ( ! [X19,X20] : (sP4(X19,X20,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) ) | $spl290),
  inference(cnf_transformation,[],[f2508_D])).
fof(f2508_D,plain,(
  ( ! [X19,X20] : (sP4(X19,X20,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) ) <=> ~$spl290),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl290])])).
fof(f2421,plain,(
  ( ! [X19,X20] : (~p(s(bool,'T_0')) | sP4(X19,X20,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) ) | $spl12),
  inference(superposition,[],[f384,f485])).
fof(f2507,plain,(
  ( ! [X18] : (~sP1('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X18) | ~p(s(bool,X18))) ) | $spl288),
  inference(cnf_transformation,[],[f2507_D])).
fof(f2507_D,plain,(
  ( ! [X18] : (~sP1('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X18) | ~p(s(bool,X18))) ) <=> ~$spl288),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl288])])).
fof(f2420,plain,(
  ( ! [X18] : (~p(s(bool,'T_0')) | ~p(s(bool,X18)) | ~sP1('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X18)) ) | $spl12),
  inference(superposition,[],[f366,f485])).
fof(f2505,plain,(
  ( ! [X14] : (sP1('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X14) | p(s(bool,X14))) ) | $spl284),
  inference(cnf_transformation,[],[f2505_D])).
fof(f2505_D,plain,(
  ( ! [X14] : (sP1('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X14) | p(s(bool,X14))) ) <=> ~$spl284),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl284])])).
fof(f2419,plain,(
  ( ! [X17] : (~p(s(bool,'T_0')) | p(s(bool,X17)) | sP1('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X17)) ) | $spl12),
  inference(superposition,[],[f365,f485])).
fof(f2506,plain,(
  ( ! [X15] : (sP1(X15,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | p(s(bool,X15))) ) | $spl286),
  inference(cnf_transformation,[],[f2506_D])).
fof(f2506_D,plain,(
  ( ! [X15] : (sP1(X15,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | p(s(bool,X15))) ) <=> ~$spl286),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl286])])).
fof(f2418,plain,(
  ( ! [X16] : (~p(s(bool,'T_0')) | sP1(X16,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | p(s(bool,X16))) ) | $spl12),
  inference(superposition,[],[f364,f485])).
fof(f2417,plain,(
  ( ! [X15] : (~p(s(bool,'T_0')) | p(s(bool,X15)) | sP1(X15,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) ) | $spl12),
  inference(superposition,[],[f363,f485])).
fof(f2416,plain,(
  ( ! [X14] : (~p(s(bool,'T_0')) | p(s(bool,X14)) | sP1('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X14)) ) | $spl12),
  inference(superposition,[],[f360,f485])).
fof(f2504,plain,(
  ( ! [X13] : (~sP1(X13,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | ~p(s(bool,X13))) ) | $spl282),
  inference(cnf_transformation,[],[f2504_D])).
fof(f2504_D,plain,(
  ( ! [X13] : (~sP1(X13,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | ~p(s(bool,X13))) ) <=> ~$spl282),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl282])])).
fof(f2415,plain,(
  ( ! [X13] : (~p(s(bool,'T_0')) | ~p(s(bool,X13)) | ~sP1(X13,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) ) | $spl12),
  inference(superposition,[],[f359,f485])).
fof(f2503,plain,(
  ( ! [X11] : (sP0('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X11)) ) | $spl280),
  inference(cnf_transformation,[],[f2503_D])).
fof(f2503_D,plain,(
  ( ! [X11] : (sP0('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X11)) ) <=> ~$spl280),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl280])])).
fof(f2414,plain,(
  ( ! [X12] : (~p(s(bool,'T_0')) | sP0('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X12)) ) | $spl12),
  inference(superposition,[],[f350,f485])).
fof(f2413,plain,(
  ( ! [X11] : (~p(s(bool,'T_0')) | sP0('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),X11)) ) | $spl12),
  inference(superposition,[],[f345,f485])).
fof(f2502,plain,(
  ( ! [X10] : (sP0(X10,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) ) | $spl278),
  inference(cnf_transformation,[],[f2502_D])).
fof(f2502_D,plain,(
  ( ! [X10] : (sP0(X10,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) ) <=> ~$spl278),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl278])])).
fof(f2412,plain,(
  ( ! [X10] : (~p(s(bool,'T_0')) | sP0(X10,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) ) | $spl12),
  inference(superposition,[],[f344,f485])).
fof(f847,plain,(
  sP5('T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl120),
  inference(cnf_transformation,[],[f847_D])).
fof(f847_D,plain,(
  sP5('T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f2352,plain,(
  sP5('T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl12),
  inference(unit_resulting_resolution,[],[f263,f485,f396])).
fof(f2353,plain,(
  sP5('T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl12),
  inference(unit_resulting_resolution,[],[f262,f485,f396])).
fof(f849,plain,(
  ~sP5('F_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl123),
  inference(cnf_transformation,[],[f849_D])).
fof(f849_D,plain,(
  ~sP5('F_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') <=> ~$spl123),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl123])])).
fof(f2354,plain,(
  ~sP5('F_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl12),
  inference(unit_resulting_resolution,[],[f268,f485,f398])).
fof(f843,plain,(
  sP5('T_0','T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl116),
  inference(cnf_transformation,[],[f843_D])).
fof(f843_D,plain,(
  sP5('T_0','T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f2371,plain,(
  sP5('T_0','T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl12),
  inference(unit_resulting_resolution,[],[f263,f485,f396])).
fof(f2372,plain,(
  sP5('T_0','T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl12),
  inference(unit_resulting_resolution,[],[f262,f485,f396])).
fof(f845,plain,(
  ~sP5('F_0','T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl119),
  inference(cnf_transformation,[],[f845_D])).
fof(f845_D,plain,(
  ~sP5('F_0','T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) <=> ~$spl119),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl119])])).
fof(f2373,plain,(
  ~sP5('F_0','T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl12),
  inference(unit_resulting_resolution,[],[f268,f485,f398])).
fof(f2347,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | p(s(bool,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) | ~p(s(bool,'F_0'))),
  introduced(theory_axiom,[])).
fof(f2284,plain,(
  ( ! [X85,X84] : (p(s(bool,'F_0')) | sP10('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X84,X85)) ) | $spl14),
  inference(superposition,[],[f443,f486])).
fof(f2341,plain,(
  ( ! [X80,X81] : (sP10(X80,X81,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) ) | $spl274),
  inference(cnf_transformation,[],[f2341_D])).
fof(f2341_D,plain,(
  ( ! [X80,X81] : (sP10(X80,X81,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) ) <=> ~$spl274),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl274])])).
fof(f2282,plain,(
  ( ! [X80,X81] : (p(s(bool,'F_0')) | sP10(X80,X81,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) ) | $spl14),
  inference(superposition,[],[f441,f486])).
fof(f2340,plain,(
  ( ! [X74,X75] : (sP9(X74,X75,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | p(s(bool,X74)) | p(s(bool,X75))) ) | $spl272),
  inference(cnf_transformation,[],[f2340_D])).
fof(f2340_D,plain,(
  ( ! [X74,X75] : (sP9(X74,X75,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | p(s(bool,X74)) | p(s(bool,X75))) ) <=> ~$spl272),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl272])])).
fof(f2279,plain,(
  ( ! [X74,X75] : (p(s(bool,'F_0')) | sP9(X74,X75,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | p(s(bool,X75)) | p(s(bool,X74))) ) | $spl14),
  inference(superposition,[],[f427,f486])).
fof(f2339,plain,(
  ( ! [X72,X73] : (sP9('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X72,X73) | p(s(bool,X73)) | p(s(bool,X72))) ) | $spl270),
  inference(cnf_transformation,[],[f2339_D])).
fof(f2339_D,plain,(
  ( ! [X72,X73] : (sP9('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X72,X73) | p(s(bool,X73)) | p(s(bool,X72))) ) <=> ~$spl270),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl270])])).
fof(f2278,plain,(
  ( ! [X72,X73] : (p(s(bool,'F_0')) | sP9('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X72,X73) | p(s(bool,X72)) | p(s(bool,X73))) ) | $spl14),
  inference(superposition,[],[f427,f486])).
fof(f2338,plain,(
  ( ! [X70,X71] : (sP9(X70,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X71) | p(s(bool,X70)) | p(s(bool,X71))) ) | $spl268),
  inference(cnf_transformation,[],[f2338_D])).
fof(f2338_D,plain,(
  ( ! [X70,X71] : (sP9(X70,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X71) | p(s(bool,X70)) | p(s(bool,X71))) ) <=> ~$spl268),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl268])])).
fof(f2277,plain,(
  ( ! [X70,X71] : (p(s(bool,'F_0')) | sP9(X70,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X71) | p(s(bool,X71)) | p(s(bool,X70))) ) | $spl14),
  inference(superposition,[],[f427,f486])).
fof(f2337,plain,(
  ( ! [X68,X69] : (sP8('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X68,X69)) ) | $spl266),
  inference(cnf_transformation,[],[f2337_D])).
fof(f2337_D,plain,(
  ( ! [X68,X69] : (sP8('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X68,X69)) ) <=> ~$spl266),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl266])])).
fof(f2276,plain,(
  ( ! [X68,X69] : (p(s(bool,'F_0')) | sP8('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X68,X69)) ) | $spl14),
  inference(superposition,[],[f426,f486])).
fof(f2336,plain,(
  ( ! [X61,X60] : (sP7('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X60,X61) | p(s(bool,X60))) ) | $spl264),
  inference(cnf_transformation,[],[f2336_D])).
fof(f2336_D,plain,(
  ( ! [X61,X60] : (sP7('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X60,X61) | p(s(bool,X60))) ) <=> ~$spl264),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl264])])).
fof(f2272,plain,(
  ( ! [X61,X60] : (p(s(bool,'F_0')) | sP7('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X60,X61) | p(s(bool,X60))) ) | $spl14),
  inference(superposition,[],[f411,f486])).
fof(f2335,plain,(
  ( ! [X59,X58] : (sP7(X58,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X59) | p(s(bool,X58))) ) | $spl262),
  inference(cnf_transformation,[],[f2335_D])).
fof(f2335_D,plain,(
  ( ! [X59,X58] : (sP7(X58,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X59) | p(s(bool,X58))) ) <=> ~$spl262),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl262])])).
fof(f2271,plain,(
  ( ! [X59,X58] : (p(s(bool,'F_0')) | sP7(X58,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X59) | p(s(bool,X58))) ) | $spl14),
  inference(superposition,[],[f411,f486])).
fof(f2334,plain,(
  ( ! [X57,X56] : (sP7(X56,X57,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | p(s(bool,X56))) ) | $spl260),
  inference(cnf_transformation,[],[f2334_D])).
fof(f2334_D,plain,(
  ( ! [X57,X56] : (sP7(X56,X57,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | p(s(bool,X56))) ) <=> ~$spl260),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl260])])).
fof(f2270,plain,(
  ( ! [X57,X56] : (p(s(bool,'F_0')) | sP7(X56,X57,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | p(s(bool,X56))) ) | $spl14),
  inference(superposition,[],[f410,f486])).
fof(f2333,plain,(
  ( ! [X54,X55] : (sP7('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X54,X55) | p(s(bool,X55))) ) | $spl258),
  inference(cnf_transformation,[],[f2333_D])).
fof(f2333_D,plain,(
  ( ! [X54,X55] : (sP7('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X54,X55) | p(s(bool,X55))) ) <=> ~$spl258),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl258])])).
fof(f2269,plain,(
  ( ! [X54,X55] : (p(s(bool,'F_0')) | sP7('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X54,X55) | p(s(bool,X55))) ) | $spl14),
  inference(superposition,[],[f410,f486])).
fof(f2332,plain,(
  ( ! [X52,X53] : (sP6('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X52,X53)) ) | $spl256),
  inference(cnf_transformation,[],[f2332_D])).
fof(f2332_D,plain,(
  ( ! [X52,X53] : (sP6('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X52,X53)) ) <=> ~$spl256),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl256])])).
fof(f2268,plain,(
  ( ! [X52,X53] : (p(s(bool,'F_0')) | sP6('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X52,X53)) ) | $spl14),
  inference(superposition,[],[f409,f486])).
fof(f2331,plain,(
  ( ! [X50,X51] : (sP6(X50,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X51)) ) | $spl254),
  inference(cnf_transformation,[],[f2331_D])).
fof(f2331_D,plain,(
  ( ! [X50,X51] : (sP6(X50,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X51)) ) <=> ~$spl254),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl254])])).
fof(f2267,plain,(
  ( ! [X50,X51] : (p(s(bool,'F_0')) | sP6(X50,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X51)) ) | $spl14),
  inference(superposition,[],[f408,f486])).
fof(f2330,plain,(
  ( ! [X39,X38] : (sP5('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X38,X39) | s(bool,X38) = s(bool,X39)) ) | $spl252),
  inference(cnf_transformation,[],[f2330_D])).
fof(f2330_D,plain,(
  ( ! [X39,X38] : (sP5('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X38,X39) | s(bool,X38) = s(bool,X39)) ) <=> ~$spl252),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl252])])).
fof(f2261,plain,(
  ( ! [X39,X38] : (p(s(bool,'F_0')) | sP5('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X38,X39) | s(bool,X38) = s(bool,X39)) ) | $spl14),
  inference(superposition,[],[f395,f486])).
fof(f2329,plain,(
  ( ! [X30,X31] : (sP3('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X30,X31)) ) | $spl250),
  inference(cnf_transformation,[],[f2329_D])).
fof(f2329_D,plain,(
  ( ! [X30,X31] : (sP3('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X30,X31)) ) <=> ~$spl250),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl250])])).
fof(f2257,plain,(
  ( ! [X30,X31] : (p(s(bool,'F_0')) | sP3('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X30,X31)) ) | $spl14),
  inference(superposition,[],[f390,f486])).
fof(f2328,plain,(
  ( ! [X28,X29] : (sP3(X28,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X29)) ) | $spl248),
  inference(cnf_transformation,[],[f2328_D])).
fof(f2328_D,plain,(
  ( ! [X28,X29] : (sP3(X28,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X29)) ) <=> ~$spl248),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl248])])).
fof(f2256,plain,(
  ( ! [X28,X29] : (p(s(bool,'F_0')) | sP3(X28,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X29)) ) | $spl14),
  inference(superposition,[],[f389,f486])).
fof(f2327,plain,(
  ( ! [X24,X25] : (sP4('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X24,X25)) ) | $spl246),
  inference(cnf_transformation,[],[f2327_D])).
fof(f2327_D,plain,(
  ( ! [X24,X25] : (sP4('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X24,X25)) ) <=> ~$spl246),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl246])])).
fof(f2254,plain,(
  ( ! [X24,X25] : (p(s(bool,'F_0')) | sP4('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X24,X25)) ) | $spl14),
  inference(superposition,[],[f386,f486])).
fof(f2326,plain,(
  ( ! [X23,X22] : (sP4(X22,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X23)) ) | $spl244),
  inference(cnf_transformation,[],[f2326_D])).
fof(f2326_D,plain,(
  ( ! [X23,X22] : (sP4(X22,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X23)) ) <=> ~$spl244),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl244])])).
fof(f2253,plain,(
  ( ! [X23,X22] : (p(s(bool,'F_0')) | sP4(X22,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X23)) ) | $spl14),
  inference(superposition,[],[f385,f486])).
fof(f1997,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | $spl24),
  inference(unit_resulting_resolution,[],[f268,f512,f398])).
fof(f2183,plain,(
  ( ! [X90,X91] : (~p(s(bool,X90)) | sP11('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X90,X91)) ) | $spl242),
  inference(cnf_transformation,[],[f2183_D])).
fof(f2183_D,plain,(
  ( ! [X90,X91] : (~p(s(bool,X90)) | sP11('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X90,X91)) ) <=> ~$spl242),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl242])])).
fof(f2110,plain,(
  ( ! [X90,X91] : (~p(s(bool,'T_0')) | ~p(s(bool,X90)) | sP11('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X90,X91)) ) | $spl26),
  inference(superposition,[],[f446,f514])).
fof(f2182,plain,(
  ( ! [X88,X89] : (p(s(bool,X88)) | sP11('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X89,X88)) ) | $spl240),
  inference(cnf_transformation,[],[f2182_D])).
fof(f2182_D,plain,(
  ( ! [X88,X89] : (p(s(bool,X88)) | sP11('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X89,X88)) ) <=> ~$spl240),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl240])])).
fof(f2109,plain,(
  ( ! [X88,X89] : (~p(s(bool,'T_0')) | p(s(bool,X88)) | sP11('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X89,X88)) ) | $spl26),
  inference(superposition,[],[f445,f514])).
fof(f2181,plain,(
  ( ! [X87,X86] : (p(s(bool,X86)) | p(s(bool,X87)) | sP11(X87,X86,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) ) | $spl238),
  inference(cnf_transformation,[],[f2181_D])).
fof(f2181_D,plain,(
  ( ! [X87,X86] : (p(s(bool,X86)) | p(s(bool,X87)) | sP11(X87,X86,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) ) <=> ~$spl238),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl238])])).
fof(f2108,plain,(
  ( ! [X87,X86] : (~p(s(bool,'T_0')) | p(s(bool,X86)) | sP11(X87,X86,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | p(s(bool,X87))) ) | $spl26),
  inference(superposition,[],[f444,f514])).
fof(f2180,plain,(
  ( ! [X83,X82] : (sP10(X82,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X83)) ) | $spl236),
  inference(cnf_transformation,[],[f2180_D])).
fof(f2180_D,plain,(
  ( ! [X83,X82] : (sP10(X82,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X83)) ) <=> ~$spl236),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl236])])).
fof(f2106,plain,(
  ( ! [X83,X82] : (~p(s(bool,'T_0')) | sP10(X82,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X83)) ) | $spl26),
  inference(superposition,[],[f442,f514])).
fof(f2179,plain,(
  ( ! [X78,X79] : (~p(s(bool,X78)) | sP9('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X78,X79)) ) | $spl234),
  inference(cnf_transformation,[],[f2179_D])).
fof(f2179_D,plain,(
  ( ! [X78,X79] : (~p(s(bool,X78)) | sP9('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X78,X79)) ) <=> ~$spl234),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl234])])).
fof(f2104,plain,(
  ( ! [X78,X79] : (~p(s(bool,'T_0')) | ~p(s(bool,X78)) | sP9('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X78,X79)) ) | $spl26),
  inference(superposition,[],[f429,f514])).
fof(f2178,plain,(
  ( ! [X76,X77] : (~p(s(bool,X76)) | sP9('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X77,X76)) ) | $spl232),
  inference(cnf_transformation,[],[f2178_D])).
fof(f2178_D,plain,(
  ( ! [X76,X77] : (~p(s(bool,X76)) | sP9('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X77,X76)) ) <=> ~$spl232),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl232])])).
fof(f2103,plain,(
  ( ! [X76,X77] : (~p(s(bool,'T_0')) | ~p(s(bool,X76)) | sP9('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X77,X76)) ) | $spl26),
  inference(superposition,[],[f428,f514])).
fof(f2177,plain,(
  ( ! [X66,X67] : (sP8(X66,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X67)) ) | $spl230),
  inference(cnf_transformation,[],[f2177_D])).
fof(f2177_D,plain,(
  ( ! [X66,X67] : (sP8(X66,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X67)) ) <=> ~$spl230),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl230])])).
fof(f2098,plain,(
  ( ! [X66,X67] : (~p(s(bool,'T_0')) | sP8(X66,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X67)) ) | $spl26),
  inference(superposition,[],[f425,f514])).
fof(f2176,plain,(
  ( ! [X64,X65] : (sP8(X64,X65,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) ) | $spl228),
  inference(cnf_transformation,[],[f2176_D])).
fof(f2176_D,plain,(
  ( ! [X64,X65] : (sP8(X64,X65,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) ) <=> ~$spl228),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl228])])).
fof(f2097,plain,(
  ( ! [X64,X65] : (~p(s(bool,'T_0')) | sP8(X64,X65,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) ) | $spl26),
  inference(superposition,[],[f424,f514])).
fof(f2175,plain,(
  ( ! [X62,X63] : (~p(s(bool,X62)) | sP7('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X62,X63) | ~p(s(bool,X63))) ) | $spl226),
  inference(cnf_transformation,[],[f2175_D])).
fof(f2175_D,plain,(
  ( ! [X62,X63] : (~p(s(bool,X62)) | sP7('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X62,X63) | ~p(s(bool,X63))) ) <=> ~$spl226),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl226])])).
fof(f2096,plain,(
  ( ! [X62,X63] : (~p(s(bool,'T_0')) | ~p(s(bool,X62)) | ~p(s(bool,X63)) | sP7('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X62,X63)) ) | $spl26),
  inference(superposition,[],[f412,f514])).
fof(f2174,plain,(
  ( ! [X48,X49] : (sP6(X48,X49,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) ) | $spl224),
  inference(cnf_transformation,[],[f2174_D])).
fof(f2174_D,plain,(
  ( ! [X48,X49] : (sP6(X48,X49,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) ) <=> ~$spl224),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl224])])).
fof(f2089,plain,(
  ( ! [X48,X49] : (~p(s(bool,'T_0')) | sP6(X48,X49,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) ) | $spl26),
  inference(superposition,[],[f407,f514])).
fof(f2173,plain,(
  ( ! [X37,X36] : (sP2('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X36,X37)) ) | $spl222),
  inference(cnf_transformation,[],[f2173_D])).
fof(f2173_D,plain,(
  ( ! [X37,X36] : (sP2('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X36,X37)) ) <=> ~$spl222),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl222])])).
fof(f2083,plain,(
  ( ! [X37,X36] : (~p(s(bool,'T_0')) | sP2('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X36,X37)) ) | $spl26),
  inference(superposition,[],[f394,f514])).
fof(f2172,plain,(
  ( ! [X35,X34] : (sP2(X34,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X35)) ) | $spl220),
  inference(cnf_transformation,[],[f2172_D])).
fof(f2172_D,plain,(
  ( ! [X35,X34] : (sP2(X34,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X35)) ) <=> ~$spl220),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl220])])).
fof(f2082,plain,(
  ( ! [X35,X34] : (~p(s(bool,'T_0')) | sP2(X34,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X35)) ) | $spl26),
  inference(superposition,[],[f393,f514])).
fof(f2171,plain,(
  ( ! [X33,X32] : (sP2(X32,X33,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) ) | $spl218),
  inference(cnf_transformation,[],[f2171_D])).
fof(f2171_D,plain,(
  ( ! [X33,X32] : (sP2(X32,X33,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) ) <=> ~$spl218),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl218])])).
fof(f2081,plain,(
  ( ! [X33,X32] : (~p(s(bool,'T_0')) | sP2(X32,X33,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) ) | $spl26),
  inference(superposition,[],[f392,f514])).
fof(f2170,plain,(
  ( ! [X26,X27] : (sP3(X26,X27,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) ) | $spl216),
  inference(cnf_transformation,[],[f2170_D])).
fof(f2170_D,plain,(
  ( ! [X26,X27] : (sP3(X26,X27,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) ) <=> ~$spl216),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl216])])).
fof(f2078,plain,(
  ( ! [X26,X27] : (~p(s(bool,'T_0')) | sP3(X26,X27,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) ) | $spl26),
  inference(superposition,[],[f388,f514])).
fof(f2169,plain,(
  ( ! [X21,X20] : (sP4(X20,X21,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) ) | $spl214),
  inference(cnf_transformation,[],[f2169_D])).
fof(f2169_D,plain,(
  ( ! [X21,X20] : (sP4(X20,X21,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) ) <=> ~$spl214),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl214])])).
fof(f2075,plain,(
  ( ! [X21,X20] : (~p(s(bool,'T_0')) | sP4(X20,X21,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) ) | $spl26),
  inference(superposition,[],[f384,f514])).
fof(f2168,plain,(
  ( ! [X19] : (~p(s(bool,X19)) | ~sP1('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X19)) ) | $spl212),
  inference(cnf_transformation,[],[f2168_D])).
fof(f2168_D,plain,(
  ( ! [X19] : (~p(s(bool,X19)) | ~sP1('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X19)) ) <=> ~$spl212),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl212])])).
fof(f2074,plain,(
  ( ! [X19] : (~p(s(bool,'T_0')) | ~p(s(bool,X19)) | ~sP1('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X19)) ) | $spl26),
  inference(superposition,[],[f366,f514])).
fof(f2166,plain,(
  ( ! [X15] : (p(s(bool,X15)) | sP1('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X15)) ) | $spl208),
  inference(cnf_transformation,[],[f2166_D])).
fof(f2166_D,plain,(
  ( ! [X15] : (p(s(bool,X15)) | sP1('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X15)) ) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f2073,plain,(
  ( ! [X18] : (~p(s(bool,'T_0')) | p(s(bool,X18)) | sP1('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X18)) ) | $spl26),
  inference(superposition,[],[f365,f514])).
fof(f2167,plain,(
  ( ! [X16] : (p(s(bool,X16)) | sP1(X16,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) ) | $spl210),
  inference(cnf_transformation,[],[f2167_D])).
fof(f2167_D,plain,(
  ( ! [X16] : (p(s(bool,X16)) | sP1(X16,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) ) <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f2072,plain,(
  ( ! [X17] : (~p(s(bool,'T_0')) | sP1(X17,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | p(s(bool,X17))) ) | $spl26),
  inference(superposition,[],[f364,f514])).
fof(f2071,plain,(
  ( ! [X16] : (~p(s(bool,'T_0')) | p(s(bool,X16)) | sP1(X16,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) ) | $spl26),
  inference(superposition,[],[f363,f514])).
fof(f2070,plain,(
  ( ! [X15] : (~p(s(bool,'T_0')) | p(s(bool,X15)) | sP1('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X15)) ) | $spl26),
  inference(superposition,[],[f360,f514])).
fof(f2165,plain,(
  ( ! [X14] : (~p(s(bool,X14)) | ~sP1(X14,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) ) | $spl206),
  inference(cnf_transformation,[],[f2165_D])).
fof(f2165_D,plain,(
  ( ! [X14] : (~p(s(bool,X14)) | ~sP1(X14,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) ) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f2069,plain,(
  ( ! [X14] : (~p(s(bool,'T_0')) | ~p(s(bool,X14)) | ~sP1(X14,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) ) | $spl26),
  inference(superposition,[],[f359,f514])).
fof(f2163,plain,(
  ( ! [X11] : (sP0('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X11)) ) | $spl204),
  inference(cnf_transformation,[],[f2163_D])).
fof(f2163_D,plain,(
  ( ! [X11] : (sP0('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X11)) ) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f2068,plain,(
  ( ! [X13] : (~p(s(bool,'T_0')) | sP0('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X13)) ) | $spl26),
  inference(superposition,[],[f350,f514])).
fof(f2160,plain,(
  ( ! [X10] : (sP0(X10,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) ) | $spl200),
  inference(cnf_transformation,[],[f2160_D])).
fof(f2160_D,plain,(
  ( ! [X10] : (sP0(X10,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) ) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f2164,plain,(
  ( ! [X12] : (~p(s(bool,'T_0')) | sP0(X12,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) ) | $spl26),
  inference(global_subsumption,[],[f261,f260,f514,f2029,f2028,f2027,f2012,f2011,f2010,f2065,f2066,f2067])).
fof(f2067,plain,(
  ( ! [X12] : (~p(s(bool,'T_0')) | p(s(bool,'F_0')) | sP0(X12,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) ) | $spl26),
  inference(superposition,[],[f349,f514])).
fof(f2066,plain,(
  ( ! [X11] : (~p(s(bool,'T_0')) | sP0('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),X11)) ) | $spl26),
  inference(superposition,[],[f345,f514])).
fof(f2065,plain,(
  ( ! [X10] : (~p(s(bool,'T_0')) | sP0(X10,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) ) | $spl26),
  inference(superposition,[],[f344,f514])).
fof(f2150,plain,(
  sP5('T_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl198),
  inference(cnf_transformation,[],[f2150_D])).
fof(f2150_D,plain,(
  sP5('T_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f2010,plain,(
  sP5('T_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl26),
  inference(unit_resulting_resolution,[],[f263,f514,f396])).
fof(f2011,plain,(
  sP5('T_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl26),
  inference(unit_resulting_resolution,[],[f262,f514,f396])).
fof(f2148,plain,(
  ~sP5('F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl197),
  inference(cnf_transformation,[],[f2148_D])).
fof(f2148_D,plain,(
  ~sP5('F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'T_0') <=> ~$spl197),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl197])])).
fof(f2012,plain,(
  ~sP5('F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'T_0') | $spl26),
  inference(unit_resulting_resolution,[],[f268,f514,f398])).
fof(f2136,plain,(
  sP5('T_0','T_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | $spl194),
  inference(cnf_transformation,[],[f2136_D])).
fof(f2136_D,plain,(
  sP5('T_0','T_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f2027,plain,(
  sP5('T_0','T_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | $spl26),
  inference(unit_resulting_resolution,[],[f263,f514,f396])).
fof(f2028,plain,(
  sP5('T_0','T_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | $spl26),
  inference(unit_resulting_resolution,[],[f262,f514,f396])).
fof(f2134,plain,(
  ~sP5('F_0','T_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | $spl193),
  inference(cnf_transformation,[],[f2134_D])).
fof(f2134_D,plain,(
  ~sP5('F_0','T_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) <=> ~$spl193),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl193])])).
fof(f2029,plain,(
  ~sP5('F_0','T_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | $spl26),
  inference(unit_resulting_resolution,[],[f268,f514,f398])).
fof(f1973,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | ($spl1 | $spl124)),
  inference(backward_demodulation,[],[f852,f466])).
fof(f1968,plain,(
  p(s(bool,'F_0')) | ($spl14 | $spl16)),
  inference(backward_demodulation,[],[f486,f504])).
fof(f842,plain,(
  ~p(s(bool,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) | $spl115),
  inference(cnf_transformation,[],[f842_D])).
fof(f842_D,plain,(
  ~p(s(bool,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) <=> ~$spl115),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl115])])).
fof(f1966,plain,(
  p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')))) | $spl188),
  inference(cnf_transformation,[],[f1966_D])).
fof(f1966_D,plain,(
  p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')))) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f1887,plain,(
  p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,'type.num.num'),'const.combin.o_2'(s(fun(sK12,'type.num.num'),sK16('const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)),sK12)),s(fun(sK13,sK12),sK15))),s(fun(sK13,bool),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')))) | ~p(s(bool,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) | $spl2),
  inference(resolution,[],[f682,f318])).
fof(f318,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,'type.num.num'),sK16(X1,X0)),s(fun(X0,bool),X1),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')))) | ~p(s(bool,'const.pred_set.countable_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f198])).
fof(f1924,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | $spl19),
  inference(unit_resulting_resolution,[],[f263,f506,f396])).
fof(f1927,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | $spl20),
  inference(unit_resulting_resolution,[],[f268,f508,f398])).
fof(f1944,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | $spl23),
  inference(unit_resulting_resolution,[],[f263,f510,f396])).
fof(f1964,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | s(bool,'T_0') != s(bool,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | s(bool,'F_0') != s(bool,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | s(bool,'F_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))),
  introduced(theory_axiom,[])).
fof(f1963,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | s(bool,'T_0') != s(bool,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) | p(s(bool,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))))),
  introduced(theory_axiom,[])).
fof(f1957,plain,(
  sP4('F_0','F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | $spl182),
  inference(cnf_transformation,[],[f1957_D])).
fof(f1957_D,plain,(
  sP4('F_0','F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f1954,plain,(
  sP4('F_0','F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | $spl24),
  inference(resolution,[],[f512,f378])).
fof(f1959,plain,(
  sP3('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0','F_0') | $spl184),
  inference(cnf_transformation,[],[f1959_D])).
fof(f1959_D,plain,(
  sP3('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0','F_0') <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f1953,plain,(
  sP3('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0','F_0') | $spl24),
  inference(resolution,[],[f512,f377])).
fof(f1961,plain,(
  sP2('F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0') | $spl186),
  inference(cnf_transformation,[],[f1961_D])).
fof(f1961_D,plain,(
  sP2('F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0') <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f1952,plain,(
  sP2('F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0') | $spl24),
  inference(resolution,[],[f512,f376])).
fof(f1948,plain,(
  sP2('F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0') | $spl24),
  inference(unit_resulting_resolution,[],[f512,f376])).
fof(f1949,plain,(
  sP3('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0','F_0') | $spl24),
  inference(unit_resulting_resolution,[],[f512,f377])).
fof(f1950,plain,(
  sP4('F_0','F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | $spl24),
  inference(unit_resulting_resolution,[],[f512,f378])).
fof(f1938,plain,(
  sP4('F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0') | $spl176),
  inference(cnf_transformation,[],[f1938_D])).
fof(f1938_D,plain,(
  sP4('F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0') <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f1934,plain,(
  sP4('F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0') | $spl20),
  inference(resolution,[],[f508,f378])).
fof(f1940,plain,(
  sP3('F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0') | $spl178),
  inference(cnf_transformation,[],[f1940_D])).
fof(f1940_D,plain,(
  sP3('F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0') <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f1933,plain,(
  sP3('F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0') | $spl20),
  inference(resolution,[],[f508,f377])).
fof(f1942,plain,(
  sP2('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0','F_0') | $spl180),
  inference(cnf_transformation,[],[f1942_D])).
fof(f1942_D,plain,(
  sP2('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0','F_0') <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f1932,plain,(
  sP2('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0','F_0') | $spl20),
  inference(resolution,[],[f508,f376])).
fof(f1928,plain,(
  sP2('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0','F_0') | $spl20),
  inference(unit_resulting_resolution,[],[f508,f376])).
fof(f1929,plain,(
  sP3('F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0') | $spl20),
  inference(unit_resulting_resolution,[],[f508,f377])).
fof(f1930,plain,(
  sP4('F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0') | $spl20),
  inference(unit_resulting_resolution,[],[f508,f378])).
fof(f1870,plain,(
  ~sP5('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl175),
  inference(cnf_transformation,[],[f1870_D])).
fof(f1870_D,plain,(
  ~sP5('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) <=> ~$spl175),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl175])])).
fof(f1866,plain,(
  ~sP5('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl83),
  inference(unit_resulting_resolution,[],[f770,f378])).
fof(f1858,plain,(
  ~sP5('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl173),
  inference(cnf_transformation,[],[f1858_D])).
fof(f1858_D,plain,(
  ~sP5('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) <=> ~$spl173),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl173])])).
fof(f1854,plain,(
  ~sP5('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl77),
  inference(unit_resulting_resolution,[],[f764,f377])).
fof(f1834,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl170),
  inference(cnf_transformation,[],[f1834_D])).
fof(f1834_D,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f1827,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl64),
  inference(resolution,[],[f752,f399])).
fof(f399,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f237])).
fof(f1821,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl64),
  inference(unit_resulting_resolution,[],[f752,f399])).
fof(f1585,plain,(
  sP5('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl168),
  inference(cnf_transformation,[],[f1585_D])).
fof(f1585_D,plain,(
  sP5('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f1456,plain,(
  sP5('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl104),
  inference(unit_resulting_resolution,[],[f263,f792,f396])).
fof(f1457,plain,(
  sP5('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl104),
  inference(unit_resulting_resolution,[],[f262,f792,f396])).
fof(f1575,plain,(
  sP5('T_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl166),
  inference(cnf_transformation,[],[f1575_D])).
fof(f1575_D,plain,(
  sP5('T_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f1471,plain,(
  sP5('T_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl104),
  inference(unit_resulting_resolution,[],[f263,f792,f396])).
fof(f1472,plain,(
  sP5('T_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl104),
  inference(unit_resulting_resolution,[],[f262,f792,f396])).
fof(f1409,plain,(
  ~sP5('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0','T_0') | $spl165),
  inference(cnf_transformation,[],[f1409_D])).
fof(f1409_D,plain,(
  ~sP5('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0','T_0') <=> ~$spl165),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl165])])).
fof(f1405,plain,(
  ~sP5('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0','T_0') | $spl87),
  inference(unit_resulting_resolution,[],[f774,f378])).
fof(f1399,plain,(
  ~sP5('T_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl163),
  inference(cnf_transformation,[],[f1399_D])).
fof(f1399_D,plain,(
  ~sP5('T_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) <=> ~$spl163),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl163])])).
fof(f1395,plain,(
  ~sP5('T_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl85),
  inference(unit_resulting_resolution,[],[f772,f378])).
fof(f1389,plain,(
  ~sP5('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl161),
  inference(cnf_transformation,[],[f1389_D])).
fof(f1389_D,plain,(
  ~sP5('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') <=> ~$spl161),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl161])])).
fof(f1385,plain,(
  ~sP5('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl81),
  inference(unit_resulting_resolution,[],[f768,f377])).
fof(f1379,plain,(
  ~sP5('F_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl159),
  inference(cnf_transformation,[],[f1379_D])).
fof(f1379_D,plain,(
  ~sP5('F_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) <=> ~$spl159),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl159])])).
fof(f1375,plain,(
  ~sP5('F_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl79),
  inference(unit_resulting_resolution,[],[f766,f377])).
fof(f1350,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','T_0') | $spl156),
  inference(cnf_transformation,[],[f1350_D])).
fof(f1350_D,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','T_0') <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f1345,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','T_0') | $spl68),
  inference(resolution,[],[f756,f399])).
fof(f1339,plain,(
  sP6('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','T_0') | $spl68),
  inference(unit_resulting_resolution,[],[f756,f399])).
fof(f1333,plain,(
  sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl154),
  inference(cnf_transformation,[],[f1333_D])).
fof(f1333_D,plain,(
  sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f1328,plain,(
  sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl66),
  inference(resolution,[],[f754,f399])).
fof(f1322,plain,(
  sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl66),
  inference(unit_resulting_resolution,[],[f754,f399])).
fof(f1254,plain,(
  sP10('F_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl152),
  inference(cnf_transformation,[],[f1254_D])).
fof(f1254_D,plain,(
  sP10('F_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f1252,plain,(
  sP10('F_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl38),
  inference(resolution,[],[f726,f435])).
fof(f1246,plain,(
  sP10('F_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl38),
  inference(unit_resulting_resolution,[],[f726,f435])).
fof(f1203,plain,(
  sP4('F_0','T_0','F_0') | $spl146),
  inference(cnf_transformation,[],[f1203_D])).
fof(f1203_D,plain,(
  sP4('F_0','T_0','F_0') <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f1199,plain,(
  sP4('F_0','T_0','F_0') | $spl112),
  inference(resolution,[],[f826,f378])).
fof(f1205,plain,(
  sP3('F_0','T_0','F_0') | $spl148),
  inference(cnf_transformation,[],[f1205_D])).
fof(f1205_D,plain,(
  sP3('F_0','T_0','F_0') <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f1198,plain,(
  sP3('F_0','T_0','F_0') | $spl112),
  inference(resolution,[],[f826,f377])).
fof(f1207,plain,(
  sP2('T_0','F_0','F_0') | $spl150),
  inference(cnf_transformation,[],[f1207_D])).
fof(f1207_D,plain,(
  sP2('T_0','F_0','F_0') <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f1197,plain,(
  sP2('T_0','F_0','F_0') | $spl112),
  inference(resolution,[],[f826,f376])).
fof(f1193,plain,(
  sP2('T_0','F_0','F_0') | $spl112),
  inference(unit_resulting_resolution,[],[f826,f376])).
fof(f1194,plain,(
  sP3('F_0','T_0','F_0') | $spl112),
  inference(unit_resulting_resolution,[],[f826,f377])).
fof(f1195,plain,(
  sP4('F_0','T_0','F_0') | $spl112),
  inference(unit_resulting_resolution,[],[f826,f378])).
fof(f976,plain,(
  sP4('F_0','F_0','T_0') | $spl140),
  inference(cnf_transformation,[],[f976_D])).
fof(f976_D,plain,(
  sP4('F_0','F_0','T_0') <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f973,plain,(
  sP4('F_0','F_0','T_0') | $spl108),
  inference(resolution,[],[f822,f378])).
fof(f978,plain,(
  sP3('T_0','F_0','F_0') | $spl142),
  inference(cnf_transformation,[],[f978_D])).
fof(f978_D,plain,(
  sP3('T_0','F_0','F_0') <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f972,plain,(
  sP3('T_0','F_0','F_0') | $spl108),
  inference(resolution,[],[f822,f377])).
fof(f980,plain,(
  sP2('F_0','T_0','F_0') | $spl144),
  inference(cnf_transformation,[],[f980_D])).
fof(f980_D,plain,(
  sP2('F_0','T_0','F_0') <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f971,plain,(
  sP2('F_0','T_0','F_0') | $spl108),
  inference(resolution,[],[f822,f376])).
fof(f966,plain,(
  sP2('F_0','T_0','F_0') | $spl108),
  inference(unit_resulting_resolution,[],[f822,f376])).
fof(f967,plain,(
  sP3('T_0','F_0','F_0') | $spl108),
  inference(unit_resulting_resolution,[],[f822,f377])).
fof(f968,plain,(
  sP4('F_0','F_0','T_0') | $spl108),
  inference(unit_resulting_resolution,[],[f822,f378])).
fof(f961,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl138),
  inference(cnf_transformation,[],[f961_D])).
fof(f961_D,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f957,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f263,f575,f582,f583,f380])).
fof(f958,plain,(
  sP5('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f262,f575,f582,f583,f380])).
fof(f889,plain,(
  sP4('F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl132),
  inference(cnf_transformation,[],[f889_D])).
fof(f889_D,plain,(
  sP4('F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f886,plain,(
  sP4('F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl10),
  inference(resolution,[],[f483,f378])).
fof(f891,plain,(
  sP3('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0') | $spl134),
  inference(cnf_transformation,[],[f891_D])).
fof(f891_D,plain,(
  sP3('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0') <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f885,plain,(
  sP3('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0') | $spl10),
  inference(resolution,[],[f483,f377])).
fof(f893,plain,(
  sP2('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') | $spl136),
  inference(cnf_transformation,[],[f893_D])).
fof(f893_D,plain,(
  sP2('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f884,plain,(
  sP2('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') | $spl10),
  inference(resolution,[],[f483,f376])).
fof(f880,plain,(
  sP2('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') | $spl10),
  inference(unit_resulting_resolution,[],[f483,f376])).
fof(f881,plain,(
  sP3('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0') | $spl10),
  inference(unit_resulting_resolution,[],[f483,f377])).
fof(f882,plain,(
  sP4('F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl10),
  inference(unit_resulting_resolution,[],[f483,f378])).
fof(f868,plain,(
  sP4('F_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | $spl126),
  inference(cnf_transformation,[],[f868_D])).
fof(f868_D,plain,(
  sP4('F_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f865,plain,(
  sP4('F_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | $spl6),
  inference(resolution,[],[f479,f378])).
fof(f870,plain,(
  sP3('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') | $spl128),
  inference(cnf_transformation,[],[f870_D])).
fof(f870_D,plain,(
  sP3('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f864,plain,(
  sP3('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') | $spl6),
  inference(resolution,[],[f479,f377])).
fof(f872,plain,(
  sP2('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0') | $spl130),
  inference(cnf_transformation,[],[f872_D])).
fof(f872_D,plain,(
  sP2('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0') <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f863,plain,(
  sP2('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0') | $spl6),
  inference(resolution,[],[f479,f376])).
fof(f859,plain,(
  sP2('const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0') | $spl6),
  inference(unit_resulting_resolution,[],[f479,f376])).
fof(f860,plain,(
  sP3('const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') | $spl6),
  inference(unit_resulting_resolution,[],[f479,f377])).
fof(f861,plain,(
  sP4('F_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | $spl6),
  inference(unit_resulting_resolution,[],[f479,f378])).
fof(f827,plain,(
  ~p(s(bool,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) | $spl13),
  inference(unit_resulting_resolution,[],[f484,f278])).
fof(f828,plain,(
  ~p(s(bool,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) | $spl13),
  inference(unit_resulting_resolution,[],[f484,f276])).
fof(f852,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl124),
  inference(cnf_transformation,[],[f852_D])).
fof(f852_D,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f829,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl13),
  inference(unit_resulting_resolution,[],[f484,f269])).
fof(f850,plain,(
  sP5('F_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl122),
  inference(cnf_transformation,[],[f850_D])).
fof(f850_D,plain,(
  sP5('F_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f830,plain,(
  sP5('F_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl13),
  inference(unit_resulting_resolution,[],[f268,f484,f395])).
fof(f848,plain,(
  ~sP5('T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl121),
  inference(cnf_transformation,[],[f848_D])).
fof(f848_D,plain,(
  ~sP5('T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') <=> ~$spl121),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl121])])).
fof(f831,plain,(
  ~sP5('T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl13),
  inference(unit_resulting_resolution,[],[f263,f484,f397])).
fof(f832,plain,(
  ~sP5('T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl13),
  inference(unit_resulting_resolution,[],[f262,f484,f397])).
fof(f833,plain,(
  ~p(s(bool,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) | $spl13),
  inference(unit_resulting_resolution,[],[f263,f484,f342])).
fof(f834,plain,(
  ~p(s(bool,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) | $spl13),
  inference(unit_resulting_resolution,[],[f263,f484,f334])).
fof(f846,plain,(
  sP5('F_0','T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl118),
  inference(cnf_transformation,[],[f846_D])).
fof(f846_D,plain,(
  sP5('F_0','T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f835,plain,(
  sP5('F_0','T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl13),
  inference(unit_resulting_resolution,[],[f268,f484,f395])).
fof(f844,plain,(
  ~sP5('T_0','T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl117),
  inference(cnf_transformation,[],[f844_D])).
fof(f844_D,plain,(
  ~sP5('T_0','T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) <=> ~$spl117),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl117])])).
fof(f836,plain,(
  ~sP5('T_0','T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl13),
  inference(unit_resulting_resolution,[],[f263,f484,f397])).
fof(f837,plain,(
  ~sP5('T_0','T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl13),
  inference(unit_resulting_resolution,[],[f262,f484,f397])).
fof(f838,plain,(
  ~p(s(bool,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) | $spl13),
  inference(unit_resulting_resolution,[],[f263,f484,f342])).
fof(f839,plain,(
  ~p(s(bool,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) | $spl13),
  inference(unit_resulting_resolution,[],[f263,f484,f334])).
fof(f826,plain,(
  sP5('F_0','F_0','T_0') | $spl112),
  inference(cnf_transformation,[],[f826_D])).
fof(f826_D,plain,(
  sP5('F_0','F_0','T_0') <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f800,plain,(
  sP5('F_0','F_0','T_0') | $spl29),
  inference(unit_resulting_resolution,[],[f268,f518,f395])).
fof(f824,plain,(
  ~sP5('T_0','F_0','T_0') | $spl111),
  inference(cnf_transformation,[],[f824_D])).
fof(f824_D,plain,(
  ~sP5('T_0','F_0','T_0') <=> ~$spl111),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl111])])).
fof(f801,plain,(
  ~sP5('T_0','F_0','T_0') | $spl29),
  inference(unit_resulting_resolution,[],[f263,f518,f397])).
fof(f802,plain,(
  ~sP5('T_0','F_0','T_0') | $spl29),
  inference(unit_resulting_resolution,[],[f262,f518,f397])).
fof(f822,plain,(
  sP5('F_0','T_0','F_0') | $spl108),
  inference(cnf_transformation,[],[f822_D])).
fof(f822_D,plain,(
  sP5('F_0','T_0','F_0') <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f810,plain,(
  sP5('F_0','T_0','F_0') | $spl29),
  inference(unit_resulting_resolution,[],[f268,f518,f395])).
fof(f820,plain,(
  ~sP5('T_0','T_0','F_0') | $spl107),
  inference(cnf_transformation,[],[f820_D])).
fof(f820_D,plain,(
  ~sP5('T_0','T_0','F_0') <=> ~$spl107),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl107])])).
fof(f811,plain,(
  ~sP5('T_0','T_0','F_0') | $spl29),
  inference(unit_resulting_resolution,[],[f263,f518,f397])).
fof(f812,plain,(
  ~sP5('T_0','T_0','F_0') | $spl29),
  inference(unit_resulting_resolution,[],[f262,f518,f397])).
fof(f792,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl104),
  inference(cnf_transformation,[],[f792_D])).
fof(f792_D,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f793,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl2),
  inference(global_subsumption,[],[f261,f260,f468,f681,f680,f677,f676,f675,f674,f673,f672,f671,f670,f669,f668,f667,f666,f663,f661,f660,f659,f658,f657,f656,f655,f654,f653,f651,f650,f649,f648,f647,f644,f642,f641,f638,f636,f635,f634,f631,f630,f629,f628,f627,f626,f625,f624,f623,f622,f621,f620,f619,f618,f617,f616,f615,f614,f613,f612,f611,f610,f609,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f594,f593,f592,f591,f590,f587,f586,f585,f584,f583,f582,f581,f580,f579,f578,f577,f576,f575,f574,f573,f572,f571,f570,f569,f568,f567,f566,f565,f564,f563,f562,f561,f560,f559,f558,f557,f556,f555,f554,f553,f552,f551,f550,f549,f548,f547,f546,f545,f541,f540,f538,f537,f536,f535,f533,f532,f530,f529,f523,f522,f521,f520,f682,f683,f684,f689,f695,f696,f697,f700,f707,f710,f711,f713,f714,f715,f716])).
fof(f684,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl2),
  inference(resolution,[],[f468,f278])).
fof(f683,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl2),
  inference(resolution,[],[f468,f276])).
fof(f520,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl2),
  inference(unit_resulting_resolution,[],[f468,f276])).
fof(f521,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl2),
  inference(unit_resulting_resolution,[],[f468,f278])).
fof(f718,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl31),
  inference(cnf_transformation,[],[f718_D])).
fof(f718_D,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) <=> ~$spl31),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl31])])).
fof(f522,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl2),
  inference(unit_resulting_resolution,[],[f468,f279])).
fof(f523,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl2),
  inference(unit_resulting_resolution,[],[f468,f281])).
fof(f529,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl2),
  inference(unit_resulting_resolution,[],[f263,f468,f334])).
fof(f530,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl2),
  inference(unit_resulting_resolution,[],[f262,f468,f334])).
fof(f532,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl2),
  inference(unit_resulting_resolution,[],[f263,f468,f334])).
fof(f533,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl2),
  inference(unit_resulting_resolution,[],[f262,f468,f334])).
fof(f535,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl2),
  inference(unit_resulting_resolution,[],[f268,f468,f337])).
fof(f536,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl2),
  inference(unit_resulting_resolution,[],[f268,f468,f338])).
fof(f537,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl2),
  inference(unit_resulting_resolution,[],[f263,f468,f342])).
fof(f538,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl2),
  inference(unit_resulting_resolution,[],[f262,f468,f342])).
fof(f540,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl2),
  inference(unit_resulting_resolution,[],[f263,f468,f342])).
fof(f541,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl2),
  inference(unit_resulting_resolution,[],[f262,f468,f342])).
fof(f784,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl97),
  inference(cnf_transformation,[],[f784_D])).
fof(f784_D,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') <=> ~$spl97),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl97])])).
fof(f547,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f263,f468,f359])).
fof(f548,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f262,f468,f359])).
fof(f782,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl95),
  inference(cnf_transformation,[],[f782_D])).
fof(f782_D,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) <=> ~$spl95),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl95])])).
fof(f549,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl2),
  inference(unit_resulting_resolution,[],[f468,f468,f359])).
fof(f786,plain,(
  ~sP1('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl99),
  inference(cnf_transformation,[],[f786_D])).
fof(f786_D,plain,(
  ~sP1('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) <=> ~$spl99),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl99])])).
fof(f550,plain,(
  ~sP1('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl2),
  inference(unit_resulting_resolution,[],[f263,f468,f359])).
fof(f551,plain,(
  ~sP1('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl2),
  inference(unit_resulting_resolution,[],[f262,f468,f359])).
fof(f552,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl2),
  inference(unit_resulting_resolution,[],[f468,f468,f359])).
fof(f788,plain,(
  sP1('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') | $spl100),
  inference(cnf_transformation,[],[f788_D])).
fof(f788_D,plain,(
  sP1('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f553,plain,(
  sP1('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f268,f468,f360])).
fof(f790,plain,(
  sP1('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl102),
  inference(cnf_transformation,[],[f790_D])).
fof(f790_D,plain,(
  sP1('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f554,plain,(
  sP1('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl2),
  inference(unit_resulting_resolution,[],[f268,f468,f363])).
fof(f555,plain,(
  sP1('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl2),
  inference(unit_resulting_resolution,[],[f268,f468,f364])).
fof(f556,plain,(
  sP1('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f268,f468,f365])).
fof(f557,plain,(
  ~sP1('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl2),
  inference(unit_resulting_resolution,[],[f263,f468,f366])).
fof(f558,plain,(
  ~sP1('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl2),
  inference(unit_resulting_resolution,[],[f262,f468,f366])).
fof(f559,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl2),
  inference(unit_resulting_resolution,[],[f468,f468,f366])).
fof(f560,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f263,f468,f366])).
fof(f561,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f262,f468,f366])).
fof(f562,plain,(
  ~sP1('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl2),
  inference(unit_resulting_resolution,[],[f468,f468,f366])).
fof(f780,plain,(
  ~sP5('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','F_0') | $spl93),
  inference(cnf_transformation,[],[f780_D])).
fof(f780_D,plain,(
  ~sP5('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','F_0') <=> ~$spl93),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl93])])).
fof(f563,plain,(
  ~sP5('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f263,f268,f468,f379])).
fof(f564,plain,(
  ~sP5('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f262,f268,f468,f379])).
fof(f776,plain,(
  ~sP5('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') | $spl89),
  inference(cnf_transformation,[],[f776_D])).
fof(f776_D,plain,(
  ~sP5('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') <=> ~$spl89),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl89])])).
fof(f565,plain,(
  ~sP5('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f468,f268,f468,f379])).
fof(f778,plain,(
  ~sP5('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') | $spl91),
  inference(cnf_transformation,[],[f778_D])).
fof(f778_D,plain,(
  ~sP5('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') <=> ~$spl91),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl91])])).
fof(f566,plain,(
  ~sP5('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f263,f268,f468,f379])).
fof(f567,plain,(
  ~sP5('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f262,f268,f468,f379])).
fof(f568,plain,(
  ~sP5('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f468,f268,f468,f379])).
fof(f774,plain,(
  ~sP4('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','F_0') | $spl87),
  inference(cnf_transformation,[],[f774_D])).
fof(f774_D,plain,(
  ~sP4('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','F_0') <=> ~$spl87),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl87])])).
fof(f569,plain,(
  ~sP4('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f263,f268,f468,f383])).
fof(f570,plain,(
  ~sP4('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f262,f268,f468,f383])).
fof(f770,plain,(
  ~sP4('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') | $spl83),
  inference(cnf_transformation,[],[f770_D])).
fof(f770_D,plain,(
  ~sP4('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') <=> ~$spl83),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl83])])).
fof(f571,plain,(
  ~sP4('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f468,f268,f468,f383])).
fof(f772,plain,(
  ~sP4('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') | $spl85),
  inference(cnf_transformation,[],[f772_D])).
fof(f772_D,plain,(
  ~sP4('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') <=> ~$spl85),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl85])])).
fof(f572,plain,(
  ~sP4('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f263,f268,f468,f383])).
fof(f573,plain,(
  ~sP4('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f262,f268,f468,f383])).
fof(f574,plain,(
  ~sP4('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f468,f268,f468,f383])).
fof(f768,plain,(
  ~sP3('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','F_0') | $spl81),
  inference(cnf_transformation,[],[f768_D])).
fof(f768_D,plain,(
  ~sP3('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','F_0') <=> ~$spl81),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl81])])).
fof(f576,plain,(
  ~sP3('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f263,f268,f468,f387])).
fof(f577,plain,(
  ~sP3('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f262,f268,f468,f387])).
fof(f764,plain,(
  ~sP3('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') | $spl77),
  inference(cnf_transformation,[],[f764_D])).
fof(f764_D,plain,(
  ~sP3('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') <=> ~$spl77),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl77])])).
fof(f578,plain,(
  ~sP3('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f468,f268,f468,f387])).
fof(f766,plain,(
  ~sP3('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') | $spl79),
  inference(cnf_transformation,[],[f766_D])).
fof(f766_D,plain,(
  ~sP3('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') <=> ~$spl79),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl79])])).
fof(f579,plain,(
  ~sP3('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f263,f268,f468,f387])).
fof(f580,plain,(
  ~sP3('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f262,f268,f468,f387])).
fof(f581,plain,(
  ~sP3('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f468,f268,f468,f387])).
fof(f762,plain,(
  ~sP6('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','F_0') | $spl75),
  inference(cnf_transformation,[],[f762_D])).
fof(f762_D,plain,(
  ~sP6('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','F_0') <=> ~$spl75),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl75])])).
fof(f590,plain,(
  ~sP6('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f263,f268,f468,f406])).
fof(f591,plain,(
  ~sP6('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f262,f268,f468,f406])).
fof(f758,plain,(
  ~sP6('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') | $spl71),
  inference(cnf_transformation,[],[f758_D])).
fof(f758_D,plain,(
  ~sP6('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') <=> ~$spl71),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl71])])).
fof(f592,plain,(
  ~sP6('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f468,f268,f468,f406])).
fof(f760,plain,(
  ~sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') | $spl73),
  inference(cnf_transformation,[],[f760_D])).
fof(f760_D,plain,(
  ~sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') <=> ~$spl73),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl73])])).
fof(f593,plain,(
  ~sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f263,f268,f468,f406])).
fof(f594,plain,(
  ~sP6('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f262,f268,f468,f406])).
fof(f595,plain,(
  ~sP6('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f468,f268,f468,f406])).
fof(f756,plain,(
  sP7('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl68),
  inference(cnf_transformation,[],[f756_D])).
fof(f756_D,plain,(
  sP7('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f597,plain,(
  sP7('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f263,f263,f468,f412])).
fof(f598,plain,(
  sP7('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f262,f263,f468,f412])).
fof(f748,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl60),
  inference(cnf_transformation,[],[f748_D])).
fof(f748_D,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f599,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f468,f263,f468,f412])).
fof(f600,plain,(
  sP7('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f263,f262,f468,f412])).
fof(f601,plain,(
  sP7('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f262,f262,f468,f412])).
fof(f602,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f468,f262,f468,f412])).
fof(f752,plain,(
  sP7('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl64),
  inference(cnf_transformation,[],[f752_D])).
fof(f752_D,plain,(
  sP7('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f603,plain,(
  sP7('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl2),
  inference(unit_resulting_resolution,[],[f263,f468,f468,f412])).
fof(f604,plain,(
  sP7('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl2),
  inference(unit_resulting_resolution,[],[f262,f468,f468,f412])).
fof(f744,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl56),
  inference(cnf_transformation,[],[f744_D])).
fof(f744_D,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f605,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl2),
  inference(unit_resulting_resolution,[],[f468,f468,f468,f412])).
fof(f754,plain,(
  sP7('T_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl66),
  inference(cnf_transformation,[],[f754_D])).
fof(f754_D,plain,(
  sP7('T_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f606,plain,(
  sP7('T_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl2),
  inference(unit_resulting_resolution,[],[f263,f263,f468,f412])).
fof(f607,plain,(
  sP7('T_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl2),
  inference(unit_resulting_resolution,[],[f262,f263,f468,f412])).
fof(f608,plain,(
  sP7('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl2),
  inference(unit_resulting_resolution,[],[f468,f263,f468,f412])).
fof(f609,plain,(
  sP7('T_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl2),
  inference(unit_resulting_resolution,[],[f263,f262,f468,f412])).
fof(f610,plain,(
  sP7('T_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl2),
  inference(unit_resulting_resolution,[],[f262,f262,f468,f412])).
fof(f611,plain,(
  sP7('T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl2),
  inference(unit_resulting_resolution,[],[f468,f262,f468,f412])).
fof(f746,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl58),
  inference(cnf_transformation,[],[f746_D])).
fof(f746_D,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f612,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl2),
  inference(unit_resulting_resolution,[],[f263,f468,f468,f412])).
fof(f613,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl2),
  inference(unit_resulting_resolution,[],[f262,f468,f468,f412])).
fof(f614,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl2),
  inference(unit_resulting_resolution,[],[f468,f468,f468,f412])).
fof(f750,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','T_0') | $spl62),
  inference(cnf_transformation,[],[f750_D])).
fof(f750_D,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','T_0') <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f615,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f263,f263,f468,f412])).
fof(f616,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f262,f263,f468,f412])).
fof(f617,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f468,f263,f468,f412])).
fof(f618,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f263,f262,f468,f412])).
fof(f619,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f262,f262,f468,f412])).
fof(f620,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f468,f262,f468,f412])).
fof(f621,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl2),
  inference(unit_resulting_resolution,[],[f263,f468,f468,f412])).
fof(f622,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl2),
  inference(unit_resulting_resolution,[],[f262,f468,f468,f412])).
fof(f623,plain,(
  sP7('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl2),
  inference(unit_resulting_resolution,[],[f468,f468,f468,f412])).
fof(f742,plain,(
  ~sP7('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl55),
  inference(cnf_transformation,[],[f742_D])).
fof(f742_D,plain,(
  ~sP7('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') <=> ~$spl55),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl55])])).
fof(f626,plain,(
  ~sP7('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f268,f263,f468,f415])).
fof(f627,plain,(
  ~sP7('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f268,f262,f468,f415])).
fof(f738,plain,(
  ~sP7('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl51),
  inference(cnf_transformation,[],[f738_D])).
fof(f738_D,plain,(
  ~sP7('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) <=> ~$spl51),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl51])])).
fof(f628,plain,(
  ~sP7('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl2),
  inference(unit_resulting_resolution,[],[f268,f468,f468,f415])).
fof(f740,plain,(
  ~sP7('F_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl53),
  inference(cnf_transformation,[],[f740_D])).
fof(f740_D,plain,(
  ~sP7('F_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) <=> ~$spl53),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl53])])).
fof(f629,plain,(
  ~sP7('F_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl2),
  inference(unit_resulting_resolution,[],[f268,f263,f468,f415])).
fof(f630,plain,(
  ~sP7('F_0','T_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl2),
  inference(unit_resulting_resolution,[],[f268,f262,f468,f415])).
fof(f631,plain,(
  ~sP7('F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl2),
  inference(unit_resulting_resolution,[],[f268,f468,f468,f415])).
fof(f736,plain,(
  ~sP8('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0','F_0') | $spl49),
  inference(cnf_transformation,[],[f736_D])).
fof(f736_D,plain,(
  ~sP8('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0','F_0') <=> ~$spl49),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl49])])).
fof(f634,plain,(
  ~sP8('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f268,f268,f468,f423])).
fof(f734,plain,(
  ~sP9('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0','F_0') | $spl47),
  inference(cnf_transformation,[],[f734_D])).
fof(f734_D,plain,(
  ~sP9('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0','F_0') <=> ~$spl47),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl47])])).
fof(f649,plain,(
  ~sP9('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f268,f268,f468,f430])).
fof(f732,plain,(
  ~sP10('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0','T_0') | $spl45),
  inference(cnf_transformation,[],[f732_D])).
fof(f732_D,plain,(
  ~sP10('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0','T_0') <=> ~$spl45),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl45])])).
fof(f653,plain,(
  ~sP10('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f268,f263,f468,f440])).
fof(f654,plain,(
  ~sP10('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f268,f262,f468,f440])).
fof(f728,plain,(
  ~sP10('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl41),
  inference(cnf_transformation,[],[f728_D])).
fof(f728_D,plain,(
  ~sP10('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) <=> ~$spl41),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl41])])).
fof(f655,plain,(
  ~sP10('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl2),
  inference(unit_resulting_resolution,[],[f268,f468,f468,f440])).
fof(f730,plain,(
  ~sP10('T_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl43),
  inference(cnf_transformation,[],[f730_D])).
fof(f730_D,plain,(
  ~sP10('T_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) <=> ~$spl43),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl43])])).
fof(f656,plain,(
  ~sP10('T_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl2),
  inference(unit_resulting_resolution,[],[f263,f268,f468,f440])).
fof(f657,plain,(
  ~sP10('T_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl2),
  inference(unit_resulting_resolution,[],[f262,f268,f468,f440])).
fof(f658,plain,(
  ~sP10('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl2),
  inference(unit_resulting_resolution,[],[f468,f268,f468,f440])).
fof(f726,plain,(
  sP11('F_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl38),
  inference(cnf_transformation,[],[f726_D])).
fof(f726_D,plain,(
  sP11('F_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f660,plain,(
  sP11('F_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl2),
  inference(unit_resulting_resolution,[],[f268,f268,f468,f444])).
fof(f724,plain,(
  ~sP11('T_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl37),
  inference(cnf_transformation,[],[f724_D])).
fof(f724_D,plain,(
  ~sP11('T_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) <=> ~$spl37),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl37])])).
fof(f668,plain,(
  ~sP11('T_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl2),
  inference(unit_resulting_resolution,[],[f268,f263,f468,f447])).
fof(f669,plain,(
  ~sP11('T_0','F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl2),
  inference(unit_resulting_resolution,[],[f268,f262,f468,f447])).
fof(f720,plain,(
  ~sP11('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl33),
  inference(cnf_transformation,[],[f720_D])).
fof(f720_D,plain,(
  ~sP11('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f670,plain,(
  ~sP11('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl2),
  inference(unit_resulting_resolution,[],[f268,f468,f468,f447])).
fof(f722,plain,(
  ~sP11('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0','T_0') | $spl35),
  inference(cnf_transformation,[],[f722_D])).
fof(f722_D,plain,(
  ~sP11('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0','T_0') <=> ~$spl35),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl35])])).
fof(f671,plain,(
  ~sP11('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f268,f263,f468,f447])).
fof(f672,plain,(
  ~sP11('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f268,f262,f468,f447])).
fof(f673,plain,(
  ~sP11('const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'F_0','const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl2),
  inference(unit_resulting_resolution,[],[f268,f468,f468,f447])).
fof(f675,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl2),
  inference(unit_resulting_resolution,[],[f268,f268,f263,f468,f457])).
fof(f676,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl2),
  inference(unit_resulting_resolution,[],[f268,f268,f262,f468,f457])).
fof(f677,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl2),
  inference(unit_resulting_resolution,[],[f268,f268,f468,f468,f457])).
fof(f680,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl2),
  inference(unit_resulting_resolution,[],[f268,f268,f468,f468,f457])).
fof(f681,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl2),
  inference(unit_resulting_resolution,[],[f268,f268,f468,f464])).
fof(f484,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl13),
  inference(cnf_transformation,[],[f484_D])).
fof(f484_D,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f516,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | ($spl1 | $spl26)),
  inference(backward_demodulation,[],[f514,f466])).
fof(f518,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl29),
  inference(cnf_transformation,[],[f518_D])).
fof(f518_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f515,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ($spl15 | $spl26)),
  inference(backward_demodulation,[],[f514,f487])).
fof(f514,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | $spl26),
  inference(cnf_transformation,[],[f514_D])).
fof(f514_D,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f502,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | $spl15),
  inference(trivial_inequality_removal,[],[f501])).
fof(f501,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | $spl15),
  inference(superposition,[],[f487,f269])).
fof(f488,plain,(
  p(s(bool,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) | $spl15),
  inference(unit_resulting_resolution,[],[f487,f282])).
fof(f282,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f173])).
fof(f489,plain,(
  p(s(bool,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) | $spl15),
  inference(unit_resulting_resolution,[],[f487,f280])).
fof(f490,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | $spl15),
  inference(unit_resulting_resolution,[],[f487,f269])).
fof(f491,plain,(
  sP5('F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0') | $spl15),
  inference(unit_resulting_resolution,[],[f268,f487,f395])).
fof(f510,plain,(
  ~sP5('T_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0') | $spl23),
  inference(cnf_transformation,[],[f510_D])).
fof(f510_D,plain,(
  ~sP5('T_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0') <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f492,plain,(
  ~sP5('T_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0') | $spl15),
  inference(unit_resulting_resolution,[],[f263,f487,f397])).
fof(f493,plain,(
  ~sP5('T_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'F_0') | $spl15),
  inference(unit_resulting_resolution,[],[f262,f487,f397])).
fof(f494,plain,(
  p(s(bool,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) | $spl15),
  inference(unit_resulting_resolution,[],[f268,f487,f339])).
fof(f339,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f209])).
fof(f495,plain,(
  p(s(bool,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) | $spl15),
  inference(unit_resulting_resolution,[],[f268,f487,f331])).
fof(f508,plain,(
  sP5('F_0','F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | $spl20),
  inference(cnf_transformation,[],[f508_D])).
fof(f508_D,plain,(
  sP5('F_0','F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f496,plain,(
  sP5('F_0','F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | $spl15),
  inference(unit_resulting_resolution,[],[f268,f487,f395])).
fof(f506,plain,(
  ~sP5('T_0','F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | $spl19),
  inference(cnf_transformation,[],[f506_D])).
fof(f506_D,plain,(
  ~sP5('T_0','F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f497,plain,(
  ~sP5('T_0','F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | $spl15),
  inference(unit_resulting_resolution,[],[f263,f487,f397])).
fof(f498,plain,(
  ~sP5('T_0','F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | $spl15),
  inference(unit_resulting_resolution,[],[f262,f487,f397])).
fof(f499,plain,(
  p(s(bool,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) | $spl15),
  inference(unit_resulting_resolution,[],[f268,f487,f339])).
fof(f500,plain,(
  p(s(bool,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14)))) | $spl15),
  inference(unit_resulting_resolution,[],[f268,f487,f331])).
fof(f485,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl12),
  inference(cnf_transformation,[],[f485_D])).
fof(f485_D,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f475,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | s(bool,'T_0') = s(bool,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl1),
  inference(superposition,[],[f466,f269])).
fof(f483,plain,(
  sP5('F_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | $spl10),
  inference(cnf_transformation,[],[f483_D])).
fof(f483_D,plain,(
  sP5('F_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f469,plain,(
  sP5('F_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f268,f466,f395])).
fof(f481,plain,(
  ~sP5('T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | $spl9),
  inference(cnf_transformation,[],[f481_D])).
fof(f481_D,plain,(
  ~sP5('T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f470,plain,(
  ~sP5('T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f263,f466,f397])).
fof(f471,plain,(
  ~sP5('T_0','const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))),'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f262,f466,f397])).
fof(f479,plain,(
  sP5('F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl6),
  inference(cnf_transformation,[],[f479_D])).
fof(f479_D,plain,(
  sP5('F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f472,plain,(
  sP5('F_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl1),
  inference(unit_resulting_resolution,[],[f268,f466,f395])).
fof(f477,plain,(
  ~sP5('T_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl5),
  inference(cnf_transformation,[],[f477_D])).
fof(f477_D,plain,(
  ~sP5('T_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f473,plain,(
  ~sP5('T_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl1),
  inference(unit_resulting_resolution,[],[f263,f466,f397])).
fof(f474,plain,(
  ~sP5('T_0','const.pred_set.countable_1'(s(fun(sK13,bool),sK14)),'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl1),
  inference(unit_resulting_resolution,[],[f262,f466,f397])).
fof(f468,plain,(
  p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) | $spl2),
  inference(cnf_transformation,[],[f468_D])).
fof(f468_D,plain,(
  p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f260,plain,(
  p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))))),
  inference(cnf_transformation,[],[f167])).
fof(f167,plain,(
  p(s(bool,'const.pred_set.INJ_3'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14),s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))) & s(bool,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) != s(bool,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13,sK14,sK15])],[f118])).
fof(f118,plain,(
  ? [X0,X1,X2,X3] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X3),s(fun(X1,bool),X2),s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X1,X0),X3),s(fun(X1,bool),X2)))))) & s(bool,'const.pred_set.countable_1'(s(fun(X1,bool),X2))) != s(bool,'const.pred_set.countable_1'(s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X1,X0),X3),s(fun(X1,bool),X2))))))),
  inference(ennf_transformation,[],[f50])).
fof(f50,plain,(
  ~! [X0,X1,X2,X3] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X3),s(fun(X1,bool),X2),s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X1,X0),X3),s(fun(X1,bool),X2)))))) => s(bool,'const.pred_set.countable_1'(s(fun(X1,bool),X2))) = s(bool,'const.pred_set.countable_1'(s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X1,X0),X3),s(fun(X1,bool),X2))))))),
  inference(flattening,[],[f49])).
fof(f49,plain,(
  ~! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X3),s(fun(X1,bool),X2),s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X1,X0),X3),s(fun(X1,bool),X2)))))) => s(bool,'const.pred_set.countable_1'(s(fun(X1,bool),X2))) = s(bool,'const.pred_set.countable_1'(s(fun(X0,bool),'const.pred_set.IMAGE_2'(s(fun(X1,X0),X3),s(fun(X1,bool),X2))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X7,X6),X9),s(fun(X7,bool),X8),s(fun(X6,bool),'const.pred_set.IMAGE_2'(s(fun(X7,X6),X9),s(fun(X7,bool),X8)))))) => s(bool,'const.pred_set.countable_1'(s(fun(X7,bool),X8))) = s(bool,'const.pred_set.countable_1'(s(fun(X6,bool),'const.pred_set.IMAGE_2'(s(fun(X7,X6),X9),s(fun(X7,bool),X8))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X7,X6),X9),s(fun(X7,bool),X8),s(fun(X6,bool),'const.pred_set.IMAGE_2'(s(fun(X7,X6),X9),s(fun(X7,bool),X8)))))) => s(bool,'const.pred_set.countable_1'(s(fun(X7,bool),X8))) = s(bool,'const.pred_set.countable_1'(s(fun(X6,bool),'const.pred_set.IMAGE_2'(s(fun(X7,X6),X9),s(fun(X7,bool),X8))))))),
  file('thm.pred_set.inj_image_countable_IFF','thm.pred_set.inj_image_countable_IFF')).
fof(f466,plain,(
  s(bool,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) != s(bool,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) | $spl1),
  inference(cnf_transformation,[],[f466_D])).
fof(f466_D,plain,(
  s(bool,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) != s(bool,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f261,plain,(
  s(bool,'const.pred_set.countable_1'(s(fun(sK13,bool),sK14))) != s(bool,'const.pred_set.countable_1'(s(fun(sK12,bool),'const.pred_set.IMAGE_2'(s(fun(sK13,sK12),sK15),s(fun(sK13,bool),sK14)))))),
  inference(cnf_transformation,[],[f167])).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 348
% Passive clauses: 552
% Generated clauses: 5295
% Final active clauses: 311
% Final passive clauses: 269
% Input formulas: 47
% Initial clauses: 199
% 
% Duplicate literals: 94
% Trivial inequalities: 74
% Fw demodulations: 851
% Bw demodulations: 22
% Global subsumptions: 330
% 
% Simple tautologies: 218
% Equational tautologies: 14
% Forward subsumptions: 2137
% Backward subsumptions: 23
% 
% Binary resolution: 760
% Unit resulting resolution: 2011
% Forward superposition: 101
% Backward superposition: 677
% Equality resolution: 12
% 
% Split clauses: 97
% Split components: 198
% SAT solver clauses: 2602
% SAT solver unit clauses: 204
% SAT solver binary clauses: 1539
% 
% Sat splits: 97
% Sat splitting refutations: 440
% 
% Memory used [KB]: 5245
% Time elapsed: 0.234 s
% ------------------------------
% ------------------------------
% Success in time 18.258 s
