Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.extreal
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.extreal
% SZS output start Proof for thm.extreal
fof(f1166,plain,(
  $false),
  inference(subsumption_resolution,[],[f1138,f433])).
fof(f433,plain,(
  ( ! [X2] : (s('type.extreal.extreal',X2) = s('type.extreal.extreal','const.extreal.extreal_min_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X2)))) )),
  inference(superposition,[],[f120,f111])).
fof(f111,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0,X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X2)))),
  inference(flattening,[],[f50])).
fof(f50,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X2)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X10] : ! [X12,X9] : s(X10,X9) = s(X10,'const.bool.COND_3'(s(bool,X12),s(X10,X9),s(X10,X9)))),
  file('thm.extreal.min_infty|split|3','thm.bool.COND_ID')).
fof(f120,plain,(
  ( ! [X0,X1] : (s('type.extreal.extreal','const.extreal.extreal_min_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))),s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) )),
  inference(cnf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0,X1] : s('type.extreal.extreal','const.extreal.extreal_min_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))),s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X6,X11] : s('type.extreal.extreal','const.extreal.extreal_min_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X11))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X11))),s('type.extreal.extreal',X6),s('type.extreal.extreal',X11)))),
  file('thm.extreal.min_infty|split|3','thm.extreal.extreal_min_def')).
fof(f1138,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.extreal_min_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))),
  inference(backward_demodulation,[],[f1125,f86])).
fof(f86,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.extreal_min_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal','const.extreal.NegInf_0')))),
  inference(cnf_transformation,[],[f74])).
fof(f74,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.extreal_min_2'(s('type.extreal.extreal',sK0),s('type.extreal.extreal','const.extreal.NegInf_0')))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f65])).
fof(f65,plain,(
  ? [X0] : s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.extreal_min_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0')))),
  inference(ennf_transformation,[],[f29])).
fof(f29,plain,(
  ~! [X0] : s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_min_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0')))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_min_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal','const.extreal.NegInf_0')))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_min_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal','const.extreal.NegInf_0')))),
  file('thm.extreal.min_infty|split|3','thm.extreal.min_infty|split|3')).
fof(f1125,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK0)),
  inference(trivial_inequality_removal,[],[f1098])).
fof(f1098,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.NegInf_0') | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK0)),
  inference(superposition,[],[f86,f610])).
fof(f610,plain,(
  ( ! [X1] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_min_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X1)) )),
  inference(forward_demodulation,[],[f586,f110])).
fof(f110,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0,X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f48])).
fof(f48,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X10] : ! [X7,X8] : s(X10,X8) = s(X10,'const.bool.COND_3'(s(bool,'F_0'),s(X10,X7),s(X10,X8)))),
  file('thm.extreal.min_infty|split|3','thm.bool.COND_CLAUSES|split|1')).
fof(f586,plain,(
  ( ! [X1] : (s('type.extreal.extreal','const.extreal.extreal_min_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal','const.extreal.NegInf_0'))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,'F_0'),s('type.extreal.extreal',X1),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X1)) )),
  inference(superposition,[],[f120,f332])).
fof(f332,plain,(
  ( ! [X4] : (s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X4),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X4)) )),
  inference(resolution,[],[f105,f100])).
fof(f100,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f39])).
fof(f39,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f38])).
fof(f38,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X9] : (s(bool,'F_0') = s(bool,X9) <=> ~p(s(bool,X9)))),
  file('thm.extreal.min_infty|split|3','thm.bool.EQ_CLAUSES|split|3')).
fof(f105,plain,(
  ( ! [X0] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0)) )),
  inference(cnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0] : ((~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0)) & (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0'))))))),
  inference(nnf_transformation,[],[f43])).
fof(f43,plain,(
  ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0')))) <=> s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X6] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal','const.extreal.NegInf_0')))) <=> s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X6))),
  file('thm.extreal.min_infty|split|3','thm.extreal.le_infty|split|2')).
% SZS output end Proof for thm.extreal
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 46
% Passive clauses: 220
% Generated clauses: 1053
% Final active clauses: 42
% Final passive clauses: 152
% Input formulas: 27
% Initial clauses: 46
% 
% Duplicate literals: 22
% Trivial inequalities: 13
% Fw subsumption resolutions: 65
% Fw demodulations: 147
% Bw demodulations: 3
% 
% Simple tautologies: 29
% Equational tautologies: 7
% Forward subsumptions: 530
% 
% Binary resolution: 59
% Factoring: 1
% Forward superposition: 294
% Backward superposition: 359
% Self superposition: 23
% Equality factoring: 6
% Equality resolution: 7
% 
% Split clauses: 8
% Split components: 16
% SAT solver clauses: 10
% SAT solver unit clauses: 2
% SAT solver binary clauses: 8
% 
% Sat splits: 8
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 10
% TWLsolver calls for satisfiability: 5
% 
% Memory used [KB]: 1279
% Time elapsed: 0.031 s
% ------------------------------
% ------------------------------
% Success in time 0.032 s
