Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.arithmetic
% (77114)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1056
% Passive clauses: 9546
% Generated clauses: 27115
% Final active clauses: 958
% Final passive clauses: 7884
% Input formulas: 40
% Initial clauses: 171
% 
% Duplicate literals: 764
% Trivial inequalities: 44
% Fw subsumption resolutions: 3131
% Fw demodulations: 1484
% Bw demodulations: 96
% 
% Simple tautologies: 292
% Equational tautologies: 58
% Forward subsumptions: 10748
% 
% Binary resolution: 11684
% Factoring: 1586
% Forward superposition: 4530
% Backward superposition: 3166
% Self superposition: 30
% Equality factoring: 32
% Equality resolution: 15
% 
% Split clauses: 868
% Split components: 1753
% SAT solver clauses: 1631
% SAT solver unit clauses: 22
% SAT solver binary clauses: 143
% 
% Sat splits: 868
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 1628
% TWLsolver calls for satisfiability: 701
% 
% Memory used [KB]: 12025
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.arithmetic
% (77135)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 519
% Passive clauses: 3794
% Generated clauses: 15970
% Final active clauses: 217
% Final passive clauses: 762
% Input formulas: 40
% Initial clauses: 171
% 
% Duplicate literals: 256
% Trivial inequalities: 69
% Fw subsumption resolutions: 1079
% Fw demodulations: 1510
% Bw demodulations: 2511
% Condensations: 217
% Global subsumptions: 13
% 
% Simple tautologies: 248
% Equational tautologies: 15
% Forward subsumptions: 4324
% Bw demodulations to eq. taut.: 291
% 
% Binary resolution: 3428
% Factoring: 106
% Forward superposition: 1052
% Backward superposition: 1280
% Self superposition: 22
% Equality factoring: 26
% Equality resolution: 7
% 
% Split clauses: 149
% Split components: 303
% SAT solver clauses: 3300
% SAT solver unit clauses: 190
% SAT solver binary clauses: 93
% SAT solver learnt clauses: 26
% SAT solver learnt literals: 10
% 
% Sat splits: 149
% Sat splitting refutations: 20
% 
% TWLsolver clauses: 3248
% TWLsolver calls for satisfiability: 1792
% 
% Memory used [KB]: 12025
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.arithmetic
% (77152)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 3400
% Passive clauses: 10561
% Generated clauses: 210645
% Final active clauses: 1804
% Final passive clauses: 1912
% Input formulas: 40
% Initial clauses: 171
% 
% Duplicate literals: 8152
% Trivial inequalities: 392
% Fw subsumption resolutions: 8060
% Fw demodulations: 34089
% Bw demodulations: 2854
% Condensations: 11228
% Global subsumptions: 674
% 
% Simple tautologies: 381
% Equational tautologies: 5709
% Forward subsumptions: 128914
% Bw demodulations to eq. taut.: 402
% 
% Unit resulting resolution: 89659
% Forward superposition: 36537
% Backward superposition: 11625
% Self superposition: 214
% Equality factoring: 184
% Equality resolution: 40
% 
% Split clauses: 519
% Split components: 1064
% SAT solver clauses: 14356
% SAT solver unit clauses: 3814
% SAT solver binary clauses: 2802
% 
% Sat splits: 519
% Sat splitting refutations: 2081
% 
% Memory used [KB]: 71128
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.arithmetic
% (77499)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 3108
% Passive clauses: 49836
% Generated clauses: 224037
% Final active clauses: 1991
% Final passive clauses: 33984
% Input formulas: 40
% Initial clauses: 170
% 
% Selected by SInE selection: 39
% SInE iterations: 1
% 
% Duplicate literals: 5124
% Trivial inequalities: 271
% Fw subsumption resolutions: 29049
% Fw demodulations: 28705
% Bw demodulations: 463
% 
% Simple tautologies: 9579
% Equational tautologies: 182
% Forward subsumptions: 95084
% 
% Binary resolution: 118095
% Factoring: 160
% Forward superposition: 28997
% Backward superposition: 6423
% Equality resolution: 9
% 
% Split clauses: 8045
% Split components: 16998
% SAT solver clauses: 8507
% SAT solver unit clauses: 32
% SAT solver binary clauses: 1495
% 
% Sat splits: 8045
% Sat splitting refutations: 125
% 
% Memory used [KB]: 54242
% Time elapsed: 4.0000 s
% ------------------------------
% ------------------------------
% remaining time: 142 next slice time: 6
lrs+1002_2:3_br=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:sas=minisat:stl=30:sac=on:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none:sp=reverse_arity:urr=on_4 on thm.arithmetic
% (77590)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 466
% Passive clauses: 2546
% Generated clauses: 17574
% Final active clauses: 432
% Final passive clauses: 1731
% Input formulas: 40
% Initial clauses: 171
% 
% Duplicate literals: 497
% Trivial inequalities: 247
% Fw subsumption resolutions: 2128
% Fw demodulations: 457
% Bw demodulations: 164
% Global subsumptions: 14
% 
% Simple tautologies: 519
% Equational tautologies: 40
% Forward subsumptions: 10159
% Bw demodulations to eq. taut.: 2
% 
% Unit resulting resolution: 5107
% Factoring: 18
% Forward superposition: 4562
% Backward superposition: 3927
% Self superposition: 80
% Equality factoring: 6
% Equality resolution: 22
% 
% Split clauses: 11
% Split components: 22
% SAT solver clauses: 3171
% SAT solver unit clauses: 873
% SAT solver binary clauses: 337
% 
% Sat splits: 11
% Sat splitting refutations: 598
% 
% Memory used [KB]: 7291
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 135 next slice time: 15
lrs+1011_12_bs=on:bsr=unit_only:cond=on:gs=on:gsaa=from_current:gsssp=full:nwc=1.1:sas=minisat:stl=60:sos=all:sac=on:sdd=off:sser=off:sfr=on:ssfp=100000:ssfq=1.2:ssnc=none:sp=reverse_arity:updr=off_14 on thm.arithmetic
% (77613)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1128
% Passive clauses: 8249
% Generated clauses: 47085
% Final active clauses: 1118
% Final passive clauses: 6887
% Input formulas: 40
% Initial clauses: 171
% 
% Duplicate literals: 1728
% Trivial inequalities: 64
% Fw subsumption resolutions: 3335
% Bw subsumption resolutions: 10
% Fw demodulations: 9524
% Bw demodulations: 9
% Condensations: 3778
% 
% Simple tautologies: 2500
% Forward subsumptions: 17991
% Backward subsumptions: 363
% 
% Binary resolution: 15823
% Forward superposition: 11101
% Backward superposition: 1772
% Equality resolution: 11
% 
% SAT solver clauses: 8252
% SAT solver unit clauses: 317
% SAT solver binary clauses: 1151
% 
% Memory used [KB]: 19317
% Time elapsed: 1.600 s
% ------------------------------
% ------------------------------
% remaining time: 119 next slice time: 9
lrs+10_8:1_bd=preordered:bs=on:ccuc=first:fsr=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:nicw=on:sas=minisat:stl=120:sos=on:sscc=on:sser=off:sfr=on:ssfp=4000:ssfq=1.0:ssnc=none:sp=reverse_arity:urr=on_7 on thm.arithmetic
% (77669)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1076
% Passive clauses: 2992
% Generated clauses: 27591
% Final active clauses: 838
% Final passive clauses: 1147
% Input formulas: 40
% Initial clauses: 171
% 
% Duplicate literals: 349
% Trivial inequalities: 202
% Fw demodulations: 4102
% Bw demodulations: 13
% Global subsumptions: 436
% 
% Simple tautologies: 678
% Equational tautologies: 339
% Forward subsumptions: 14719
% Backward subsumptions: 1131
% 
% Binary resolution: 7511
% Unit resulting resolution: 12262
% Factoring: 6
% Forward superposition: 321
% Backward superposition: 1797
% Self superposition: 24
% Equality factoring: 26
% Equality resolution: 21
% 
% Split clauses: 234
% Split components: 468
% SAT solver clauses: 6404
% SAT solver unit clauses: 29
% SAT solver binary clauses: 2963
% 
% Sat splits: 234
% Sat splitting refutations: 693
% 
% Memory used [KB]: 17910
% Time elapsed: 1.0000 s
% ------------------------------
% ------------------------------
% remaining time: 109 next slice time: 85
lrs+1_5:4_cond=on:fsr=off:fde=none:gs=on:gsem=on:lwlo=on:nm=64:nwc=1:stl=60:sos=all:spl=off_80 on thm.arithmetic
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.arithmetic
% SZS output start Proof for thm.arithmetic
fof(f15136,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f15072])).
fof(f15072,plain,(
  s(bool,'F_0') != s(bool,'F_0')),
  inference(resolution,[],[f14950,f198])).
fof(f198,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f136])).
fof(f136,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f56])).
fof(f56,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X13] : (s(bool,'F_0') = s(bool,X13) <=> ~p(s(bool,X13)))),
  file('thm.arithmetic.arithmetic_3|49|4_prove_ltb','thm.bool.EQ_CLAUSES|split|3')).
fof(f14950,plain,(
  p(s(bool,'F_0'))),
  inference(global_subsumption,[],[f14948,f14715])).
fof(f14715,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK2))))),
  inference(superposition,[],[f227,f14676])).
fof(f14676,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK2)))),
  inference(global_subsumption,[],[f14600,f12637])).
fof(f12637,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK4),s('type.num.num',sK1))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK3),s('type.num.num',sK1))))))),
  inference(resolution,[],[f1136,f899])).
fof(f899,plain,(
  ( ! [X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK4),s('type.num.num',sK2))),s('type.num.num',X1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK3),s('type.num.num',sK1))))))) )),
  inference(resolution,[],[f177,f257])).
fof(f257,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X2)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X2)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X2)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X2)))))),
  inference(flattening,[],[f125])).
fof(f125,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X2))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X2)))))),
  inference(ennf_transformation,[],[f88])).
fof(f88,plain,(
  ! [X0,X1,X2] : ((p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) & p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X2))))) => p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X2)))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X10,X11,X12] : ((p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X10),s('type.num.num',X11)))) & p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num',X12))))) => p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X10),s('type.num.num',X12)))))),
  file('thm.arithmetic.arithmetic_3|49|4_prove_ltb','thm.arithmetic.LESS_EQ_LESS_TRANS')).
fof(f177,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK4),s('type.num.num',sK2))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK3),s('type.num.num',sK1))))))),
  inference(cnf_transformation,[],[f130])).
fof(f130,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK4),s('type.num.num',sK3)))) & p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK2),s('type.num.num',sK1)))) & ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK4),s('type.num.num',sK2))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK3),s('type.num.num',sK1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK1,sK2,sK3,sK4])],[f101])).
fof(f101,plain,(
  ? [X0,X1,X2,X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num',X2)))) & p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))) & ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X3),s('type.num.num',X1))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num',X0)))))))),
  inference(flattening,[],[f100])).
fof(f100,plain,(
  ? [X0,X1,X2,X3] : ((p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num',X2)))) & p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))))) & ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X3),s('type.num.num',X1))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num',X0)))))))),
  inference(ennf_transformation,[],[f42])).
fof(f42,plain,(
  ~! [X0,X1,X2,X3] : ((p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num',X2)))) & p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))))) => p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X3),s('type.num.num',X1))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num',X0)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8,X9] : ((p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num',X8)))) & p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num',X6))))) => p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X9),s('type.num.num',X7))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X8),s('type.num.num',X6)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8,X9] : ((p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num',X8)))) & p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num',X6))))) => p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X9),s('type.num.num',X7))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X8),s('type.num.num',X6)))))))),
  file('thm.arithmetic.arithmetic_3|49|4_prove_ltb','thm.arithmetic.arithmetic_3|49|4_prove_ltb')).
fof(f1136,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',sK2))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',sK1))))))) )),
  inference(resolution,[],[f751,f263])).
fof(f263,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2))))) )),
  inference(cnf_transformation,[],[f162])).
fof(f162,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X2)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2))))) & ((s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) & ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X2))))))))),
  inference(flattening,[],[f161])).
fof(f161,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X2)))))) | (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2)))))) & ((s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) & ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X2))))))))),
  inference(nnf_transformation,[],[f90])).
fof(f90,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X2)))))) <=> (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2))))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X10,X11,X12] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X10),s('type.num.num',X11))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X10),s('type.num.num',X12)))))) <=> (s('type.num.num','const.num.|48|_0') = s('type.num.num',X10) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X11),s('type.num.num',X12))))))),
  file('thm.arithmetic.arithmetic_3|49|4_prove_ltb','thm.arithmetic.LE_MULT_LCANCEL')).
fof(f751,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK2),s('type.num.num',sK1))))),
  inference(resolution,[],[f176,f227])).
fof(f176,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK2),s('type.num.num',sK1))))),
  inference(cnf_transformation,[],[f130])).
fof(f14600,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK2))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK4),s('type.num.num',sK1))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK3),s('type.num.num',sK1))))))),
  inference(resolution,[],[f2534,f650])).
fof(f650,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK4),s('type.num.num',X0))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK3),s('type.num.num',X0))))))) )),
  inference(resolution,[],[f175,f260])).
fof(f260,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num',X1)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X2)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f160])).
fof(f160,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num',X1)))))) | (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) & p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X2)))))) & (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X2)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num',X1))))))))),
  inference(flattening,[],[f159])).
fof(f159,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num',X1)))))) | (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) & p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X2)))))) & ((~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X2))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num',X1))))))))),
  inference(nnf_transformation,[],[f89])).
fof(f89,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num',X1)))))) <=> (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) & p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X2))))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X10,X11,X12] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X10),s('type.num.num',X11))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X12),s('type.num.num',X11)))))) <=> (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X11)))) & p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X10),s('type.num.num',X12))))))),
  file('thm.arithmetic.arithmetic_3|49|4_prove_ltb','thm.arithmetic.LT_MULT_RCANCEL')).
fof(f175,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK4),s('type.num.num',sK3))))),
  inference(cnf_transformation,[],[f130])).
fof(f2534,plain,(
  ( ! [X6] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num',sK1)))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num',sK2)))) )),
  inference(resolution,[],[f750,f197])).
fof(f197,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f135])).
fof(f135,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f54])).
fof(f54,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X13] : (s(bool,'F_0') = s(bool,X13) <=> ~p(s(bool,X13)))),
  file('thm.arithmetic.arithmetic_3|49|4_prove_ltb','thm.bool.EQ_CLAUSES|split|2')).
fof(f750,plain,(
  ( ! [X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',sK2)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',sK1))))) )),
  inference(resolution,[],[f176,f257])).
fof(f227,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f151])).
fof(f151,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num',X0) = s('type.num.num',X1)) & ((~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) & s('type.num.num',X0) != s('type.num.num',X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(flattening,[],[f150])).
fof(f150,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num',X0) = s('type.num.num',X1))) & ((~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) & s('type.num.num',X0) != s('type.num.num',X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(nnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num',X0) = s('type.num.num',X1)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X10,X11] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X10),s('type.num.num',X11)))) <=> (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X10),s('type.num.num',X11)))) | s('type.num.num',X10) = s('type.num.num',X11)))),
  file('thm.arithmetic.arithmetic_3|49|4_prove_ltb','thm.arithmetic.LESS_OR_EQ')).
fof(f14948,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK2))))),
  inference(trivial_inequality_removal,[],[f14947])).
fof(f14947,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.|48|_0') | p(s(bool,'F_0')) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK2))))),
  inference(superposition,[],[f14716,f215])).
fof(f215,plain,(
  ( ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X0] : ((s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))) & (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0)))),
  inference(nnf_transformation,[],[f65])).
fof(f65,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) <=> p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))))),
  inference(flattening,[],[f64])).
fof(f64,plain,(
  ! [X0] : (~s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) <=> p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X11] : (~s('type.num.num','const.num.|48|_0') = s('type.num.num',X11) <=> p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X11)))))),
  file('thm.arithmetic.arithmetic_3|49|4_prove_ltb','thm.arithmetic.NOT_ZERO_LT_ZERO')).
fof(f14716,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK2) | p(s(bool,'F_0'))),
  inference(superposition,[],[f228,f14676])).
fof(f228,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num',X0) != s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f151])).
% SZS output end Proof for thm.arithmetic
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 234
% Passive clauses: 2555
% Generated clauses: 14472
% Final active clauses: 214
% Final passive clauses: 2188
% Input formulas: 40
% Initial clauses: 302
% 
% Duplicate literals: 758
% Trivial inequalities: 111
% Fw demodulations: 1539
% Bw demodulations: 207
% Condensations: 1154
% Global subsumptions: 45
% 
% Simple tautologies: 907
% Equational tautologies: 8
% Forward subsumptions: 7815
% Bw demodulations to eq. taut.: 7
% 
% Binary resolution: 5405
% Factoring: 7
% Forward superposition: 4116
% Backward superposition: 1351
% Equality resolution: 14
% 
% SAT solver clauses: 2758
% SAT solver unit clauses: 158
% SAT solver binary clauses: 266
% SAT solver learnt clauses: 92
% SAT solver learnt literals: 194
% 
% TWLsolver clauses: 2600
% TWLsolver calls for satisfiability: 66951
% 
% Memory used [KB]: 9210
% Time elapsed: 0.501 s
% ------------------------------
% ------------------------------
% Success in time 19.51 s
