Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.ringNorm
% (77355)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 68
% Passive clauses: 1945
% Generated clauses: 11554
% Final active clauses: 66
% Final passive clauses: 1820
% Input formulas: 25
% Initial clauses: 26
% 
% Duplicate literals: 719
% Trivial inequalities: 105
% Fw subsumption resolutions: 939
% Fw demodulations: 473
% 
% Simple tautologies: 229
% Equational tautologies: 95
% Forward subsumptions: 5864
% 
% Binary resolution: 25
% Forward superposition: 4870
% Backward superposition: 3923
% Self superposition: 352
% Equality factoring: 262
% Equality resolution: 6
% 
% Split clauses: 405
% Split components: 820
% SAT solver clauses: 1243
% SAT solver unit clauses: 412
% SAT solver binary clauses: 393
% 
% Sat splits: 405
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1239
% TWLsolver calls for satisfiability: 44
% 
% Memory used [KB]: 12025
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.ringNorm
% (77372)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 37
% Passive clauses: 229
% Generated clauses: 6139
% Final active clauses: 34
% Final passive clauses: 190
% Input formulas: 25
% Initial clauses: 26
% 
% Duplicate literals: 908
% Trivial inequalities: 58
% Fw subsumption resolutions: 141
% Fw demodulations: 36
% Bw demodulations: 1
% Condensations: 1757
% 
% Simple tautologies: 119
% Equational tautologies: 77
% Forward subsumptions: 2108
% 
% Binary resolution: 16
% Forward superposition: 1713
% Backward superposition: 1511
% Self superposition: 268
% Equality factoring: 222
% Equality resolution: 5
% 
% Split clauses: 27
% Split components: 54
% SAT solver clauses: 601
% SAT solver unit clauses: 46
% SAT solver binary clauses: 66
% 
% Sat splits: 27
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 601
% TWLsolver calls for satisfiability: 57
% 
% Memory used [KB]: 5884
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.ringNorm
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.ringNorm
% SZS output start Proof for thm.ringNorm
fof(f15591,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f15548])).
fof(f15548,plain,(
  s('type.canonical.canonical_sum'(sK0),sK2) != s('type.canonical.canonical_sum'(sK0),sK2)),
  inference(superposition,[],[f77,f15401])).
fof(f15401,plain,(
  ( ! [X6,X4,X5] : (s('type.canonical.canonical_sum'(X4),X6) = s('type.canonical.canonical_sum'(X4),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(X4),X5),s('type.canonical.canonical_sum'(X4),X6),s('type.canonical.canonical_sum'(X4),'const.canonical.Nil_monom_0')))) )),
  inference(superposition,[],[f10553,f83])).
fof(f83,plain,(
  ( ! [X2,X0,X1] : (s('type.canonical.canonical_sum'(X0),X2) = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_merge_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))) )),
  inference(cnf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0,X1,X2] : s('type.canonical.canonical_sum'(X0),X2) = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_merge_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))),
  inference(flattening,[],[f34])).
fof(f34,plain,(
  ! [X0] : ! [X1,X2] : s('type.canonical.canonical_sum'(X0),X2) = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_merge_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6] : ! [X13,X8] : s('type.canonical.canonical_sum'(X6),X8) = s('type.canonical.canonical_sum'(X6),'const.canonical.canonical_sum_merge_3'(s('type.semi_ring.semi_ring'(X6),X13),s('type.canonical.canonical_sum'(X6),X8),s('type.canonical.canonical_sum'(X6),'const.canonical.Nil_monom_0')))),
  file('thm.ringNorm.canonical_sum_merge_def|split|4','thm.canonical.canonical_sum_merge_def|split|4')).
fof(f10553,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_merge_3'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3)))) )),
  inference(forward_demodulation,[],[f10504,f91])).
fof(f91,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_merge_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))) )),
  inference(cnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0,X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_merge_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))),
  inference(flattening,[],[f48])).
fof(f48,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_merge_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X6] : ! [X27,X28,X29] : s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(X6),X27),s('type.canonical.canonical_sum'(X6),X28),s('type.canonical.canonical_sum'(X6),X29))) = s('type.canonical.canonical_sum'(X6),happ(s(fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6)),happ(s(fun('type.canonical.canonical_sum'(X6),fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6))),'const.ringNorm.r_canonical_sum_merge_1'(s('type.ring.ring'(X6),X27))),s('type.canonical.canonical_sum'(X6),X28))),s('type.canonical.canonical_sum'(X6),X29)))),
  file('thm.ringNorm.canonical_sum_merge_def|split|4','ARITY_1_3_const.ringNorm.r_canonical_sum_merge')).
fof(f10504,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_merge_3'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_merge_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))) )),
  inference(superposition,[],[f92,f81])).
fof(f81,plain,(
  ( ! [X0,X1] : (s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_merge_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_merge_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f31])).
fof(f31,plain,(
  ! [X0,X1] : s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_merge_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_merge_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f30])).
fof(f30,plain,(
  ! [X0] : ! [X1] : s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_merge_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_merge_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X6] : ! [X7] : s(fun('type.canonical.canonical_sum'(X6),fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6))),'const.ringNorm.r_canonical_sum_merge_1'(s('type.ring.ring'(X6),X7))) = s(fun('type.canonical.canonical_sum'(X6),fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6))),'const.canonical.canonical_sum_merge_1'(s('type.semi_ring.semi_ring'(X6),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X6),X7)))))),
  file('thm.ringNorm.canonical_sum_merge_def|split|4','thm.ringNorm.r_canonical_sum_merge_def')).
fof(f92,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_merge_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_merge_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))) )),
  inference(cnf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0,X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_merge_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_merge_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))),
  inference(flattening,[],[f50])).
fof(f50,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_merge_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_merge_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X6] : ! [X27,X28,X29] : s('type.canonical.canonical_sum'(X6),'const.canonical.canonical_sum_merge_3'(s('type.semi_ring.semi_ring'(X6),X27),s('type.canonical.canonical_sum'(X6),X28),s('type.canonical.canonical_sum'(X6),X29))) = s('type.canonical.canonical_sum'(X6),happ(s(fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6)),happ(s(fun('type.canonical.canonical_sum'(X6),fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6))),'const.canonical.canonical_sum_merge_1'(s('type.semi_ring.semi_ring'(X6),X27))),s('type.canonical.canonical_sum'(X6),X28))),s('type.canonical.canonical_sum'(X6),X29)))),
  file('thm.ringNorm.canonical_sum_merge_def|split|4','ARITY_1_3_const.canonical.canonical_sum_merge')).
fof(f77,plain,(
  s('type.canonical.canonical_sum'(sK0),sK2) != s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(sK0),sK1),s('type.canonical.canonical_sum'(sK0),sK2),s('type.canonical.canonical_sum'(sK0),'const.canonical.Nil_monom_0')))),
  inference(cnf_transformation,[],[f74])).
fof(f74,plain,(
  s('type.canonical.canonical_sum'(sK0),sK2) != s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(sK0),sK1),s('type.canonical.canonical_sum'(sK0),sK2),s('type.canonical.canonical_sum'(sK0),'const.canonical.Nil_monom_0')))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f72])).
fof(f72,plain,(
  ? [X0,X1,X2] : s('type.canonical.canonical_sum'(X0),X2) != s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))),
  inference(ennf_transformation,[],[f28])).
fof(f28,plain,(
  ~! [X0,X1,X2] : s('type.canonical.canonical_sum'(X0),X2) = s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))),
  inference(flattening,[],[f27])).
fof(f27,plain,(
  ~! [X0] : ! [X1,X2] : s('type.canonical.canonical_sum'(X0),X2) = s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : s('type.canonical.canonical_sum'(X6),X8) = s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(X6),X7),s('type.canonical.canonical_sum'(X6),X8),s('type.canonical.canonical_sum'(X6),'const.canonical.Nil_monom_0')))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : s('type.canonical.canonical_sum'(X6),X8) = s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(X6),X7),s('type.canonical.canonical_sum'(X6),X8),s('type.canonical.canonical_sum'(X6),'const.canonical.Nil_monom_0')))),
  file('thm.ringNorm.canonical_sum_merge_def|split|4','thm.ringNorm.canonical_sum_merge_def|split|4')).
% SZS output end Proof for thm.ringNorm
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 66
% Passive clauses: 732
% Generated clauses: 15383
% Final active clauses: 64
% Final passive clauses: 595
% Input formulas: 25
% Initial clauses: 26
% 
% Duplicate literals: 1854
% Trivial inequalities: 93
% Fw subsumption resolutions: 502
% Fw demodulations: 480
% Bw demodulations: 67
% Condensations: 3577
% 
% Simple tautologies: 205
% Equational tautologies: 138
% Forward subsumptions: 7908
% Bw demodulations to eq. taut.: 4
% 
% Unit resulting resolution: 81
% Forward superposition: 4389
% Backward superposition: 3835
% Self superposition: 345
% Equality factoring: 254
% Equality resolution: 6
% 
% Split clauses: 24
% Split components: 48
% SAT solver clauses: 783
% SAT solver unit clauses: 59
% SAT solver binary clauses: 61
% 
% Sat splits: 24
% Sat splitting refutations: 1
% 
% Memory used [KB]: 15351
% Time elapsed: 2.055 s
% ------------------------------
% ------------------------------
% Success in time 3.385 s
