Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.canonical
% (99196)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 115
% Passive clauses: 2964
% Generated clauses: 20694
% Final active clauses: 115
% Final passive clauses: 2772
% Input formulas: 69
% Initial clauses: 124
% 
% Duplicate literals: 425
% Trivial inequalities: 235
% Fw subsumption resolutions: 2776
% Fw demodulations: 3085
% 
% Simple tautologies: 268
% Equational tautologies: 60
% Forward subsumptions: 10111
% 
% Binary resolution: 210
% Factoring: 1
% Forward superposition: 6903
% Backward superposition: 6401
% Self superposition: 300
% Equality factoring: 139
% Equality resolution: 36
% 
% Split clauses: 360
% Split components: 724
% SAT solver clauses: 569
% SAT solver unit clauses: 15
% SAT solver binary clauses: 115
% 
% Sat splits: 360
% 
% TWLsolver clauses: 569
% TWLsolver calls for satisfiability: 85
% 
% Memory used [KB]: 14456
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.canonical
% (99215)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 94
% Passive clauses: 692
% Generated clauses: 12242
% Final active clauses: 93
% Final passive clauses: 567
% Input formulas: 69
% Initial clauses: 124
% 
% Duplicate literals: 455
% Trivial inequalities: 150
% Fw subsumption resolutions: 959
% Fw demodulations: 1217
% Bw demodulations: 2
% Condensations: 1574
% Global subsumptions: 12
% 
% Simple tautologies: 183
% Equational tautologies: 57
% Forward subsumptions: 5678
% Bw demodulations to eq. taut.: 30
% 
% Binary resolution: 79
% Factoring: 1
% Forward superposition: 3111
% Backward superposition: 3530
% Self superposition: 199
% Equality factoring: 116
% Equality resolution: 30
% 
% Split clauses: 44
% Split components: 88
% SAT solver clauses: 1415
% SAT solver unit clauses: 158
% SAT solver binary clauses: 51
% SAT solver learnt clauses: 24
% 
% Sat splits: 44
% 
% TWLsolver clauses: 1367
% TWLsolver calls for satisfiability: 369
% 
% Memory used [KB]: 11769
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.canonical
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.canonical
% SZS output start Proof for thm.canonical
fof(f99650,plain,(
  $false),
  inference(subsumption_resolution,[],[f99649,f322])).
fof(f322,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.spolynom'(X0),'const.canonical.SPconst_1'(s(X0,X3))) != s('type.canonical.spolynom'(X0),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(X0),X2),s('type.canonical.spolynom'(X0),X1)))) )),
  inference(cnf_transformation,[],[f139])).
fof(f139,plain,(
  ! [X0,X1,X2,X3] : s('type.canonical.spolynom'(X0),'const.canonical.SPconst_1'(s(X0,X3))) != s('type.canonical.spolynom'(X0),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(X0),X2),s('type.canonical.spolynom'(X0),X1)))),
  inference(flattening,[],[f138])).
fof(f138,plain,(
  ! [X0] : ! [X1,X2,X3] : ~s('type.canonical.spolynom'(X0),'const.canonical.SPconst_1'(s(X0,X3))) = s('type.canonical.spolynom'(X0),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(X0),X2),s('type.canonical.spolynom'(X0),X1)))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X6] : ! [X7,X8,X9] : ~s('type.canonical.spolynom'(X6),'const.canonical.SPconst_1'(s(X6,X9))) = s('type.canonical.spolynom'(X6),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(X6),X8),s('type.canonical.spolynom'(X6),X7)))),
  file('thm.canonical.spolynom_distinct|split|4','thm.canonical.spolynom_distinct|split|3')).
fof(f99649,plain,(
  s('type.canonical.spolynom'(sK0),'const.canonical.SPconst_1'(s(sK0,sK3))) = s('type.canonical.spolynom'(sK0),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(sK0),sK2),s('type.canonical.spolynom'(sK0),sK1)))),
  inference(forward_demodulation,[],[f98478,f99648])).
fof(f99648,plain,(
  ( ! [X138] : (s('type.canonical.spolynom'(sK0),'const.canonical.SPconst_1'(s(sK0,sK3))) = s('type.canonical.spolynom'(sK0),happ(s(fun(sK0,'type.canonical.spolynom'(sK0)),X138),s(sK0,sK3)))) )),
  inference(forward_demodulation,[],[f98477,f244])).
fof(f244,plain,(
  s('type.canonical.spolynom'(sK0),'const.canonical.SPconst_1'(s(sK0,sK3))) = s('type.canonical.spolynom'(sK0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(sK0),sK2),s('type.canonical.spolynom'(sK0),sK1)))),
  inference(cnf_transformation,[],[f201])).
fof(f201,plain,(
  s('type.canonical.spolynom'(sK0),'const.canonical.SPconst_1'(s(sK0,sK3))) = s('type.canonical.spolynom'(sK0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(sK0),sK2),s('type.canonical.spolynom'(sK0),sK1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f184])).
fof(f184,plain,(
  ? [X0,X1,X2,X3] : s('type.canonical.spolynom'(X0),'const.canonical.SPconst_1'(s(X0,X3))) = s('type.canonical.spolynom'(X0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(X0),X2),s('type.canonical.spolynom'(X0),X1)))),
  inference(ennf_transformation,[],[f72])).
fof(f72,plain,(
  ~! [X0,X1,X2,X3] : s('type.canonical.spolynom'(X0),'const.canonical.SPconst_1'(s(X0,X3))) != s('type.canonical.spolynom'(X0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(X0),X2),s('type.canonical.spolynom'(X0),X1)))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ~! [X0] : ! [X1,X2,X3] : ~s('type.canonical.spolynom'(X0),'const.canonical.SPconst_1'(s(X0,X3))) = s('type.canonical.spolynom'(X0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(X0),X2),s('type.canonical.spolynom'(X0),X1)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : ~s('type.canonical.spolynom'(X6),'const.canonical.SPconst_1'(s(X6,X9))) = s('type.canonical.spolynom'(X6),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(X6),X8),s('type.canonical.spolynom'(X6),X7)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : ~s('type.canonical.spolynom'(X6),'const.canonical.SPconst_1'(s(X6,X9))) = s('type.canonical.spolynom'(X6),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(X6),X8),s('type.canonical.spolynom'(X6),X7)))),
  file('thm.canonical.spolynom_distinct|split|4','thm.canonical.spolynom_distinct|split|4')).
fof(f98477,plain,(
  ( ! [X138] : (s('type.canonical.spolynom'(sK0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(sK0),sK2),s('type.canonical.spolynom'(sK0),sK1))) = s('type.canonical.spolynom'(sK0),happ(s(fun(sK0,'type.canonical.spolynom'(sK0)),X138),s(sK0,sK3)))) )),
  inference(superposition,[],[f38461,f306])).
fof(f306,plain,(
  ( ! [X2,X0,X1] : (s('type.canonical.spolynom'(X0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.spolynom'(X0),happ(s(fun('type.canonical.spolynom'(X0),'type.canonical.spolynom'(X0)),happ(s(fun('type.canonical.spolynom'(X0),fun('type.canonical.spolynom'(X0),'type.canonical.spolynom'(X0))),'const.canonical.SPmult_0'),s('type.canonical.spolynom'(X0),X1))),s('type.canonical.spolynom'(X0),X2)))) )),
  inference(cnf_transformation,[],[f115])).
fof(f115,plain,(
  ! [X0,X1,X2] : s('type.canonical.spolynom'(X0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.spolynom'(X0),happ(s(fun('type.canonical.spolynom'(X0),'type.canonical.spolynom'(X0)),happ(s(fun('type.canonical.spolynom'(X0),fun('type.canonical.spolynom'(X0),'type.canonical.spolynom'(X0))),'const.canonical.SPmult_0'),s('type.canonical.spolynom'(X0),X1))),s('type.canonical.spolynom'(X0),X2)))),
  inference(flattening,[],[f114])).
fof(f114,plain,(
  ! [X0] : ! [X1,X2] : s('type.canonical.spolynom'(X0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.spolynom'(X0),happ(s(fun('type.canonical.spolynom'(X0),'type.canonical.spolynom'(X0)),happ(s(fun('type.canonical.spolynom'(X0),fun('type.canonical.spolynom'(X0),'type.canonical.spolynom'(X0))),'const.canonical.SPmult_0'),s('type.canonical.spolynom'(X0),X1))),s('type.canonical.spolynom'(X0),X2)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6] : ! [X10,X11] : s('type.canonical.spolynom'(X6),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(X6),X10),s('type.canonical.spolynom'(X6),X11))) = s('type.canonical.spolynom'(X6),happ(s(fun('type.canonical.spolynom'(X6),'type.canonical.spolynom'(X6)),happ(s(fun('type.canonical.spolynom'(X6),fun('type.canonical.spolynom'(X6),'type.canonical.spolynom'(X6))),'const.canonical.SPmult_0'),s('type.canonical.spolynom'(X6),X10))),s('type.canonical.spolynom'(X6),X11)))),
  file('thm.canonical.spolynom_distinct|split|4','ARITY_0_2_const.canonical.SPmult')).
fof(f38461,plain,(
  ( ! [X7,X5,X9] : (s(X5,happ(s(fun(sK0,X5),X7),s(sK0,sK3))) = s(X5,happ(s(fun('type.canonical.spolynom'(sK0),X5),happ(s(fun('type.canonical.spolynom'(sK0),fun('type.canonical.spolynom'(sK0),X5)),X9),s('type.canonical.spolynom'(sK0),sK2))),s('type.canonical.spolynom'(sK0),sK1)))) )),
  inference(superposition,[],[f15432,f342])).
fof(f342,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),X4),s(X1,X2))) = s(X0,'const.canonical.spolynom_CASE_5'(s('type.canonical.spolynom'(X1),'const.canonical.SPconst_1'(s(X1,X2))),s(fun('type.quote.index',X0),X3),s(fun(X1,X0),X4),s(fun('type.canonical.spolynom'(X1),fun('type.canonical.spolynom'(X1),X0)),X5),s(fun('type.canonical.spolynom'(X1),fun('type.canonical.spolynom'(X1),X0)),X6)))) )),
  inference(cnf_transformation,[],[f163])).
fof(f163,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(X0,happ(s(fun(X1,X0),X4),s(X1,X2))) = s(X0,'const.canonical.spolynom_CASE_5'(s('type.canonical.spolynom'(X1),'const.canonical.SPconst_1'(s(X1,X2))),s(fun('type.quote.index',X0),X3),s(fun(X1,X0),X4),s(fun('type.canonical.spolynom'(X1),fun('type.canonical.spolynom'(X1),X0)),X5),s(fun('type.canonical.spolynom'(X1),fun('type.canonical.spolynom'(X1),X0)),X6)))),
  inference(flattening,[],[f162])).
fof(f162,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6] : s(X0,happ(s(fun(X1,X0),X4),s(X1,X2))) = s(X0,'const.canonical.spolynom_CASE_5'(s('type.canonical.spolynom'(X1),'const.canonical.SPconst_1'(s(X1,X2))),s(fun('type.quote.index',X0),X3),s(fun(X1,X0),X4),s(fun('type.canonical.spolynom'(X1),fun('type.canonical.spolynom'(X1),X0)),X5),s(fun('type.canonical.spolynom'(X1),fun('type.canonical.spolynom'(X1),X0)),X6)))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X29,X6] : ! [X9,X30,X31,X32,X33] : s(X29,happ(s(fun(X6,X29),X31),s(X6,X9))) = s(X29,'const.canonical.spolynom_CASE_5'(s('type.canonical.spolynom'(X6),'const.canonical.SPconst_1'(s(X6,X9))),s(fun('type.quote.index',X29),X30),s(fun(X6,X29),X31),s(fun('type.canonical.spolynom'(X6),fun('type.canonical.spolynom'(X6),X29)),X32),s(fun('type.canonical.spolynom'(X6),fun('type.canonical.spolynom'(X6),X29)),X33)))),
  file('thm.canonical.spolynom_distinct|split|4','thm.canonical.spolynom_case_def|split|1')).
fof(f15432,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,happ(s(fun('type.canonical.spolynom'(sK0),X0),happ(s(fun('type.canonical.spolynom'(sK0),fun('type.canonical.spolynom'(sK0),X0)),X1),s('type.canonical.spolynom'(sK0),sK2))),s('type.canonical.spolynom'(sK0),sK1))) = s(X0,'const.canonical.spolynom_CASE_5'(s('type.canonical.spolynom'(sK0),'const.canonical.SPconst_1'(s(sK0,sK3))),s(fun('type.quote.index',X0),X2),s(fun(sK0,X0),X3),s(fun('type.canonical.spolynom'(sK0),fun('type.canonical.spolynom'(sK0),X0)),X4),s(fun('type.canonical.spolynom'(sK0),fun('type.canonical.spolynom'(sK0),X0)),X1)))) )),
  inference(superposition,[],[f345,f244])).
fof(f345,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun('type.canonical.spolynom'(X1),X0),happ(s(fun('type.canonical.spolynom'(X1),fun('type.canonical.spolynom'(X1),X0)),X7),s('type.canonical.spolynom'(X1),X2))),s('type.canonical.spolynom'(X1),X3))) = s(X0,'const.canonical.spolynom_CASE_5'(s('type.canonical.spolynom'(X1),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(X1),X2),s('type.canonical.spolynom'(X1),X3))),s(fun('type.quote.index',X0),X4),s(fun(X1,X0),X5),s(fun('type.canonical.spolynom'(X1),fun('type.canonical.spolynom'(X1),X0)),X6),s(fun('type.canonical.spolynom'(X1),fun('type.canonical.spolynom'(X1),X0)),X7)))) )),
  inference(cnf_transformation,[],[f169])).
fof(f169,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun('type.canonical.spolynom'(X1),X0),happ(s(fun('type.canonical.spolynom'(X1),fun('type.canonical.spolynom'(X1),X0)),X7),s('type.canonical.spolynom'(X1),X2))),s('type.canonical.spolynom'(X1),X3))) = s(X0,'const.canonical.spolynom_CASE_5'(s('type.canonical.spolynom'(X1),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(X1),X2),s('type.canonical.spolynom'(X1),X3))),s(fun('type.quote.index',X0),X4),s(fun(X1,X0),X5),s(fun('type.canonical.spolynom'(X1),fun('type.canonical.spolynom'(X1),X0)),X6),s(fun('type.canonical.spolynom'(X1),fun('type.canonical.spolynom'(X1),X0)),X7)))),
  inference(flattening,[],[f168])).
fof(f168,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun('type.canonical.spolynom'(X1),X0),happ(s(fun('type.canonical.spolynom'(X1),fun('type.canonical.spolynom'(X1),X0)),X7),s('type.canonical.spolynom'(X1),X2))),s('type.canonical.spolynom'(X1),X3))) = s(X0,'const.canonical.spolynom_CASE_5'(s('type.canonical.spolynom'(X1),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(X1),X2),s('type.canonical.spolynom'(X1),X3))),s(fun('type.quote.index',X0),X4),s(fun(X1,X0),X5),s(fun('type.canonical.spolynom'(X1),fun('type.canonical.spolynom'(X1),X0)),X6),s(fun('type.canonical.spolynom'(X1),fun('type.canonical.spolynom'(X1),X0)),X7)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X29,X6] : ! [X8,X7,X30,X31,X32,X33] : s(X29,happ(s(fun('type.canonical.spolynom'(X6),X29),happ(s(fun('type.canonical.spolynom'(X6),fun('type.canonical.spolynom'(X6),X29)),X33),s('type.canonical.spolynom'(X6),X8))),s('type.canonical.spolynom'(X6),X7))) = s(X29,'const.canonical.spolynom_CASE_5'(s('type.canonical.spolynom'(X6),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(X6),X8),s('type.canonical.spolynom'(X6),X7))),s(fun('type.quote.index',X29),X30),s(fun(X6,X29),X31),s(fun('type.canonical.spolynom'(X6),fun('type.canonical.spolynom'(X6),X29)),X32),s(fun('type.canonical.spolynom'(X6),fun('type.canonical.spolynom'(X6),X29)),X33)))),
  file('thm.canonical.spolynom_distinct|split|4','thm.canonical.spolynom_case_def|split|3')).
fof(f98478,plain,(
  ( ! [X139] : (s('type.canonical.spolynom'(sK0),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(sK0),sK2),s('type.canonical.spolynom'(sK0),sK1))) = s('type.canonical.spolynom'(sK0),happ(s(fun(sK0,'type.canonical.spolynom'(sK0)),X139),s(sK0,sK3)))) )),
  inference(superposition,[],[f38461,f307])).
fof(f307,plain,(
  ( ! [X2,X0,X1] : (s('type.canonical.spolynom'(X0),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.spolynom'(X0),happ(s(fun('type.canonical.spolynom'(X0),'type.canonical.spolynom'(X0)),happ(s(fun('type.canonical.spolynom'(X0),fun('type.canonical.spolynom'(X0),'type.canonical.spolynom'(X0))),'const.canonical.SPplus_0'),s('type.canonical.spolynom'(X0),X1))),s('type.canonical.spolynom'(X0),X2)))) )),
  inference(cnf_transformation,[],[f117])).
fof(f117,plain,(
  ! [X0,X1,X2] : s('type.canonical.spolynom'(X0),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.spolynom'(X0),happ(s(fun('type.canonical.spolynom'(X0),'type.canonical.spolynom'(X0)),happ(s(fun('type.canonical.spolynom'(X0),fun('type.canonical.spolynom'(X0),'type.canonical.spolynom'(X0))),'const.canonical.SPplus_0'),s('type.canonical.spolynom'(X0),X1))),s('type.canonical.spolynom'(X0),X2)))),
  inference(flattening,[],[f116])).
fof(f116,plain,(
  ! [X0] : ! [X1,X2] : s('type.canonical.spolynom'(X0),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.spolynom'(X0),happ(s(fun('type.canonical.spolynom'(X0),'type.canonical.spolynom'(X0)),happ(s(fun('type.canonical.spolynom'(X0),fun('type.canonical.spolynom'(X0),'type.canonical.spolynom'(X0))),'const.canonical.SPplus_0'),s('type.canonical.spolynom'(X0),X1))),s('type.canonical.spolynom'(X0),X2)))),
  inference(rectify,[],[f49])).
fof(f49,axiom,(
  ! [X6] : ! [X10,X11] : s('type.canonical.spolynom'(X6),'const.canonical.SPplus_2'(s('type.canonical.spolynom'(X6),X10),s('type.canonical.spolynom'(X6),X11))) = s('type.canonical.spolynom'(X6),happ(s(fun('type.canonical.spolynom'(X6),'type.canonical.spolynom'(X6)),happ(s(fun('type.canonical.spolynom'(X6),fun('type.canonical.spolynom'(X6),'type.canonical.spolynom'(X6))),'const.canonical.SPplus_0'),s('type.canonical.spolynom'(X6),X10))),s('type.canonical.spolynom'(X6),X11)))),
  file('thm.canonical.spolynom_distinct|split|4','ARITY_0_2_const.canonical.SPplus')).
% SZS output end Proof for thm.canonical
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 368
% Passive clauses: 6889
% Generated clauses: 103233
% Final active clauses: 253
% Final passive clauses: 1571
% Input formulas: 69
% Initial clauses: 124
% Discarded non-redundant clauses: 3243
% 
% Duplicate literals: 2133
% Trivial inequalities: 944
% Fw subsumption resolutions: 6194
% Fw demodulations: 18316
% Bw demodulations: 3318
% Condensations: 5886
% Global subsumptions: 266
% 
% Simple tautologies: 693
% Equational tautologies: 480
% Forward subsumptions: 56423
% Fw demodulations to eq. taut.: 2804
% Bw demodulations to eq. taut.: 253
% 
% Unit resulting resolution: 12346
% Factoring: 4
% Forward superposition: 27530
% Backward superposition: 20617
% Self superposition: 453
% Equality factoring: 184
% Equality resolution: 72
% 
% Split clauses: 173
% Split components: 366
% SAT solver clauses: 7785
% SAT solver unit clauses: 398
% SAT solver binary clauses: 2099
% 
% Sat splits: 173
% Sat splitting refutations: 282
% 
% Memory used [KB]: 73431
% Time elapsed: 6.077 s
% ------------------------------
% ------------------------------
% Success in time 7.41 s
