Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.complex
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.complex
% SZS output start Proof for thm.complex
fof(f9176,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f9168])).
fof(f9168,plain,(
  s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_scalar_lmul_2'(s('type.realax.real',sK10),s('type.pair.prod'('type.realax.real','type.realax.real'),sK11))) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_scalar_lmul_2'(s('type.realax.real',sK10),s('type.pair.prod'('type.realax.real','type.realax.real'),sK11)))),
  inference(superposition,[],[f254,f9130])).
fof(f9130,plain,(
  ( ! [X0,X1] : (s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_scalar_lmul_2'(s('type.realax.real',X1),s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_scalar_rmul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.realax.real',X1)))) )),
  inference(forward_demodulation,[],[f4036,f4034])).
fof(f4034,plain,(
  ( ! [X0,X1] : (s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_scalar_lmul_2'(s('type.realax.real',X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1))),s('type.realax.real',X0)))))) )),
  inference(superposition,[],[f336,f335])).
fof(f335,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X0)))) )),
  inference(cnf_transformation,[],[f100])).
fof(f100,plain,(
  ! [X0,X1] : s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X0)))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X12,X16] : s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X12),s('type.realax.real',X16))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X16),s('type.realax.real',X12)))),
  file('thm.complex.COMPLEX_SCALAR_MUL_COMM','thm.realax.REAL_MUL_SYM')).
fof(f336,plain,(
  ( ! [X0,X1] : (s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_scalar_lmul_2'(s('type.realax.real',X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1)))))))) )),
  inference(cnf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0,X1] : s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_scalar_lmul_2'(s('type.realax.real',X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1)))))))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X6,X7] : s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_scalar_lmul_2'(s('type.realax.real',X6),s('type.pair.prod'('type.realax.real','type.realax.real'),X7))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X6),s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X7))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X6),s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X7)))))))),
  file('thm.complex.COMPLEX_SCALAR_MUL_COMM','thm.complex.complex_scalar_lmul')).
fof(f4036,plain,(
  ( ! [X0,X1] : (s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_scalar_rmul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.realax.real',X1))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))),s('type.realax.real',X1)))))) )),
  inference(superposition,[],[f337,f335])).
fof(f337,plain,(
  ( ! [X0,X1] : (s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_scalar_rmul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.realax.real',X1))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))),s('type.realax.real',X1)))))) )),
  inference(cnf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0,X1] : s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_scalar_rmul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.realax.real',X1))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))),s('type.realax.real',X1)))))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X7,X6] : s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_scalar_rmul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X7),s('type.realax.real',X6))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X7))),s('type.realax.real',X6))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X7))),s('type.realax.real',X6)))))),
  file('thm.complex.COMPLEX_SCALAR_MUL_COMM','thm.complex.complex_scalar_rmul')).
fof(f254,plain,(
  s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_scalar_lmul_2'(s('type.realax.real',sK10),s('type.pair.prod'('type.realax.real','type.realax.real'),sK11))) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_scalar_rmul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK11),s('type.realax.real',sK10)))),
  inference(cnf_transformation,[],[f164])).
fof(f164,plain,(
  s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_scalar_lmul_2'(s('type.realax.real',sK10),s('type.pair.prod'('type.realax.real','type.realax.real'),sK11))) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_scalar_rmul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK11),s('type.realax.real',sK10)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK10,sK11])],[f124])).
fof(f124,plain,(
  ? [X0,X1] : s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_scalar_lmul_2'(s('type.realax.real',X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_scalar_rmul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1),s('type.realax.real',X0)))),
  inference(ennf_transformation,[],[f54])).
fof(f54,plain,(
  ~! [X0,X1] : s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_scalar_lmul_2'(s('type.realax.real',X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_scalar_rmul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1),s('type.realax.real',X0)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_scalar_lmul_2'(s('type.realax.real',X6),s('type.pair.prod'('type.realax.real','type.realax.real'),X7))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_scalar_rmul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X7),s('type.realax.real',X6)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_scalar_lmul_2'(s('type.realax.real',X6),s('type.pair.prod'('type.realax.real','type.realax.real'),X7))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_scalar_rmul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X7),s('type.realax.real',X6)))),
  file('thm.complex.COMPLEX_SCALAR_MUL_COMM','thm.complex.COMPLEX_SCALAR_MUL_COMM')).
% SZS output end Proof for thm.complex
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 489
% Passive clauses: 3114
% Generated clauses: 8899
% Final active clauses: 450
% Final passive clauses: 2286
% Input formulas: 52
% Initial clauses: 189
% 
% Duplicate literals: 232
% Trivial inequalities: 70
% Fw subsumption resolutions: 943
% Fw demodulations: 42
% Bw demodulations: 38
% 
% Simple tautologies: 187
% Equational tautologies: 60
% Forward subsumptions: 4102
% 
% Binary resolution: 4415
% Factoring: 902
% Forward superposition: 925
% Backward superposition: 942
% Self superposition: 33
% Equality factoring: 26
% Equality resolution: 29
% 
% Split clauses: 228
% Split components: 462
% SAT solver clauses: 439
% SAT solver unit clauses: 9
% SAT solver binary clauses: 110
% 
% Sat splits: 228
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 437
% TWLsolver calls for satisfiability: 269
% 
% Memory used [KB]: 3965
% Time elapsed: 0.240 s
% ------------------------------
% ------------------------------
% Success in time 0.243 s
