Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.pred_set
% (82469)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 711
% Passive clauses: 8710
% Generated clauses: 18357
% Final active clauses: 655
% Final passive clauses: 7536
% Input formulas: 34
% Initial clauses: 156
% 
% Duplicate literals: 508
% Trivial inequalities: 94
% Fw subsumption resolutions: 1572
% Fw demodulations: 231
% Bw demodulations: 55
% 
% Simple tautologies: 192
% Equational tautologies: 66
% Forward subsumptions: 6624
% 
% Binary resolution: 10439
% Factoring: 1331
% Forward superposition: 2583
% Backward superposition: 1090
% Self superposition: 24
% Equality factoring: 26
% Equality resolution: 45
% 
% Split clauses: 372
% Split components: 750
% SAT solver clauses: 721
% SAT solver unit clauses: 14
% SAT solver binary clauses: 183
% 
% Sat splits: 372
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 719
% TWLsolver calls for satisfiability: 444
% 
% Memory used [KB]: 8699
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.pred_set
% (82476)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 373
% Passive clauses: 2430
% Generated clauses: 9948
% Final active clauses: 357
% Final passive clauses: 1778
% Input formulas: 34
% Initial clauses: 156
% 
% Duplicate literals: 206
% Trivial inequalities: 18
% Fw subsumption resolutions: 848
% Fw demodulations: 79
% Bw demodulations: 285
% Condensations: 237
% Global subsumptions: 13
% 
% Simple tautologies: 195
% Equational tautologies: 24
% Forward subsumptions: 3015
% Bw demodulations to eq. taut.: 4
% 
% Binary resolution: 3683
% Factoring: 342
% Forward superposition: 906
% Backward superposition: 666
% Self superposition: 26
% Equality factoring: 26
% Equality resolution: 8
% 
% Split clauses: 48
% Split components: 101
% SAT solver clauses: 3044
% SAT solver unit clauses: 174
% SAT solver binary clauses: 80
% SAT solver learnt clauses: 26
% SAT solver learnt literals: 10
% 
% Sat splits: 48
% Sat splitting refutations: 12
% 
% TWLsolver clauses: 2992
% TWLsolver calls for satisfiability: 1700
% 
% Memory used [KB]: 11129
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f16019,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1160,f1083,f1162,f1087,f1165,f1091,f1167,f1094,f1169,f1097,f1171,f1100,f1173,f1103,f1175,f1106,f1177,f1109,f1179,f1112,f1181,f1115,f1183,f1118,f1185,f1121,f1187,f1124,f1189,f1127,f1191,f1130,f1193,f1133,f1213,f1214,f1215,f1216,f1218,f1221,f1597,f1598,f1604,f1602,f1599,f1600,f1607,f1601,f1609,f1606,f1650,f1651,f1657,f1655,f1652,f1653,f1660,f1654,f1662,f1659,f1718,f1719,f1726,f1723,f1720,f1721,f1729,f1722,f1731,f1724,f1728,f1741,f1742,f1754,f1752,f2256,f2257,f2258,f2259,f2266,f2260,f2268,f2265,f2358,f2359,f2360,f2361,f2368,f2362,f2370,f2367,f2788,f2791,f1084,f2789,f4032,f4033,f4038,f4039,f4129,f4130,f4131,f4132,f4134,f4136,f4137,f4138,f4140,f4142,f4144,f4145,f4146,f4148,f4150,f4152,f4153,f4154,f4263,f4264,f4265,f4270,f4272,f4529,f4530,f4531,f4532,f4533,f4534,f4646,f4647,f4648,f4651,f4652,f4653,f4654,f4655,f4656,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4667,f4668,f4669,f4670,f4671,f4672,f4673,f4674,f4675,f4676,f4677,f4678,f4679,f4680,f4681,f4682,f4683,f4684,f4685,f4686,f4687,f4688,f4689,f4690,f4691,f4692,f4693,f4694,f4695,f4696,f4697,f4698,f4699,f4700,f4701,f4702,f4703,f4704,f4705,f4706,f4707,f4708,f4709,f4710,f4711,f4712,f4713,f4714,f4715,f4716,f4717,f4718,f4719,f4720,f4721,f4722,f4723,f4724,f4725,f4726,f4727,f4728,f4729,f4730,f4731,f4732,f4733,f4734,f4735,f4736,f4737,f4738,f4739,f4740,f4741,f6436,f6437,f6993,f6996,f6994,f7404,f7405,f7406,f7407,f7614,f7615,f7616,f7617,f7618,f7619,f7620,f7621,f7622,f7624,f7625,f7626,f7627,f7628,f7629,f7630,f7631,f7632,f7634,f7636,f7637,f7638,f7639,f7640,f7642,f7644,f7646,f7648,f7649,f7650,f7652,f7654,f7656,f7658,f7659,f7660,f7778,f7779,f7781,f7783,f7786,f7791,f7801,f7449,f7446,f2263,f2261,f9588,f9589,f9590,f4034,f9591,f9771,f9772,f9773,f9774,f9775,f9776,f9777,f9778,f9779,f9781,f9782,f9783,f9784,f9785,f9786,f9787,f9788,f9789,f9791,f9793,f9794,f9795,f9797,f9798,f9799,f9801,f9803,f9805,f9806,f9807,f9809,f9811,f9813,f9814,f9815,f9933,f9934,f9936,f9938,f9941,f9946,f9956,f9622,f9619,f2365,f2363,f12868,f12887,f12888,f12889,f12890,f12891,f12892,f12900,f12930,f12953,f12955,f12956,f12957,f12958,f12959,f13072,f13073,f13075,f13077,f13080,f13094,f13096,f13102,f15980])).
fof(f15980,plain,(
  $false | ($spl76 | $spl80)),
  inference(subsumption_resolution,[],[f15979,f364])).
fof(f364,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f203,f200])).
fof(f200,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.UNION_DIFF_2','thm.bool.TRUTH')).
fof(f203,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f128])).
fof(f128,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f38])).
fof(f38,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f17])).
fof(f17,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.pred_set.UNION_DIFF_2','thm.bool.NOT_CLAUSES|split|1')).
fof(f15979,plain,(
  p(s(bool,'F_0')) | ($spl76 | $spl80)),
  inference(forward_demodulation,[],[f15978,f2267])).
fof(f2267,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | $spl80),
  inference(cnf_transformation,[],[f2267_D])).
fof(f2267_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f15978,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)))) | $spl76),
  inference(subsumption_resolution,[],[f15977,f245])).
fof(f245,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f150])).
fof(f150,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))) & (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))))),
  inference(flattening,[],[f149])).
fof(f149,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))) & ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))))),
  inference(nnf_transformation,[],[f68])).
fof(f68,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(flattening,[],[f67])).
fof(f67,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X6] : ! [X7,X8,X11] : (p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),'const.pred_set.DIFF_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X7)))) & ~p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X8))))))),
  file('thm.pred_set.UNION_DIFF_2','thm.pred_set.IN_DIFF')).
fof(f15977,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)))) | p(s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))) | $spl76),
  inference(subsumption_resolution,[],[f15905,f200])).
fof(f15905,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)))) | p(s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))) | $spl76),
  inference(superposition,[],[f248,f2265])).
fof(f248,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f152])).
fof(f152,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) & ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))))),
  inference(flattening,[],[f151])).
fof(f151,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))) & ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))))),
  inference(nnf_transformation,[],[f70])).
fof(f70,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(flattening,[],[f69])).
fof(f69,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X6] : ! [X7,X8,X11] : (p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X7)))) | p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X8))))))),
  file('thm.pred_set.UNION_DIFF_2','thm.pred_set.IN_UNION')).
fof(f13102,plain,(
  $false | ($spl76 | $spl78 | $spl106)),
  inference(subsumption_resolution,[],[f13101,f12960])).
fof(f12960,plain,(
  ( ! [X91] : (s(bool,'T_0') = s(bool,X91)) ) | ($spl76 | $spl78 | $spl106)),
  inference(global_subsumption,[],[f7405,f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1718,f1719,f1720,f1721,f1722,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2269,f2314,f2313,f2312,f2309,f2308,f2323,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f3197,f3209,f3223,f3866,f4029,f4032,f721,f720,f2303,f6433,f6436,f7404,f7406,f2266,f5332,f7449,f7446,f1576,f1574,f1556,f1538,f1515,f1512,f1492,f1489,f1474,f1471,f1451,f1449,f1579,f1559,f1554,f1536,f1510,f1487,f1469,f1454,f8521,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f1756,f896,f797,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f9622,f9619,f12379,f12383,f12400,f12399,f12420,f12419,f12444,f12443,f12579,f12578,f2265])).
fof(f12578,plain,(
  sP3(f(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12))),g(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12))),'F_0')),
  inference(unit_resulting_resolution,[],[f1648,f282])).
fof(f12579,plain,(
  sP2(g(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12))),f(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12))),'F_0')),
  inference(unit_resulting_resolution,[],[f1648,f281])).
fof(f12443,plain,(
  sP3(g(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12))),f(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12))),'F_0')),
  inference(unit_resulting_resolution,[],[f1646,f282])).
fof(f12444,plain,(
  sP2(f(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12))),g(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12))),'F_0')),
  inference(unit_resulting_resolution,[],[f1646,f281])).
fof(f12419,plain,(
  sP3(f(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))),g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))),'F_0')),
  inference(unit_resulting_resolution,[],[f1595,f282])).
fof(f12420,plain,(
  sP2(g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))),f(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))),'F_0')),
  inference(unit_resulting_resolution,[],[f1595,f281])).
fof(f12399,plain,(
  sP3(g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))),f(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))),'F_0')),
  inference(unit_resulting_resolution,[],[f1593,f282])).
fof(f12400,plain,(
  sP2(f(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))),g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))),'F_0')),
  inference(unit_resulting_resolution,[],[f1593,f281])).
fof(f12383,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X2)) | ~sP11(X0,X2,'T_0') | s(bool,'T_0') = s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f12382])).
fof(f12382,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X2)) | ~sP11(X0,X2,'T_0') | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) )),
  inference(condensation,[],[f12381])).
fof(f12381,plain,(
  ( ! [X30,X28,X31,X29,X32] : (p(s(bool,X29)) | ~sP11(X30,X29,'T_0') | s(bool,X30) = s(bool,X31) | s(bool,'T_0') = s(bool,X31) | s(bool,X28) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(subsumption_resolution,[],[f12364,f518])).
fof(f12364,plain,(
  ( ! [X30,X28,X31,X29,X32] : (~p(s(bool,X28)) | p(s(bool,X29)) | ~sP11(X30,X29,'T_0') | s(bool,X30) = s(bool,X31) | s(bool,'T_0') = s(bool,X31) | s(bool,X28) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(superposition,[],[f1152,f472])).
fof(f12379,plain,(
  ( ! [X6,X4,X8,X7,X5] : (p(s(bool,X5)) | ~sP11(X6,X5,X4) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7) | s(bool,X4) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) )),
  inference(subsumption_resolution,[],[f12359,f200])).
fof(f12359,plain,(
  ( ! [X6,X4,X8,X7,X5] : (~p(s(bool,'T_0')) | p(s(bool,X5)) | ~sP11(X6,X5,X4) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7) | s(bool,X4) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f1152,f472])).
fof(f220,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f132])).
fof(f132,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f51])).
fof(f51,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X8] : (s(bool,'F_0') = s(bool,X8) <=> ~p(s(bool,X8)))),
  file('thm.pred_set.UNION_DIFF_2','thm.bool.EQ_CLAUSES|split|2')).
fof(f221,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f132])).
fof(f722,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK16(X0,X1,X2,bool))) )),
  inference(superposition,[],[f280,f207])).
fof(f723,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK16(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK16(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f280,f207])).
fof(f1448,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X1),s(fun(bool,bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f245,f207])).
fof(f1453,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X20),s(fun(bool,bool),X21)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X19),s(fun(bool,bool),X20)))) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f245,f207])).
fof(f1468,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X2),s(fun(bool,bool),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),X1)))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f246,f207])).
fof(f1473,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X20)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X19),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X21),s(fun(bool,bool),X20)))))) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f246,f207])).
fof(f1486,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X1),s(fun(bool,bool),X2)))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f249,f207])).
fof(f1491,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X20)))) | p(s(bool,'const.bool.IN_2'(s(bool,X19),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X20),s(fun(bool,bool),X21)))))) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f249,f207])).
fof(f1509,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X2),s(fun(bool,bool),X1)))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f250,f207])).
fof(f1514,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X20)))) | p(s(bool,'const.bool.IN_2'(s(bool,X19),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X21),s(fun(bool,bool),X20)))))) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f250,f207])).
fof(f1535,plain,(
  ( ! [X17,X18,X16] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(X16,sK15(X17,X18,X16)),s(fun(X16,bool),X17))) | s(fun(X16,bool),X17) = s(fun(X16,bool),X18) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X16,sK15(X17,X18,X16)),s(fun(X16,bool),X18)))) )),
  inference(superposition,[],[f244,f207])).
fof(f1540,plain,(
  ( ! [X12,X13] : (s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X12))) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X13))) | s(fun(bool,bool),X12) = s(fun(bool,bool),X13) | s(bool,'T_0') = s(bool,sK15(X12,X13,bool))) )),
  inference(superposition,[],[f244,f207])).
fof(f1545,plain,(
  ( ! [X17,X18,X16] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(X16,sK15(X17,X18,X16)),s(fun(X16,bool),X18))) | s(fun(X16,bool),X17) = s(fun(X16,bool),X18) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X16,sK15(X17,X18,X16)),s(fun(X16,bool),X17)))) )),
  inference(superposition,[],[f244,f207])).
fof(f1553,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X1),s(fun(bool,bool),X2)))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f247,f207])).
fof(f1558,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X20)))) | p(s(bool,'const.bool.IN_2'(s(bool,X19),s(fun(bool,bool),X21)))) | p(s(bool,'const.bool.IN_2'(s(bool,X19),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X20),s(fun(bool,bool),X21)))))) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f247,f207])).
fof(f1573,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X1),s(fun(bool,bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X2)))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f248,f207])).
fof(f1578,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X20),s(fun(bool,bool),X21)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X19),s(fun(bool,bool),X20)))) | p(s(bool,'const.bool.IN_2'(s(bool,X19),s(fun(bool,bool),X21)))) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f248,f207])).
fof(f1826,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))) | s(bool,'T_0') = s(bool,X0) | s(bool,X3) = s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X1),s(fun(bool,bool),X2))))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f1816,f200])).
fof(f1816,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))) | s(bool,'T_0') = s(bool,X0) | s(bool,X3) = s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X1),s(fun(bool,bool),X2))))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1448,f472])).
fof(f1865,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(bool,X2),s(fun(bool,bool),X0)))) | s(bool,'T_0') = s(bool,X2) | s(bool,X3) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X0),s(fun(bool,bool),X1))))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f1860,f200])).
fof(f1860,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(bool,X2),s(fun(bool,bool),X0)))) | s(bool,'T_0') = s(bool,X2) | s(bool,X3) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X0),s(fun(bool,bool),X1))))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1453,f472])).
fof(f1901,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,X2),s(fun(bool,bool),X1)))) | s(bool,'T_0') = s(bool,X2) | s(bool,X3) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X0),s(fun(bool,bool),X1))))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f1896,f200])).
fof(f1896,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(bool,X2),s(fun(bool,bool),X1)))) | s(bool,'T_0') = s(bool,X2) | s(bool,X3) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X0),s(fun(bool,bool),X1))))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1468,f472])).
fof(f1936,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,X1),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X2),s(fun(bool,bool),X0)))))) | s(bool,'T_0') = s(bool,X1) | s(bool,X3) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f1931,f200])).
fof(f1931,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(bool,X1),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X2),s(fun(bool,bool),X0)))))) | s(bool,'T_0') = s(bool,X1) | s(bool,X3) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1473,f472])).
fof(f2067,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X1),s(fun(bool,bool),X2)))))) | s(bool,'T_0') = s(bool,X0) | s(bool,X3) = s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),X1))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f2057,f200])).
fof(f2057,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X1),s(fun(bool,bool),X2)))))) | s(bool,'T_0') = s(bool,X0) | s(bool,X3) = s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),X1))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1486,f472])).
fof(f2102,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(bool,X1),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X0),s(fun(bool,bool),X2)))))) | s(bool,'T_0') = s(bool,X1) | s(bool,X3) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f2097,f200])).
fof(f2097,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(bool,X1),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X0),s(fun(bool,bool),X2)))))) | s(bool,'T_0') = s(bool,X1) | s(bool,X3) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1491,f472])).
fof(f2146,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X2),s(fun(bool,bool),X1)))))) | s(bool,'T_0') = s(bool,X0) | s(bool,X3) = s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),X1))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f2136,f200])).
fof(f2136,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X2),s(fun(bool,bool),X1)))))) | s(bool,'T_0') = s(bool,X0) | s(bool,X3) = s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),X1))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1509,f472])).
fof(f2186,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(bool,X1),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X2),s(fun(bool,bool),X0)))))) | s(bool,'T_0') = s(bool,X1) | s(bool,X3) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f2181,f200])).
fof(f2181,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(bool,X1),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X2),s(fun(bool,bool),X0)))))) | s(bool,'T_0') = s(bool,X1) | s(bool,X3) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1514,f472])).
fof(f2214,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X1))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK15(X0,X1,bool)) | s(bool,X2) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1540,f472])).
fof(f2216,plain,(
  ( ! [X6,X8,X7] : (s(bool,X7) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X8))) | s(fun(bool,bool),X6) = s(fun(bool,bool),X8) | s(bool,'T_0') = s(bool,sK15(X6,X8,bool)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X6))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f1540,f472])).
fof(f2218,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X13))) | s(fun(bool,bool),X12) = s(fun(bool,bool),X13) | s(bool,'T_0') = s(bool,sK15(X12,X13,bool)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X12)))) )),
  inference(superposition,[],[f1540,f207])).
fof(f2225,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X1))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK15(X1,X0,bool)) | s(bool,X2) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1540,f472])).
fof(f2227,plain,(
  ( ! [X6,X8,X7] : (s(bool,X7) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X8))) | s(fun(bool,bool),X6) = s(fun(bool,bool),X8) | s(bool,'T_0') = s(bool,sK15(X8,X6,bool)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X6))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f1540,f472])).
fof(f2229,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X13))) | s(fun(bool,bool),X12) = s(fun(bool,bool),X13) | s(bool,'T_0') = s(bool,sK15(X13,X12,bool)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X12)))) )),
  inference(superposition,[],[f1540,f207])).
fof(f2462,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X0))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,sK15(X0,X1,bool)),s(fun(bool,bool),X1))) | s(bool,X2) = s(bool,sK15(X0,X1,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1535,f472])).
fof(f2464,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(bool,X8),s(fun(bool,bool),X6))) | s(fun(bool,bool),X6) = s(fun(bool,bool),X7) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,sK15(X6,X7,bool)),s(fun(bool,bool),X7))) | s(bool,'T_0') = s(bool,sK15(X6,X7,bool)) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f1535,f472])).
fof(f2466,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X12))) | s(fun(bool,bool),X12) = s(fun(bool,bool),X13) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,sK15(X12,X13,bool)),s(fun(bool,bool),X13))) | s(bool,'T_0') = s(bool,sK15(X12,X13,bool))) )),
  inference(superposition,[],[f1535,f207])).
fof(f2501,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X1))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,sK15(X0,X1,bool)),s(fun(bool,bool),X0))) | s(bool,X2) = s(bool,sK15(X0,X1,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1545,f472])).
fof(f2503,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(bool,X8),s(fun(bool,bool),X7))) | s(fun(bool,bool),X6) = s(fun(bool,bool),X7) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,sK15(X6,X7,bool)),s(fun(bool,bool),X6))) | s(bool,'T_0') = s(bool,sK15(X6,X7,bool)) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f1545,f472])).
fof(f2505,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X13))) | s(fun(bool,bool),X12) = s(fun(bool,bool),X13) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,sK15(X12,X13,bool)),s(fun(bool,bool),X12))) | s(bool,'T_0') = s(bool,sK15(X12,X13,bool))) )),
  inference(superposition,[],[f1545,f207])).
fof(f2556,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X1),s(fun(bool,bool),X2)))))) | s(bool,'T_0') = s(bool,X0) | s(bool,X3) = s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),X1))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f2546,f200])).
fof(f2546,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X1),s(fun(bool,bool),X2)))))) | s(bool,'T_0') = s(bool,X0) | s(bool,X3) = s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),X1))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1553,f472])).
fof(f2590,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(bool,X1),s(fun(bool,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(bool,X1),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X0),s(fun(bool,bool),X2)))))) | s(bool,'T_0') = s(bool,X1) | s(bool,X3) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f2585,f200])).
fof(f2585,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(bool,X1),s(fun(bool,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(bool,X1),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X0),s(fun(bool,bool),X2)))))) | s(bool,'T_0') = s(bool,X1) | s(bool,X3) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1558,f472])).
fof(f2634,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X2)))) | s(bool,'T_0') = s(bool,X0) | s(bool,X3) = s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X1),s(fun(bool,bool),X2))))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f2624,f200])).
fof(f2624,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X2)))) | s(bool,'T_0') = s(bool,X0) | s(bool,X3) = s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X1),s(fun(bool,bool),X2))))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1573,f472])).
fof(f2668,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(bool,X2),s(fun(bool,bool),X0)))) | p(s(bool,'const.bool.IN_2'(s(bool,X2),s(fun(bool,bool),X1)))) | s(bool,'T_0') = s(bool,X2) | s(bool,X3) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X0),s(fun(bool,bool),X1))))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f2663,f200])).
fof(f2663,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(bool,X2),s(fun(bool,bool),X0)))) | p(s(bool,'const.bool.IN_2'(s(bool,X2),s(fun(bool,bool),X1)))) | s(bool,'T_0') = s(bool,X2) | s(bool,X3) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X0),s(fun(bool,bool),X1))))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1578,f472])).
fof(f797,plain,(
  ( ! [X111,X112,X110,X113] : (s(bool,X113) != s(bool,f(s(X110,sK16(X111,X112,bool,X110)))) | s(fun(X110,bool),X111) = s(fun(X110,bool),X112) | s(bool,'T_0') = s(bool,X113) | s(bool,'T_0') = s(bool,g(s(X110,sK16(X111,X112,bool,X110))))) )),
  inference(superposition,[],[f280,f472])).
fof(f896,plain,(
  ( ! [X111,X109,X110,X108] : (s(bool,'T_0') != s(bool,f(s(X108,sK16(X109,X110,bool,X108)))) | s(fun(X108,bool),X109) = s(fun(X108,bool),X110) | s(bool,X111) = s(bool,g(s(X108,sK16(X109,X110,bool,X108)))) | s(bool,'T_0') = s(bool,X111)) )),
  inference(superposition,[],[f280,f472])).
fof(f1756,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,g(s(X0,sK16(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,f(s(X0,sK16(X1,X2,bool,X0))))) )),
  inference(condensation,[],[f1755])).
fof(f1755,plain,(
  ( ! [X10,X8,X11,X9] : (s(fun(X8,bool),X9) = s(fun(X8,bool),X10) | s(bool,'T_0') = s(bool,g(s(X8,sK16(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,f(s(X8,sK16(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(subsumption_resolution,[],[f1748,f207])).
fof(f1748,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'F_0') != s(bool,X11) | s(fun(X8,bool),X9) = s(fun(X8,bool),X10) | s(bool,'T_0') = s(bool,g(s(X8,sK16(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,f(s(X8,sK16(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f723,f472])).
fof(f244,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.bool.IN_2'(s(X0,sK15(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK15(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) )),
  inference(cnf_transformation,[],[f148])).
fof(f148,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (s(bool,'const.bool.IN_2'(s(X0,sK15(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK15(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK15])],[f147])).
fof(f147,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(rectify,[],[f146])).
fof(f146,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(nnf_transformation,[],[f66])).
fof(f66,plain,(
  ! [X0,X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(flattening,[],[f65])).
fof(f65,plain,(
  ! [X0] : ! [X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X6] : ! [X7,X8] : (s(fun(X6,bool),X7) = s(fun(X6,bool),X8) <=> ! [X11] : s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X7))) = s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X8))))),
  file('thm.pred_set.UNION_DIFF_2','thm.pred_set.EXTENSION')).
fof(f243,plain,(
  ( ! [X4,X2,X0,X1] : (s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) )),
  inference(cnf_transformation,[],[f148])).
fof(f246,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))) )),
  inference(cnf_transformation,[],[f150])).
fof(f247,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))) )),
  inference(cnf_transformation,[],[f150])).
fof(f249,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))) )),
  inference(cnf_transformation,[],[f152])).
fof(f250,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))) )),
  inference(cnf_transformation,[],[f152])).
fof(f1461,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) | s(bool,X4) = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f1456,f200])).
fof(f1456,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) | s(bool,X4) = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f245,f472])).
fof(f1481,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X3),s(fun(X0,bool),X2)))))) | s(bool,X4) = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f1476,f200])).
fof(f1476,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X3),s(fun(X0,bool),X2)))))) | s(bool,X4) = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f246,f472])).
fof(f1499,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))))) | s(bool,X4) = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f1494,f200])).
fof(f1494,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))))) | s(bool,X4) = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f249,f472])).
fof(f1522,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X3),s(fun(X0,bool),X2)))))) | s(bool,X4) = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f1517,f200])).
fof(f1517,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X3),s(fun(X0,bool),X2)))))) | s(bool,X4) = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f250,f472])).
fof(f1566,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))))) | s(bool,X4) = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f1561,f200])).
fof(f1561,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))))) | s(bool,X4) = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f247,f472])).
fof(f1586,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3)))) | s(bool,X4) = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f1581,f200])).
fof(f1581,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3)))) | s(bool,X4) = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f248,f472])).
fof(f8521,plain,(
  s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) != s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))))),
  inference(unit_resulting_resolution,[],[f199,f244])).
fof(f1454,plain,(
  ( ! [X24,X23,X25,X22] : (~p(s(bool,'const.bool.IN_2'(s(bool,X22),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X23),s(fun(bool,bool),X24)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X23)))) | s(bool,X22) = s(bool,X25) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f245,f472])).
fof(f1469,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X4)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X3),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X5),s(fun(bool,bool),X4)))))) | s(bool,X3) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f246,f472])).
fof(f1487,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X4)))) | p(s(bool,'const.bool.IN_2'(s(bool,X3),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X4),s(fun(bool,bool),X5)))))) | s(bool,X3) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f249,f472])).
fof(f1510,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X4)))) | p(s(bool,'const.bool.IN_2'(s(bool,X3),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X5),s(fun(bool,bool),X4)))))) | s(bool,X3) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f250,f472])).
fof(f1536,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X0))) != s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X1))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,X2) = s(bool,sK15(X0,X1,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f244,f472])).
fof(f1554,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X4)))) | p(s(bool,'const.bool.IN_2'(s(bool,X3),s(fun(bool,bool),X5)))) | p(s(bool,'const.bool.IN_2'(s(bool,X3),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X4),s(fun(bool,bool),X5)))))) | s(bool,X3) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f247,f472])).
fof(f1559,plain,(
  ( ! [X24,X23,X25,X22] : (~p(s(bool,'const.bool.IN_2'(s(bool,X22),s(fun(bool,bool),X23)))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X24)))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X23),s(fun(bool,bool),X24)))))) | s(bool,X22) = s(bool,X25) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f247,f472])).
fof(f1579,plain,(
  ( ! [X24,X23,X25,X22] : (~p(s(bool,'const.bool.IN_2'(s(bool,X22),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X23),s(fun(bool,bool),X24)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X23)))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X24)))) | s(bool,X22) = s(bool,X25) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f248,f472])).
fof(f1449,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X4),s(fun(bool,bool),X5)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X3),s(fun(bool,bool),X4)))) | s(bool,X3) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f245,f472])).
fof(f1451,plain,(
  ( ! [X14,X12,X13,X11] : (~p(s(bool,'const.bool.IN_2'(s(bool,X12),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X13),s(fun(bool,bool),X14)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X11),s(fun(bool,bool),X13)))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f245,f472])).
fof(f1471,plain,(
  ( ! [X14,X12,X13,X11] : (~p(s(bool,'const.bool.IN_2'(s(bool,X12),s(fun(bool,bool),X13)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X11),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X14),s(fun(bool,bool),X13)))))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f246,f472])).
fof(f1474,plain,(
  ( ! [X24,X23,X25,X22] : (~p(s(bool,'const.bool.IN_2'(s(bool,X22),s(fun(bool,bool),X23)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X24),s(fun(bool,bool),X23)))))) | s(bool,X22) = s(bool,X25) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f246,f472])).
fof(f1489,plain,(
  ( ! [X14,X12,X13,X11] : (~p(s(bool,'const.bool.IN_2'(s(bool,X12),s(fun(bool,bool),X13)))) | p(s(bool,'const.bool.IN_2'(s(bool,X11),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X13),s(fun(bool,bool),X14)))))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f249,f472])).
fof(f1492,plain,(
  ( ! [X24,X23,X25,X22] : (~p(s(bool,'const.bool.IN_2'(s(bool,X22),s(fun(bool,bool),X23)))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X23),s(fun(bool,bool),X24)))))) | s(bool,X22) = s(bool,X25) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f249,f472])).
fof(f1512,plain,(
  ( ! [X14,X12,X13,X11] : (~p(s(bool,'const.bool.IN_2'(s(bool,X12),s(fun(bool,bool),X13)))) | p(s(bool,'const.bool.IN_2'(s(bool,X11),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X14),s(fun(bool,bool),X13)))))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f250,f472])).
fof(f1515,plain,(
  ( ! [X24,X23,X25,X22] : (~p(s(bool,'const.bool.IN_2'(s(bool,X22),s(fun(bool,bool),X23)))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X24),s(fun(bool,bool),X23)))))) | s(bool,X22) = s(bool,X25) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f250,f472])).
fof(f1538,plain,(
  ( ! [X6,X8,X7] : (s(bool,'const.bool.IN_2'(s(bool,X8),s(fun(bool,bool),X6))) != s(bool,'const.bool.IN_2'(s(bool,X8),s(fun(bool,bool),X7))) | s(fun(bool,bool),X6) = s(fun(bool,bool),X7) | s(bool,'T_0') = s(bool,sK15(X6,X7,bool)) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f244,f472])).
fof(f1556,plain,(
  ( ! [X14,X12,X13,X11] : (~p(s(bool,'const.bool.IN_2'(s(bool,X12),s(fun(bool,bool),X13)))) | p(s(bool,'const.bool.IN_2'(s(bool,X11),s(fun(bool,bool),X14)))) | p(s(bool,'const.bool.IN_2'(s(bool,X11),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X13),s(fun(bool,bool),X14)))))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f247,f472])).
fof(f1574,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X4),s(fun(bool,bool),X5)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X3),s(fun(bool,bool),X4)))) | p(s(bool,'const.bool.IN_2'(s(bool,X3),s(fun(bool,bool),X5)))) | s(bool,X3) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f248,f472])).
fof(f1576,plain,(
  ( ! [X14,X12,X13,X11] : (~p(s(bool,'const.bool.IN_2'(s(bool,X12),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X13),s(fun(bool,bool),X14)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X11),s(fun(bool,bool),X13)))) | p(s(bool,'const.bool.IN_2'(s(bool,X11),s(fun(bool,bool),X14)))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f248,f472])).
fof(f5332,plain,(
  s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) != s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))))),
  inference(unit_resulting_resolution,[],[f199,f244])).
fof(f6433,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f6419,f364])).
fof(f6419,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | sP1(X0,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1139,f207])).
fof(f2303,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | $spl76),
  inference(unit_resulting_resolution,[],[f407,f2269,f1136])).
fof(f720,plain,(
  s(bool,f(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12)))) != s(bool,g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))))),
  inference(unit_resulting_resolution,[],[f199,f280])).
fof(f721,plain,(
  s(bool,f(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) != s(bool,g(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12))))),
  inference(unit_resulting_resolution,[],[f199,f280])).
fof(f4029,plain,(
  ( ! [X2,X0,X1] : (sP1(X1,X0) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f4013,f364])).
fof(f4013,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | sP1(X1,X0) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1138,f207])).
fof(f3866,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1134,f207])).
fof(f3223,plain,(
  sP3('T_0','const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)),'F_0') | $spl76),
  inference(unit_resulting_resolution,[],[f2313,f282])).
fof(f3209,plain,(
  ~sP4('T_0','T_0','const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | $spl76),
  inference(unit_resulting_resolution,[],[f200,f376,f371,f2309,f285])).
fof(f3197,plain,(
  sP3('const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)),'T_0','F_0') | $spl76),
  inference(unit_resulting_resolution,[],[f2308,f282])).
fof(f2475,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,sK15(X1,X2,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,sK15(X1,X2,X0)),s(fun(X0,bool),X1)))) )),
  inference(condensation,[],[f2474])).
fof(f2474,plain,(
  ( ! [X10,X8,X11,X9] : (s(fun(X8,bool),X9) = s(fun(X8,bool),X10) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X8,sK15(X9,X10,X8)),s(fun(X8,bool),X10))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X8,sK15(X9,X10,X8)),s(fun(X8,bool),X9))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(subsumption_resolution,[],[f2469,f207])).
fof(f2469,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'F_0') != s(bool,X11) | s(fun(X8,bool),X9) = s(fun(X8,bool),X10) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X8,sK15(X9,X10,X8)),s(fun(X8,bool),X10))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X8,sK15(X9,X10,X8)),s(fun(X8,bool),X9))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f1535,f472])).
fof(f2353,plain,(
  sP5('F_0','const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))),'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13)))),
  inference(unit_resulting_resolution,[],[f364,f1530,f300])).
fof(f2354,plain,(
  ~sP5('T_0','const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))),'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13)))),
  inference(unit_resulting_resolution,[],[f200,f1530,f302])).
fof(f2356,plain,(
  sP5('F_0','const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13)),'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))))),
  inference(unit_resulting_resolution,[],[f364,f1530,f300])).
fof(f2357,plain,(
  ~sP5('T_0','const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13)),'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))))),
  inference(unit_resulting_resolution,[],[f200,f1530,f302])).
fof(f2323,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | s(bool,'T_0') = s(bool,X0)) ) | $spl76),
  inference(trivial_inequality_removal,[],[f2317])).
fof(f2317,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | s(bool,'T_0') = s(bool,X0)) ) | $spl76),
  inference(superposition,[],[f2269,f472])).
fof(f2308,plain,(
  sP5('F_0','const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)),'T_0') | $spl76),
  inference(unit_resulting_resolution,[],[f364,f2269,f300])).
fof(f2309,plain,(
  ~sP5('T_0','const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)),'T_0') | $spl76),
  inference(unit_resulting_resolution,[],[f200,f2269,f302])).
fof(f2312,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)))) | $spl76),
  inference(unit_resulting_resolution,[],[f200,f2269,f254])).
fof(f2313,plain,(
  sP5('F_0','T_0','const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | $spl76),
  inference(unit_resulting_resolution,[],[f364,f2269,f300])).
fof(f2314,plain,(
  ~sP5('T_0','T_0','const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | $spl76),
  inference(unit_resulting_resolution,[],[f200,f2269,f302])).
fof(f2269,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | $spl76),
  inference(backward_demodulation,[],[f2265,f1529])).
fof(f2251,plain,(
  sP5('F_0','const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))),'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)))),
  inference(unit_resulting_resolution,[],[f364,f1529,f300])).
fof(f2252,plain,(
  ~sP5('T_0','const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))),'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)))),
  inference(unit_resulting_resolution,[],[f200,f1529,f302])).
fof(f2254,plain,(
  sP5('F_0','const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)),'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))))),
  inference(unit_resulting_resolution,[],[f364,f1529,f300])).
fof(f2255,plain,(
  ~sP5('T_0','const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)),'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))))),
  inference(unit_resulting_resolution,[],[f200,f1529,f302])).
fof(f1646,plain,(
  sP5('F_0',g(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12))),f(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12))))),
  inference(unit_resulting_resolution,[],[f364,f721,f300])).
fof(f1647,plain,(
  ~sP5('T_0',g(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12))),f(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12))))),
  inference(unit_resulting_resolution,[],[f200,f721,f302])).
fof(f1648,plain,(
  sP5('F_0',f(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12))),g(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12))))),
  inference(unit_resulting_resolution,[],[f364,f721,f300])).
fof(f1649,plain,(
  ~sP5('T_0',f(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12))),g(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12))))),
  inference(unit_resulting_resolution,[],[f200,f721,f302])).
fof(f1593,plain,(
  sP5('F_0',g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))),f(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))))),
  inference(unit_resulting_resolution,[],[f364,f720,f300])).
fof(f1594,plain,(
  ~sP5('T_0',g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))),f(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))))),
  inference(unit_resulting_resolution,[],[f200,f720,f302])).
fof(f1595,plain,(
  sP5('F_0',f(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))),g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))))),
  inference(unit_resulting_resolution,[],[f364,f720,f300])).
fof(f1596,plain,(
  ~sP5('T_0',f(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))),g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))))),
  inference(unit_resulting_resolution,[],[f200,f720,f302])).
fof(f1543,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X11) != s(bool,'const.bool.IN_2'(s(X8,sK15(X9,X10,X8)),s(fun(X8,bool),X10))) | s(fun(X8,bool),X9) = s(fun(X8,bool),X10) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X8,sK15(X9,X10,X8)),s(fun(X8,bool),X9))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f244,f472])).
fof(f1541,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X0,sK15(X1,X2,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,X3) = s(bool,'const.bool.IN_2'(s(X0,sK15(X1,X2,X0)),s(fun(X0,bool),X1))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f244,f472])).
fof(f1533,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X11) != s(bool,'const.bool.IN_2'(s(X8,sK15(X9,X10,X8)),s(fun(X8,bool),X9))) | s(fun(X8,bool),X9) = s(fun(X8,bool),X10) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X8,sK15(X9,X10,X8)),s(fun(X8,bool),X10))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f244,f472])).
fof(f1531,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X0,sK15(X1,X2,X0)),s(fun(X0,bool),X1))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,X3) = s(bool,'const.bool.IN_2'(s(X0,sK15(X1,X2,X0)),s(fun(X0,bool),X2))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f244,f472])).
fof(f1529,plain,(
  s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) != s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))))),
  inference(unit_resulting_resolution,[],[f199,f244])).
fof(f1530,plain,(
  s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) != s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))))),
  inference(unit_resulting_resolution,[],[f199,f244])).
fof(f895,plain,(
  ( ! [X107,X105,X106,X104] : (s(X106,f(s(bool,'T_0'))) != s(X106,g(s(bool,'T_0'))) | s(fun(bool,X106),X104) = s(fun(bool,X106),X105) | s(bool,X107) = s(bool,sK16(X104,X105,X106,bool)) | s(bool,'T_0') = s(bool,X107)) )),
  inference(superposition,[],[f280,f472])).
fof(f1154,plain,(
  ( ! [X94,X97,X95,X96] : (~sP8(X94,X95,X96) | sP9(X94,X95,X96) | s(bool,X94) = s(bool,X97) | s(bool,'T_0') = s(bool,X97)) )),
  inference(subsumption_resolution,[],[f893,f200])).
fof(f893,plain,(
  ( ! [X94,X97,X95,X96] : (~p(s(bool,'T_0')) | ~sP8(X94,X95,X96) | sP9(X94,X95,X96) | s(bool,X94) = s(bool,X97) | s(bool,'T_0') = s(bool,X97)) )),
  inference(superposition,[],[f356,f472])).
fof(f1153,plain,(
  ( ! [X92,X90,X93,X91] : (~sP10(X90,X91,X92) | sP11(X90,X91,X92) | s(bool,X90) = s(bool,X93) | s(bool,'T_0') = s(bool,X93)) )),
  inference(subsumption_resolution,[],[f892,f200])).
fof(f892,plain,(
  ( ! [X92,X90,X93,X91] : (~p(s(bool,'T_0')) | ~sP10(X90,X91,X92) | sP11(X90,X91,X92) | s(bool,X90) = s(bool,X93) | s(bool,'T_0') = s(bool,X93)) )),
  inference(superposition,[],[f355,f472])).
fof(f1152,plain,(
  ( ! [X88,X87,X89,X86] : (~p(s(bool,X87)) | p(s(bool,X88)) | ~sP11(X86,X88,X87) | s(bool,X86) = s(bool,X89) | s(bool,'T_0') = s(bool,X89)) )),
  inference(subsumption_resolution,[],[f891,f200])).
fof(f891,plain,(
  ( ! [X88,X87,X89,X86] : (~p(s(bool,'T_0')) | ~p(s(bool,X87)) | p(s(bool,X88)) | ~sP11(X86,X88,X87) | s(bool,X86) = s(bool,X89) | s(bool,'T_0') = s(bool,X89)) )),
  inference(superposition,[],[f352,f472])).
fof(f1151,plain,(
  ( ! [X85,X83,X84,X82] : (sP11(X82,X83,X84) | ~p(s(bool,X83)) | s(bool,X82) = s(bool,X85) | s(bool,'T_0') = s(bool,X85)) )),
  inference(subsumption_resolution,[],[f890,f200])).
fof(f890,plain,(
  ( ! [X85,X83,X84,X82] : (~p(s(bool,'T_0')) | ~p(s(bool,X83)) | sP11(X82,X83,X84) | s(bool,X82) = s(bool,X85) | s(bool,'T_0') = s(bool,X85)) )),
  inference(superposition,[],[f351,f472])).
fof(f1150,plain,(
  ( ! [X80,X78,X81,X79] : (sP11(X78,X80,X79) | p(s(bool,X79)) | s(bool,X78) = s(bool,X81) | s(bool,'T_0') = s(bool,X81)) )),
  inference(subsumption_resolution,[],[f889,f200])).
fof(f889,plain,(
  ( ! [X80,X78,X81,X79] : (~p(s(bool,'T_0')) | p(s(bool,X79)) | sP11(X78,X80,X79) | s(bool,X78) = s(bool,X81) | s(bool,'T_0') = s(bool,X81)) )),
  inference(superposition,[],[f350,f472])).
fof(f1149,plain,(
  ( ! [X76,X74,X77,X75] : (~sP10(X76,X75,X74) | ~p(s(bool,X76)) | p(s(bool,X75)) | s(bool,X74) = s(bool,X77) | s(bool,'T_0') = s(bool,X77)) )),
  inference(subsumption_resolution,[],[f888,f200])).
fof(f888,plain,(
  ( ! [X76,X74,X77,X75] : (~p(s(bool,'T_0')) | p(s(bool,X75)) | ~p(s(bool,X76)) | ~sP10(X76,X75,X74) | s(bool,X74) = s(bool,X77) | s(bool,'T_0') = s(bool,X77)) )),
  inference(superposition,[],[f345,f472])).
fof(f1148,plain,(
  ( ! [X70,X72,X71,X73] : (~sP9(X70,X72,X71) | p(s(bool,X72)) | p(s(bool,X71)) | s(bool,X70) = s(bool,X73) | s(bool,'T_0') = s(bool,X73)) )),
  inference(subsumption_resolution,[],[f887,f200])).
fof(f887,plain,(
  ( ! [X70,X72,X71,X73] : (~p(s(bool,'T_0')) | p(s(bool,X71)) | p(s(bool,X72)) | ~sP9(X70,X72,X71) | s(bool,X70) = s(bool,X73) | s(bool,'T_0') = s(bool,X73)) )),
  inference(superposition,[],[f335,f472])).
fof(f1147,plain,(
  ( ! [X68,X66,X69,X67] : (sP9(X66,X67,X68) | ~p(s(bool,X67)) | s(bool,X66) = s(bool,X69) | s(bool,'T_0') = s(bool,X69)) )),
  inference(subsumption_resolution,[],[f886,f200])).
fof(f886,plain,(
  ( ! [X68,X66,X69,X67] : (~p(s(bool,'T_0')) | ~p(s(bool,X67)) | sP9(X66,X67,X68) | s(bool,X66) = s(bool,X69) | s(bool,'T_0') = s(bool,X69)) )),
  inference(superposition,[],[f334,f472])).
fof(f1146,plain,(
  ( ! [X64,X62,X65,X63] : (sP9(X62,X64,X63) | ~p(s(bool,X63)) | s(bool,X62) = s(bool,X65) | s(bool,'T_0') = s(bool,X65)) )),
  inference(subsumption_resolution,[],[f885,f200])).
fof(f885,plain,(
  ( ! [X64,X62,X65,X63] : (~p(s(bool,'T_0')) | ~p(s(bool,X63)) | sP9(X62,X64,X63) | s(bool,X62) = s(bool,X65) | s(bool,'T_0') = s(bool,X65)) )),
  inference(superposition,[],[f333,f472])).
fof(f1145,plain,(
  ( ! [X61,X59,X60,X58] : (~sP8(X58,X59,X60) | p(s(bool,X60)) | p(s(bool,X59)) | s(bool,X58) = s(bool,X61) | s(bool,'T_0') = s(bool,X61)) )),
  inference(subsumption_resolution,[],[f884,f200])).
fof(f884,plain,(
  ( ! [X61,X59,X60,X58] : (~p(s(bool,'T_0')) | p(s(bool,X59)) | p(s(bool,X60)) | ~sP8(X58,X59,X60) | s(bool,X58) = s(bool,X61) | s(bool,'T_0') = s(bool,X61)) )),
  inference(superposition,[],[f328,f472])).
fof(f1144,plain,(
  ( ! [X57,X54,X56,X55] : (~sP7(X56,X55,X54) | p(s(bool,X56)) | ~p(s(bool,X55)) | s(bool,X54) = s(bool,X57) | s(bool,'T_0') = s(bool,X57)) )),
  inference(subsumption_resolution,[],[f883,f200])).
fof(f883,plain,(
  ( ! [X57,X54,X56,X55] : (~p(s(bool,'T_0')) | ~p(s(bool,X55)) | p(s(bool,X56)) | ~sP7(X56,X55,X54) | s(bool,X54) = s(bool,X57) | s(bool,'T_0') = s(bool,X57)) )),
  inference(superposition,[],[f320,f472])).
fof(f1143,plain,(
  ( ! [X52,X50,X53,X51] : (~sP6(X51,X50,X52) | p(s(bool,X52)) | ~p(s(bool,X51)) | s(bool,X50) = s(bool,X53) | s(bool,'T_0') = s(bool,X53)) )),
  inference(subsumption_resolution,[],[f882,f200])).
fof(f882,plain,(
  ( ! [X52,X50,X53,X51] : (~p(s(bool,'T_0')) | ~p(s(bool,X51)) | p(s(bool,X52)) | ~sP6(X51,X50,X52) | s(bool,X50) = s(bool,X53) | s(bool,'T_0') = s(bool,X53)) )),
  inference(superposition,[],[f311,f472])).
fof(f1142,plain,(
  ( ! [X47,X48,X46,X49] : (~sP3(X47,X46,X48) | p(s(bool,X48)) | ~p(s(bool,X47)) | s(bool,X46) = s(bool,X49) | s(bool,'T_0') = s(bool,X49)) )),
  inference(subsumption_resolution,[],[f881,f200])).
fof(f881,plain,(
  ( ! [X47,X48,X46,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,X47)) | p(s(bool,X48)) | ~sP3(X47,X46,X48) | s(bool,X46) = s(bool,X49) | s(bool,'T_0') = s(bool,X49)) )),
  inference(superposition,[],[f292,f472])).
fof(f1141,plain,(
  ( ! [X45,X43,X44,X42] : (~sP4(X43,X42,X44) | p(s(bool,X44)) | ~p(s(bool,X43)) | s(bool,X42) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) )),
  inference(subsumption_resolution,[],[f880,f200])).
fof(f880,plain,(
  ( ! [X45,X43,X44,X42] : (~p(s(bool,'T_0')) | ~p(s(bool,X43)) | p(s(bool,X44)) | ~sP4(X43,X42,X44) | s(bool,X42) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) )),
  inference(superposition,[],[f288,f472])).
fof(f1140,plain,(
  ( ! [X39,X41,X38,X40] : (sP5(X39,X40,X38) | ~sP4(X39,X38,X40) | ~sP3(X40,X38,X39) | ~sP2(X38,X40,X39) | s(bool,X38) = s(bool,X41) | s(bool,'T_0') = s(bool,X41)) )),
  inference(subsumption_resolution,[],[f879,f200])).
fof(f879,plain,(
  ( ! [X39,X41,X38,X40] : (~p(s(bool,'T_0')) | sP5(X39,X40,X38) | ~sP4(X39,X38,X40) | ~sP3(X40,X38,X39) | ~sP2(X38,X40,X39) | s(bool,X38) = s(bool,X41) | s(bool,'T_0') = s(bool,X41)) )),
  inference(superposition,[],[f285,f472])).
fof(f1139,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,X36)) | sP1(X36,X35) | s(bool,X35) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) )),
  inference(subsumption_resolution,[],[f878,f200])).
fof(f878,plain,(
  ( ! [X37,X35,X36] : (~p(s(bool,'T_0')) | p(s(bool,X36)) | sP1(X36,X35) | s(bool,X35) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) )),
  inference(superposition,[],[f275,f472])).
fof(f1138,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,X33)) | sP1(X32,X33) | s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f877,f200])).
fof(f877,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,'T_0')) | p(s(bool,X33)) | sP1(X32,X33) | s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f272,f472])).
fof(f1137,plain,(
  ( ! [X30,X31,X29] : (sP0(X29,X30) | s(bool,X29) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) )),
  inference(subsumption_resolution,[],[f876,f200])).
fof(f876,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,'T_0')) | sP0(X29,X30) | s(bool,X29) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) )),
  inference(superposition,[],[f257,f472])).
fof(f1136,plain,(
  ( ! [X28,X26,X27] : (sP0(X27,X26) | s(bool,X26) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f875,f200])).
fof(f875,plain,(
  ( ! [X28,X26,X27] : (~p(s(bool,'T_0')) | sP0(X27,X26) | s(bool,X26) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f256,f472])).
fof(f1135,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,X24)) | s(bool,X23) = s(bool,X24) | s(bool,X23) = s(bool,X25) | s(bool,'T_0') = s(bool,X25)) )),
  inference(subsumption_resolution,[],[f874,f200])).
fof(f874,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,'T_0')) | ~p(s(bool,X24)) | s(bool,X23) = s(bool,X24) | s(bool,X23) = s(bool,X25) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f254,f472])).
fof(f1134,plain,(
  ( ! [X14,X13] : (s(bool,'F_0') != s(bool,X13) | s(bool,X13) = s(bool,X14) | s(bool,'T_0') = s(bool,X14)) )),
  inference(subsumption_resolution,[],[f870,f200])).
fof(f870,plain,(
  ( ! [X14,X13] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X13) | s(bool,X13) = s(bool,X14) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f220,f472])).
fof(f833,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X2)) | s(bool,X1) != s(bool,X2) | s(bool,X1) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f216,f472])).
fof(f796,plain,(
  ( ! [X109,X107,X108,X106] : (s(X108,f(s(bool,X109))) != s(X108,g(s(bool,X109))) | s(fun(bool,X108),X106) = s(fun(bool,X108),X107) | s(bool,'T_0') = s(bool,X109) | s(bool,'T_0') = s(bool,sK16(X106,X107,X108,bool))) )),
  inference(superposition,[],[f280,f472])).
fof(f743,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f472,f472])).
fof(f701,plain,(
  sP5('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f374,f371,f369,f364,f286])).
fof(f679,plain,(
  sP5('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f376,f371,f368,f200,f285])).
fof(f658,plain,(
  ~sP11('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f652,f340])).
fof(f652,plain,(
  ~sP10('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f200,f364,f200,f345])).
fof(f648,plain,(
  sP11('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f385,f200,f364,f344])).
fof(f622,plain,(
  ~sP9('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f616,f323])).
fof(f616,plain,(
  ~sP8('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f364,f364,f200,f328])).
fof(f612,plain,(
  sP9('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f382,f364,f364,f327])).
fof(f588,plain,(
  ~sP7('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f582,f304])).
fof(f582,plain,(
  ~sP6('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f364,f200,f200,f311])).
fof(f576,plain,(
  sP7('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f377,f200,f200,f307])).
fof(f573,plain,(
  ~sP5('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f564,f281])).
fof(f569,plain,(
  ~sP5('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f559,f282])).
fof(f564,plain,(
  ~sP2('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f364,f364,f364,f296])).
fof(f559,plain,(
  ~sP3('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f364,f200,f200,f292])).
fof(f549,plain,(
  ~sP4('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f364,f200,f200,f288])).
fof(f535,plain,(
  ~sP5('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f200,f386,f302])).
fof(f537,plain,(
  ~sP5('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f200,f386,f302])).
fof(f526,plain,(
  sP5('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f364,f386,f300])).
fof(f527,plain,(
  sP5('F_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f364,f386,f300])).
fof(f518,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f515,f364])).
fof(f515,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f251,f207])).
fof(f472,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f207,f207])).
fof(f448,plain,(
  ( ! [X0] : (sP11('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f200,f200,f351])).
fof(f441,plain,(
  ( ! [X0] : (sP11('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f364,f200,f350])).
fof(f439,plain,(
  ( ! [X0] : (~sP11('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f200,f364,f339])).
fof(f438,plain,(
  ( ! [X0] : (~sP11('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f364,f364,f338])).
fof(f431,plain,(
  ( ! [X0] : (sP9('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f200,f200,f334])).
fof(f428,plain,(
  ( ! [X0] : (sP9('T_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f200,f200,f333])).
fof(f425,plain,(
  ( ! [X0] : (~sP9('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f200,f364,f322])).
fof(f421,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f200,f364,f321])).
fof(f420,plain,(
  ( ! [X0] : (sP7('F_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f364,f364,f316])).
fof(f414,plain,(
  ( ! [X0] : (sP7('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f364,f364,f315])).
fof(f411,plain,(
  ( ! [X0] : (~sP7('T_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f200,f364,f306])).
fof(f410,plain,(
  ( ! [X0] : (~sP7('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f200,f364,f305])).
fof(f407,plain,(
  ~sP0('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f364,f364,f365])).
fof(f400,plain,(
  sP1('F_0','T_0')),
  inference(unit_resulting_resolution,[],[f364,f200,f275])).
fof(f397,plain,(
  sP1('T_0','F_0')),
  inference(unit_resulting_resolution,[],[f364,f200,f272])).
fof(f394,plain,(
  ~sP1('T_0','T_0')),
  inference(unit_resulting_resolution,[],[f200,f200,f271])).
fof(f393,plain,(
  ~sP1('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f364,f364,f270])).
fof(f386,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f364,f216])).
fof(f385,plain,(
  ( ! [X0,X1] : (sP10('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f364,f348])).
fof(f384,plain,(
  ( ! [X0,X1] : (sP10(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f200,f347])).
fof(f383,plain,(
  ( ! [X0,X1] : (sP10(X0,X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f364,f346])).
fof(f382,plain,(
  ( ! [X0,X1] : (sP8('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f364,f331])).
fof(f381,plain,(
  ( ! [X0,X1] : (sP8(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f200,f330])).
fof(f380,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f200,f329])).
fof(f379,plain,(
  ( ! [X0,X1] : (sP6('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f364,f314])).
fof(f378,plain,(
  ( ! [X0,X1] : (sP6(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f364,f313])).
fof(f377,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f200,f312])).
fof(f376,plain,(
  ( ! [X0,X1] : (sP2('T_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f200,f299])).
fof(f375,plain,(
  ( ! [X0,X1] : (sP2(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f200,f298])).
fof(f374,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f200,f297])).
fof(f373,plain,(
  ( ! [X0,X1] : (sP3('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f364,f295])).
fof(f372,plain,(
  ( ! [X0,X1] : (sP3(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f364,f294])).
fof(f371,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f200,f293])).
fof(f370,plain,(
  ( ! [X0,X1] : (sP4('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f364,f291])).
fof(f369,plain,(
  ( ! [X0,X1] : (sP4(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f364,f290])).
fof(f368,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f200,f289])).
fof(f367,plain,(
  ( ! [X0] : (sP0('T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f200,f257])).
fof(f366,plain,(
  ( ! [X0] : (sP0(X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f200,f256])).
fof(f349,plain,(
  ( ! [X2,X0,X1] : (sP11(X2,X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f198])).
fof(f198,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f197])).
fof(f197,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP11(X2,X0,X1))),
  inference(definition_folding,[],[f108,f125,f124])).
fof(f124,plain,(
  ! [X2,X0,X1] : (sP10(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f125,plain,(
  ! [X2,X0,X1] : (sP11(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f108,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f87])).
fof(f87,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f86])).
fof(f86,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X13,X14,X15] : ((p(s(bool,X15)) <=> (p(s(bool,X14)) => p(s(bool,X13)))) <=> ((p(s(bool,X15)) | p(s(bool,X14))) & (p(s(bool,X15)) | ~p(s(bool,X13))) & (~p(s(bool,X14)) | p(s(bool,X13)) | ~p(s(bool,X15)))))),
  file('thm.pred_set.UNION_DIFF_2','thm.sat.sat_27_prove_ltb')).
fof(f350,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f198])).
fof(f351,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f198])).
fof(f352,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f198])).
fof(f345,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f196])).
fof(f196,plain,(
  ! [X0,X1,X2] : ((~sP10(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP10(X0,X1,X2)))),
  inference(rectify,[],[f195])).
fof(f195,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(flattening,[],[f194])).
fof(f194,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(nnf_transformation,[],[f124])).
fof(f346,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f196])).
fof(f347,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f196])).
fof(f348,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f196])).
fof(f338,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f193])).
fof(f193,plain,(
  ! [X0,X1,X2] : ((~sP11(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP10(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP10(X0,X1,X2) | sP11(X0,X1,X2)))),
  inference(rectify,[],[f192])).
fof(f192,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1) | sP11(X2,X0,X1)))),
  inference(flattening,[],[f191])).
fof(f191,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1)) | sP11(X2,X0,X1)))),
  inference(nnf_transformation,[],[f125])).
fof(f339,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f193])).
fof(f340,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f193])).
fof(f355,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP10(X0,X1,X2) | sP11(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f341])).
fof(f341,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f193])).
fof(f342,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f193])).
fof(f343,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f193])).
fof(f344,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f193])).
fof(f332,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f190])).
fof(f190,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f189])).
fof(f189,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP9(X2,X0,X1))),
  inference(definition_folding,[],[f85,f122,f121])).
fof(f121,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f122,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f85,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f84])).
fof(f84,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X13,X14,X15] : ((p(s(bool,X15)) <=> (p(s(bool,X14)) | p(s(bool,X13)))) <=> ((p(s(bool,X15)) | ~p(s(bool,X14))) & (p(s(bool,X15)) | ~p(s(bool,X13))) & (p(s(bool,X14)) | p(s(bool,X13)) | ~p(s(bool,X15)))))),
  file('thm.pred_set.UNION_DIFF_2','thm.sat.sat_25_prove_ltb')).
fof(f333,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f190])).
fof(f334,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f190])).
fof(f335,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)) | ~sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f190])).
fof(f328,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f188])).
fof(f188,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f187])).
fof(f187,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f186])).
fof(f186,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f121])).
fof(f329,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f188])).
fof(f330,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f188])).
fof(f331,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f188])).
fof(f321,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f185])).
fof(f185,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP8(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f184])).
fof(f184,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1) | sP9(X2,X0,X1)))),
  inference(flattening,[],[f183])).
fof(f183,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1)) | sP9(X2,X0,X1)))),
  inference(nnf_transformation,[],[f122])).
fof(f322,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f185])).
fof(f323,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f185])).
fof(f356,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f324])).
fof(f324,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f185])).
fof(f325,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f185])).
fof(f326,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f185])).
fof(f327,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f185])).
fof(f315,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f182])).
fof(f182,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f181])).
fof(f181,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP7(X2,X0,X1))),
  inference(definition_folding,[],[f83,f119,f118])).
fof(f118,plain,(
  ! [X0,X1,X2] : (sP6(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f119,plain,(
  ! [X2,X0,X1] : (sP7(X2,X0,X1) <=> (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f83,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f82])).
fof(f82,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X13,X14,X15] : ((p(s(bool,X15)) <=> (p(s(bool,X14)) & p(s(bool,X13)))) <=> ((p(s(bool,X15)) | ~p(s(bool,X14)) | ~p(s(bool,X13))) & (p(s(bool,X14)) | ~p(s(bool,X15))) & (p(s(bool,X13)) | ~p(s(bool,X15)))))),
  file('thm.pred_set.UNION_DIFF_2','thm.sat.sat_23_prove_ltb')).
fof(f316,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f182])).
fof(f317,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f182])).
fof(f320,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP7(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f182])).
fof(f311,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f180])).
fof(f180,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(flattening,[],[f179])).
fof(f179,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(nnf_transformation,[],[f118])).
fof(f312,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f180])).
fof(f313,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f180])).
fof(f314,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f180])).
fof(f304,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f178])).
fof(f178,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | (sP6(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP6(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f177])).
fof(f177,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP7(X2,X0,X1)))),
  inference(flattening,[],[f176])).
fof(f176,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP7(X2,X0,X1)))),
  inference(nnf_transformation,[],[f119])).
fof(f305,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f178])).
fof(f306,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f178])).
fof(f307,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f178])).
fof(f357,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | ~sP6(X1,X2,X0)) )),
  inference(duplicate_literal_removal,[],[f310])).
fof(f310,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X0)) | ~sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f178])).
fof(f300,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f175])).
fof(f175,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP5(X2,X1,X0)) & (~sP5(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f117])).
fof(f117,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP5(X2,X1,X0))),
  inference(definition_folding,[],[f81,f116,f115,f114,f113])).
fof(f113,plain,(
  ! [X0,X1,X2] : (sP2(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f114,plain,(
  ! [X1,X0,X2] : (sP3(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f115,plain,(
  ! [X2,X0,X1] : (sP4(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f116,plain,(
  ! [X2,X1,X0] : (sP5(X2,X1,X0) <=> (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f81,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f80])).
fof(f80,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X13,X14,X15] : ((p(s(bool,X15)) <=> s(bool,X13) = s(bool,X14)) <=> ((p(s(bool,X15)) | p(s(bool,X14)) | p(s(bool,X13))) & (p(s(bool,X15)) | ~p(s(bool,X13)) | ~p(s(bool,X14))) & (p(s(bool,X14)) | ~p(s(bool,X13)) | ~p(s(bool,X15))) & (p(s(bool,X13)) | ~p(s(bool,X14)) | ~p(s(bool,X15)))))),
  file('thm.pred_set.UNION_DIFF_2','thm.sat.sat_2|49|_prove_ltb')).
fof(f301,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f175])).
fof(f302,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | ~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f175])).
fof(f303,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f175])).
fof(f296,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f174])).
fof(f174,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(flattening,[],[f173])).
fof(f173,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(nnf_transformation,[],[f113])).
fof(f297,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f174])).
fof(f298,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f174])).
fof(f299,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f174])).
fof(f292,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f172])).
fof(f172,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f171])).
fof(f171,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(flattening,[],[f170])).
fof(f170,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(nnf_transformation,[],[f114])).
fof(f293,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f172])).
fof(f294,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f172])).
fof(f295,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f172])).
fof(f288,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f169])).
fof(f169,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f168])).
fof(f168,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(flattening,[],[f167])).
fof(f167,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(nnf_transformation,[],[f115])).
fof(f289,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f169])).
fof(f290,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f169])).
fof(f291,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f169])).
fof(f281,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f166])).
fof(f166,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | (sP2(X2,X1,X0) & sP3(X1,X2,X0) & sP4(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP2(X2,X1,X0) | ~sP3(X1,X2,X0) | ~sP4(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f165])).
fof(f165,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP5(X2,X1,X0)))),
  inference(flattening,[],[f164])).
fof(f164,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP5(X2,X1,X0)))),
  inference(nnf_transformation,[],[f116])).
fof(f282,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP3(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f166])).
fof(f283,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP4(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f166])).
fof(f284,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f166])).
fof(f285,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP5(X0,X1,X2) | ~sP4(X0,X2,X1) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f166])).
fof(f286,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X2,X1) | p(s(bool,X1)) | sP5(X0,X1,X2) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f166])).
fof(f287,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X2,X1) | p(s(bool,X0)) | sP5(X0,X1,X2) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f166])).
fof(f280,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK16(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK16(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f163])).
fof(f163,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK16(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK16(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK16])],[f107])).
fof(f107,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f78])).
fof(f78,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.pred_set.UNION_DIFF_2',aHL_EXT)).
fof(f270,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f161])).
fof(f161,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1)))),
  inference(rectify,[],[f160])).
fof(f160,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0)))),
  inference(flattening,[],[f159])).
fof(f159,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP1(X1,X0)))),
  inference(nnf_transformation,[],[f111])).
fof(f111,plain,(
  ! [X1,X0] : (sP1(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f271,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f161])).
fof(f272,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f161])).
fof(f275,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f161])).
fof(f365,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f255,f364])).
fof(f255,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f154])).
fof(f154,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f153])).
fof(f153,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f256,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f154])).
fof(f257,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f154])).
fof(f251,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f101])).
fof(f101,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X9,X10] : ((p(s(bool,X9)) => p(s(bool,X10))) => ((p(s(bool,X10)) => p(s(bool,X9))) => s(bool,X9) = s(bool,X10)))),
  file('thm.pred_set.UNION_DIFF_2','thm.bool.IMP_ANTISYM_AX')).
fof(f254,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f102])).
fof(f216,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f130])).
fof(f130,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X8] : (s(bool,'T_0') = s(bool,X8) <=> p(s(bool,X8)))),
  file('thm.pred_set.UNION_DIFF_2','thm.bool.EQ_CLAUSES|split|1')).
fof(f217,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f130])).
fof(f207,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.pred_set.UNION_DIFF_2',aHL_BOOL_CASES)).
fof(f199,plain,(
  s(fun(sK12,bool),sK13) != s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))),
  inference(cnf_transformation,[],[f127])).
fof(f127,plain,(
  s(fun(sK12,bool),sK13) != s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13,sK14])],[f88])).
fof(f88,plain,(
  ? [X0,X1,X2] : s(fun(X0,bool),X1) != s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))),
  inference(ennf_transformation,[],[f37])).
fof(f37,plain,(
  ~! [X0,X1,X2] : s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))),
  inference(flattening,[],[f36])).
fof(f36,plain,(
  ~! [X0] : ! [X1,X2] : s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : s(fun(X6,bool),X7) = s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),X7),s(fun(X6,bool),'const.pred_set.DIFF_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : s(fun(X6,bool),X7) = s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),X7),s(fun(X6,bool),'const.pred_set.DIFF_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))))),
  file('thm.pred_set.UNION_DIFF_2','thm.pred_set.UNION_DIFF_2')).
fof(f13101,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | $spl76),
  inference(forward_demodulation,[],[f5332,f2265])).
fof(f13096,plain,(
  $false | ($spl76 | $spl78 | $spl106)),
  inference(subsumption_resolution,[],[f13095,f12960])).
fof(f13095,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | $spl76),
  inference(forward_demodulation,[],[f10785,f2265])).
fof(f10785,plain,(
  s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) != s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))))),
  inference(unit_resulting_resolution,[],[f199,f244])).
fof(f13094,plain,(
  $false | ($spl58 | $spl76 | $spl78 | $spl106)),
  inference(global_subsumption,[],[f13093,f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1718,f1719,f1720,f1721,f1722,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2269,f2314,f2313,f2312,f2309,f2308,f2323,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f3197,f3209,f3223,f3866,f4029,f4032,f721,f720,f2303,f6433,f6436,f7404,f7406,f2266,f5332,f7449,f7446,f1576,f1574,f1556,f1538,f1515,f1512,f1492,f1489,f1474,f1471,f1451,f1449,f1579,f1559,f1554,f1536,f1510,f1487,f1469,f1454,f8521,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f1756,f896,f797,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f9622,f9619,f12379,f12383,f12400,f12399,f12420,f12419,f12444,f12443,f12579,f12578,f2265,f7405,f12960,f13071,f13076])).
fof(f13076,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | ($spl76 | $spl78 | $spl106)),
  inference(subsumption_resolution,[],[f13013,f12960])).
fof(f13013,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X0,sK15(X2,X1,X0)),s(fun(X0,bool),X1))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | ($spl76 | $spl78 | $spl106)),
  inference(backward_demodulation,[],[f12960,f244])).
fof(f13071,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl76 | $spl78 | $spl106)),
  inference(subsumption_resolution,[],[f12961,f12960])).
fof(f12961,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | ($spl76 | $spl78 | $spl106)),
  inference(backward_demodulation,[],[f12960,f220])).
fof(f13093,plain,(
  ( ! [X159,X160] : (s(fun(bool,bool),X159) = s(fun(bool,bool),X160)) ) | ($spl58 | $spl76 | $spl78 | $spl106)),
  inference(subsumption_resolution,[],[f13092,f200])).
fof(f13092,plain,(
  ( ! [X159,X160] : (~p(s(bool,'T_0')) | s(fun(bool,bool),X159) = s(fun(bool,bool),X160)) ) | ($spl58 | $spl76 | $spl78 | $spl106)),
  inference(forward_demodulation,[],[f13091,f12960])).
fof(f13091,plain,(
  ( ! [X161,X159,X162,X160] : (~p(s(bool,'const.bool.IN_2'(s(bool,sK16(X159,X160,bool,bool)),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X162),s(fun(bool,bool),X161)))))) | s(fun(bool,bool),X159) = s(fun(bool,bool),X160)) ) | ($spl58 | $spl76 | $spl78 | $spl106)),
  inference(subsumption_resolution,[],[f13038,f200])).
fof(f13038,plain,(
  ( ! [X161,X159,X162,X160] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(bool,sK16(X159,X160,bool,bool)),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X162),s(fun(bool,bool),X161)))))) | s(fun(bool,bool),X159) = s(fun(bool,bool),X160)) ) | ($spl58 | $spl76 | $spl78 | $spl106)),
  inference(backward_demodulation,[],[f12960,f6958])).
fof(f6958,plain,(
  ( ! [X161,X159,X162,X160] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X161)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,sK16(X159,X160,bool,bool)),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X162),s(fun(bool,bool),X161)))))) | s(fun(bool,bool),X159) = s(fun(bool,bool),X160)) ) | $spl58),
  inference(superposition,[],[f246,f1724])).
fof(f13080,plain,(
  $false | ($spl76 | $spl78 | $spl106)),
  inference(subsumption_resolution,[],[f13079,f200])).
fof(f13079,plain,(
  ~p(s(bool,'T_0')) | ($spl76 | $spl78 | $spl106)),
  inference(forward_demodulation,[],[f13078,f12960])).
fof(f13078,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))) ) | ($spl76 | $spl78 | $spl106)),
  inference(subsumption_resolution,[],[f13016,f200])).
fof(f13016,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))) ) | ($spl76 | $spl78 | $spl106)),
  inference(backward_demodulation,[],[f12960,f246])).
fof(f13077,plain,(
  $false | ($spl76 | $spl78 | $spl106)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1718,f1719,f1720,f1721,f1722,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2269,f2314,f2313,f2312,f2309,f2308,f2323,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f3197,f3209,f3223,f3866,f4029,f4032,f721,f720,f2303,f6433,f6436,f7404,f7406,f2266,f5332,f7449,f7446,f1576,f1574,f1556,f1538,f1515,f1512,f1492,f1489,f1474,f1471,f1451,f1449,f1579,f1559,f1554,f1536,f1510,f1487,f1469,f1454,f8521,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f1756,f896,f797,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f9622,f9619,f12379,f12383,f12400,f12399,f12420,f12419,f12444,f12443,f12579,f12578,f2265,f7405,f12960,f13071,f13076])).
fof(f13075,plain,(
  $false | ($spl58 | $spl76 | $spl78 | $spl106)),
  inference(global_subsumption,[],[f13074,f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1718,f1719,f1720,f1721,f1722,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2269,f2314,f2313,f2312,f2309,f2308,f2323,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f3197,f3209,f3223,f3866,f4029,f4032,f721,f720,f2303,f6433,f6436,f7404,f7406,f2266,f5332,f7449,f7446,f1576,f1574,f1556,f1538,f1515,f1512,f1492,f1489,f1474,f1471,f1451,f1449,f1579,f1559,f1554,f1536,f1510,f1487,f1469,f1454,f8521,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f1756,f896,f797,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f9622,f9619,f12379,f12383,f12400,f12399,f12420,f12419,f12444,f12443,f12579,f12578,f2265,f7405,f12960,f13071])).
fof(f13074,plain,(
  ( ! [X6,X7] : (s(fun(bool,bool),X6) = s(fun(bool,bool),X7)) ) | ($spl58 | $spl76 | $spl78 | $spl106)),
  inference(subsumption_resolution,[],[f13007,f12960])).
fof(f13007,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,sK16(X6,X7,bool,bool)) | s(fun(bool,bool),X6) = s(fun(bool,bool),X7)) ) | ($spl58 | $spl76 | $spl78 | $spl106)),
  inference(backward_demodulation,[],[f12960,f6971])).
fof(f6971,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,sK16(X6,X7,bool,bool)) | s(fun(bool,bool),X6) = s(fun(bool,bool),X7)) ) | $spl58),
  inference(subsumption_resolution,[],[f6915,f200])).
fof(f6915,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,sK16(X6,X7,bool,bool)) | s(fun(bool,bool),X6) = s(fun(bool,bool),X7)) ) | $spl58),
  inference(superposition,[],[f220,f1724])).
fof(f13073,plain,(
  $false | ($spl76 | $spl78 | $spl106)),
  inference(subsumption_resolution,[],[f12963,f200])).
fof(f12963,plain,(
  ~p(s(bool,'T_0')) | ($spl76 | $spl78 | $spl106)),
  inference(backward_demodulation,[],[f12960,f364])).
fof(f13072,plain,(
  $false | ($spl76 | $spl78 | $spl106)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1718,f1719,f1720,f1721,f1722,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2269,f2314,f2313,f2312,f2309,f2308,f2323,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f3197,f3209,f3223,f3866,f4029,f4032,f721,f720,f2303,f6433,f6436,f7404,f7406,f2266,f5332,f7449,f7446,f1576,f1574,f1556,f1538,f1515,f1512,f1492,f1489,f1474,f1471,f1451,f1449,f1579,f1559,f1554,f1536,f1510,f1487,f1469,f1454,f8521,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f1756,f896,f797,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f9622,f9619,f12379,f12383,f12400,f12399,f12420,f12419,f12444,f12443,f12579,f12578,f2265,f7405,f12960,f13071])).
fof(f12959,plain,(
  $false | ($spl76 | $spl81)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1718,f1719,f1720,f1721,f1722,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2269,f2314,f2313,f2312,f2309,f2308,f2323,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f3197,f3209,f3223,f3866,f4029,f4032,f721,f720,f2303,f6433,f6436,f7404,f7406,f5332,f2268,f7449,f7446,f9249,f9247,f9245,f9244,f9239,f9238,f9237,f9236,f9234,f9232,f9231,f9210,f9209,f1576,f1574,f1556,f1538,f1515,f1512,f1492,f1489,f1474,f1471,f1451,f1449,f1579,f1559,f1554,f1536,f1510,f1487,f1469,f1454,f8521,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f1756,f896,f797,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f9622,f9619,f11695,f11894,f11896,f12379,f12383,f12400,f12399,f12420,f12419,f12444,f12443,f12457,f12456,f12455,f12454,f12535,f12536,f12537,f12538,f12544,f12545,f12546,f12547,f12548,f12549,f12550,f12551,f12552,f12553,f12554,f12555,f12556,f12557,f12558,f12559,f12560,f12561,f12562,f12573,f12579,f12578,f12583,f12586,f10803,f9258,f2265])).
fof(f9258,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))))))) | $spl81),
  inference(backward_demodulation,[],[f9234,f5332])).
fof(f10803,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))))))) | $spl81),
  inference(forward_demodulation,[],[f10785,f9234])).
fof(f12586,plain,(
  ~sP3('T_0','const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)),'F_0') | $spl81),
  inference(unit_resulting_resolution,[],[f364,f375,f370,f12456,f287])).
fof(f12583,plain,(
  sP4('T_0','T_0','const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | $spl81),
  inference(unit_resulting_resolution,[],[f12455,f283])).
fof(f12573,plain,(
  ~sP3('const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)),'T_0','F_0') | $spl81),
  inference(unit_resulting_resolution,[],[f200,f376,f370,f12454,f285])).
fof(f12562,plain,(
  ( ! [X54,X56,X55] : (p(s(bool,X54)) | ~sP11(X55,X54,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | s(bool,X55) = s(bool,X56) | s(bool,'T_0') = s(bool,X56)) ) | $spl81),
  inference(subsumption_resolution,[],[f12505,f200])).
fof(f12505,plain,(
  ( ! [X54,X56,X55] : (~p(s(bool,'T_0')) | p(s(bool,X54)) | ~sP11(X55,X54,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | s(bool,X55) = s(bool,X56) | s(bool,'T_0') = s(bool,X56)) ) | $spl81),
  inference(superposition,[],[f1152,f9234])).
fof(f12561,plain,(
  ( ! [X48,X49] : (s(bool,X48) = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | s(bool,X48) = s(bool,X49) | s(bool,'T_0') = s(bool,X49)) ) | $spl81),
  inference(subsumption_resolution,[],[f12502,f200])).
fof(f12502,plain,(
  ( ! [X48,X49] : (~p(s(bool,'T_0')) | s(bool,X48) = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | s(bool,X48) = s(bool,X49) | s(bool,'T_0') = s(bool,X49)) ) | $spl81),
  inference(superposition,[],[f1135,f9234])).
fof(f12560,plain,(
  ( ! [X41,X40] : (~sP8('const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)),X40,X41) | sP9('const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)),X40,X41)) ) | $spl81),
  inference(subsumption_resolution,[],[f12494,f200])).
fof(f12494,plain,(
  ( ! [X41,X40] : (~p(s(bool,'T_0')) | ~sP8('const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)),X40,X41) | sP9('const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)),X40,X41)) ) | $spl81),
  inference(superposition,[],[f356,f9234])).
fof(f12559,plain,(
  ( ! [X39,X38] : (~sP10('const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)),X38,X39) | sP11('const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)),X38,X39)) ) | $spl81),
  inference(subsumption_resolution,[],[f12493,f200])).
fof(f12493,plain,(
  ( ! [X39,X38] : (~p(s(bool,'T_0')) | ~sP10('const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)),X38,X39) | sP11('const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)),X38,X39)) ) | $spl81),
  inference(superposition,[],[f355,f9234])).
fof(f12558,plain,(
  ( ! [X37,X36] : (~p(s(bool,X36)) | p(s(bool,X37)) | ~sP11('const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)),X37,X36)) ) | $spl81),
  inference(subsumption_resolution,[],[f12492,f200])).
fof(f12492,plain,(
  ( ! [X37,X36] : (~p(s(bool,'T_0')) | ~p(s(bool,X36)) | p(s(bool,X37)) | ~sP11('const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)),X37,X36)) ) | $spl81),
  inference(superposition,[],[f352,f9234])).
fof(f12557,plain,(
  ( ! [X35,X34] : (~p(s(bool,X34)) | sP11('const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)),X34,X35)) ) | $spl81),
  inference(subsumption_resolution,[],[f12491,f200])).
fof(f12491,plain,(
  ( ! [X35,X34] : (~p(s(bool,'T_0')) | ~p(s(bool,X34)) | sP11('const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)),X34,X35)) ) | $spl81),
  inference(superposition,[],[f351,f9234])).
fof(f12556,plain,(
  ( ! [X33,X32] : (p(s(bool,X32)) | sP11('const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)),X33,X32)) ) | $spl81),
  inference(subsumption_resolution,[],[f12490,f200])).
fof(f12490,plain,(
  ( ! [X33,X32] : (~p(s(bool,'T_0')) | p(s(bool,X32)) | sP11('const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)),X33,X32)) ) | $spl81),
  inference(superposition,[],[f350,f9234])).
fof(f12555,plain,(
  ( ! [X30,X31] : (p(s(bool,X30)) | ~p(s(bool,X31)) | ~sP10(X31,X30,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)))) ) | $spl81),
  inference(subsumption_resolution,[],[f12489,f200])).
fof(f12489,plain,(
  ( ! [X30,X31] : (~p(s(bool,'T_0')) | p(s(bool,X30)) | ~p(s(bool,X31)) | ~sP10(X31,X30,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)))) ) | $spl81),
  inference(superposition,[],[f345,f9234])).
fof(f12554,plain,(
  ( ! [X28,X29] : (p(s(bool,X28)) | p(s(bool,X29)) | ~sP9('const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)),X29,X28)) ) | $spl81),
  inference(subsumption_resolution,[],[f12488,f200])).
fof(f12488,plain,(
  ( ! [X28,X29] : (~p(s(bool,'T_0')) | p(s(bool,X28)) | p(s(bool,X29)) | ~sP9('const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)),X29,X28)) ) | $spl81),
  inference(superposition,[],[f335,f9234])).
fof(f12553,plain,(
  ( ! [X26,X27] : (~p(s(bool,X26)) | sP9('const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)),X26,X27)) ) | $spl81),
  inference(subsumption_resolution,[],[f12487,f200])).
fof(f12487,plain,(
  ( ! [X26,X27] : (~p(s(bool,'T_0')) | ~p(s(bool,X26)) | sP9('const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)),X26,X27)) ) | $spl81),
  inference(superposition,[],[f334,f9234])).
fof(f12552,plain,(
  ( ! [X24,X25] : (~p(s(bool,X24)) | sP9('const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)),X25,X24)) ) | $spl81),
  inference(subsumption_resolution,[],[f12486,f200])).
fof(f12486,plain,(
  ( ! [X24,X25] : (~p(s(bool,'T_0')) | ~p(s(bool,X24)) | sP9('const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)),X25,X24)) ) | $spl81),
  inference(superposition,[],[f333,f9234])).
fof(f12551,plain,(
  ( ! [X23,X22] : (p(s(bool,X22)) | p(s(bool,X23)) | ~sP8('const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)),X22,X23)) ) | $spl81),
  inference(subsumption_resolution,[],[f12485,f200])).
fof(f12485,plain,(
  ( ! [X23,X22] : (~p(s(bool,'T_0')) | p(s(bool,X22)) | p(s(bool,X23)) | ~sP8('const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)),X22,X23)) ) | $spl81),
  inference(superposition,[],[f328,f9234])).
fof(f12550,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | p(s(bool,X21)) | ~sP7(X21,X20,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)))) ) | $spl81),
  inference(subsumption_resolution,[],[f12484,f200])).
fof(f12484,plain,(
  ( ! [X21,X20] : (~p(s(bool,'T_0')) | ~p(s(bool,X20)) | p(s(bool,X21)) | ~sP7(X21,X20,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)))) ) | $spl81),
  inference(superposition,[],[f320,f9234])).
fof(f12549,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | p(s(bool,X19)) | ~sP6(X18,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)),X19)) ) | $spl81),
  inference(subsumption_resolution,[],[f12483,f200])).
fof(f12483,plain,(
  ( ! [X19,X18] : (~p(s(bool,'T_0')) | ~p(s(bool,X18)) | p(s(bool,X19)) | ~sP6(X18,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)),X19)) ) | $spl81),
  inference(superposition,[],[f311,f9234])).
fof(f12548,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | p(s(bool,X17)) | ~sP3(X16,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)),X17)) ) | $spl81),
  inference(subsumption_resolution,[],[f12482,f200])).
fof(f12482,plain,(
  ( ! [X17,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X16)) | p(s(bool,X17)) | ~sP3(X16,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)),X17)) ) | $spl81),
  inference(superposition,[],[f292,f9234])).
fof(f12547,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | p(s(bool,X15)) | ~sP4(X14,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)),X15)) ) | $spl81),
  inference(subsumption_resolution,[],[f12481,f200])).
fof(f12481,plain,(
  ( ! [X14,X15] : (~p(s(bool,'T_0')) | ~p(s(bool,X14)) | p(s(bool,X15)) | ~sP4(X14,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)),X15)) ) | $spl81),
  inference(superposition,[],[f288,f9234])).
fof(f12546,plain,(
  ( ! [X12,X13] : (sP5(X12,X13,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | ~sP4(X12,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)),X13) | ~sP3(X13,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)),X12) | ~sP2('const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)),X13,X12)) ) | $spl81),
  inference(subsumption_resolution,[],[f12480,f200])).
fof(f12480,plain,(
  ( ! [X12,X13] : (~p(s(bool,'T_0')) | sP5(X12,X13,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | ~sP4(X12,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)),X13) | ~sP3(X13,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)),X12) | ~sP2('const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)),X13,X12)) ) | $spl81),
  inference(superposition,[],[f285,f9234])).
fof(f12545,plain,(
  ( ! [X11] : (p(s(bool,X11)) | sP1(X11,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)))) ) | $spl81),
  inference(subsumption_resolution,[],[f12479,f200])).
fof(f12479,plain,(
  ( ! [X11] : (~p(s(bool,'T_0')) | p(s(bool,X11)) | sP1(X11,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)))) ) | $spl81),
  inference(superposition,[],[f275,f9234])).
fof(f12544,plain,(
  ( ! [X10] : (p(s(bool,X10)) | sP1('const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)),X10)) ) | $spl81),
  inference(subsumption_resolution,[],[f12478,f200])).
fof(f12478,plain,(
  ( ! [X10] : (~p(s(bool,'T_0')) | p(s(bool,X10)) | sP1('const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)),X10)) ) | $spl81),
  inference(superposition,[],[f272,f9234])).
fof(f12538,plain,(
  ( ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),X4),s(fun(sK12,bool),sK13))))))) ) | $spl81),
  inference(subsumption_resolution,[],[f12467,f200])).
fof(f12467,plain,(
  ( ! [X4] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),X4),s(fun(sK12,bool),sK13))))))) ) | $spl81),
  inference(superposition,[],[f246,f9234])).
fof(f12537,plain,(
  ( ! [X3] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),X3))))))) ) | $spl81),
  inference(subsumption_resolution,[],[f12466,f200])).
fof(f12466,plain,(
  ( ! [X3] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),X3))))))) ) | $spl81),
  inference(superposition,[],[f249,f9234])).
fof(f12536,plain,(
  ( ! [X2] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),X2),s(fun(sK12,bool),sK13))))))) ) | $spl81),
  inference(subsumption_resolution,[],[f12465,f200])).
fof(f12465,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),X2),s(fun(sK12,bool),sK13))))))) ) | $spl81),
  inference(superposition,[],[f250,f9234])).
fof(f12535,plain,(
  ( ! [X1] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),X1))))))) ) | $spl81),
  inference(subsumption_resolution,[],[f12464,f200])).
fof(f12464,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),X1))))))) ) | $spl81),
  inference(superposition,[],[f247,f9234])).
fof(f12454,plain,(
  ~sP5('F_0','const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)),'T_0') | $spl81),
  inference(unit_resulting_resolution,[],[f364,f9234,f303])).
fof(f12455,plain,(
  sP5('T_0','const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)),'T_0') | $spl81),
  inference(unit_resulting_resolution,[],[f200,f9234,f301])).
fof(f12456,plain,(
  ~sP5('F_0','T_0','const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | $spl81),
  inference(unit_resulting_resolution,[],[f364,f9234,f303])).
fof(f12457,plain,(
  sP5('T_0','T_0','const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | $spl81),
  inference(unit_resulting_resolution,[],[f200,f9234,f301])).
fof(f11896,plain,(
  ~sP4('T_0','const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)),'F_0') | $spl81),
  inference(unit_resulting_resolution,[],[f364,f374,f373,f9232,f286])).
fof(f11894,plain,(
  sP2('const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)),'F_0','F_0') | $spl81),
  inference(unit_resulting_resolution,[],[f9231,f281])).
fof(f11695,plain,(
  sP2('F_0','const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)),'F_0') | $spl81),
  inference(unit_resulting_resolution,[],[f9209,f281])).
fof(f9209,plain,(
  sP5('F_0','const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)),'F_0') | $spl81),
  inference(unit_resulting_resolution,[],[f364,f2268,f300])).
fof(f9210,plain,(
  ~sP5('T_0','const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)),'F_0') | $spl81),
  inference(unit_resulting_resolution,[],[f200,f2268,f302])).
fof(f9231,plain,(
  sP5('F_0','F_0','const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | $spl81),
  inference(unit_resulting_resolution,[],[f364,f2268,f300])).
fof(f9232,plain,(
  ~sP5('T_0','F_0','const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | $spl81),
  inference(unit_resulting_resolution,[],[f200,f2268,f302])).
fof(f9234,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | $spl81),
  inference(unit_resulting_resolution,[],[f364,f2268,f518])).
fof(f9236,plain,(
  ( ! [X0] : (sP0(X0,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)))) ) | $spl81),
  inference(unit_resulting_resolution,[],[f386,f2268,f1136])).
fof(f9237,plain,(
  ( ! [X0] : (sP0('const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)),X0)) ) | $spl81),
  inference(unit_resulting_resolution,[],[f386,f2268,f1137])).
fof(f9238,plain,(
  sP1('const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)),'F_0') | $spl81),
  inference(unit_resulting_resolution,[],[f364,f386,f2268,f1138])).
fof(f9239,plain,(
  sP1('F_0','const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | $spl81),
  inference(unit_resulting_resolution,[],[f364,f386,f2268,f1139])).
fof(f9244,plain,(
  ( ! [X0] : (sP11('const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)),X0,'F_0')) ) | $spl81),
  inference(unit_resulting_resolution,[],[f386,f383,f2268,f1153])).
fof(f9245,plain,(
  ( ! [X0] : (sP11('const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)),'T_0',X0)) ) | $spl81),
  inference(unit_resulting_resolution,[],[f386,f384,f2268,f1153])).
fof(f9247,plain,(
  ( ! [X0] : (sP9('const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)),X0,'T_0')) ) | $spl81),
  inference(unit_resulting_resolution,[],[f386,f380,f2268,f1154])).
fof(f9249,plain,(
  ( ! [X0] : (sP9('const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)),'T_0',X0)) ) | $spl81),
  inference(unit_resulting_resolution,[],[f386,f381,f2268,f1154])).
fof(f12958,plain,(
  $false | ($spl76 | $spl81 | $spl106)),
  inference(global_subsumption,[],[f7405,f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1718,f1719,f1720,f1721,f1722,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2265,f2269,f2314,f2313,f2312,f2309,f2308,f2323,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f3197,f3209,f3223,f3866,f4029,f4032,f721,f720,f2303,f6433,f6436,f7404,f7406,f5332,f2268,f7449,f7446,f9249,f9247,f9245,f9244,f9239,f9238,f9237,f9236,f9234,f9232,f9231,f9210,f9209,f1576,f1574,f1556,f1538,f1515,f1512,f1492,f1489,f1474,f1471,f1451,f1449,f1579,f1559,f1554,f1536,f1510,f1487,f1469,f1454,f8521,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f1756,f896,f797,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f9622,f9619,f11695,f11894,f11896,f12379,f12383,f12400,f12399,f12420,f12419,f12444,f12443,f12457,f12456,f12455,f12454,f12535,f12536,f12537,f12538,f12544,f12545,f12546,f12547,f12548,f12549,f12550,f12551,f12552,f12553,f12554,f12555,f12556,f12557,f12558,f12559,f12560,f12561,f12562,f12573,f12579,f12578,f12583,f12586,f10803,f9258])).
fof(f12957,plain,(
  $false | ($spl76 | $spl81 | $spl108)),
  inference(global_subsumption,[],[f7407,f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1718,f1719,f1720,f1721,f1722,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2265,f2269,f2314,f2313,f2312,f2309,f2308,f2323,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f3197,f3209,f3223,f3866,f4029,f4032,f721,f720,f2303,f6433,f6436,f7404,f7406,f5332,f2268,f7449,f7446,f9249,f9247,f9245,f9244,f9239,f9238,f9237,f9236,f9234,f9232,f9231,f9210,f9209,f1576,f1574,f1556,f1538,f1515,f1512,f1492,f1489,f1474,f1471,f1451,f1449,f1579,f1559,f1554,f1536,f1510,f1487,f1469,f1454,f8521,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f1756,f896,f797,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f9622,f9619,f11695,f11894,f11896,f12379,f12383,f12400,f12399,f12420,f12419,f12444,f12443,f12457,f12456,f12455,f12454,f12535,f12536,f12537,f12538,f12544,f12545,f12546,f12547,f12548,f12549,f12550,f12551,f12552,f12553,f12554,f12555,f12556,f12557,f12558,f12559,f12560,f12561,f12562,f12573,f12579,f12578,f12583,f12586,f10803,f9258])).
fof(f12956,plain,(
  $false | ($spl76 | $spl81)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1718,f1719,f1720,f1721,f1722,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2265,f2269,f2314,f2313,f2312,f2309,f2308,f2323,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f3197,f3209,f3223,f3866,f4029,f4032,f721,f720,f2303,f6433,f6436,f7404,f7406,f5332,f2268,f7449,f7446,f9249,f9247,f9245,f9244,f9239,f9238,f9237,f9236,f9234,f9232,f9231,f9210,f9209,f1576,f1574,f1556,f1538,f1515,f1512,f1492,f1489,f1474,f1471,f1451,f1449,f1579,f1559,f1554,f1536,f1510,f1487,f1469,f1454,f8521,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f1756,f896,f797,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f9622,f9619,f11695,f11894,f11896,f12379,f12383,f12400,f12399,f12420,f12419,f12444,f12443,f12457,f12456,f12455,f12454,f12535,f12536,f12537,f12538,f12544,f12545,f12546,f12547,f12548,f12549,f12550,f12551,f12552,f12553,f12554,f12555,f12556,f12557,f12558,f12559,f12560,f12561,f12562,f12573,f12579,f12578,f12583,f12586,f10803,f9258])).
fof(f12955,plain,(
  $false | ($spl76 | $spl81)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1718,f1719,f1720,f1721,f1722,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2265,f2269,f2314,f2313,f2312,f2309,f2308,f2323,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f3197,f3209,f3223,f3866,f4029,f4032,f721,f720,f2303,f6433,f6436,f7404,f7406,f5332,f2268,f7449,f7446,f9249,f9247,f9245,f9244,f9239,f9238,f9237,f9236,f9234,f9232,f9231,f9210,f9209,f1576,f1574,f1556,f1538,f1515,f1512,f1492,f1489,f1474,f1471,f1451,f1449,f1579,f1559,f1554,f1536,f1510,f1487,f1469,f1454,f8521,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f1756,f896,f797,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f9622,f9619,f11695,f11894,f11896,f12379,f12383,f12400,f12399,f12420,f12419,f12444,f12443,f12457,f12456,f12455,f12454,f12535,f12536,f12537,f12538,f12544,f12545,f12546,f12547,f12548,f12549,f12550,f12551,f12552,f12553,f12554,f12555,f12556,f12557,f12558,f12559,f12560,f12561,f12562,f12573,f12579,f12578,f12583,f12586,f10803])).
fof(f12953,plain,(
  $false | ($spl77 | $spl81)),
  inference(subsumption_resolution,[],[f12605,f12537])).
fof(f12605,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))))) | $spl77),
  inference(unit_resulting_resolution,[],[f2264,f217])).
fof(f2264,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))))))) | $spl77),
  inference(cnf_transformation,[],[f2264_D])).
fof(f2264_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))))))) <=> ~$spl77),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl77])])).
fof(f12930,plain,(
  $false | ($spl77 | $spl81)),
  inference(subsumption_resolution,[],[f12642,f12537])).
fof(f12642,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))))) | $spl77),
  inference(unit_resulting_resolution,[],[f2264,f2264,f2264,f1135])).
fof(f12900,plain,(
  $false | ($spl77 | $spl81)),
  inference(subsumption_resolution,[],[f12755,f12537])).
fof(f12755,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))))) | $spl77),
  inference(unit_resulting_resolution,[],[f200,f2264,f254])).
fof(f12892,plain,(
  $false | ($spl77 | $spl81)),
  inference(subsumption_resolution,[],[f12770,f12537])).
fof(f12770,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))))) | $spl77),
  inference(unit_resulting_resolution,[],[f2264,f2264,f2264,f1135])).
fof(f12891,plain,(
  $false | ($spl77 | $spl81)),
  inference(subsumption_resolution,[],[f12771,f12537])).
fof(f12771,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))))) | $spl77),
  inference(unit_resulting_resolution,[],[f386,f386,f2264,f1135])).
fof(f12890,plain,(
  $false | ($spl49 | $spl77 | $spl81)),
  inference(subsumption_resolution,[],[f12772,f12537])).
fof(f12772,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))))) | ($spl49 | $spl77)),
  inference(unit_resulting_resolution,[],[f1657,f1657,f2264,f1135])).
fof(f12889,plain,(
  $false | ($spl39 | $spl77 | $spl81)),
  inference(subsumption_resolution,[],[f12773,f12537])).
fof(f12773,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))))) | ($spl39 | $spl77)),
  inference(unit_resulting_resolution,[],[f1604,f1604,f2264,f1135])).
fof(f12888,plain,(
  $false | ($spl77 | $spl81 | $spl105)),
  inference(subsumption_resolution,[],[f12774,f12537])).
fof(f12774,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))))) | ($spl77 | $spl105)),
  inference(unit_resulting_resolution,[],[f6996,f6996,f2264,f1135])).
fof(f12887,plain,(
  $false | ($spl77 | $spl81)),
  inference(subsumption_resolution,[],[f12775,f12537])).
fof(f12775,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))))) | $spl77),
  inference(unit_resulting_resolution,[],[f2264,f2264,f2264,f1135])).
fof(f12868,plain,(
  $false | ($spl77 | $spl81)),
  inference(subsumption_resolution,[],[f12845,f12537])).
fof(f12845,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))))) | $spl77),
  inference(unit_resulting_resolution,[],[f200,f2264,f254])).
fof(f2363,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl82),
  inference(cnf_transformation,[],[f2363_D])).
fof(f2363_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))))))) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f2365,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | $spl85),
  inference(cnf_transformation,[],[f2365_D])).
fof(f2365_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) <=> ~$spl85),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl85])])).
fof(f9619,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f8521,f472])).
fof(f9622,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | s(bool,X3) = s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))))))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f8521,f472])).
fof(f9956,plain,(
  $false | ($spl58 | $spl85 | $spl88)),
  inference(global_subsumption,[],[f9955,f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1718,f1719,f1720,f1721,f1722,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f3866,f4029,f4032,f721,f720,f6433,f6436,f1576,f1574,f1556,f1538,f1515,f1512,f1492,f1489,f1474,f1471,f1451,f1449,f1579,f1559,f1554,f1536,f1510,f1487,f1469,f1454,f5332,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f1756,f896,f797,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f7449,f7446,f8521,f2365,f2368,f9757,f9756,f9755,f9724,f9723,f9717,f9715,f9704,f9703,f9669,f9664,f9659,f9656,f9653,f9650,f9816,f9932,f9937])).
fof(f9937,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | ($spl85 | $spl88)),
  inference(subsumption_resolution,[],[f9871,f9816])).
fof(f9871,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X0,sK15(X2,X1,X0)),s(fun(X0,bool),X1))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | ($spl85 | $spl88)),
  inference(backward_demodulation,[],[f9816,f244])).
fof(f9932,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl85 | $spl88)),
  inference(subsumption_resolution,[],[f9817,f9816])).
fof(f9817,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | ($spl85 | $spl88)),
  inference(backward_demodulation,[],[f9816,f220])).
fof(f9816,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl85 | $spl88)),
  inference(subsumption_resolution,[],[f9766,f2368])).
fof(f9766,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | s(bool,'T_0') = s(bool,X0)) ) | $spl85),
  inference(trivial_inequality_removal,[],[f9761])).
fof(f9761,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | s(bool,'T_0') = s(bool,X0)) ) | $spl85),
  inference(superposition,[],[f2365,f472])).
fof(f9650,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))))) | $spl85),
  inference(unit_resulting_resolution,[],[f8521,f2365,f518])).
fof(f9653,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))))))) | $spl85),
  inference(unit_resulting_resolution,[],[f8521,f2365,f1134])).
fof(f9656,plain,(
  ( ! [X0] : (sP0(X0,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))))) ) | $spl85),
  inference(unit_resulting_resolution,[],[f8521,f2365,f1136])).
fof(f9659,plain,(
  ( ! [X0] : (sP0('const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))),X0)) ) | $spl85),
  inference(unit_resulting_resolution,[],[f8521,f2365,f1137])).
fof(f9664,plain,(
  sP1('const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))),'F_0') | $spl85),
  inference(unit_resulting_resolution,[],[f364,f8521,f2365,f1138])).
fof(f9669,plain,(
  sP1('F_0','const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))))))) | $spl85),
  inference(unit_resulting_resolution,[],[f364,f8521,f2365,f1139])).
fof(f9703,plain,(
  ( ! [X0] : (sP11('const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))),X0,'F_0')) ) | $spl85),
  inference(unit_resulting_resolution,[],[f383,f8521,f2365,f1153])).
fof(f9704,plain,(
  ( ! [X0] : (sP11('const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))),'T_0',X0)) ) | $spl85),
  inference(unit_resulting_resolution,[],[f384,f8521,f2365,f1153])).
fof(f9715,plain,(
  ( ! [X0] : (sP9('const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))),X0,'T_0')) ) | $spl85),
  inference(unit_resulting_resolution,[],[f380,f8521,f2365,f1154])).
fof(f9717,plain,(
  ( ! [X0] : (sP9('const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))),'T_0',X0)) ) | $spl85),
  inference(unit_resulting_resolution,[],[f381,f8521,f2365,f1154])).
fof(f9723,plain,(
  sP5('F_0','const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13)),'T_0') | $spl85),
  inference(unit_resulting_resolution,[],[f364,f2365,f300])).
fof(f9724,plain,(
  ~sP5('T_0','const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13)),'T_0') | $spl85),
  inference(unit_resulting_resolution,[],[f200,f2365,f302])).
fof(f9755,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13)))) | $spl85),
  inference(unit_resulting_resolution,[],[f200,f2365,f254])).
fof(f9756,plain,(
  sP5('F_0','T_0','const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | $spl85),
  inference(unit_resulting_resolution,[],[f364,f2365,f300])).
fof(f9757,plain,(
  ~sP5('T_0','T_0','const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | $spl85),
  inference(unit_resulting_resolution,[],[f200,f2365,f302])).
fof(f9955,plain,(
  ( ! [X159,X160] : (s(fun(bool,bool),X159) = s(fun(bool,bool),X160)) ) | ($spl58 | $spl85 | $spl88)),
  inference(subsumption_resolution,[],[f9954,f200])).
fof(f9954,plain,(
  ( ! [X159,X160] : (~p(s(bool,'T_0')) | s(fun(bool,bool),X159) = s(fun(bool,bool),X160)) ) | ($spl58 | $spl85 | $spl88)),
  inference(forward_demodulation,[],[f9953,f9816])).
fof(f9953,plain,(
  ( ! [X161,X159,X162,X160] : (~p(s(bool,'const.bool.IN_2'(s(bool,sK16(X159,X160,bool,bool)),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X162),s(fun(bool,bool),X161)))))) | s(fun(bool,bool),X159) = s(fun(bool,bool),X160)) ) | ($spl58 | $spl85 | $spl88)),
  inference(subsumption_resolution,[],[f9899,f200])).
fof(f9899,plain,(
  ( ! [X161,X159,X162,X160] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(bool,sK16(X159,X160,bool,bool)),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X162),s(fun(bool,bool),X161)))))) | s(fun(bool,bool),X159) = s(fun(bool,bool),X160)) ) | ($spl58 | $spl85 | $spl88)),
  inference(backward_demodulation,[],[f9816,f6958])).
fof(f9946,plain,(
  $false | ($spl85 | $spl88)),
  inference(subsumption_resolution,[],[f9887,f200])).
fof(f9887,plain,(
  ~p(s(bool,'T_0')) | ($spl85 | $spl88)),
  inference(backward_demodulation,[],[f9816,f9755])).
fof(f9941,plain,(
  $false | ($spl85 | $spl88)),
  inference(subsumption_resolution,[],[f9940,f200])).
fof(f9940,plain,(
  ~p(s(bool,'T_0')) | ($spl85 | $spl88)),
  inference(forward_demodulation,[],[f9939,f9816])).
fof(f9939,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))) ) | ($spl85 | $spl88)),
  inference(subsumption_resolution,[],[f9874,f200])).
fof(f9874,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))) ) | ($spl85 | $spl88)),
  inference(backward_demodulation,[],[f9816,f246])).
fof(f9938,plain,(
  $false | ($spl85 | $spl88)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1718,f1719,f1720,f1721,f1722,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f3866,f4029,f4032,f721,f720,f6433,f6436,f1576,f1574,f1556,f1538,f1515,f1512,f1492,f1489,f1474,f1471,f1451,f1449,f1579,f1559,f1554,f1536,f1510,f1487,f1469,f1454,f5332,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f1756,f896,f797,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f7449,f7446,f8521,f2365,f2368,f9757,f9756,f9755,f9724,f9723,f9717,f9715,f9704,f9703,f9669,f9664,f9659,f9656,f9653,f9650,f9816,f9932,f9937])).
fof(f9936,plain,(
  $false | ($spl58 | $spl85 | $spl88)),
  inference(global_subsumption,[],[f9935,f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1718,f1719,f1720,f1721,f1722,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f3866,f4029,f4032,f721,f720,f6433,f6436,f1576,f1574,f1556,f1538,f1515,f1512,f1492,f1489,f1474,f1471,f1451,f1449,f1579,f1559,f1554,f1536,f1510,f1487,f1469,f1454,f5332,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f1756,f896,f797,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f7449,f7446,f8521,f2365,f2368,f9757,f9756,f9755,f9724,f9723,f9717,f9715,f9704,f9703,f9669,f9664,f9659,f9656,f9653,f9650,f9816,f9932])).
fof(f9935,plain,(
  ( ! [X6,X7] : (s(fun(bool,bool),X6) = s(fun(bool,bool),X7)) ) | ($spl58 | $spl85 | $spl88)),
  inference(subsumption_resolution,[],[f9864,f9816])).
fof(f9864,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,sK16(X6,X7,bool,bool)) | s(fun(bool,bool),X6) = s(fun(bool,bool),X7)) ) | ($spl58 | $spl85 | $spl88)),
  inference(backward_demodulation,[],[f9816,f6971])).
fof(f9934,plain,(
  $false | ($spl85 | $spl88)),
  inference(subsumption_resolution,[],[f9819,f200])).
fof(f9819,plain,(
  ~p(s(bool,'T_0')) | ($spl85 | $spl88)),
  inference(backward_demodulation,[],[f9816,f364])).
fof(f9933,plain,(
  $false | ($spl85 | $spl88)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1718,f1719,f1720,f1721,f1722,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f3866,f4029,f4032,f721,f720,f6433,f6436,f1576,f1574,f1556,f1538,f1515,f1512,f1492,f1489,f1474,f1471,f1451,f1449,f1579,f1559,f1554,f1536,f1510,f1487,f1469,f1454,f5332,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f1756,f896,f797,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f7449,f7446,f8521,f2365,f2368,f9757,f9756,f9755,f9724,f9723,f9717,f9715,f9704,f9703,f9669,f9664,f9659,f9656,f9653,f9650,f9816,f9932])).
fof(f9815,plain,(
  $false | ($spl85 | $spl91)),
  inference(subsumption_resolution,[],[f9632,f2370])).
fof(f9632,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | $spl85),
  inference(unit_resulting_resolution,[],[f2365,f207])).
fof(f9814,plain,(
  $false | ($spl85 | $spl91)),
  inference(subsumption_resolution,[],[f9634,f2370])).
fof(f9634,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | $spl85),
  inference(unit_resulting_resolution,[],[f386,f2365,f472])).
fof(f9813,plain,(
  $false | ($spl49 | $spl85 | $spl91)),
  inference(subsumption_resolution,[],[f9812,f2370])).
fof(f9812,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | ($spl49 | $spl85)),
  inference(forward_demodulation,[],[f9635,f4071])).
fof(f4071,plain,(
  s(bool,'F_0') = s(bool,f(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) | $spl49),
  inference(unit_resulting_resolution,[],[f1657,f207])).
fof(f9635,plain,(
  s(bool,f(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) = s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | ($spl49 | $spl85)),
  inference(unit_resulting_resolution,[],[f1657,f2365,f472])).
fof(f9811,plain,(
  $false | ($spl39 | $spl85 | $spl91)),
  inference(subsumption_resolution,[],[f9810,f2370])).
fof(f9810,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | ($spl39 | $spl85)),
  inference(forward_demodulation,[],[f9636,f6220])).
fof(f6220,plain,(
  s(bool,'F_0') = s(bool,f(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12)))) | $spl39),
  inference(unit_resulting_resolution,[],[f393,f364,f1604,f1138])).
fof(f9636,plain,(
  s(bool,f(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12)))) = s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | ($spl39 | $spl85)),
  inference(unit_resulting_resolution,[],[f1604,f2365,f472])).
fof(f9809,plain,(
  $false | ($spl85 | $spl91 | $spl92 | $spl105)),
  inference(subsumption_resolution,[],[f9808,f2370])).
fof(f9808,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | ($spl85 | $spl92 | $spl105)),
  inference(forward_demodulation,[],[f9637,f7060])).
fof(f7060,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | ($spl92 | $spl105)),
  inference(unit_resulting_resolution,[],[f382,f5107,f6996,f1154])).
fof(f5107,plain,(
  ( ! [X0] : (~sP9('F_0',f(s(bool,'T_0')),X0)) ) | $spl92),
  inference(unit_resulting_resolution,[],[f364,f4986,f322])).
fof(f4986,plain,(
  p(s(bool,f(s(bool,'T_0')))) | $spl92),
  inference(unit_resulting_resolution,[],[f364,f200,f4975,f349])).
fof(f4975,plain,(
  ~sP11(f(s(bool,'T_0')),'F_0','T_0') | $spl92),
  inference(unit_resulting_resolution,[],[f364,f2789])).
fof(f9637,plain,(
  s(bool,g(s(bool,'T_0'))) = s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | ($spl85 | $spl105)),
  inference(unit_resulting_resolution,[],[f6996,f2365,f472])).
fof(f9807,plain,(
  $false | ($spl85 | $spl87)),
  inference(subsumption_resolution,[],[f9640,f2366])).
fof(f2366,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))))))) | $spl87),
  inference(cnf_transformation,[],[f2366_D])).
fof(f2366_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))))))) <=> ~$spl87),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl87])])).
fof(f9640,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))))))) | $spl85),
  inference(unit_resulting_resolution,[],[f8521,f2365,f472])).
fof(f9806,plain,(
  $false | ($spl85 | $spl91)),
  inference(subsumption_resolution,[],[f9641,f2370])).
fof(f9641,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | $spl85),
  inference(unit_resulting_resolution,[],[f386,f2365,f472])).
fof(f9805,plain,(
  $false | ($spl49 | $spl85 | $spl91)),
  inference(subsumption_resolution,[],[f9804,f2370])).
fof(f9804,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | ($spl49 | $spl85)),
  inference(forward_demodulation,[],[f9642,f4071])).
fof(f9642,plain,(
  s(bool,f(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) = s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | ($spl49 | $spl85)),
  inference(unit_resulting_resolution,[],[f1657,f2365,f472])).
fof(f9803,plain,(
  $false | ($spl39 | $spl85 | $spl91)),
  inference(subsumption_resolution,[],[f9802,f2370])).
fof(f9802,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | ($spl39 | $spl85)),
  inference(forward_demodulation,[],[f9643,f6220])).
fof(f9643,plain,(
  s(bool,f(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12)))) = s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | ($spl39 | $spl85)),
  inference(unit_resulting_resolution,[],[f1604,f2365,f472])).
fof(f9801,plain,(
  $false | ($spl85 | $spl91 | $spl92 | $spl105)),
  inference(subsumption_resolution,[],[f9800,f2370])).
fof(f9800,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | ($spl85 | $spl92 | $spl105)),
  inference(forward_demodulation,[],[f9644,f7060])).
fof(f9644,plain,(
  s(bool,g(s(bool,'T_0'))) = s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | ($spl85 | $spl105)),
  inference(unit_resulting_resolution,[],[f6996,f2365,f472])).
fof(f9799,plain,(
  $false | ($spl85 | $spl87)),
  inference(subsumption_resolution,[],[f9647,f2366])).
fof(f9647,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))))))) | $spl85),
  inference(unit_resulting_resolution,[],[f8521,f2365,f472])).
fof(f9798,plain,(
  $false | ($spl85 | $spl91)),
  inference(subsumption_resolution,[],[f9648,f2370])).
fof(f9648,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | $spl85),
  inference(unit_resulting_resolution,[],[f364,f2365,f518])).
fof(f9797,plain,(
  $false | ($spl85 | $spl91 | $spl92 | $spl105)),
  inference(subsumption_resolution,[],[f9796,f2370])).
fof(f9796,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | ($spl85 | $spl92 | $spl105)),
  inference(forward_demodulation,[],[f9651,f7060])).
fof(f9651,plain,(
  s(bool,g(s(bool,'T_0'))) = s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | ($spl85 | $spl92 | $spl105)),
  inference(unit_resulting_resolution,[],[f7060,f2365,f1134])).
fof(f9795,plain,(
  $false | ($spl85 | $spl91)),
  inference(subsumption_resolution,[],[f9654,f2370])).
fof(f9654,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | $spl85),
  inference(unit_resulting_resolution,[],[f407,f2365,f1136])).
fof(f9794,plain,(
  $false | ($spl85 | $spl91)),
  inference(subsumption_resolution,[],[f9657,f2370])).
fof(f9657,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | $spl85),
  inference(unit_resulting_resolution,[],[f407,f2365,f1137])).
fof(f9793,plain,(
  $false | ($spl85 | $spl91)),
  inference(subsumption_resolution,[],[f9660,f2370])).
fof(f9660,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | $spl85),
  inference(unit_resulting_resolution,[],[f393,f364,f2365,f1138])).
fof(f9791,plain,(
  $false | ($spl85 | $spl91)),
  inference(subsumption_resolution,[],[f9665,f2370])).
fof(f9665,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | $spl85),
  inference(unit_resulting_resolution,[],[f393,f364,f2365,f1139])).
fof(f9789,plain,(
  $false | ($spl85 | $spl91)),
  inference(subsumption_resolution,[],[f9670,f2370])).
fof(f9670,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | $spl85),
  inference(unit_resulting_resolution,[],[f421,f200,f2365,f1146])).
fof(f9788,plain,(
  $false | ($spl85 | $spl91)),
  inference(subsumption_resolution,[],[f9671,f2370])).
fof(f9671,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | $spl85),
  inference(unit_resulting_resolution,[],[f425,f200,f2365,f1146])).
fof(f9787,plain,(
  $false | ($spl85 | $spl91 | $spl92)),
  inference(subsumption_resolution,[],[f9672,f2370])).
fof(f9672,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | ($spl85 | $spl92)),
  inference(unit_resulting_resolution,[],[f5107,f200,f2365,f1146])).
fof(f9786,plain,(
  $false | ($spl85 | $spl91)),
  inference(subsumption_resolution,[],[f9676,f2370])).
fof(f9676,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | $spl85),
  inference(unit_resulting_resolution,[],[f425,f200,f2365,f1147])).
fof(f9785,plain,(
  $false | ($spl85 | $spl91)),
  inference(subsumption_resolution,[],[f9677,f2370])).
fof(f9677,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | $spl85),
  inference(unit_resulting_resolution,[],[f421,f200,f2365,f1147])).
fof(f9784,plain,(
  $false | ($spl85 | $spl91 | $spl92)),
  inference(subsumption_resolution,[],[f9678,f2370])).
fof(f9678,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | ($spl85 | $spl92)),
  inference(unit_resulting_resolution,[],[f5105,f200,f2365,f1147])).
fof(f5105,plain,(
  ( ! [X0] : (~sP9('F_0',X0,f(s(bool,'T_0')))) ) | $spl92),
  inference(unit_resulting_resolution,[],[f364,f4986,f321])).
fof(f9783,plain,(
  $false | ($spl85 | $spl91)),
  inference(subsumption_resolution,[],[f9682,f2370])).
fof(f9682,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | $spl85),
  inference(unit_resulting_resolution,[],[f438,f364,f2365,f1150])).
fof(f9782,plain,(
  $false | ($spl85 | $spl91)),
  inference(subsumption_resolution,[],[f9683,f2370])).
fof(f9683,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | $spl85),
  inference(unit_resulting_resolution,[],[f439,f364,f2365,f1150])).
fof(f9781,plain,(
  $false | ($spl85 | $spl91 | $spl92)),
  inference(subsumption_resolution,[],[f9684,f2370])).
fof(f9684,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | ($spl85 | $spl92)),
  inference(unit_resulting_resolution,[],[f5127,f364,f2365,f1150])).
fof(f5127,plain,(
  ( ! [X0] : (~sP11('F_0',f(s(bool,'T_0')),X0)) ) | $spl92),
  inference(unit_resulting_resolution,[],[f364,f4986,f339])).
fof(f9779,plain,(
  $false | ($spl85 | $spl91)),
  inference(subsumption_resolution,[],[f9689,f2370])).
fof(f9689,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | $spl85),
  inference(unit_resulting_resolution,[],[f439,f200,f2365,f1151])).
fof(f9778,plain,(
  $false | ($spl85 | $spl91)),
  inference(subsumption_resolution,[],[f9690,f2370])).
fof(f9690,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | $spl85),
  inference(unit_resulting_resolution,[],[f438,f200,f2365,f1151])).
fof(f9777,plain,(
  $false | ($spl85 | $spl91)),
  inference(subsumption_resolution,[],[f9695,f2370])).
fof(f9695,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | $spl85),
  inference(unit_resulting_resolution,[],[f385,f438,f2365,f1153])).
fof(f9776,plain,(
  $false | ($spl85 | $spl91)),
  inference(subsumption_resolution,[],[f9696,f2370])).
fof(f9696,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | $spl85),
  inference(unit_resulting_resolution,[],[f385,f439,f2365,f1153])).
fof(f9775,plain,(
  $false | ($spl85 | $spl91 | $spl92)),
  inference(subsumption_resolution,[],[f9697,f2370])).
fof(f9697,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | ($spl85 | $spl92)),
  inference(unit_resulting_resolution,[],[f385,f5127,f2365,f1153])).
fof(f9774,plain,(
  $false | ($spl85 | $spl91)),
  inference(subsumption_resolution,[],[f9706,f2370])).
fof(f9706,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | $spl85),
  inference(unit_resulting_resolution,[],[f382,f421,f2365,f1154])).
fof(f9773,plain,(
  $false | ($spl85 | $spl91 | $spl92)),
  inference(subsumption_resolution,[],[f9707,f2370])).
fof(f9707,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | ($spl85 | $spl92)),
  inference(unit_resulting_resolution,[],[f382,f5105,f2365,f1154])).
fof(f9772,plain,(
  $false | ($spl85 | $spl91)),
  inference(subsumption_resolution,[],[f9708,f2370])).
fof(f9708,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | $spl85),
  inference(unit_resulting_resolution,[],[f382,f425,f2365,f1154])).
fof(f9771,plain,(
  $false | ($spl85 | $spl91 | $spl92)),
  inference(subsumption_resolution,[],[f9709,f2370])).
fof(f9709,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | ($spl85 | $spl92)),
  inference(unit_resulting_resolution,[],[f382,f5107,f2365,f1154])).
fof(f9591,plain,(
  ( ! [X94,X93] : (sP1('const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13)),X93) | s(bool,'T_0') = s(bool,X94) | s(bool,X93) = s(bool,X94)) ) | $spl112),
  inference(cnf_transformation,[],[f9591_D])).
fof(f9591_D,plain,(
  ( ! [X94,X93] : (sP1('const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13)),X93) | s(bool,'T_0') = s(bool,X94) | s(bool,X93) = s(bool,X94)) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f4034,plain,(
  ( ! [X33] : (p(s(bool,X33))) ) | $spl96),
  inference(cnf_transformation,[],[f4034_D])).
fof(f4034_D,plain,(
  ( ! [X33] : (p(s(bool,X33))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f9590,plain,(
  ( ! [X94,X92,X93] : (p(s(bool,X92)) | sP1('const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13)),X93) | s(bool,X93) = s(bool,X94) | s(bool,'T_0') = s(bool,X94)) ) | $spl86),
  inference(subsumption_resolution,[],[f9470,f216])).
fof(f9470,plain,(
  ( ! [X94,X92,X93] : (p(s(bool,X92)) | sP1('const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13)),X93) | s(bool,X93) = s(bool,X94) | s(bool,'T_0') = s(bool,X94) | s(bool,'T_0') = s(bool,X92)) ) | $spl86),
  inference(superposition,[],[f1139,f2430])).
fof(f2430,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | s(bool,'T_0') = s(bool,X0)) ) | $spl86),
  inference(trivial_inequality_removal,[],[f2424])).
fof(f2424,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | s(bool,'T_0') = s(bool,X0)) ) | $spl86),
  inference(superposition,[],[f2371,f472])).
fof(f2371,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | $spl86),
  inference(backward_demodulation,[],[f2367,f1530])).
fof(f9589,plain,(
  ( ! [X90,X91] : (sP1(X90,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | s(bool,'T_0') = s(bool,X91) | s(bool,X90) = s(bool,X91)) ) | $spl110),
  inference(cnf_transformation,[],[f9589_D])).
fof(f9589_D,plain,(
  ( ! [X90,X91] : (sP1(X90,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | s(bool,'T_0') = s(bool,X91) | s(bool,X90) = s(bool,X91)) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f9588,plain,(
  ( ! [X90,X91,X89] : (p(s(bool,X89)) | sP1(X90,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | s(bool,X90) = s(bool,X91) | s(bool,'T_0') = s(bool,X91)) ) | $spl86),
  inference(subsumption_resolution,[],[f9469,f216])).
fof(f9469,plain,(
  ( ! [X90,X91,X89] : (p(s(bool,X89)) | sP1(X90,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | s(bool,X90) = s(bool,X91) | s(bool,'T_0') = s(bool,X91) | s(bool,'T_0') = s(bool,X89)) ) | $spl86),
  inference(superposition,[],[f1138,f2430])).
fof(f2261,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl72),
  inference(cnf_transformation,[],[f2261_D])).
fof(f2261_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))))))) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f2263,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | $spl75),
  inference(cnf_transformation,[],[f2263_D])).
fof(f2263_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) <=> ~$spl75),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl75])])).
fof(f7446,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f5332,f472])).
fof(f7449,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | s(bool,X3) = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))))))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f5332,f472])).
fof(f7801,plain,(
  $false | ($spl58 | $spl75 | $spl78)),
  inference(global_subsumption,[],[f7800,f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1718,f1719,f1720,f1721,f1722,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f3866,f4029,f4032,f721,f720,f220,f221,f364,f722,f723,f1448,f1453,f1468,f1473,f1486,f1491,f1509,f1514,f1535,f1540,f1545,f1553,f1558,f1573,f1578,f1826,f1865,f1901,f1936,f2067,f2102,f2146,f2186,f2214,f2216,f2218,f2225,f2227,f2229,f2462,f2464,f2466,f2501,f2503,f2505,f2556,f2590,f2634,f2668,f1756,f896,f797,f244,f243,f245,f246,f247,f248,f249,f250,f1461,f1481,f1499,f1522,f1566,f1586,f1487,f1489,f1510,f1512,f1576,f1579,f1449,f1451,f1471,f1474,f1492,f1515,f1538,f1554,f1556,f1559,f1574,f1454,f1469,f1536,f6433,f6436,f5332,f2263,f2266,f7601,f7600,f7599,f7568,f7567,f7562,f7560,f7549,f7548,f7514,f7509,f7504,f7501,f7498,f7496,f7661,f7777,f7782])).
fof(f7782,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | ($spl75 | $spl78)),
  inference(subsumption_resolution,[],[f7716,f7661])).
fof(f7716,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X0,sK15(X2,X1,X0)),s(fun(X0,bool),X1))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | ($spl75 | $spl78)),
  inference(backward_demodulation,[],[f7661,f244])).
fof(f7777,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl75 | $spl78)),
  inference(subsumption_resolution,[],[f7662,f7661])).
fof(f7662,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | ($spl75 | $spl78)),
  inference(backward_demodulation,[],[f7661,f220])).
fof(f7661,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl75 | $spl78)),
  inference(subsumption_resolution,[],[f7610,f2266])).
fof(f7610,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | s(bool,'T_0') = s(bool,X0)) ) | $spl75),
  inference(trivial_inequality_removal,[],[f7605])).
fof(f7605,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | s(bool,'T_0') = s(bool,X0)) ) | $spl75),
  inference(superposition,[],[f2263,f472])).
fof(f7496,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))))) | $spl75),
  inference(unit_resulting_resolution,[],[f5332,f2263,f518])).
fof(f7498,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))))))) | $spl75),
  inference(unit_resulting_resolution,[],[f5332,f2263,f1134])).
fof(f7501,plain,(
  ( ! [X0] : (sP0(X0,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))))) ) | $spl75),
  inference(unit_resulting_resolution,[],[f5332,f2263,f1136])).
fof(f7504,plain,(
  ( ! [X0] : (sP0('const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))),X0)) ) | $spl75),
  inference(unit_resulting_resolution,[],[f5332,f2263,f1137])).
fof(f7509,plain,(
  sP1('const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))),'F_0') | $spl75),
  inference(unit_resulting_resolution,[],[f364,f5332,f2263,f1138])).
fof(f7514,plain,(
  sP1('F_0','const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))))))) | $spl75),
  inference(unit_resulting_resolution,[],[f364,f5332,f2263,f1139])).
fof(f7548,plain,(
  ( ! [X0] : (sP11('const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))),X0,'F_0')) ) | $spl75),
  inference(unit_resulting_resolution,[],[f383,f5332,f2263,f1153])).
fof(f7549,plain,(
  ( ! [X0] : (sP11('const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))),'T_0',X0)) ) | $spl75),
  inference(unit_resulting_resolution,[],[f384,f5332,f2263,f1153])).
fof(f7560,plain,(
  ( ! [X0] : (sP9('const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))),X0,'T_0')) ) | $spl75),
  inference(unit_resulting_resolution,[],[f380,f5332,f2263,f1154])).
fof(f7562,plain,(
  ( ! [X0] : (sP9('const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))),'T_0',X0)) ) | $spl75),
  inference(unit_resulting_resolution,[],[f381,f5332,f2263,f1154])).
fof(f7567,plain,(
  sP5('F_0','const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)),'T_0') | $spl75),
  inference(unit_resulting_resolution,[],[f364,f2263,f300])).
fof(f7568,plain,(
  ~sP5('T_0','const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)),'T_0') | $spl75),
  inference(unit_resulting_resolution,[],[f200,f2263,f302])).
fof(f7599,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)))) | $spl75),
  inference(unit_resulting_resolution,[],[f200,f2263,f254])).
fof(f7600,plain,(
  sP5('F_0','T_0','const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | $spl75),
  inference(unit_resulting_resolution,[],[f364,f2263,f300])).
fof(f7601,plain,(
  ~sP5('T_0','T_0','const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | $spl75),
  inference(unit_resulting_resolution,[],[f200,f2263,f302])).
fof(f7800,plain,(
  ( ! [X159,X160] : (s(fun(bool,bool),X159) = s(fun(bool,bool),X160)) ) | ($spl58 | $spl75 | $spl78)),
  inference(subsumption_resolution,[],[f7799,f200])).
fof(f7799,plain,(
  ( ! [X159,X160] : (~p(s(bool,'T_0')) | s(fun(bool,bool),X159) = s(fun(bool,bool),X160)) ) | ($spl58 | $spl75 | $spl78)),
  inference(forward_demodulation,[],[f7798,f7661])).
fof(f7798,plain,(
  ( ! [X161,X159,X162,X160] : (~p(s(bool,'const.bool.IN_2'(s(bool,sK16(X159,X160,bool,bool)),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X162),s(fun(bool,bool),X161)))))) | s(fun(bool,bool),X159) = s(fun(bool,bool),X160)) ) | ($spl58 | $spl75 | $spl78)),
  inference(subsumption_resolution,[],[f7744,f200])).
fof(f7744,plain,(
  ( ! [X161,X159,X162,X160] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(bool,sK16(X159,X160,bool,bool)),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X162),s(fun(bool,bool),X161)))))) | s(fun(bool,bool),X159) = s(fun(bool,bool),X160)) ) | ($spl58 | $spl75 | $spl78)),
  inference(backward_demodulation,[],[f7661,f6958])).
fof(f7791,plain,(
  $false | ($spl75 | $spl78)),
  inference(subsumption_resolution,[],[f7732,f200])).
fof(f7732,plain,(
  ~p(s(bool,'T_0')) | ($spl75 | $spl78)),
  inference(backward_demodulation,[],[f7661,f7599])).
fof(f7786,plain,(
  $false | ($spl75 | $spl78)),
  inference(subsumption_resolution,[],[f7785,f200])).
fof(f7785,plain,(
  ~p(s(bool,'T_0')) | ($spl75 | $spl78)),
  inference(forward_demodulation,[],[f7784,f7661])).
fof(f7784,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))) ) | ($spl75 | $spl78)),
  inference(subsumption_resolution,[],[f7719,f200])).
fof(f7719,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))) ) | ($spl75 | $spl78)),
  inference(backward_demodulation,[],[f7661,f246])).
fof(f7783,plain,(
  $false | ($spl75 | $spl78)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1718,f1719,f1720,f1721,f1722,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f3866,f4029,f4032,f721,f720,f220,f221,f364,f722,f723,f1448,f1453,f1468,f1473,f1486,f1491,f1509,f1514,f1535,f1540,f1545,f1553,f1558,f1573,f1578,f1826,f1865,f1901,f1936,f2067,f2102,f2146,f2186,f2214,f2216,f2218,f2225,f2227,f2229,f2462,f2464,f2466,f2501,f2503,f2505,f2556,f2590,f2634,f2668,f1756,f896,f797,f244,f243,f245,f246,f247,f248,f249,f250,f1461,f1481,f1499,f1522,f1566,f1586,f1487,f1489,f1510,f1512,f1576,f1579,f1449,f1451,f1471,f1474,f1492,f1515,f1538,f1554,f1556,f1559,f1574,f1454,f1469,f1536,f6433,f6436,f5332,f2263,f2266,f7601,f7600,f7599,f7568,f7567,f7562,f7560,f7549,f7548,f7514,f7509,f7504,f7501,f7498,f7496,f7661,f7777,f7782])).
fof(f7781,plain,(
  $false | ($spl58 | $spl75 | $spl78)),
  inference(global_subsumption,[],[f7780,f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1718,f1719,f1720,f1721,f1722,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f3866,f4029,f4032,f721,f720,f220,f221,f364,f722,f723,f1448,f1453,f1468,f1473,f1486,f1491,f1509,f1514,f1535,f1540,f1545,f1553,f1558,f1573,f1578,f1826,f1865,f1901,f1936,f2067,f2102,f2146,f2186,f2214,f2216,f2218,f2225,f2227,f2229,f2462,f2464,f2466,f2501,f2503,f2505,f2556,f2590,f2634,f2668,f1756,f896,f797,f244,f243,f245,f246,f247,f248,f249,f250,f1461,f1481,f1499,f1522,f1566,f1586,f1487,f1489,f1510,f1512,f1576,f1579,f1449,f1451,f1471,f1474,f1492,f1515,f1538,f1554,f1556,f1559,f1574,f1454,f1469,f1536,f6433,f6436,f5332,f2263,f2266,f7601,f7600,f7599,f7568,f7567,f7562,f7560,f7549,f7548,f7514,f7509,f7504,f7501,f7498,f7496,f7661,f7777])).
fof(f7780,plain,(
  ( ! [X6,X7] : (s(fun(bool,bool),X6) = s(fun(bool,bool),X7)) ) | ($spl58 | $spl75 | $spl78)),
  inference(subsumption_resolution,[],[f7710,f7661])).
fof(f7710,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,sK16(X6,X7,bool,bool)) | s(fun(bool,bool),X6) = s(fun(bool,bool),X7)) ) | ($spl58 | $spl75 | $spl78)),
  inference(backward_demodulation,[],[f7661,f6971])).
fof(f7779,plain,(
  $false | ($spl75 | $spl78)),
  inference(subsumption_resolution,[],[f7664,f200])).
fof(f7664,plain,(
  ~p(s(bool,'T_0')) | ($spl75 | $spl78)),
  inference(backward_demodulation,[],[f7661,f364])).
fof(f7778,plain,(
  $false | ($spl75 | $spl78)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1718,f1719,f1720,f1721,f1722,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f3866,f4029,f4032,f721,f720,f220,f221,f364,f722,f723,f1448,f1453,f1468,f1473,f1486,f1491,f1509,f1514,f1535,f1540,f1545,f1553,f1558,f1573,f1578,f1826,f1865,f1901,f1936,f2067,f2102,f2146,f2186,f2214,f2216,f2218,f2225,f2227,f2229,f2462,f2464,f2466,f2501,f2503,f2505,f2556,f2590,f2634,f2668,f1756,f896,f797,f244,f243,f245,f246,f247,f248,f249,f250,f1461,f1481,f1499,f1522,f1566,f1586,f1487,f1489,f1510,f1512,f1576,f1579,f1449,f1451,f1471,f1474,f1492,f1515,f1538,f1554,f1556,f1559,f1574,f1454,f1469,f1536,f6433,f6436,f5332,f2263,f2266,f7601,f7600,f7599,f7568,f7567,f7562,f7560,f7549,f7548,f7514,f7509,f7504,f7501,f7498,f7496,f7661,f7777])).
fof(f7660,plain,(
  $false | ($spl75 | $spl81)),
  inference(subsumption_resolution,[],[f7476,f2268])).
fof(f7476,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | $spl75),
  inference(unit_resulting_resolution,[],[f2263,f207])).
fof(f7659,plain,(
  $false | ($spl75 | $spl81)),
  inference(subsumption_resolution,[],[f7478,f2268])).
fof(f7478,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | $spl75),
  inference(unit_resulting_resolution,[],[f386,f2263,f472])).
fof(f7658,plain,(
  $false | ($spl49 | $spl75 | $spl81)),
  inference(subsumption_resolution,[],[f7657,f2268])).
fof(f7657,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | ($spl49 | $spl75)),
  inference(forward_demodulation,[],[f7479,f4071])).
fof(f7479,plain,(
  s(bool,f(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | ($spl49 | $spl75)),
  inference(unit_resulting_resolution,[],[f1657,f2263,f472])).
fof(f7656,plain,(
  $false | ($spl39 | $spl75 | $spl81)),
  inference(subsumption_resolution,[],[f7655,f2268])).
fof(f7655,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | ($spl39 | $spl75)),
  inference(forward_demodulation,[],[f7480,f6220])).
fof(f7480,plain,(
  s(bool,f(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12)))) = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | ($spl39 | $spl75)),
  inference(unit_resulting_resolution,[],[f1604,f2263,f472])).
fof(f7654,plain,(
  $false | ($spl75 | $spl81 | $spl92 | $spl105)),
  inference(subsumption_resolution,[],[f7653,f2268])).
fof(f7653,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | ($spl75 | $spl92 | $spl105)),
  inference(forward_demodulation,[],[f7481,f7060])).
fof(f7481,plain,(
  s(bool,g(s(bool,'T_0'))) = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | ($spl75 | $spl105)),
  inference(unit_resulting_resolution,[],[f6996,f2263,f472])).
fof(f7652,plain,(
  $false | ($spl75 | $spl81 | $spl86)),
  inference(subsumption_resolution,[],[f7651,f2268])).
fof(f7651,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | ($spl75 | $spl86)),
  inference(forward_demodulation,[],[f7482,f2410])).
fof(f2410,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | $spl86),
  inference(unit_resulting_resolution,[],[f407,f2371,f1136])).
fof(f7482,plain,(
  s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) = s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | ($spl75 | $spl86)),
  inference(unit_resulting_resolution,[],[f2371,f2263,f472])).
fof(f7650,plain,(
  $false | ($spl75 | $spl77)),
  inference(subsumption_resolution,[],[f7485,f2264])).
fof(f7485,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))))))) | $spl75),
  inference(unit_resulting_resolution,[],[f5332,f2263,f472])).
fof(f7649,plain,(
  $false | ($spl75 | $spl81)),
  inference(subsumption_resolution,[],[f7486,f2268])).
fof(f7486,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | $spl75),
  inference(unit_resulting_resolution,[],[f386,f2263,f472])).
fof(f7648,plain,(
  $false | ($spl49 | $spl75 | $spl81)),
  inference(subsumption_resolution,[],[f7647,f2268])).
fof(f7647,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | ($spl49 | $spl75)),
  inference(forward_demodulation,[],[f7487,f4071])).
fof(f7487,plain,(
  s(bool,f(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | ($spl49 | $spl75)),
  inference(unit_resulting_resolution,[],[f1657,f2263,f472])).
fof(f7646,plain,(
  $false | ($spl39 | $spl75 | $spl81)),
  inference(subsumption_resolution,[],[f7645,f2268])).
fof(f7645,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | ($spl39 | $spl75)),
  inference(forward_demodulation,[],[f7488,f6220])).
fof(f7488,plain,(
  s(bool,f(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12)))) = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | ($spl39 | $spl75)),
  inference(unit_resulting_resolution,[],[f1604,f2263,f472])).
fof(f7644,plain,(
  $false | ($spl75 | $spl81 | $spl92 | $spl105)),
  inference(subsumption_resolution,[],[f7643,f2268])).
fof(f7643,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | ($spl75 | $spl92 | $spl105)),
  inference(forward_demodulation,[],[f7489,f7060])).
fof(f7489,plain,(
  s(bool,g(s(bool,'T_0'))) = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | ($spl75 | $spl105)),
  inference(unit_resulting_resolution,[],[f6996,f2263,f472])).
fof(f7642,plain,(
  $false | ($spl75 | $spl81 | $spl86)),
  inference(subsumption_resolution,[],[f7641,f2268])).
fof(f7641,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | ($spl75 | $spl86)),
  inference(forward_demodulation,[],[f7490,f2410])).
fof(f7490,plain,(
  s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) = s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | ($spl75 | $spl86)),
  inference(unit_resulting_resolution,[],[f2371,f2263,f472])).
fof(f7640,plain,(
  $false | ($spl75 | $spl77)),
  inference(subsumption_resolution,[],[f7493,f2264])).
fof(f7493,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))))))) | $spl75),
  inference(unit_resulting_resolution,[],[f5332,f2263,f472])).
fof(f7639,plain,(
  $false | ($spl75 | $spl81)),
  inference(subsumption_resolution,[],[f7494,f2268])).
fof(f7494,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | $spl75),
  inference(unit_resulting_resolution,[],[f364,f2263,f518])).
fof(f7638,plain,(
  $false | ($spl75 | $spl81)),
  inference(subsumption_resolution,[],[f7499,f2268])).
fof(f7499,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | $spl75),
  inference(unit_resulting_resolution,[],[f407,f2263,f1136])).
fof(f7637,plain,(
  $false | ($spl75 | $spl81)),
  inference(subsumption_resolution,[],[f7502,f2268])).
fof(f7502,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | $spl75),
  inference(unit_resulting_resolution,[],[f407,f2263,f1137])).
fof(f7636,plain,(
  $false | ($spl75 | $spl81)),
  inference(subsumption_resolution,[],[f7505,f2268])).
fof(f7505,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | $spl75),
  inference(unit_resulting_resolution,[],[f393,f364,f2263,f1138])).
fof(f7634,plain,(
  $false | ($spl75 | $spl81)),
  inference(subsumption_resolution,[],[f7510,f2268])).
fof(f7510,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | $spl75),
  inference(unit_resulting_resolution,[],[f393,f364,f2263,f1139])).
fof(f7632,plain,(
  $false | ($spl75 | $spl81)),
  inference(subsumption_resolution,[],[f7515,f2268])).
fof(f7515,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | $spl75),
  inference(unit_resulting_resolution,[],[f421,f200,f2263,f1146])).
fof(f7631,plain,(
  $false | ($spl75 | $spl81)),
  inference(subsumption_resolution,[],[f7516,f2268])).
fof(f7516,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | $spl75),
  inference(unit_resulting_resolution,[],[f425,f200,f2263,f1146])).
fof(f7630,plain,(
  $false | ($spl75 | $spl81 | $spl92)),
  inference(subsumption_resolution,[],[f7517,f2268])).
fof(f7517,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | ($spl75 | $spl92)),
  inference(unit_resulting_resolution,[],[f5107,f200,f2263,f1146])).
fof(f7629,plain,(
  $false | ($spl75 | $spl81)),
  inference(subsumption_resolution,[],[f7521,f2268])).
fof(f7521,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | $spl75),
  inference(unit_resulting_resolution,[],[f425,f200,f2263,f1147])).
fof(f7628,plain,(
  $false | ($spl75 | $spl81)),
  inference(subsumption_resolution,[],[f7522,f2268])).
fof(f7522,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | $spl75),
  inference(unit_resulting_resolution,[],[f421,f200,f2263,f1147])).
fof(f7627,plain,(
  $false | ($spl75 | $spl81 | $spl92)),
  inference(subsumption_resolution,[],[f7523,f2268])).
fof(f7523,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | ($spl75 | $spl92)),
  inference(unit_resulting_resolution,[],[f5105,f200,f2263,f1147])).
fof(f7626,plain,(
  $false | ($spl75 | $spl81)),
  inference(subsumption_resolution,[],[f7527,f2268])).
fof(f7527,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | $spl75),
  inference(unit_resulting_resolution,[],[f438,f364,f2263,f1150])).
fof(f7625,plain,(
  $false | ($spl75 | $spl81)),
  inference(subsumption_resolution,[],[f7528,f2268])).
fof(f7528,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | $spl75),
  inference(unit_resulting_resolution,[],[f439,f364,f2263,f1150])).
fof(f7624,plain,(
  $false | ($spl75 | $spl81 | $spl92)),
  inference(subsumption_resolution,[],[f7529,f2268])).
fof(f7529,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | ($spl75 | $spl92)),
  inference(unit_resulting_resolution,[],[f5127,f364,f2263,f1150])).
fof(f7622,plain,(
  $false | ($spl75 | $spl81)),
  inference(subsumption_resolution,[],[f7534,f2268])).
fof(f7534,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | $spl75),
  inference(unit_resulting_resolution,[],[f439,f200,f2263,f1151])).
fof(f7621,plain,(
  $false | ($spl75 | $spl81)),
  inference(subsumption_resolution,[],[f7535,f2268])).
fof(f7535,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | $spl75),
  inference(unit_resulting_resolution,[],[f438,f200,f2263,f1151])).
fof(f7620,plain,(
  $false | ($spl75 | $spl81)),
  inference(subsumption_resolution,[],[f7540,f2268])).
fof(f7540,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | $spl75),
  inference(unit_resulting_resolution,[],[f385,f438,f2263,f1153])).
fof(f7619,plain,(
  $false | ($spl75 | $spl81)),
  inference(subsumption_resolution,[],[f7541,f2268])).
fof(f7541,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | $spl75),
  inference(unit_resulting_resolution,[],[f385,f439,f2263,f1153])).
fof(f7618,plain,(
  $false | ($spl75 | $spl81 | $spl92)),
  inference(subsumption_resolution,[],[f7542,f2268])).
fof(f7542,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | ($spl75 | $spl92)),
  inference(unit_resulting_resolution,[],[f385,f5127,f2263,f1153])).
fof(f7617,plain,(
  $false | ($spl75 | $spl81)),
  inference(subsumption_resolution,[],[f7551,f2268])).
fof(f7551,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | $spl75),
  inference(unit_resulting_resolution,[],[f382,f421,f2263,f1154])).
fof(f7616,plain,(
  $false | ($spl75 | $spl81 | $spl92)),
  inference(subsumption_resolution,[],[f7552,f2268])).
fof(f7552,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | ($spl75 | $spl92)),
  inference(unit_resulting_resolution,[],[f382,f5105,f2263,f1154])).
fof(f7615,plain,(
  $false | ($spl75 | $spl81)),
  inference(subsumption_resolution,[],[f7553,f2268])).
fof(f7553,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | $spl75),
  inference(unit_resulting_resolution,[],[f382,f425,f2263,f1154])).
fof(f7614,plain,(
  $false | ($spl75 | $spl81 | $spl92)),
  inference(subsumption_resolution,[],[f7554,f2268])).
fof(f7554,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | ($spl75 | $spl92)),
  inference(unit_resulting_resolution,[],[f382,f5107,f2263,f1154])).
fof(f7407,plain,(
  ( ! [X94,X93] : (sP1('const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)),X93) | s(bool,'T_0') = s(bool,X94) | s(bool,X93) = s(bool,X94)) ) | $spl108),
  inference(cnf_transformation,[],[f7407_D])).
fof(f7407_D,plain,(
  ( ! [X94,X93] : (sP1('const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)),X93) | s(bool,'T_0') = s(bool,X94) | s(bool,X93) = s(bool,X94)) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f7406,plain,(
  ( ! [X94,X92,X93] : (p(s(bool,X92)) | sP1('const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)),X93) | s(bool,X93) = s(bool,X94) | s(bool,'T_0') = s(bool,X94)) ) | $spl76),
  inference(subsumption_resolution,[],[f7290,f216])).
fof(f7290,plain,(
  ( ! [X94,X92,X93] : (p(s(bool,X92)) | sP1('const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13)),X93) | s(bool,X93) = s(bool,X94) | s(bool,'T_0') = s(bool,X94) | s(bool,'T_0') = s(bool,X92)) ) | $spl76),
  inference(superposition,[],[f1139,f2323])).
fof(f7405,plain,(
  ( ! [X90,X91] : (sP1(X90,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | s(bool,'T_0') = s(bool,X91) | s(bool,X90) = s(bool,X91)) ) | $spl106),
  inference(cnf_transformation,[],[f7405_D])).
fof(f7405_D,plain,(
  ( ! [X90,X91] : (sP1(X90,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | s(bool,'T_0') = s(bool,X91) | s(bool,X90) = s(bool,X91)) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f7404,plain,(
  ( ! [X90,X91,X89] : (p(s(bool,X89)) | sP1(X90,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | s(bool,X90) = s(bool,X91) | s(bool,'T_0') = s(bool,X91)) ) | $spl76),
  inference(subsumption_resolution,[],[f7289,f216])).
fof(f7289,plain,(
  ( ! [X90,X91,X89] : (p(s(bool,X89)) | sP1(X90,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | s(bool,X90) = s(bool,X91) | s(bool,'T_0') = s(bool,X91) | s(bool,'T_0') = s(bool,X89)) ) | $spl76),
  inference(superposition,[],[f1138,f2323])).
fof(f6994,plain,(
  ( ! [X142,X141] : (s(fun(bool,bool),X141) = s(fun(bool,bool),X142)) ) | $spl102),
  inference(cnf_transformation,[],[f6994_D])).
fof(f6994_D,plain,(
  ( ! [X142,X141] : (s(fun(bool,bool),X141) = s(fun(bool,bool),X142)) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f6996,plain,(
  s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | $spl105),
  inference(cnf_transformation,[],[f6996_D])).
fof(f6996_D,plain,(
  s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) <=> ~$spl105),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl105])])).
fof(f6993,plain,(
  ( ! [X142,X141] : (s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X141) = s(fun(bool,bool),X142)) ) | ($spl58 | $spl92)),
  inference(forward_demodulation,[],[f6960,f5024])).
fof(f5024,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | $spl92),
  inference(unit_resulting_resolution,[],[f200,f4986,f254])).
fof(f6960,plain,(
  ( ! [X142,X141] : (s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X141) = s(fun(bool,bool),X142)) ) | $spl58),
  inference(duplicate_literal_removal,[],[f6953])).
fof(f6953,plain,(
  ( ! [X142,X141] : (s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X141) = s(fun(bool,bool),X142) | s(fun(bool,bool),X141) = s(fun(bool,bool),X142)) ) | $spl58),
  inference(superposition,[],[f280,f1724])).
fof(f6437,plain,(
  ( ! [X28,X29] : (sP1('F_0',X28) | s(bool,'T_0') = s(bool,X29) | s(bool,X28) = s(bool,X29)) ) | $spl100),
  inference(cnf_transformation,[],[f6437_D])).
fof(f6437_D,plain,(
  ( ! [X28,X29] : (sP1('F_0',X28) | s(bool,'T_0') = s(bool,X29) | s(bool,X28) = s(bool,X29)) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f6436,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | sP1('F_0',X28) | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29)) )),
  inference(subsumption_resolution,[],[f6426,f216])).
fof(f6426,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | sP1('F_0',X28) | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f1139,f207])).
fof(f4741,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f1621,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f797,f896,f1756,f720,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1609,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f1606])).
fof(f1621,plain,(
  s(bool,'F_0') = s(bool,f(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12)))) | $spl40),
  inference(unit_resulting_resolution,[],[f407,f1610,f1136])).
fof(f4068,plain,(
  sP4('T_0',f(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))),'F_0') | $spl40),
  inference(unit_resulting_resolution,[],[f3922,f283])).
fof(f4063,plain,(
  ~sP2(f(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))),'F_0','F_0') | $spl40),
  inference(unit_resulting_resolution,[],[f364,f373,f370,f3921,f286])).
fof(f4040,plain,(
  ~sP2('F_0',f(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))),'F_0') | $spl40),
  inference(unit_resulting_resolution,[],[f364,f372,f370,f3919,f287])).
fof(f4036,plain,(
  ( ! [X39,X38] : (sP1(X38,f(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12)))) | s(bool,X38) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) ) | $spl40),
  inference(subsumption_resolution,[],[f4024,f364])).
fof(f4024,plain,(
  ( ! [X39,X38] : (p(s(bool,'F_0')) | sP1(X38,f(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12)))) | s(bool,X38) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) ) | $spl40),
  inference(superposition,[],[f1138,f1621])).
fof(f3919,plain,(
  ~sP5('F_0',f(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))),'F_0') | $spl40),
  inference(unit_resulting_resolution,[],[f364,f1621,f303])).
fof(f3920,plain,(
  sP5('T_0',f(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))),'F_0') | $spl40),
  inference(unit_resulting_resolution,[],[f200,f1621,f301])).
fof(f3921,plain,(
  ~sP5('F_0','F_0',f(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12)))) | $spl40),
  inference(unit_resulting_resolution,[],[f364,f1621,f303])).
fof(f3922,plain,(
  sP5('T_0','F_0',f(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12)))) | $spl40),
  inference(unit_resulting_resolution,[],[f200,f1621,f301])).
fof(f3261,plain,(
  ~sP3('T_0',g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))),'F_0') | $spl40),
  inference(unit_resulting_resolution,[],[f364,f375,f370,f2955,f287])).
fof(f3258,plain,(
  sP4('T_0','T_0',g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12)))) | $spl40),
  inference(unit_resulting_resolution,[],[f2954,f283])).
fof(f3248,plain,(
  ~sP3(g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))),'T_0','F_0') | $spl40),
  inference(unit_resulting_resolution,[],[f200,f376,f370,f2953,f285])).
fof(f3037,plain,(
  ( ! [X35,X36] : (~sP8(g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))),X35,X36) | sP9(g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))),X35,X36)) ) | $spl40),
  inference(subsumption_resolution,[],[f2986,f200])).
fof(f2986,plain,(
  ( ! [X35,X36] : (~p(s(bool,'T_0')) | ~sP8(g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))),X35,X36) | sP9(g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))),X35,X36)) ) | $spl40),
  inference(superposition,[],[f356,f1606])).
fof(f3036,plain,(
  ( ! [X33,X34] : (~sP10(g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))),X33,X34) | sP11(g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))),X33,X34)) ) | $spl40),
  inference(subsumption_resolution,[],[f2985,f200])).
fof(f2985,plain,(
  ( ! [X33,X34] : (~p(s(bool,'T_0')) | ~sP10(g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))),X33,X34) | sP11(g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))),X33,X34)) ) | $spl40),
  inference(superposition,[],[f355,f1606])).
fof(f3035,plain,(
  ( ! [X31,X32] : (~p(s(bool,X31)) | p(s(bool,X32)) | ~sP11(g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))),X32,X31)) ) | $spl40),
  inference(subsumption_resolution,[],[f2984,f200])).
fof(f2984,plain,(
  ( ! [X31,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X31)) | p(s(bool,X32)) | ~sP11(g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))),X32,X31)) ) | $spl40),
  inference(superposition,[],[f352,f1606])).
fof(f3034,plain,(
  ( ! [X30,X29] : (~p(s(bool,X29)) | sP11(g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))),X29,X30)) ) | $spl40),
  inference(subsumption_resolution,[],[f2983,f200])).
fof(f2983,plain,(
  ( ! [X30,X29] : (~p(s(bool,'T_0')) | ~p(s(bool,X29)) | sP11(g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))),X29,X30)) ) | $spl40),
  inference(superposition,[],[f351,f1606])).
fof(f3033,plain,(
  ( ! [X28,X27] : (p(s(bool,X27)) | sP11(g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))),X28,X27)) ) | $spl40),
  inference(subsumption_resolution,[],[f2982,f200])).
fof(f2982,plain,(
  ( ! [X28,X27] : (~p(s(bool,'T_0')) | p(s(bool,X27)) | sP11(g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))),X28,X27)) ) | $spl40),
  inference(superposition,[],[f350,f1606])).
fof(f3032,plain,(
  ( ! [X26,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | ~sP10(X26,X25,g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))))) ) | $spl40),
  inference(subsumption_resolution,[],[f2981,f200])).
fof(f2981,plain,(
  ( ! [X26,X25] : (~p(s(bool,'T_0')) | p(s(bool,X25)) | ~p(s(bool,X26)) | ~sP10(X26,X25,g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))))) ) | $spl40),
  inference(superposition,[],[f345,f1606])).
fof(f3031,plain,(
  ( ! [X24,X23] : (p(s(bool,X23)) | p(s(bool,X24)) | ~sP9(g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))),X24,X23)) ) | $spl40),
  inference(subsumption_resolution,[],[f2980,f200])).
fof(f2980,plain,(
  ( ! [X24,X23] : (~p(s(bool,'T_0')) | p(s(bool,X23)) | p(s(bool,X24)) | ~sP9(g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))),X24,X23)) ) | $spl40),
  inference(superposition,[],[f335,f1606])).
fof(f3030,plain,(
  ( ! [X21,X22] : (~p(s(bool,X21)) | sP9(g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))),X21,X22)) ) | $spl40),
  inference(subsumption_resolution,[],[f2979,f200])).
fof(f2979,plain,(
  ( ! [X21,X22] : (~p(s(bool,'T_0')) | ~p(s(bool,X21)) | sP9(g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))),X21,X22)) ) | $spl40),
  inference(superposition,[],[f334,f1606])).
fof(f3029,plain,(
  ( ! [X19,X20] : (~p(s(bool,X19)) | sP9(g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))),X20,X19)) ) | $spl40),
  inference(subsumption_resolution,[],[f2978,f200])).
fof(f2978,plain,(
  ( ! [X19,X20] : (~p(s(bool,'T_0')) | ~p(s(bool,X19)) | sP9(g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))),X20,X19)) ) | $spl40),
  inference(superposition,[],[f333,f1606])).
fof(f3028,plain,(
  ( ! [X17,X18] : (p(s(bool,X17)) | p(s(bool,X18)) | ~sP8(g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))),X17,X18)) ) | $spl40),
  inference(subsumption_resolution,[],[f2977,f200])).
fof(f2977,plain,(
  ( ! [X17,X18] : (~p(s(bool,'T_0')) | p(s(bool,X17)) | p(s(bool,X18)) | ~sP8(g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))),X17,X18)) ) | $spl40),
  inference(superposition,[],[f328,f1606])).
fof(f3027,plain,(
  ( ! [X15,X16] : (~p(s(bool,X15)) | p(s(bool,X16)) | ~sP7(X16,X15,g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))))) ) | $spl40),
  inference(subsumption_resolution,[],[f2976,f200])).
fof(f2976,plain,(
  ( ! [X15,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X15)) | p(s(bool,X16)) | ~sP7(X16,X15,g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))))) ) | $spl40),
  inference(superposition,[],[f320,f1606])).
fof(f3026,plain,(
  ( ! [X14,X13] : (~p(s(bool,X13)) | p(s(bool,X14)) | ~sP6(X13,g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))),X14)) ) | $spl40),
  inference(subsumption_resolution,[],[f2975,f200])).
fof(f2975,plain,(
  ( ! [X14,X13] : (~p(s(bool,'T_0')) | ~p(s(bool,X13)) | p(s(bool,X14)) | ~sP6(X13,g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))),X14)) ) | $spl40),
  inference(superposition,[],[f311,f1606])).
fof(f3025,plain,(
  ( ! [X12,X11] : (~p(s(bool,X11)) | p(s(bool,X12)) | ~sP3(X11,g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))),X12)) ) | $spl40),
  inference(subsumption_resolution,[],[f2974,f200])).
fof(f2974,plain,(
  ( ! [X12,X11] : (~p(s(bool,'T_0')) | ~p(s(bool,X11)) | p(s(bool,X12)) | ~sP3(X11,g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))),X12)) ) | $spl40),
  inference(superposition,[],[f292,f1606])).
fof(f3024,plain,(
  ( ! [X10,X9] : (~p(s(bool,X9)) | p(s(bool,X10)) | ~sP4(X9,g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))),X10)) ) | $spl40),
  inference(subsumption_resolution,[],[f2973,f200])).
fof(f2973,plain,(
  ( ! [X10,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,X9)) | p(s(bool,X10)) | ~sP4(X9,g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))),X10)) ) | $spl40),
  inference(superposition,[],[f288,f1606])).
fof(f3023,plain,(
  ( ! [X8,X7] : (sP5(X7,X8,g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12)))) | ~sP4(X7,g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))),X8) | ~sP3(X8,g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))),X7) | ~sP2(g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))),X8,X7)) ) | $spl40),
  inference(subsumption_resolution,[],[f2972,f200])).
fof(f2972,plain,(
  ( ! [X8,X7] : (~p(s(bool,'T_0')) | sP5(X7,X8,g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12)))) | ~sP4(X7,g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))),X8) | ~sP3(X8,g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))),X7) | ~sP2(g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))),X8,X7)) ) | $spl40),
  inference(superposition,[],[f285,f1606])).
fof(f3022,plain,(
  ( ! [X6] : (p(s(bool,X6)) | sP1(X6,g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))))) ) | $spl40),
  inference(subsumption_resolution,[],[f2971,f200])).
fof(f2971,plain,(
  ( ! [X6] : (~p(s(bool,'T_0')) | p(s(bool,X6)) | sP1(X6,g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))))) ) | $spl40),
  inference(superposition,[],[f275,f1606])).
fof(f3021,plain,(
  ( ! [X5] : (p(s(bool,X5)) | sP1(g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))),X5)) ) | $spl40),
  inference(subsumption_resolution,[],[f2970,f200])).
fof(f2970,plain,(
  ( ! [X5] : (~p(s(bool,'T_0')) | p(s(bool,X5)) | sP1(g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))),X5)) ) | $spl40),
  inference(superposition,[],[f272,f1606])).
fof(f3020,plain,(
  ( ! [X4] : (sP0(g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))),X4)) ) | $spl40),
  inference(subsumption_resolution,[],[f2969,f200])).
fof(f2969,plain,(
  ( ! [X4] : (~p(s(bool,'T_0')) | sP0(g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))),X4)) ) | $spl40),
  inference(superposition,[],[f257,f1606])).
fof(f3019,plain,(
  ( ! [X3] : (sP0(X3,g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))))) ) | $spl40),
  inference(subsumption_resolution,[],[f2968,f200])).
fof(f2968,plain,(
  ( ! [X3] : (~p(s(bool,'T_0')) | sP0(X3,g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))))) ) | $spl40),
  inference(superposition,[],[f256,f1606])).
fof(f2953,plain,(
  ~sP5('F_0',g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))),'T_0') | $spl40),
  inference(unit_resulting_resolution,[],[f364,f1606,f303])).
fof(f2954,plain,(
  sP5('T_0',g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))),'T_0') | $spl40),
  inference(unit_resulting_resolution,[],[f200,f1606,f301])).
fof(f2955,plain,(
  ~sP5('F_0','T_0',g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12)))) | $spl40),
  inference(unit_resulting_resolution,[],[f364,f1606,f303])).
fof(f2956,plain,(
  sP5('T_0','T_0',g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12)))) | $spl40),
  inference(unit_resulting_resolution,[],[f200,f1606,f301])).
fof(f1739,plain,(
  sP3('T_0',f(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))),'F_0') | $spl40),
  inference(unit_resulting_resolution,[],[f1631,f282])).
fof(f1733,plain,(
  ~sP4('T_0','T_0',f(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12)))) | $spl40),
  inference(unit_resulting_resolution,[],[f200,f376,f371,f1627,f285])).
fof(f1669,plain,(
  sP3(f(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))),'T_0','F_0') | $spl40),
  inference(unit_resulting_resolution,[],[f1626,f282])).
fof(f1641,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12)))) | s(bool,'T_0') = s(bool,X0)) ) | $spl40),
  inference(trivial_inequality_removal,[],[f1635])).
fof(f1635,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,f(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12)))) | s(bool,'T_0') = s(bool,X0)) ) | $spl40),
  inference(superposition,[],[f1610,f472])).
fof(f1626,plain,(
  sP5('F_0',f(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))),'T_0') | $spl40),
  inference(unit_resulting_resolution,[],[f364,f1610,f300])).
fof(f1627,plain,(
  ~sP5('T_0',f(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))),'T_0') | $spl40),
  inference(unit_resulting_resolution,[],[f200,f1610,f302])).
fof(f1630,plain,(
  ~p(s(bool,f(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))))) | $spl40),
  inference(unit_resulting_resolution,[],[f200,f1610,f254])).
fof(f1631,plain,(
  sP5('F_0','T_0',f(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12)))) | $spl40),
  inference(unit_resulting_resolution,[],[f364,f1610,f300])).
fof(f1632,plain,(
  ~sP5('T_0','T_0',f(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12)))) | $spl40),
  inference(unit_resulting_resolution,[],[f200,f1610,f302])).
fof(f1610,plain,(
  s(bool,'T_0') != s(bool,f(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12)))) | $spl40),
  inference(backward_demodulation,[],[f1606,f720])).
fof(f4740,plain,(
  $false | ($spl40 | $spl45 | $spl98)),
  inference(global_subsumption,[],[f4039,f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f1621,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f797,f896,f1756,f720,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1609,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220])).
fof(f4739,plain,(
  $false | ($spl40 | $spl45 | $spl92)),
  inference(global_subsumption,[],[f2789,f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f1621,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f797,f896,f1756,f720,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1609,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220])).
fof(f4738,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f1621,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f797,f896,f1756,f720,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1609,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220])).
fof(f4737,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f1621,f220,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f797,f896,f1756,f720,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1609,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221])).
fof(f4736,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f1621,f221,f220,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f797,f896,f1756,f720,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1609,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364])).
fof(f4735,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f1621,f364,f221,f220,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f797,f896,f1756,f720,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1609,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722])).
fof(f4734,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f1621,f722,f364,f221,f220,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f797,f896,f1756,f720,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1609,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723])).
fof(f4733,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f1621,f723,f722,f364,f221,f220,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f797,f896,f1756,f720,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1609,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448])).
fof(f4732,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f1621,f1448,f723,f722,f364,f221,f220,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f797,f896,f1756,f720,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1609,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453])).
fof(f4731,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f1621,f1453,f1448,f723,f722,f364,f221,f220,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f797,f896,f1756,f720,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1609,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468])).
fof(f4730,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f1621,f1468,f1453,f1448,f723,f722,f364,f221,f220,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f797,f896,f1756,f720,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1609,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473])).
fof(f4729,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f1621,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f797,f896,f1756,f720,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1609,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486])).
fof(f4728,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f1621,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f797,f896,f1756,f720,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1609,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491])).
fof(f4727,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f1621,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f797,f896,f1756,f720,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1609,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509])).
fof(f4726,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f1621,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f797,f896,f1756,f720,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1609,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514])).
fof(f4725,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f1621,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f797,f896,f1756,f720,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1609,f1578,f1573,f1558,f1553,f1545,f1540,f1535])).
fof(f4724,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f1621,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f797,f896,f1756,f720,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1609,f1578,f1573,f1558,f1553,f1545,f1540])).
fof(f4723,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f1621,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f797,f896,f1756,f720,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1609,f1578,f1573,f1558,f1553,f1545])).
fof(f4722,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f1621,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f797,f896,f1756,f720,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1609,f1578,f1573,f1558,f1553])).
fof(f4721,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f1621,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f797,f896,f1756,f720,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1609,f1578,f1573,f1558])).
fof(f4720,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f1621,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f797,f896,f1756,f720,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1609,f1578,f1573])).
fof(f4719,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f1621,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f797,f896,f1756,f720,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1609,f1578])).
fof(f4718,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f797,f896,f1756,f720,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1609])).
fof(f4717,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f1609,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f797,f896,f1756,f720,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826])).
fof(f4716,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f1609,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f797,f896,f1756,f720,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865])).
fof(f4715,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f1865,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f1609,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f797,f896,f1756,f720,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901])).
fof(f4714,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f1901,f1865,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f1609,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f797,f896,f1756,f720,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936])).
fof(f4713,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f1936,f1901,f1865,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f1609,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f797,f896,f1756,f720,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067])).
fof(f4712,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f2067,f1936,f1901,f1865,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f1609,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f797,f896,f1756,f720,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102])).
fof(f4711,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f2102,f2067,f1936,f1901,f1865,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f1609,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f797,f896,f1756,f720,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146])).
fof(f4710,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f1609,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f797,f896,f1756,f720,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186])).
fof(f4709,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f1609,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f797,f896,f1756,f720,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214])).
fof(f4708,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f1609,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f797,f896,f1756,f720,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216])).
fof(f4707,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f1609,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f797,f896,f1756,f720,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218])).
fof(f4706,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f1609,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f797,f896,f1756,f720,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225])).
fof(f4705,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f1609,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f797,f896,f1756,f720,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227])).
fof(f4704,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f1609,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f797,f896,f1756,f720,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229])).
fof(f4703,plain,(
  $false | ($spl40 | $spl45 | $spl76)),
  inference(global_subsumption,[],[f2303,f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f1609,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f797,f896,f1756,f720,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462])).
fof(f4702,plain,(
  $false | ($spl40 | $spl45 | $spl86)),
  inference(global_subsumption,[],[f2410,f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f1609,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f797,f896,f1756,f720,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462])).
fof(f4701,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f1609,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f797,f896,f1756,f720,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462])).
fof(f4700,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f1609,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f797,f896,f1756,f720,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464])).
fof(f4699,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f1609,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f797,f896,f1756,f720,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466])).
fof(f4698,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f1609,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f797,f896,f1756,f720,f2668,f2634,f2590,f2556,f2505,f2503,f2501])).
fof(f4697,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f1609,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f797,f896,f1756,f720,f2668,f2634,f2590,f2556,f2505,f2503])).
fof(f4696,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f1609,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f797,f896,f1756,f720,f2668,f2634,f2590,f2556,f2505])).
fof(f4695,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f1609,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f797,f896,f1756,f720,f2668,f2634,f2590,f2556])).
fof(f4694,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f1609,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f797,f896,f1756,f720,f2668,f2634,f2590])).
fof(f4693,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f1609,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f797,f896,f1756,f720,f2668,f2634])).
fof(f4692,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f1609,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f797,f896,f1756,f720,f2668])).
fof(f4691,plain,(
  $false | ($spl40 | $spl45 | $spl49)),
  inference(global_subsumption,[],[f4071,f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f1609,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f797,f896,f1756,f720])).
fof(f4690,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f1609,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f797,f896,f1756,f720])).
fof(f4689,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f720,f1609,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f797,f896,f1756])).
fof(f4688,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f1756,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f720,f1609,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f797,f896])).
fof(f4687,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f896,f1756,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f720,f1609,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244,f797])).
fof(f4686,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f797,f896,f1756,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f720,f1609,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243,f244])).
fof(f4685,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f244,f797,f896,f1756,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f720,f1609,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245,f243])).
fof(f4684,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f244,f243,f797,f896,f1756,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f720,f1609,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246,f245])).
fof(f4683,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f244,f245,f243,f797,f896,f1756,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f720,f1609,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247,f246])).
fof(f4682,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f244,f246,f245,f243,f797,f896,f1756,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f720,f1609,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248,f247])).
fof(f4681,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f244,f247,f246,f245,f243,f797,f896,f1756,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f720,f1609,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249,f248])).
fof(f4680,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f244,f247,f246,f245,f243,f248,f797,f896,f1756,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f720,f1609,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481,f1461,f250,f249])).
fof(f4679,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f244,f247,f246,f245,f243,f249,f248,f797,f896,f1756,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f720,f1609,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481,f1461,f250])).
fof(f4678,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f244,f247,f246,f245,f243,f250,f249,f248,f797,f896,f1756,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f720,f1609,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481,f1461])).
fof(f4677,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f244,f1461,f247,f246,f245,f243,f250,f249,f248,f797,f896,f1756,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f720,f1609,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499,f1481])).
fof(f4676,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f244,f1481,f1461,f247,f246,f245,f243,f250,f249,f248,f797,f896,f1756,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f720,f1609,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522,f1499])).
fof(f4675,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f244,f1481,f1461,f247,f246,f245,f243,f250,f249,f248,f1499,f797,f896,f1756,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f720,f1609,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566,f1522])).
fof(f4674,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f244,f1481,f1461,f247,f246,f245,f243,f250,f249,f248,f1499,f1522,f797,f896,f1756,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f720,f1609,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586,f1566])).
fof(f4673,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f244,f1566,f1481,f1461,f247,f246,f245,f243,f250,f249,f248,f1499,f1522,f797,f896,f1756,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f720,f1609,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487,f1586])).
fof(f4672,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f244,f1566,f1481,f1461,f247,f246,f245,f243,f250,f249,f248,f1499,f1522,f1586,f797,f896,f1756,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f720,f1609,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489,f1487])).
fof(f4671,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f1487,f244,f1566,f1481,f1461,f247,f246,f245,f243,f250,f249,f248,f1499,f1522,f1586,f797,f896,f1756,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f720,f1609,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510,f1489])).
fof(f4670,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f1489,f1487,f244,f1566,f1481,f1461,f247,f246,f245,f243,f250,f249,f248,f1499,f1522,f1586,f797,f896,f1756,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f720,f1609,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512,f1510])).
fof(f4669,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f1489,f1510,f1487,f244,f1566,f1481,f1461,f247,f246,f245,f243,f250,f249,f248,f1499,f1522,f1586,f797,f896,f1756,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f720,f1609,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576,f1512])).
fof(f4668,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f1512,f1489,f1510,f1487,f244,f1566,f1481,f1461,f247,f246,f245,f243,f250,f249,f248,f1499,f1522,f1586,f797,f896,f1756,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f720,f1609,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579,f1576])).
fof(f4667,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f1576,f1512,f1489,f1510,f1487,f244,f1566,f1481,f1461,f247,f246,f245,f243,f250,f249,f248,f1499,f1522,f1586,f797,f896,f1756,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f720,f1609,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449,f1579])).
fof(f4666,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f1576,f1512,f1489,f1579,f1510,f1487,f244,f1566,f1481,f1461,f247,f246,f245,f243,f250,f249,f248,f1499,f1522,f1586,f797,f896,f1756,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f720,f1609,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451,f1449])).
fof(f4665,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f1576,f1512,f1489,f1449,f1579,f1510,f1487,f244,f1566,f1481,f1461,f247,f246,f245,f243,f250,f249,f248,f1499,f1522,f1586,f797,f896,f1756,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f720,f1609,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471,f1451])).
fof(f4664,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f1576,f1512,f1489,f1451,f1449,f1579,f1510,f1487,f244,f1566,f1481,f1461,f247,f246,f245,f243,f250,f249,f248,f1499,f1522,f1586,f797,f896,f1756,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f720,f1609,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474,f1471])).
fof(f4663,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f1576,f1512,f1489,f1471,f1451,f1449,f1579,f1510,f1487,f244,f1566,f1481,f1461,f247,f246,f245,f243,f250,f249,f248,f1499,f1522,f1586,f797,f896,f1756,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f720,f1609,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492,f1474])).
fof(f4662,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f1576,f1512,f1489,f1474,f1471,f1451,f1449,f1579,f1510,f1487,f244,f1566,f1481,f1461,f247,f246,f245,f243,f250,f249,f248,f1499,f1522,f1586,f797,f896,f1756,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f720,f1609,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515,f1492])).
fof(f4661,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f1576,f1512,f1492,f1489,f1474,f1471,f1451,f1449,f1579,f1510,f1487,f244,f1566,f1481,f1461,f247,f246,f245,f243,f250,f249,f248,f1499,f1522,f1586,f797,f896,f1756,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f720,f1609,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538,f1515])).
fof(f4660,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f1576,f1515,f1512,f1492,f1489,f1474,f1471,f1451,f1449,f1579,f1510,f1487,f244,f1566,f1481,f1461,f247,f246,f245,f243,f250,f249,f248,f1499,f1522,f1586,f797,f896,f1756,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f720,f1609,f1536,f1469,f1454,f1574,f1559,f1556,f1554,f1538])).
fof(f4659,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f1576,f1538,f1515,f1512,f1492,f1489,f1474,f1471,f1451,f1449,f1579,f1510,f1487,f244,f1566,f1481,f1461,f247,f246,f245,f243,f250,f249,f248,f1499,f1522,f1586,f797,f896,f1756,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f720,f1609,f1536,f1469,f1454,f1574,f1559,f1556,f1554])).
fof(f4658,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f1576,f1538,f1515,f1512,f1492,f1489,f1474,f1471,f1451,f1449,f1579,f1554,f1510,f1487,f244,f1566,f1481,f1461,f247,f246,f245,f243,f250,f249,f248,f1499,f1522,f1586,f797,f896,f1756,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f720,f1609,f1536,f1469,f1454,f1574,f1559,f1556])).
fof(f4657,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f1576,f1556,f1538,f1515,f1512,f1492,f1489,f1474,f1471,f1451,f1449,f1579,f1554,f1510,f1487,f244,f1566,f1481,f1461,f247,f246,f245,f243,f250,f249,f248,f1499,f1522,f1586,f797,f896,f1756,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f720,f1609,f1536,f1469,f1454,f1574,f1559])).
fof(f4656,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f1576,f1556,f1538,f1515,f1512,f1492,f1489,f1474,f1471,f1451,f1449,f1579,f1559,f1554,f1510,f1487,f244,f1566,f1481,f1461,f247,f246,f245,f243,f250,f249,f248,f1499,f1522,f1586,f797,f896,f1756,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f720,f1609,f1536,f1469,f1454,f1574])).
fof(f4655,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f1576,f1574,f1556,f1538,f1515,f1512,f1492,f1489,f1474,f1471,f1451,f1449,f1579,f1559,f1554,f1510,f1487,f244,f1566,f1481,f1461,f247,f246,f245,f243,f250,f249,f248,f1499,f1522,f1586,f797,f896,f1756,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f720,f1609,f1536,f1469,f1454])).
fof(f4654,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f1576,f1574,f1556,f1538,f1515,f1512,f1492,f1489,f1474,f1471,f1451,f1449,f1579,f1559,f1554,f1510,f1487,f1454,f244,f1566,f1481,f1461,f247,f246,f245,f243,f250,f249,f248,f1499,f1522,f1586,f797,f896,f1756,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f720,f1609,f1536,f1469])).
fof(f4653,plain,(
  $false | ($spl40 | $spl45)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f1576,f1574,f1556,f1538,f1515,f1512,f1492,f1489,f1474,f1471,f1451,f1449,f1579,f1559,f1554,f1510,f1487,f1469,f1454,f244,f1566,f1481,f1461,f247,f246,f245,f243,f250,f249,f248,f1499,f1522,f1586,f797,f896,f1756,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f720,f1609,f1536])).
fof(f4652,plain,(
  $false | ($spl40 | $spl45 | $spl68)),
  inference(global_subsumption,[],[f1752,f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1606,f1610,f1632,f1631,f1630,f1627,f1626,f1641,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1669,f1718,f1719,f1720,f1721,f1722,f1733,f1739,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f2956,f2955,f2954,f2953,f3019,f3020,f3021,f3022,f3023,f3024,f3025,f3026,f3027,f3028,f3029,f3030,f3031,f3032,f3033,f3034,f3035,f3036,f3037,f3248,f3258,f3261,f3866,f3922,f3921,f3920,f3919,f4029,f4032,f4036,f721,f4040,f4063,f4068,f1576,f1574,f1556,f1538,f1515,f1512,f1492,f1489,f1474,f1471,f1451,f1449,f1579,f1559,f1554,f1536,f1510,f1487,f1469,f1454,f244,f1566,f1481,f1461,f247,f246,f245,f243,f250,f249,f248,f1499,f1522,f1586,f797,f896,f1756,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1621,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f720,f1609])).
fof(f4651,plain,(
  $false | ($spl42 | $spl96)),
  inference(subsumption_resolution,[],[f4650,f4034])).
fof(f4650,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))) ) | ($spl42 | $spl96)),
  inference(subsumption_resolution,[],[f4590,f200])).
fof(f4590,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))) ) | ($spl42 | $spl96)),
  inference(backward_demodulation,[],[f4535,f246])).
fof(f4535,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,X2)) ) | ($spl42 | $spl96)),
  inference(global_subsumption,[],[f1607,f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1530,f1529,f1531,f1533,f1541,f1543,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1718,f1719,f1720,f1721,f1722,f1741,f1742,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2475,f3866,f4029,f4032,f721,f1576,f1574,f1556,f1538,f1515,f1512,f1492,f1489,f1474,f1471,f1451,f1449,f1579,f1559,f1554,f1536,f1510,f1487,f1469,f1454,f244,f1566,f1481,f1461,f247,f246,f245,f243,f250,f249,f248,f1499,f1522,f1586,f797,f896,f1756,f2668,f2634,f2590,f2556,f2505,f2503,f2501,f2466,f2464,f2462,f2229,f2227,f2225,f2218,f2216,f2214,f2186,f2146,f2102,f2067,f1936,f1901,f1865,f1826,f1578,f1573,f1558,f1553,f1545,f1540,f1535,f1514,f1509,f1491,f1486,f1473,f1468,f1453,f1448,f723,f722,f364,f221,f220,f4034,f720])).
fof(f4648,plain,(
  $false | ($spl42 | $spl96)),
  inference(subsumption_resolution,[],[f4584,f4535])).
fof(f4584,plain,(
  s(bool,'T_0') != s(bool,g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12)))) | ($spl42 | $spl96)),
  inference(backward_demodulation,[],[f4535,f720])).
fof(f4647,plain,(
  $false | ($spl42 | $spl96)),
  inference(subsumption_resolution,[],[f4538,f200])).
fof(f4538,plain,(
  ~p(s(bool,'T_0')) | ($spl42 | $spl96)),
  inference(backward_demodulation,[],[f4535,f364])).
fof(f4646,plain,(
  $false | ($spl42 | $spl96)),
  inference(subsumption_resolution,[],[f4645,f4034])).
fof(f4645,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl42 | $spl96)),
  inference(subsumption_resolution,[],[f4536,f4535])).
fof(f4536,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | ($spl42 | $spl96)),
  inference(backward_demodulation,[],[f4535,f220])).
fof(f4534,plain,(
  $false | ($spl70 | $spl96)),
  inference(subsumption_resolution,[],[f4295,f4034])).
fof(f4295,plain,(
  ~p(s(bool,f(s(bool,'T_0')))) | $spl70),
  inference(unit_resulting_resolution,[],[f2795,f200,f364,f284])).
fof(f2795,plain,(
  sP5('T_0',f(s(bool,'T_0')),'F_0') | $spl70),
  inference(unit_resulting_resolution,[],[f200,f1753,f301])).
fof(f1753,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | $spl70),
  inference(cnf_transformation,[],[f1753_D])).
fof(f1753_D,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f4533,plain,(
  $false | ($spl40 | $spl96)),
  inference(subsumption_resolution,[],[f4296,f4034])).
fof(f4296,plain,(
  ~p(s(bool,f(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))))) | $spl40),
  inference(unit_resulting_resolution,[],[f3920,f200,f364,f284])).
fof(f4532,plain,(
  $false | ($spl70 | $spl96)),
  inference(subsumption_resolution,[],[f4330,f4034])).
fof(f4330,plain,(
  ~p(s(bool,f(s(bool,'T_0')))) | $spl70),
  inference(unit_resulting_resolution,[],[f2941,f200,f364,f288])).
fof(f2941,plain,(
  sP4('T_0',f(s(bool,'T_0')),'F_0') | $spl70),
  inference(unit_resulting_resolution,[],[f2797,f283])).
fof(f2797,plain,(
  sP5('T_0','F_0',f(s(bool,'T_0'))) | $spl70),
  inference(unit_resulting_resolution,[],[f200,f1753,f301])).
fof(f4531,plain,(
  $false | ($spl40 | $spl96)),
  inference(subsumption_resolution,[],[f4336,f4034])).
fof(f4336,plain,(
  ~p(s(bool,f(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))))) | $spl40),
  inference(unit_resulting_resolution,[],[f1669,f200,f364,f292])).
fof(f4530,plain,(
  $false | ($spl40 | $spl96)),
  inference(subsumption_resolution,[],[f4337,f4034])).
fof(f4337,plain,(
  ~p(s(bool,f(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))))) | $spl40),
  inference(unit_resulting_resolution,[],[f1739,f200,f364,f292])).
fof(f4529,plain,(
  $false | ($spl10 | $spl40 | $spl96)),
  inference(subsumption_resolution,[],[f4493,f4034])).
fof(f4493,plain,(
  ~p(s(bool,f(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))))) | ($spl10 | $spl40)),
  inference(unit_resulting_resolution,[],[f1669,f364,f1097])).
fof(f4272,plain,(
  $false | ($spl49 | $spl52)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f250,f249,f248,f247,f246,f245,f243,f244,f797,f896,f720,f723,f722,f364,f221,f220,f1448,f1449,f1451,f1453,f1454,f1461,f1468,f1469,f1471,f1473,f1474,f1481,f1486,f1487,f1489,f1491,f1492,f1499,f1509,f1510,f1512,f1514,f1515,f1522,f1530,f1529,f1531,f1533,f1535,f1536,f1538,f1540,f1541,f1543,f1545,f1553,f1554,f1556,f1558,f1559,f1566,f1573,f1574,f1576,f1578,f1579,f1586,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1718,f1719,f1720,f1721,f1722,f1741,f1742,f1756,f1826,f1865,f1901,f1936,f2067,f2102,f2146,f2186,f2214,f2216,f2218,f2225,f2227,f2229,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2462,f2464,f2466,f2475,f2501,f2503,f2505,f2556,f2590,f2634,f2668,f3866,f4029,f4032,f1657,f721,f1660,f4118,f4117,f4116,f4109,f4108,f4103,f4099,f4096,f4093,f4089,f4155,f4262,f4271])).
fof(f4271,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | ($spl49 | $spl52)),
  inference(subsumption_resolution,[],[f4225,f4155])).
fof(f4225,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X0,sK15(X2,X1,X0)),s(fun(X0,bool),X1))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | ($spl49 | $spl52)),
  inference(backward_demodulation,[],[f4155,f244])).
fof(f4262,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl49 | $spl52)),
  inference(subsumption_resolution,[],[f4156,f4155])).
fof(f4156,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | ($spl49 | $spl52)),
  inference(backward_demodulation,[],[f4155,f220])).
fof(f4155,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl49 | $spl52)),
  inference(subsumption_resolution,[],[f4127,f1660])).
fof(f4127,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) | s(bool,'T_0') = s(bool,X0)) ) | $spl49),
  inference(trivial_inequality_removal,[],[f4122])).
fof(f4122,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,f(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) | s(bool,'T_0') = s(bool,X0)) ) | $spl49),
  inference(superposition,[],[f1657,f472])).
fof(f4089,plain,(
  p(s(bool,g(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12))))) | $spl49),
  inference(unit_resulting_resolution,[],[f721,f1657,f518])).
fof(f4093,plain,(
  s(bool,'F_0') != s(bool,g(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) | $spl49),
  inference(unit_resulting_resolution,[],[f721,f1657,f1134])).
fof(f4096,plain,(
  ( ! [X0] : (sP0(X0,g(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12))))) ) | $spl49),
  inference(unit_resulting_resolution,[],[f721,f1657,f1136])).
fof(f4099,plain,(
  ( ! [X0] : (sP0(g(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12))),X0)) ) | $spl49),
  inference(unit_resulting_resolution,[],[f721,f1657,f1137])).
fof(f4103,plain,(
  sP1(g(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12))),'F_0') | $spl49),
  inference(unit_resulting_resolution,[],[f364,f721,f1657,f1138])).
fof(f4108,plain,(
  sP5('F_0',f(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12))),'T_0') | $spl49),
  inference(unit_resulting_resolution,[],[f364,f1657,f300])).
fof(f4109,plain,(
  ~sP5('T_0',f(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12))),'T_0') | $spl49),
  inference(unit_resulting_resolution,[],[f200,f1657,f302])).
fof(f4116,plain,(
  ~p(s(bool,f(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12))))) | $spl49),
  inference(unit_resulting_resolution,[],[f200,f1657,f254])).
fof(f4117,plain,(
  sP5('F_0','T_0',f(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) | $spl49),
  inference(unit_resulting_resolution,[],[f364,f1657,f300])).
fof(f4118,plain,(
  ~sP5('T_0','T_0',f(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) | $spl49),
  inference(unit_resulting_resolution,[],[f200,f1657,f302])).
fof(f4270,plain,(
  $false | ($spl49 | $spl52)),
  inference(subsumption_resolution,[],[f4269,f200])).
fof(f4269,plain,(
  ~p(s(bool,'T_0')) | ($spl49 | $spl52)),
  inference(forward_demodulation,[],[f4268,f4155])).
fof(f4268,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))) ) | ($spl49 | $spl52)),
  inference(subsumption_resolution,[],[f4220,f200])).
fof(f4220,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))) ) | ($spl49 | $spl52)),
  inference(backward_demodulation,[],[f4155,f246])).
fof(f4265,plain,(
  $false | ($spl49 | $spl52)),
  inference(subsumption_resolution,[],[f4209,f200])).
fof(f4209,plain,(
  ~p(s(bool,'T_0')) | ($spl49 | $spl52)),
  inference(backward_demodulation,[],[f4155,f4116])).
fof(f4264,plain,(
  $false | ($spl49 | $spl52)),
  inference(subsumption_resolution,[],[f4158,f200])).
fof(f4158,plain,(
  ~p(s(bool,'T_0')) | ($spl49 | $spl52)),
  inference(backward_demodulation,[],[f4155,f364])).
fof(f4263,plain,(
  $false | ($spl49 | $spl52)),
  inference(global_subsumption,[],[f199,f200,f207,f217,f216,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f743,f796,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f250,f249,f248,f247,f246,f245,f243,f244,f797,f896,f720,f723,f722,f364,f221,f220,f1448,f1449,f1451,f1453,f1454,f1461,f1468,f1469,f1471,f1473,f1474,f1481,f1486,f1487,f1489,f1491,f1492,f1499,f1509,f1510,f1512,f1514,f1515,f1522,f1530,f1529,f1531,f1533,f1535,f1536,f1538,f1540,f1541,f1543,f1545,f1553,f1554,f1556,f1558,f1559,f1566,f1573,f1574,f1576,f1578,f1579,f1586,f1596,f1595,f1594,f1593,f1597,f1598,f1599,f1600,f1601,f1649,f1648,f1647,f1646,f1650,f1651,f1652,f1653,f1654,f1718,f1719,f1720,f1721,f1722,f1741,f1742,f1756,f1826,f1865,f1901,f1936,f2067,f2102,f2146,f2186,f2214,f2216,f2218,f2225,f2227,f2229,f2255,f2254,f2252,f2251,f2256,f2257,f2258,f2259,f2260,f2357,f2356,f2354,f2353,f2358,f2359,f2360,f2361,f2362,f2462,f2464,f2466,f2475,f2501,f2503,f2505,f2556,f2590,f2634,f2668,f3866,f4029,f4032,f1657,f721,f1660,f4118,f4117,f4116,f4109,f4108,f4103,f4099,f4096,f4093,f4089,f4155,f4262])).
fof(f4154,plain,(
  $false | ($spl49 | $spl55)),
  inference(subsumption_resolution,[],[f4071,f1662])).
fof(f4153,plain,(
  $false | ($spl49 | $spl55)),
  inference(subsumption_resolution,[],[f4073,f1662])).
fof(f4073,plain,(
  s(bool,'F_0') = s(bool,f(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) | $spl49),
  inference(unit_resulting_resolution,[],[f386,f1657,f472])).
fof(f4152,plain,(
  $false | ($spl40 | $spl49 | $spl55)),
  inference(subsumption_resolution,[],[f4151,f1662])).
fof(f4151,plain,(
  s(bool,'F_0') = s(bool,f(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) | ($spl40 | $spl49)),
  inference(forward_demodulation,[],[f4074,f1621])).
fof(f4074,plain,(
  s(bool,f(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12)))) = s(bool,f(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) | ($spl40 | $spl49)),
  inference(unit_resulting_resolution,[],[f1610,f1657,f472])).
fof(f4150,plain,(
  $false | ($spl49 | $spl55 | $spl76)),
  inference(subsumption_resolution,[],[f4149,f1662])).
fof(f4149,plain,(
  s(bool,'F_0') = s(bool,f(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) | ($spl49 | $spl76)),
  inference(forward_demodulation,[],[f4076,f2303])).
fof(f4076,plain,(
  s(bool,f(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | ($spl49 | $spl76)),
  inference(unit_resulting_resolution,[],[f2269,f1657,f472])).
fof(f4148,plain,(
  $false | ($spl49 | $spl55 | $spl86)),
  inference(subsumption_resolution,[],[f4147,f1662])).
fof(f4147,plain,(
  s(bool,'F_0') = s(bool,f(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) | ($spl49 | $spl86)),
  inference(forward_demodulation,[],[f4077,f2410])).
fof(f4077,plain,(
  s(bool,f(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) = s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | ($spl49 | $spl86)),
  inference(unit_resulting_resolution,[],[f2371,f1657,f472])).
fof(f4146,plain,(
  $false | ($spl49 | $spl51)),
  inference(subsumption_resolution,[],[f4079,f1658])).
fof(f1658,plain,(
  s(bool,'T_0') != s(bool,g(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) | $spl51),
  inference(cnf_transformation,[],[f1658_D])).
fof(f1658_D,plain,(
  s(bool,'T_0') != s(bool,g(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) <=> ~$spl51),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl51])])).
fof(f4079,plain,(
  s(bool,'T_0') = s(bool,g(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) | $spl49),
  inference(unit_resulting_resolution,[],[f721,f1657,f472])).
fof(f4145,plain,(
  $false | ($spl49 | $spl55)),
  inference(subsumption_resolution,[],[f4080,f1662])).
fof(f4080,plain,(
  s(bool,'F_0') = s(bool,f(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) | $spl49),
  inference(unit_resulting_resolution,[],[f386,f1657,f472])).
fof(f4144,plain,(
  $false | ($spl40 | $spl49 | $spl55)),
  inference(subsumption_resolution,[],[f4143,f1662])).
fof(f4143,plain,(
  s(bool,'F_0') = s(bool,f(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) | ($spl40 | $spl49)),
  inference(forward_demodulation,[],[f4081,f1621])).
fof(f4081,plain,(
  s(bool,f(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12)))) = s(bool,f(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) | ($spl40 | $spl49)),
  inference(unit_resulting_resolution,[],[f1610,f1657,f472])).
fof(f4142,plain,(
  $false | ($spl49 | $spl55 | $spl76)),
  inference(subsumption_resolution,[],[f4141,f1662])).
fof(f4141,plain,(
  s(bool,'F_0') = s(bool,f(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) | ($spl49 | $spl76)),
  inference(forward_demodulation,[],[f4083,f2303])).
fof(f4083,plain,(
  s(bool,f(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | ($spl49 | $spl76)),
  inference(unit_resulting_resolution,[],[f2269,f1657,f472])).
fof(f4140,plain,(
  $false | ($spl49 | $spl55 | $spl86)),
  inference(subsumption_resolution,[],[f4139,f1662])).
fof(f4139,plain,(
  s(bool,'F_0') = s(bool,f(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) | ($spl49 | $spl86)),
  inference(forward_demodulation,[],[f4084,f2410])).
fof(f4084,plain,(
  s(bool,f(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) = s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | ($spl49 | $spl86)),
  inference(unit_resulting_resolution,[],[f2371,f1657,f472])).
fof(f4138,plain,(
  $false | ($spl49 | $spl51)),
  inference(subsumption_resolution,[],[f4086,f1658])).
fof(f4086,plain,(
  s(bool,'T_0') = s(bool,g(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) | $spl49),
  inference(unit_resulting_resolution,[],[f721,f1657,f472])).
fof(f4137,plain,(
  $false | ($spl49 | $spl55)),
  inference(subsumption_resolution,[],[f4087,f1662])).
fof(f4087,plain,(
  s(bool,'F_0') = s(bool,f(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) | $spl49),
  inference(unit_resulting_resolution,[],[f364,f1657,f518])).
fof(f4136,plain,(
  $false | ($spl49 | $spl55 | $spl70)),
  inference(subsumption_resolution,[],[f4135,f1662])).
fof(f4135,plain,(
  s(bool,'F_0') = s(bool,f(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) | ($spl49 | $spl70)),
  inference(forward_demodulation,[],[f4090,f1753])).
fof(f4090,plain,(
  s(bool,f(s(bool,'T_0'))) = s(bool,f(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) | ($spl49 | $spl70)),
  inference(unit_resulting_resolution,[],[f1753,f1657,f1134])).
fof(f4134,plain,(
  $false | ($spl40 | $spl49 | $spl55)),
  inference(subsumption_resolution,[],[f4133,f1662])).
fof(f4133,plain,(
  s(bool,'F_0') = s(bool,f(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) | ($spl40 | $spl49)),
  inference(forward_demodulation,[],[f4091,f1621])).
fof(f4091,plain,(
  s(bool,f(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12)))) = s(bool,f(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) | ($spl40 | $spl49)),
  inference(unit_resulting_resolution,[],[f1621,f1657,f1134])).
fof(f4132,plain,(
  $false | ($spl49 | $spl55)),
  inference(subsumption_resolution,[],[f4094,f1662])).
fof(f4094,plain,(
  s(bool,'F_0') = s(bool,f(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) | $spl49),
  inference(unit_resulting_resolution,[],[f407,f1657,f1136])).
fof(f4131,plain,(
  $false | ($spl49 | $spl55)),
  inference(subsumption_resolution,[],[f4097,f1662])).
fof(f4097,plain,(
  s(bool,'F_0') = s(bool,f(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) | $spl49),
  inference(unit_resulting_resolution,[],[f407,f1657,f1137])).
fof(f4130,plain,(
  $false | ($spl49 | $spl55)),
  inference(subsumption_resolution,[],[f4100,f1662])).
fof(f4100,plain,(
  s(bool,'F_0') = s(bool,f(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) | $spl49),
  inference(unit_resulting_resolution,[],[f393,f364,f1657,f1138])).
fof(f4129,plain,(
  $false | ($spl40 | $spl49 | $spl55)),
  inference(subsumption_resolution,[],[f4128,f1662])).
fof(f4128,plain,(
  s(bool,'F_0') = s(bool,f(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) | ($spl40 | $spl49)),
  inference(forward_demodulation,[],[f4106,f1621])).
fof(f4106,plain,(
  s(bool,f(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12)))) = s(bool,f(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) | ($spl40 | $spl49)),
  inference(unit_resulting_resolution,[],[f1657,f1641])).
fof(f4039,plain,(
  ( ! [X45,X44] : (sP1(X44,f(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) | s(bool,'T_0') = s(bool,X45) | s(bool,X44) = s(bool,X45)) ) | $spl98),
  inference(cnf_transformation,[],[f4039_D])).
fof(f4039_D,plain,(
  ( ! [X45,X44] : (sP1(X44,f(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) | s(bool,'T_0') = s(bool,X45) | s(bool,X44) = s(bool,X45)) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f4038,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,X43)) | sP1(X44,f(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) | s(bool,X44) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) ) | $spl50),
  inference(subsumption_resolution,[],[f4026,f216])).
fof(f4026,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,X43)) | sP1(X44,f(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) | s(bool,X44) = s(bool,X45) | s(bool,'T_0') = s(bool,X45) | s(bool,'T_0') = s(bool,X43)) ) | $spl50),
  inference(superposition,[],[f1138,f1703])).
fof(f1703,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) | s(bool,'T_0') = s(bool,X0)) ) | $spl50),
  inference(trivial_inequality_removal,[],[f1697])).
fof(f1697,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,f(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) | s(bool,'T_0') = s(bool,X0)) ) | $spl50),
  inference(superposition,[],[f1663,f472])).
fof(f1663,plain,(
  s(bool,'T_0') != s(bool,f(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) | $spl50),
  inference(backward_demodulation,[],[f1659,f721])).
fof(f4033,plain,(
  ( ! [X35,X34] : (sP1(X34,'F_0') | s(bool,'T_0') = s(bool,X35) | s(bool,X34) = s(bool,X35)) ) | $spl94),
  inference(cnf_transformation,[],[f4033_D])).
fof(f4033_D,plain,(
  ( ! [X35,X34] : (sP1(X34,'F_0') | s(bool,'T_0') = s(bool,X35) | s(bool,X34) = s(bool,X35)) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f4032,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X33)) | sP1(X34,'F_0') | s(bool,X34) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) )),
  inference(subsumption_resolution,[],[f4022,f216])).
fof(f4022,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X33)) | sP1(X34,'F_0') | s(bool,X34) = s(bool,X35) | s(bool,'T_0') = s(bool,X35) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f1138,f207])).
fof(f2789,plain,(
  ( ! [X1] : (~sP11(f(s(bool,'T_0')),X1,'T_0') | p(s(bool,X1))) ) | $spl92),
  inference(cnf_transformation,[],[f2789_D])).
fof(f2789_D,plain,(
  ( ! [X1] : (~sP11(f(s(bool,'T_0')),X1,'T_0') | p(s(bool,X1))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f1084,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl2),
  inference(cnf_transformation,[],[f1084_D])).
fof(f1084_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f2791,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP11(f(s(bool,'T_0')),X1,'T_0')) ) | $spl71),
  inference(condensation,[],[f2790])).
fof(f2790,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X18)) | ~sP11(f(s(bool,'T_0')),X18,'T_0') | s(bool,X17) = s(bool,X19) | s(bool,'T_0') = s(bool,X19)) ) | $spl71),
  inference(subsumption_resolution,[],[f2783,f518])).
fof(f2783,plain,(
  ( ! [X19,X17,X18] : (~p(s(bool,X17)) | p(s(bool,X18)) | ~sP11(f(s(bool,'T_0')),X18,'T_0') | s(bool,X17) = s(bool,X19) | s(bool,'T_0') = s(bool,X19)) ) | $spl71),
  inference(superposition,[],[f2037,f472])).
fof(f2037,plain,(
  ( ! [X31,X32] : (~p(s(bool,X31)) | p(s(bool,X32)) | ~sP11(f(s(bool,'T_0')),X32,X31)) ) | $spl71),
  inference(subsumption_resolution,[],[f1988,f200])).
fof(f1988,plain,(
  ( ! [X31,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X31)) | p(s(bool,X32)) | ~sP11(f(s(bool,'T_0')),X32,X31)) ) | $spl71),
  inference(superposition,[],[f352,f1771])).
fof(f1771,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | $spl71),
  inference(unit_resulting_resolution,[],[f364,f1754,f518])).
fof(f2788,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP11(f(s(bool,'T_0')),X1,'T_0')) ) | $spl71),
  inference(condensation,[],[f2787])).
fof(f2787,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X15)) | ~sP11(f(s(bool,'T_0')),X15,'T_0') | s(bool,X14) = s(bool,X16) | s(bool,'T_0') = s(bool,X16)) ) | $spl71),
  inference(subsumption_resolution,[],[f2782,f518])).
fof(f2782,plain,(
  ( ! [X14,X15,X16] : (~p(s(bool,X14)) | p(s(bool,X15)) | ~sP11(f(s(bool,'T_0')),X15,'T_0') | s(bool,X14) = s(bool,X16) | s(bool,'T_0') = s(bool,X16)) ) | $spl71),
  inference(superposition,[],[f2037,f472])).
fof(f2367,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))))))) | $spl86),
  inference(cnf_transformation,[],[f2367_D])).
fof(f2367_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))))))) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f2370,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | $spl91),
  inference(cnf_transformation,[],[f2370_D])).
fof(f2370_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) <=> ~$spl91),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl91])])).
fof(f2362,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))))),
  inference(superposition,[],[f1530,f207])).
fof(f2368,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | s(bool,'T_0') = s(bool,X2)) ) | $spl88),
  inference(cnf_transformation,[],[f2368_D])).
fof(f2368_D,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f2361,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))))) )),
  inference(superposition,[],[f1530,f472])).
fof(f2360,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))))))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1530,f472])).
fof(f2359,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | s(bool,X1) = s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))))))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(superposition,[],[f1530,f472])).
fof(f2358,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),sK13))) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK12,sK15(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1530,f472])).
fof(f2265,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))))))) | $spl76),
  inference(cnf_transformation,[],[f2265_D])).
fof(f2265_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))))))) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f2268,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | $spl81),
  inference(cnf_transformation,[],[f2268_D])).
fof(f2268_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) <=> ~$spl81),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl81])])).
fof(f2260,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))))),
  inference(superposition,[],[f1529,f207])).
fof(f2266,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | s(bool,'T_0') = s(bool,X2)) ) | $spl78),
  inference(cnf_transformation,[],[f2266_D])).
fof(f2266_D,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f2259,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))))))) )),
  inference(superposition,[],[f1529,f472])).
fof(f2258,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))))))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1529,f472])).
fof(f2257,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | s(bool,X1) = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))))))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(superposition,[],[f1529,f472])).
fof(f2256,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),sK13))) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK12,sK15('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,sK12)),s(fun(sK12,bool),'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14))))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1529,f472])).
fof(f1752,plain,(
  ( ! [X2,X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X2) = s(bool,sK16(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK16(X0,X1,bool,bool))))) ) | $spl68),
  inference(cnf_transformation,[],[f1752_D])).
fof(f1752_D,plain,(
  ( ! [X2,X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X2) = s(bool,sK16(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK16(X0,X1,bool,bool))))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f1754,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | $spl71),
  inference(cnf_transformation,[],[f1754_D])).
fof(f1754_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl71),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl71])])).
fof(f1742,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,g(s(bool,sK16(X3,X4,bool,bool)))) | s(bool,X5) = s(bool,sK16(X3,X4,bool,bool)) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f723,f472])).
fof(f1741,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,g(s(bool,sK16(X0,X1,bool,bool)))) | s(bool,X2) = s(bool,sK16(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f723,f472])).
fof(f1728,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl62),
  inference(cnf_transformation,[],[f1728_D])).
fof(f1728_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f1724,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK16(X0,X1,bool,bool)) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1)) ) | $spl58),
  inference(cnf_transformation,[],[f1724_D])).
fof(f1724_D,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK16(X0,X1,bool,bool)) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1)) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f1731,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl67),
  inference(cnf_transformation,[],[f1731_D])).
fof(f1731_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl67),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl67])])).
fof(f1722,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X12) = s(fun(bool,bool),X13) | s(bool,'T_0') = s(bool,sK16(X12,X13,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f722,f207])).
fof(f1729,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl64),
  inference(cnf_transformation,[],[f1729_D])).
fof(f1729_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f1721,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X10) = s(fun(bool,bool),X11) | s(bool,'T_0') = s(bool,sK16(X10,X11,bool,bool)) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f722,f472])).
fof(f1720,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X7) = s(fun(bool,bool),X8) | s(bool,'T_0') = s(bool,sK16(X7,X8,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f722,f472])).
fof(f1723,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl56),
  inference(cnf_transformation,[],[f1723_D])).
fof(f1723_D,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f1726,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | $spl61),
  inference(cnf_transformation,[],[f1726_D])).
fof(f1726_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl61),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl61])])).
fof(f1719,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,sK16(X3,X4,bool,bool)) | s(bool,X5) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f722,f472])).
fof(f1718,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK16(X0,X1,bool,bool)) | s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f722,f472])).
fof(f1659,plain,(
  s(bool,'T_0') = s(bool,g(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) | $spl50),
  inference(cnf_transformation,[],[f1659_D])).
fof(f1659_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f1662,plain,(
  s(bool,'F_0') != s(bool,f(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) | $spl55),
  inference(cnf_transformation,[],[f1662_D])).
fof(f1662_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) <=> ~$spl55),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl55])])).
fof(f1654,plain,(
  s(bool,'F_0') != s(bool,f(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) | s(bool,'T_0') = s(bool,g(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12))))),
  inference(superposition,[],[f721,f207])).
fof(f1660,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) | s(bool,'T_0') = s(bool,X2)) ) | $spl52),
  inference(cnf_transformation,[],[f1660_D])).
fof(f1660_D,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f1653,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,f(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,g(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12))))) )),
  inference(superposition,[],[f721,f472])).
fof(f1652,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) | s(bool,'T_0') = s(bool,g(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f721,f472])).
fof(f1655,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) | s(bool,'T_0') = s(bool,X0)) ) | $spl46),
  inference(cnf_transformation,[],[f1655_D])).
fof(f1655_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f1657,plain,(
  s(bool,'T_0') != s(bool,f(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) | $spl49),
  inference(cnf_transformation,[],[f1657_D])).
fof(f1657_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) <=> ~$spl49),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl49])])).
fof(f1651,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,f(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) | s(bool,X1) = s(bool,g(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(superposition,[],[f721,f472])).
fof(f1650,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) | s(bool,X0) = s(bool,g(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f721,f472])).
fof(f1606,plain,(
  s(bool,'T_0') = s(bool,g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12)))) | $spl40),
  inference(cnf_transformation,[],[f1606_D])).
fof(f1606_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12)))) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f1609,plain,(
  s(bool,'F_0') != s(bool,f(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12)))) | $spl45),
  inference(cnf_transformation,[],[f1609_D])).
fof(f1609_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12)))) <=> ~$spl45),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl45])])).
fof(f1601,plain,(
  s(bool,'F_0') != s(bool,f(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12)))) | s(bool,'T_0') = s(bool,g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))))),
  inference(superposition,[],[f720,f207])).
fof(f1607,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12)))) | s(bool,'T_0') = s(bool,X2)) ) | $spl42),
  inference(cnf_transformation,[],[f1607_D])).
fof(f1607_D,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12)))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f1600,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,f(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12)))) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12))))) )),
  inference(superposition,[],[f720,f472])).
fof(f1599,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12)))) | s(bool,'T_0') = s(bool,g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12)))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f720,f472])).
fof(f1602,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12)))) | s(bool,'T_0') = s(bool,X0)) ) | $spl36),
  inference(cnf_transformation,[],[f1602_D])).
fof(f1602_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12)))) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f1604,plain,(
  s(bool,'T_0') != s(bool,f(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12)))) | $spl39),
  inference(cnf_transformation,[],[f1604_D])).
fof(f1604_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12)))) <=> ~$spl39),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl39])])).
fof(f1598,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,f(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12)))) | s(bool,X1) = s(bool,g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12)))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(superposition,[],[f720,f472])).
fof(f1597,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12)))) | s(bool,X0) = s(bool,g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12)))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f720,f472])).
fof(f1221,plain,(
  $false | $spl2),
  inference(subsumption_resolution,[],[f1220,f200])).
fof(f1220,plain,(
  ~p(s(bool,'T_0')) | $spl2),
  inference(forward_demodulation,[],[f1219,f1084])).
fof(f1219,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))) ) | $spl2),
  inference(subsumption_resolution,[],[f1206,f200])).
fof(f1206,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))) ) | $spl2),
  inference(backward_demodulation,[],[f1084,f246])).
fof(f1218,plain,(
  $false | $spl2),
  inference(global_subsumption,[],[f199,f200,f364,f207,f217,f216,f221,f220,f244,f243,f247,f246,f245,f250,f249,f248,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f721,f720,f722,f723,f743,f796,f797,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f896,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1084,f1212,f1217])).
fof(f1217,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | $spl2),
  inference(subsumption_resolution,[],[f1203,f1084])).
fof(f1203,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X0,sK15(X2,X1,X0)),s(fun(X0,bool),X1))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | $spl2),
  inference(backward_demodulation,[],[f1084,f244])).
fof(f1212,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl2),
  inference(subsumption_resolution,[],[f1194,f1084])).
fof(f1194,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl2),
  inference(backward_demodulation,[],[f1084,f220])).
fof(f1216,plain,(
  $false | $spl2),
  inference(subsumption_resolution,[],[f1200,f1084])).
fof(f1200,plain,(
  s(bool,'T_0') != s(bool,g(s(sK12,sK16('const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),sK13,bool,sK12)))) | $spl2),
  inference(backward_demodulation,[],[f1084,f720])).
fof(f1215,plain,(
  $false | $spl2),
  inference(subsumption_resolution,[],[f1199,f1084])).
fof(f1199,plain,(
  s(bool,'T_0') != s(bool,g(s(sK12,sK16(sK13,'const.pred_set.UNION_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'const.pred_set.DIFF_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),sK14)))),bool,sK12)))) | $spl2),
  inference(backward_demodulation,[],[f1084,f721])).
fof(f1214,plain,(
  $false | $spl2),
  inference(subsumption_resolution,[],[f1196,f200])).
fof(f1196,plain,(
  ~p(s(bool,'T_0')) | $spl2),
  inference(backward_demodulation,[],[f1084,f364])).
fof(f1213,plain,(
  $false | $spl2),
  inference(global_subsumption,[],[f199,f200,f364,f207,f217,f216,f221,f220,f244,f243,f247,f246,f245,f250,f249,f248,f254,f251,f257,f256,f365,f275,f272,f271,f270,f280,f287,f286,f285,f284,f283,f282,f281,f291,f290,f289,f288,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f357,f307,f306,f305,f304,f314,f313,f312,f311,f320,f317,f316,f315,f327,f326,f325,f356,f323,f322,f321,f331,f330,f329,f328,f335,f334,f333,f332,f344,f343,f342,f355,f340,f339,f338,f348,f347,f346,f345,f352,f351,f350,f349,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f376,f377,f378,f379,f380,f381,f382,f383,f384,f385,f386,f393,f394,f397,f400,f407,f410,f411,f414,f420,f421,f425,f428,f431,f438,f439,f441,f448,f472,f518,f527,f526,f537,f535,f549,f559,f564,f569,f573,f576,f582,f588,f612,f616,f622,f648,f652,f658,f679,f701,f721,f720,f722,f723,f743,f796,f797,f833,f1082,f1086,f1090,f1093,f1096,f1099,f1102,f1105,f1108,f1111,f1114,f1117,f1120,f1123,f1126,f1129,f1132,f1134,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f895,f896,f1160,f1162,f1165,f1167,f1169,f1171,f1173,f1175,f1177,f1179,f1181,f1183,f1185,f1187,f1189,f1191,f1193,f1084,f1212])).
fof(f1133,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~sP8('T_0',X1,X2)) ) | $spl34),
  inference(cnf_transformation,[],[f1133_D])).
fof(f1133_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~sP8('T_0',X1,X2)) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f1193,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8('T_0',X1,X2) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f1192])).
fof(f1192,plain,(
  ( ! [X94,X92,X93,X91] : (~sP8('T_0',X92,X93) | sP9('T_0',X92,X93) | s(bool,X91) = s(bool,X94) | s(bool,'T_0') = s(bool,X94)) )),
  inference(subsumption_resolution,[],[f924,f518])).
fof(f924,plain,(
  ( ! [X94,X92,X93,X91] : (~p(s(bool,X91)) | ~sP8('T_0',X92,X93) | sP9('T_0',X92,X93) | s(bool,X91) = s(bool,X94) | s(bool,'T_0') = s(bool,X94)) )),
  inference(superposition,[],[f356,f472])).
fof(f1130,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~sP10('T_0',X1,X2)) ) | $spl32),
  inference(cnf_transformation,[],[f1130_D])).
fof(f1130_D,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~sP10('T_0',X1,X2)) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f1191,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP10('T_0',X1,X2) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f1190])).
fof(f1190,plain,(
  ( ! [X90,X88,X87,X89] : (~sP10('T_0',X88,X89) | sP11('T_0',X88,X89) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(subsumption_resolution,[],[f923,f518])).
fof(f923,plain,(
  ( ! [X90,X88,X87,X89] : (~p(s(bool,X87)) | ~sP10('T_0',X88,X89) | sP11('T_0',X88,X89) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(superposition,[],[f355,f472])).
fof(f1127,plain,(
  ( ! [X2,X1] : (~sP11('T_0',X2,X1) | p(s(bool,X2)) | ~p(s(bool,X1))) ) | $spl30),
  inference(cnf_transformation,[],[f1127_D])).
fof(f1127_D,plain,(
  ( ! [X2,X1] : (~sP11('T_0',X2,X1) | p(s(bool,X2)) | ~p(s(bool,X1))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f1189,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f1188])).
fof(f1188,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,X84)) | p(s(bool,X85)) | ~sP11('T_0',X85,X84) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(subsumption_resolution,[],[f922,f518])).
fof(f922,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,X83)) | ~p(s(bool,X84)) | p(s(bool,X85)) | ~sP11('T_0',X85,X84) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(superposition,[],[f352,f472])).
fof(f1124,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~p(s(bool,X1))) ) | $spl28),
  inference(cnf_transformation,[],[f1124_D])).
fof(f1124_D,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~p(s(bool,X1))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f1187,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f1186])).
fof(f1186,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,X80)) | sP11('T_0',X80,X81) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(subsumption_resolution,[],[f921,f518])).
fof(f921,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,X79)) | ~p(s(bool,X80)) | sP11('T_0',X80,X81) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(superposition,[],[f351,f472])).
fof(f1121,plain,(
  ( ! [X2,X1] : (sP11('T_0',X2,X1) | p(s(bool,X1))) ) | $spl26),
  inference(cnf_transformation,[],[f1121_D])).
fof(f1121_D,plain,(
  ( ! [X2,X1] : (sP11('T_0',X2,X1) | p(s(bool,X1))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f1185,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f1184])).
fof(f1184,plain,(
  ( ! [X78,X76,X77,X75] : (p(s(bool,X76)) | sP11('T_0',X77,X76) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(subsumption_resolution,[],[f920,f518])).
fof(f920,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,X75)) | p(s(bool,X76)) | sP11('T_0',X77,X76) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f350,f472])).
fof(f1118,plain,(
  ( ! [X2,X1] : (~sP10(X2,X1,'T_0') | ~p(s(bool,X2)) | p(s(bool,X1))) ) | $spl24),
  inference(cnf_transformation,[],[f1118_D])).
fof(f1118_D,plain,(
  ( ! [X2,X1] : (~sP10(X2,X1,'T_0') | ~p(s(bool,X2)) | p(s(bool,X1))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f1183,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP10(X2,X1,'T_0')) )),
  inference(condensation,[],[f1182])).
fof(f1182,plain,(
  ( ! [X74,X72,X71,X73] : (p(s(bool,X72)) | ~p(s(bool,X73)) | ~sP10(X73,X72,'T_0') | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(subsumption_resolution,[],[f919,f518])).
fof(f919,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,X71)) | p(s(bool,X72)) | ~p(s(bool,X73)) | ~sP10(X73,X72,'T_0') | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(superposition,[],[f345,f472])).
fof(f1115,plain,(
  ( ! [X2,X1] : (~sP9('T_0',X2,X1) | p(s(bool,X2)) | p(s(bool,X1))) ) | $spl22),
  inference(cnf_transformation,[],[f1115_D])).
fof(f1115_D,plain,(
  ( ! [X2,X1] : (~sP9('T_0',X2,X1) | p(s(bool,X2)) | p(s(bool,X1))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f1181,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f1180])).
fof(f1180,plain,(
  ( ! [X70,X68,X69,X67] : (p(s(bool,X68)) | p(s(bool,X69)) | ~sP9('T_0',X69,X68) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(subsumption_resolution,[],[f918,f518])).
fof(f918,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,X67)) | p(s(bool,X68)) | p(s(bool,X69)) | ~sP9('T_0',X69,X68) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f335,f472])).
fof(f1112,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~p(s(bool,X1))) ) | $spl20),
  inference(cnf_transformation,[],[f1112_D])).
fof(f1112_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~p(s(bool,X1))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f1179,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f1178])).
fof(f1178,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,X64)) | sP9('T_0',X64,X65) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(subsumption_resolution,[],[f917,f518])).
fof(f917,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,X63)) | ~p(s(bool,X64)) | sP9('T_0',X64,X65) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f334,f472])).
fof(f1109,plain,(
  ( ! [X2,X1] : (sP9('T_0',X2,X1) | ~p(s(bool,X1))) ) | $spl18),
  inference(cnf_transformation,[],[f1109_D])).
fof(f1109_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X2,X1) | ~p(s(bool,X1))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f1177,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f1176])).
fof(f1176,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X60)) | sP9('T_0',X61,X60) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(subsumption_resolution,[],[f916,f518])).
fof(f916,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X59)) | ~p(s(bool,X60)) | sP9('T_0',X61,X60) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f333,f472])).
fof(f1106,plain,(
  ( ! [X2,X1] : (~sP8('T_0',X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) ) | $spl16),
  inference(cnf_transformation,[],[f1106_D])).
fof(f1106_D,plain,(
  ( ! [X2,X1] : (~sP8('T_0',X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f1175,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP8('T_0',X1,X2)) )),
  inference(condensation,[],[f1174])).
fof(f1174,plain,(
  ( ! [X57,X58,X56,X55] : (p(s(bool,X56)) | p(s(bool,X57)) | ~sP8('T_0',X56,X57) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(subsumption_resolution,[],[f915,f518])).
fof(f915,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,X55)) | p(s(bool,X56)) | p(s(bool,X57)) | ~sP8('T_0',X56,X57) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(superposition,[],[f328,f472])).
fof(f1103,plain,(
  ( ! [X2,X1] : (~sP7(X2,X1,'T_0') | p(s(bool,X2)) | ~p(s(bool,X1))) ) | $spl14),
  inference(cnf_transformation,[],[f1103_D])).
fof(f1103_D,plain,(
  ( ! [X2,X1] : (~sP7(X2,X1,'T_0') | p(s(bool,X2)) | ~p(s(bool,X1))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f1173,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,X1,'T_0')) )),
  inference(condensation,[],[f1172])).
fof(f1172,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X52)) | p(s(bool,X53)) | ~sP7(X53,X52,'T_0') | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(subsumption_resolution,[],[f914,f518])).
fof(f914,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X51)) | ~p(s(bool,X52)) | p(s(bool,X53)) | ~sP7(X53,X52,'T_0') | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(superposition,[],[f320,f472])).
fof(f1100,plain,(
  ( ! [X2,X1] : (~sP6(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl12),
  inference(cnf_transformation,[],[f1100_D])).
fof(f1100_D,plain,(
  ( ! [X2,X1] : (~sP6(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f1171,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X1,'T_0',X2)) )),
  inference(condensation,[],[f1170])).
fof(f1170,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X48)) | p(s(bool,X49)) | ~sP6(X48,'T_0',X49) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(subsumption_resolution,[],[f913,f518])).
fof(f913,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X47)) | ~p(s(bool,X48)) | p(s(bool,X49)) | ~sP6(X48,'T_0',X49) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f311,f472])).
fof(f1097,plain,(
  ( ! [X2,X1] : (~sP3(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl10),
  inference(cnf_transformation,[],[f1097_D])).
fof(f1097_D,plain,(
  ( ! [X2,X1] : (~sP3(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f1169,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP3(X1,'T_0',X2)) )),
  inference(condensation,[],[f1168])).
fof(f1168,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X44)) | p(s(bool,X45)) | ~sP3(X44,'T_0',X45) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f912,f518])).
fof(f912,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X43)) | ~p(s(bool,X44)) | p(s(bool,X45)) | ~sP3(X44,'T_0',X45) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f292,f472])).
fof(f1094,plain,(
  ( ! [X2,X1] : (~sP4(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl8),
  inference(cnf_transformation,[],[f1094_D])).
fof(f1094_D,plain,(
  ( ! [X2,X1] : (~sP4(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f1167,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X1,'T_0',X2)) )),
  inference(condensation,[],[f1166])).
fof(f1166,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X40)) | p(s(bool,X41)) | ~sP4(X40,'T_0',X41) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(subsumption_resolution,[],[f911,f518])).
fof(f911,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X39)) | ~p(s(bool,X40)) | p(s(bool,X41)) | ~sP4(X40,'T_0',X41) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f288,f472])).
fof(f1091,plain,(
  ( ! [X2,X1] : (~sP4(X1,'T_0',X2) | ~sP3(X2,'T_0',X1) | sP5(X1,X2,'T_0')) ) | $spl6),
  inference(cnf_transformation,[],[f1091_D])).
fof(f1091_D,plain,(
  ( ! [X2,X1] : (~sP4(X1,'T_0',X2) | ~sP3(X2,'T_0',X1) | sP5(X1,X2,'T_0')) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f1165,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP5(X1,X2,'T_0') | ~sP4(X1,'T_0',X2) | ~sP3(X2,'T_0',X1)) )),
  inference(subsumption_resolution,[],[f1164,f376])).
fof(f1164,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP5(X1,X2,'T_0') | ~sP4(X1,'T_0',X2) | ~sP3(X2,'T_0',X1) | ~sP2('T_0',X2,X1)) )),
  inference(condensation,[],[f1163])).
fof(f1163,plain,(
  ( ! [X37,X35,X38,X36] : (sP5(X36,X37,'T_0') | ~sP4(X36,'T_0',X37) | ~sP3(X37,'T_0',X36) | ~sP2('T_0',X37,X36) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(subsumption_resolution,[],[f910,f518])).
fof(f910,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X35)) | sP5(X36,X37,'T_0') | ~sP4(X36,'T_0',X37) | ~sP3(X37,'T_0',X36) | ~sP2('T_0',X37,X36) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f285,f472])).
fof(f1087,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1(X1,'T_0')) ) | $spl4),
  inference(cnf_transformation,[],[f1087_D])).
fof(f1087_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1(X1,'T_0')) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f1162,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f1161])).
fof(f1161,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,X33)) | sP1(X33,'T_0') | s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f909,f518])).
fof(f909,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,X32)) | p(s(bool,X33)) | sP1(X33,'T_0') | s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f275,f472])).
fof(f1083,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1('T_0',X1)) ) | $spl0),
  inference(cnf_transformation,[],[f1083_D])).
fof(f1083_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1('T_0',X1)) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f1160,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f1159])).
fof(f1159,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X30)) | sP1('T_0',X30) | s(bool,X29) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) )),
  inference(subsumption_resolution,[],[f908,f518])).
fof(f908,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,X29)) | p(s(bool,X30)) | sP1('T_0',X30) | s(bool,X29) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) )),
  inference(superposition,[],[f272,f472])).
fof(f1132,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8('T_0',X1,X2) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f1131])).
fof(f1131,plain,(
  ( ! [X94,X92,X93,X91] : (~sP8('T_0',X92,X93) | sP9('T_0',X92,X93) | s(bool,X91) = s(bool,X94) | s(bool,'T_0') = s(bool,X94)) )),
  inference(subsumption_resolution,[],[f860,f518])).
fof(f860,plain,(
  ( ! [X94,X92,X93,X91] : (~p(s(bool,X91)) | ~sP8('T_0',X92,X93) | sP9('T_0',X92,X93) | s(bool,X91) = s(bool,X94) | s(bool,'T_0') = s(bool,X94)) )),
  inference(superposition,[],[f356,f472])).
fof(f1129,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP10('T_0',X1,X2) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f1128])).
fof(f1128,plain,(
  ( ! [X90,X88,X87,X89] : (~sP10('T_0',X88,X89) | sP11('T_0',X88,X89) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(subsumption_resolution,[],[f859,f518])).
fof(f859,plain,(
  ( ! [X90,X88,X87,X89] : (~p(s(bool,X87)) | ~sP10('T_0',X88,X89) | sP11('T_0',X88,X89) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(superposition,[],[f355,f472])).
fof(f1126,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f1125])).
fof(f1125,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,X84)) | p(s(bool,X85)) | ~sP11('T_0',X85,X84) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(subsumption_resolution,[],[f858,f518])).
fof(f858,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,X83)) | ~p(s(bool,X84)) | p(s(bool,X85)) | ~sP11('T_0',X85,X84) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(superposition,[],[f352,f472])).
fof(f1123,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f1122])).
fof(f1122,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,X80)) | sP11('T_0',X80,X81) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(subsumption_resolution,[],[f857,f518])).
fof(f857,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,X79)) | ~p(s(bool,X80)) | sP11('T_0',X80,X81) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(superposition,[],[f351,f472])).
fof(f1120,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f1119])).
fof(f1119,plain,(
  ( ! [X78,X76,X77,X75] : (p(s(bool,X76)) | sP11('T_0',X77,X76) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(subsumption_resolution,[],[f856,f518])).
fof(f856,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,X75)) | p(s(bool,X76)) | sP11('T_0',X77,X76) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f350,f472])).
fof(f1117,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP10(X2,X1,'T_0')) )),
  inference(condensation,[],[f1116])).
fof(f1116,plain,(
  ( ! [X74,X72,X71,X73] : (p(s(bool,X72)) | ~p(s(bool,X73)) | ~sP10(X73,X72,'T_0') | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(subsumption_resolution,[],[f855,f518])).
fof(f855,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,X71)) | p(s(bool,X72)) | ~p(s(bool,X73)) | ~sP10(X73,X72,'T_0') | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(superposition,[],[f345,f472])).
fof(f1114,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f1113])).
fof(f1113,plain,(
  ( ! [X70,X68,X69,X67] : (p(s(bool,X68)) | p(s(bool,X69)) | ~sP9('T_0',X69,X68) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(subsumption_resolution,[],[f854,f518])).
fof(f854,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,X67)) | p(s(bool,X68)) | p(s(bool,X69)) | ~sP9('T_0',X69,X68) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f335,f472])).
fof(f1111,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f1110])).
fof(f1110,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,X64)) | sP9('T_0',X64,X65) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(subsumption_resolution,[],[f853,f518])).
fof(f853,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,X63)) | ~p(s(bool,X64)) | sP9('T_0',X64,X65) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f334,f472])).
fof(f1108,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f1107])).
fof(f1107,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X60)) | sP9('T_0',X61,X60) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(subsumption_resolution,[],[f852,f518])).
fof(f852,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X59)) | ~p(s(bool,X60)) | sP9('T_0',X61,X60) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f333,f472])).
fof(f1105,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP8('T_0',X1,X2)) )),
  inference(condensation,[],[f1104])).
fof(f1104,plain,(
  ( ! [X57,X58,X56,X55] : (p(s(bool,X56)) | p(s(bool,X57)) | ~sP8('T_0',X56,X57) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(subsumption_resolution,[],[f851,f518])).
fof(f851,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,X55)) | p(s(bool,X56)) | p(s(bool,X57)) | ~sP8('T_0',X56,X57) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(superposition,[],[f328,f472])).
fof(f1102,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,X1,'T_0')) )),
  inference(condensation,[],[f1101])).
fof(f1101,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X52)) | p(s(bool,X53)) | ~sP7(X53,X52,'T_0') | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(subsumption_resolution,[],[f850,f518])).
fof(f850,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X51)) | ~p(s(bool,X52)) | p(s(bool,X53)) | ~sP7(X53,X52,'T_0') | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(superposition,[],[f320,f472])).
fof(f1099,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X1,'T_0',X2)) )),
  inference(condensation,[],[f1098])).
fof(f1098,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X48)) | p(s(bool,X49)) | ~sP6(X48,'T_0',X49) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(subsumption_resolution,[],[f849,f518])).
fof(f849,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X47)) | ~p(s(bool,X48)) | p(s(bool,X49)) | ~sP6(X48,'T_0',X49) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f311,f472])).
fof(f1096,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP3(X1,'T_0',X2)) )),
  inference(condensation,[],[f1095])).
fof(f1095,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X44)) | p(s(bool,X45)) | ~sP3(X44,'T_0',X45) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f848,f518])).
fof(f848,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X43)) | ~p(s(bool,X44)) | p(s(bool,X45)) | ~sP3(X44,'T_0',X45) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f292,f472])).
fof(f1093,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X1,'T_0',X2)) )),
  inference(condensation,[],[f1092])).
fof(f1092,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X40)) | p(s(bool,X41)) | ~sP4(X40,'T_0',X41) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(subsumption_resolution,[],[f847,f518])).
fof(f847,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X39)) | ~p(s(bool,X40)) | p(s(bool,X41)) | ~sP4(X40,'T_0',X41) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f288,f472])).
fof(f1090,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP5(X1,X2,'T_0') | ~sP4(X1,'T_0',X2) | ~sP3(X2,'T_0',X1)) )),
  inference(subsumption_resolution,[],[f1089,f376])).
fof(f1089,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP5(X1,X2,'T_0') | ~sP4(X1,'T_0',X2) | ~sP3(X2,'T_0',X1) | ~sP2('T_0',X2,X1)) )),
  inference(condensation,[],[f1088])).
fof(f1088,plain,(
  ( ! [X37,X35,X38,X36] : (sP5(X36,X37,'T_0') | ~sP4(X36,'T_0',X37) | ~sP3(X37,'T_0',X36) | ~sP2('T_0',X37,X36) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(subsumption_resolution,[],[f846,f518])).
fof(f846,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X35)) | sP5(X36,X37,'T_0') | ~sP4(X36,'T_0',X37) | ~sP3(X37,'T_0',X36) | ~sP2('T_0',X37,X36) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f285,f472])).
fof(f1086,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f1085])).
fof(f1085,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,X33)) | sP1(X33,'T_0') | s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f845,f518])).
fof(f845,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,X32)) | p(s(bool,X33)) | sP1(X33,'T_0') | s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f275,f472])).
fof(f1082,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f1081])).
fof(f1081,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X30)) | sP1('T_0',X30) | s(bool,X29) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) )),
  inference(subsumption_resolution,[],[f844,f518])).
fof(f844,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,X29)) | p(s(bool,X30)) | sP1('T_0',X30) | s(bool,X29) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) )),
  inference(superposition,[],[f272,f472])).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 758
% Passive clauses: 1651
% Generated clauses: 16024
% Final active clauses: 486
% Final passive clauses: 318
% Input formulas: 34
% Initial clauses: 156
% 
% Duplicate literals: 661
% Trivial inequalities: 88
% Fw subsumption resolutions: 674
% Fw demodulations: 1875
% Bw demodulations: 506
% Condensations: 390
% Global subsumptions: 115
% 
% Simple tautologies: 100
% Equational tautologies: 367
% Forward subsumptions: 9790
% Bw demodulations to eq. taut.: 5
% 
% Unit resulting resolution: 5841
% Factoring: 1
% Forward superposition: 2998
% Backward superposition: 1926
% Self superposition: 34
% Equality factoring: 38
% Equality resolution: 19
% 
% Split clauses: 75
% Split components: 155
% SAT solver clauses: 2156
% SAT solver unit clauses: 111
% SAT solver binary clauses: 741
% 
% Sat splits: 75
% Sat splitting refutations: 238
% 
% Memory used [KB]: 6780
% Time elapsed: 0.593 s
% ------------------------------
% ------------------------------
% Success in time 1.911 s
