Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.extreal
% (48827)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1274
% Passive clauses: 7808
% Generated clauses: 38735
% Final active clauses: 1247
% Final passive clauses: 6250
% Input formulas: 261
% Initial clauses: 517
% 
% Duplicate literals: 198
% Trivial inequalities: 77
% Fw subsumption resolutions: 2206
% Fw demodulations: 3232
% Bw demodulations: 26
% 
% Simple tautologies: 294
% Equational tautologies: 34
% Forward subsumptions: 24652
% 
% Binary resolution: 7039
% Factoring: 490
% Forward superposition: 11992
% Backward superposition: 12728
% Self superposition: 48
% Equality factoring: 2
% Equality resolution: 22
% 
% Split clauses: 351
% Split components: 704
% SAT solver clauses: 522
% SAT solver unit clauses: 25
% SAT solver binary clauses: 95
% 
% Sat splits: 350
% 
% TWLsolver clauses: 522
% TWLsolver calls for satisfiability: 413
% 
% Memory used [KB]: 19317
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.extreal
% (48852)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 743
% Passive clauses: 3347
% Generated clauses: 26857
% Final active clauses: 722
% Final passive clauses: 2559
% Input formulas: 261
% Initial clauses: 517
% 
% Duplicate literals: 102
% Trivial inequalities: 53
% Fw subsumption resolutions: 797
% Fw demodulations: 3419
% Bw demodulations: 66
% Condensations: 71
% Global subsumptions: 36
% 
% Simple tautologies: 278
% Equational tautologies: 18
% Forward subsumptions: 14920
% Fw demodulations to eq. taut.: 11
% 
% Binary resolution: 2714
% Factoring: 162
% Forward superposition: 8678
% Backward superposition: 6772
% Self superposition: 59
% Equality factoring: 2
% Equality resolution: 23
% 
% Split clauses: 18
% Split components: 36
% SAT solver clauses: 4405
% SAT solver unit clauses: 711
% SAT solver binary clauses: 157
% SAT solver learnt clauses: 72
% SAT solver learnt literals: 12
% 
% Sat splits: 18
% 
% TWLsolver clauses: 4261
% TWLsolver calls for satisfiability: 4973
% 
% Memory used [KB]: 24178
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.extreal
% (48874)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 926
% Passive clauses: 7719
% Generated clauses: 115196
% Final active clauses: 889
% Final passive clauses: 3159
% Input formulas: 261
% Initial clauses: 517
% Discarded non-redundant clauses: 28378
% 
% Duplicate literals: 10
% Trivial inequalities: 47
% Fw subsumption resolutions: 1601
% Fw demodulations: 39838
% Bw demodulations: 95
% Condensations: 33
% Global subsumptions: 1894
% 
% Simple tautologies: 106
% Equational tautologies: 75
% Forward subsumptions: 36969
% Fw demodulations to eq. taut.: 557
% Bw demodulations to eq. taut.: 14
% 
% Unit resulting resolution: 54536
% Factoring: 1
% Forward superposition: 9967
% Backward superposition: 6492
% Self superposition: 79
% Equality factoring: 2
% Equality resolution: 22
% 
% Split clauses: 37
% Split components: 74
% SAT solver clauses: 13442
% SAT solver unit clauses: 6779
% SAT solver binary clauses: 5453
% 
% Sat splits: 37
% Sat splitting refutations: 3753
% 
% Memory used [KB]: 352446
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.extreal
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.extreal
% SZS output start Proof for thm.extreal
fof(f6766,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f832,f833,f834,f835,f836,f837,f839,f931,f975,f980,f983,f986,f987,f977,f991,f1112,f1108,f1213,f1214,f1215,f1216,f1217,f1218,f1258,f1275,f1274,f1273,f982,f1323,f1335,f1337,f1338,f1339,f1340,f1341,f965,f1342,f1344,f1436,f963,f1482,f1484,f1489,f1490,f1491,f1492,f1493,f1495,f1617,f1496,f1585,f1616,f1620,f1623,f1626,f1627,f1628,f1629,f1625,f1649,f1651,f1653,f1654,f1655,f1650,f1656,f1652,f1673,f1675,f1677,f1678,f1679,f1674,f1680,f1676,f1699,f1701,f1703,f1704,f1705,f1700,f1706,f1702,f1725,f1727,f1729,f1730,f1731,f1726,f1732,f1728,f1752,f1754,f1756,f1757,f1758,f1753,f1759,f1755,f1778,f1780,f1782,f1783,f1784,f1779,f1785,f1781,f1804,f1806,f1808,f1809,f1810,f1805,f1811,f1807,f1830,f1832,f1834,f1835,f1836,f1831,f1837,f1833,f1857,f1859,f1861,f1862,f1863,f1858,f1864,f1860,f1884,f1886,f1888,f1889,f1890,f1885,f1891,f1887,f1911,f1913,f1915,f1916,f1917,f1912,f1918,f1914,f2106,f2105,f2104,f2322,f2321,f2340,f2344,f2362,f1622,f2363,f2365,f2367,f2368,f2369,f2364,f2370,f2366,f2371,f2373,f2375,f2376,f2377,f2372,f2378,f2374,f2379,f2381,f2383,f2384,f2385,f2380,f2386,f2382,f2387,f2389,f2391,f2392,f2393,f2388,f2394,f2390,f2395,f2397,f2399,f2400,f2401,f2396,f2402,f2398,f2403,f2405,f2407,f2408,f2409,f2404,f2410,f2406,f2411,f2413,f2415,f2416,f2417,f2412,f2418,f2414,f2419,f2421,f2423,f2424,f2425,f2420,f2426,f2422,f2427,f2429,f990,f2431,f2432,f985,f2433,f2428,f2434,f2430,f979,f2515,f2521,f2532,f2536,f2540,f2517,f2571,f2703,f2710,f2711,f2712,f2713,f2714,f2845,f2715,f2717,f2809,f2844,f2934,f2940,f2951,f2955,f2959,f2936,f3158,f3164,f3175,f3179,f3183,f988,f3160,f3659,f3660,f3677,f3678,f3912,f3913,f4451,f4452,f4456,f4468,f4469,f4473,f2325,f4899,f4900,f4917,f4918,f5058,f5059,f5060,f5061,f5062,f5064,f5190,f5065,f5154,f5189,f6003,f6004,f6021,f6022,f2109,f6162,f6163,f6164,f6165,f964,f6166,f968,f6168,f966,f6291,f6169,f967,f6258,f969,f6290,f6765])).
fof(f6765,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2449,f928])).
fof(f928,plain,(
  s(bool,'T_0') != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0')))),
  inference(resolution,[],[f565,f763])).
fof(f763,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f542])).
fof(f542,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f321])).
fof(f321,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X22] : (s(bool,'T_0') = s(bool,X22) <=> p(s(bool,X22)))),
  file('thm.extreal.le_infty|split|1','thm.bool.EQ_CLAUSES|split|1')).
fof(f565,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0'))))),
  inference(cnf_transformation,[],[f433])).
fof(f433,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0'))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12])],[f372])).
fof(f372,plain,(
  ? [X0] : ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.PosInf_0'))))),
  inference(ennf_transformation,[],[f263])).
fof(f263,plain,(
  ~! [X0] : p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.PosInf_0'))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal','const.extreal.PosInf_0'))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal','const.extreal.PosInf_0'))))),
  file('thm.extreal.le_infty|split|1','thm.extreal.le_infty|split|1')).
fof(f2449,plain,(
  s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0'))) | $spl16),
  inference(superposition,[],[f794,f979])).
fof(f794,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0')))) )),
  inference(cnf_transformation,[],[f347])).
fof(f347,plain,(
  ! [X0] : s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0')))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X19] : s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X19))),s('type.extreal.extreal','const.extreal.PosInf_0')))),
  file('thm.extreal.le_infty|split|1','thm.extreal.extreal_le_def|split|3')).
fof(f6290,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl134),
  inference(cnf_transformation,[],[f6290_D])).
fof(f6290_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f969,plain,(
  ( ! [X15] : (p(s(bool,X15))) ) | $spl12),
  inference(cnf_transformation,[],[f969_D])).
fof(f969_D,plain,(
  ( ! [X15] : (p(s(bool,X15))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f6258,plain,(
  ( ! [X163,X162] : (p(s(bool,X162)) | p(s(bool,X163)) | s(bool,X163) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl16),
  inference(resolution,[],[f5995,f823])).
fof(f823,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f570])).
fof(f570,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f435])).
fof(f435,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f434])).
fof(f434,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f374])).
fof(f374,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f373])).
fof(f373,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f264])).
fof(f264,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f182])).
fof(f182,axiom,(
  ! [X6,X64,X12,X65] : ((s(bool,X6) = s(bool,X64) & (p(s(bool,X64)) => s(bool,X12) = s(bool,X65))) => ((p(s(bool,X6)) => p(s(bool,X12))) <=> (p(s(bool,X64)) => p(s(bool,X65)))))),
  file('thm.extreal.le_infty|split|1','thm.bool.IMP_CONG')).
fof(f5995,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl16),
  inference(subsumption_resolution,[],[f5725,f780])).
fof(f780,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f412])).
fof(f412,plain,(
  ! [X0] : (~p(s(bool,'F_0')) | p(s(bool,X0)))),
  inference(ennf_transformation,[],[f336])).
fof(f336,plain,(
  ! [X0] : (p(s(bool,'F_0')) => p(s(bool,X0)))),
  inference(rectify,[],[f54])).
fof(f54,axiom,(
  ! [X22] : (p(s(bool,'F_0')) => p(s(bool,X22)))),
  file('thm.extreal.le_infty|split|1','thm.bool.FALSITY')).
fof(f5725,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl16),
  inference(superposition,[],[f790,f2448])).
fof(f2448,plain,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12))) | $spl16),
  inference(superposition,[],[f793,f979])).
fof(f793,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))))) )),
  inference(cnf_transformation,[],[f346])).
fof(f346,plain,(
  ! [X0] : s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X29] : s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X29)))))),
  file('thm.extreal.le_infty|split|1','thm.extreal.extreal_le_def|split|6')).
fof(f790,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0))))) )),
  inference(cnf_transformation,[],[f556])).
fof(f556,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0))))) & (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))))),
  inference(nnf_transformation,[],[f343])).
fof(f343,plain,(
  ! [X0,X1] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) <=> ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))))),
  inference(flattening,[],[f342])).
fof(f342,plain,(
  ! [X0,X1] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) <=> ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X6,X12] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X12)))) <=> ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X12),s('type.extreal.extreal',X6)))))),
  file('thm.extreal.le_infty|split|1','thm.extreal.extreal_lt_def')).
fof(f967,plain,(
  ( ! [X6,X8] : (s(bool,X6) != s(bool,X8) | ~p(s(bool,X6))) ) | $spl8),
  inference(cnf_transformation,[],[f967_D])).
fof(f967_D,plain,(
  ( ! [X6,X8] : (s(bool,X6) != s(bool,X8) | ~p(s(bool,X6))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f6169,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s(bool,X22) != s(bool,X23)) ) | $spl16),
  inference(resolution,[],[f5995,f573])).
fof(f573,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f435])).
fof(f6291,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl136),
  inference(cnf_transformation,[],[f6291_D])).
fof(f6291_D,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f966,plain,(
  ( ! [X4,X3] : (~p(s(bool,X3)) | s(bool,X3) != s(bool,X4) | p(s(bool,X4))) ) | $spl6),
  inference(cnf_transformation,[],[f966_D])).
fof(f966_D,plain,(
  ( ! [X4,X3] : (~p(s(bool,X3)) | s(bool,X3) != s(bool,X4) | p(s(bool,X4))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f6168,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,X19) | s(bool,X20) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl16),
  inference(resolution,[],[f5995,f573])).
fof(f968,plain,(
  ( ! [X17] : (~p(s(bool,X17))) ) | $spl10),
  inference(cnf_transformation,[],[f968_D])).
fof(f968_D,plain,(
  ( ! [X17] : (~p(s(bool,X17))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f6166,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl16),
  inference(resolution,[],[f5995,f572])).
fof(f572,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f435])).
fof(f964,plain,(
  ( ! [X2,X0] : (s(bool,X0) != s(bool,X2) | p(s(bool,X0))) ) | $spl2),
  inference(cnf_transformation,[],[f964_D])).
fof(f964_D,plain,(
  ( ! [X2,X0] : (s(bool,X0) != s(bool,X2) | p(s(bool,X0))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f6165,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl16),
  inference(resolution,[],[f5995,f571])).
fof(f571,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f435])).
fof(f6164,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s(bool,X7) != s(bool,X8)) ) | $spl16),
  inference(resolution,[],[f5995,f569])).
fof(f569,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f435])).
fof(f6163,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X4) | s(bool,X5) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl16),
  inference(resolution,[],[f5995,f569])).
fof(f6162,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl16),
  inference(resolution,[],[f5995,f567])).
fof(f567,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f435])).
fof(f2109,plain,(
  ( ! [X29] : (s(bool,'F_0') != s(bool,X29)) ) | $spl78),
  inference(cnf_transformation,[],[f2109_D])).
fof(f2109_D,plain,(
  ( ! [X29] : (s(bool,'F_0') != s(bool,X29)) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f6022,plain,(
  ( ! [X85,X86,X84] : (s(bool,'F_0') != s(bool,X84) | p(s(bool,X85)) | ~p(s(bool,X86)) | s(bool,X85) != s(bool,X86)) ) | $spl16),
  inference(subsumption_resolution,[],[f5754,f571])).
fof(f5754,plain,(
  ( ! [X85,X86,X84] : (s(bool,'F_0') != s(bool,X84) | p(s(bool,X85)) | ~p(s(bool,X84)) | ~p(s(bool,X86)) | s(bool,X85) != s(bool,X86)) ) | $spl16),
  inference(superposition,[],[f573,f2448])).
fof(f6021,plain,(
  ( ! [X83,X81,X82] : (s(bool,'F_0') != s(bool,X81) | p(s(bool,X82)) | ~p(s(bool,X83)) | s(bool,X82) != s(bool,X83)) ) | $spl16),
  inference(subsumption_resolution,[],[f5753,f823])).
fof(f5753,plain,(
  ( ! [X83,X81,X82] : (s(bool,'F_0') != s(bool,X81) | p(s(bool,X82)) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)))) | ~p(s(bool,X83)) | s(bool,X82) != s(bool,X83)) ) | $spl16),
  inference(superposition,[],[f573,f2448])).
fof(f6004,plain,(
  ( ! [X26,X24,X25] : (s(bool,'F_0') != s(bool,X24) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X26)) ) | $spl16),
  inference(subsumption_resolution,[],[f5734,f823])).
fof(f5734,plain,(
  ( ! [X26,X24,X25] : (s(bool,'F_0') != s(bool,X24) | p(s(bool,X25)) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)))) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X26)) ) | $spl16),
  inference(superposition,[],[f569,f2448])).
fof(f6003,plain,(
  ( ! [X23,X21,X22] : (s(bool,'F_0') != s(bool,X21) | p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X23)) ) | $spl16),
  inference(subsumption_resolution,[],[f5733,f571])).
fof(f5733,plain,(
  ( ! [X23,X21,X22] : (s(bool,'F_0') != s(bool,X21) | p(s(bool,X22)) | ~p(s(bool,X21)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X23)) ) | $spl16),
  inference(superposition,[],[f569,f2448])).
fof(f5189,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK12)))) ) | $spl130),
  inference(cnf_transformation,[],[f5189_D])).
fof(f5189_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK12)))) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f5154,plain,(
  ( ! [X163,X162] : (p(s(bool,X162)) | p(s(bool,X163)) | s(bool,X163) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK12)))) ) | $spl16),
  inference(resolution,[],[f4891,f823])).
fof(f4891,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK12)))) | $spl16),
  inference(subsumption_resolution,[],[f4621,f780])).
fof(f4621,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK12)))) | $spl16),
  inference(superposition,[],[f790,f2447])).
fof(f2447,plain,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.NegInf_0'))) | $spl16),
  inference(superposition,[],[f792,f979])).
fof(f792,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.NegInf_0')))) )),
  inference(cnf_transformation,[],[f345])).
fof(f345,plain,(
  ! [X0] : s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.NegInf_0')))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X37] : s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X37))),s('type.extreal.extreal','const.extreal.NegInf_0')))),
  file('thm.extreal.le_infty|split|1','thm.extreal.extreal_le_def|split|5')).
fof(f5065,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK12))) | s(bool,X22) != s(bool,X23)) ) | $spl16),
  inference(resolution,[],[f4891,f573])).
fof(f5190,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK12)))) ) | $spl132),
  inference(cnf_transformation,[],[f5190_D])).
fof(f5190_D,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK12)))) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f5064,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,X19) | s(bool,X20) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK12)))) ) | $spl16),
  inference(resolution,[],[f4891,f573])).
fof(f5062,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK12)))) ) | $spl16),
  inference(resolution,[],[f4891,f572])).
fof(f5061,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK12)))) ) | $spl16),
  inference(resolution,[],[f4891,f571])).
fof(f5060,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK12))) | s(bool,X7) != s(bool,X8)) ) | $spl16),
  inference(resolution,[],[f4891,f569])).
fof(f5059,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X4) | s(bool,X5) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK12)))) ) | $spl16),
  inference(resolution,[],[f4891,f569])).
fof(f5058,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK12)))) ) | $spl16),
  inference(resolution,[],[f4891,f567])).
fof(f4918,plain,(
  ( ! [X85,X86,X84] : (s(bool,'F_0') != s(bool,X84) | p(s(bool,X85)) | ~p(s(bool,X86)) | s(bool,X85) != s(bool,X86)) ) | $spl16),
  inference(subsumption_resolution,[],[f4650,f571])).
fof(f4650,plain,(
  ( ! [X85,X86,X84] : (s(bool,'F_0') != s(bool,X84) | p(s(bool,X85)) | ~p(s(bool,X84)) | ~p(s(bool,X86)) | s(bool,X85) != s(bool,X86)) ) | $spl16),
  inference(superposition,[],[f573,f2447])).
fof(f4917,plain,(
  ( ! [X83,X81,X82] : (s(bool,'F_0') != s(bool,X81) | p(s(bool,X82)) | ~p(s(bool,X83)) | s(bool,X82) != s(bool,X83)) ) | $spl16),
  inference(subsumption_resolution,[],[f4649,f823])).
fof(f4649,plain,(
  ( ! [X83,X81,X82] : (s(bool,'F_0') != s(bool,X81) | p(s(bool,X82)) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.NegInf_0')))) | ~p(s(bool,X83)) | s(bool,X82) != s(bool,X83)) ) | $spl16),
  inference(superposition,[],[f573,f2447])).
fof(f4900,plain,(
  ( ! [X26,X24,X25] : (s(bool,'F_0') != s(bool,X24) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X26)) ) | $spl16),
  inference(subsumption_resolution,[],[f4630,f823])).
fof(f4630,plain,(
  ( ! [X26,X24,X25] : (s(bool,'F_0') != s(bool,X24) | p(s(bool,X25)) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.NegInf_0')))) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X26)) ) | $spl16),
  inference(superposition,[],[f569,f2447])).
fof(f4899,plain,(
  ( ! [X23,X21,X22] : (s(bool,'F_0') != s(bool,X21) | p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X23)) ) | $spl16),
  inference(subsumption_resolution,[],[f4629,f571])).
fof(f4629,plain,(
  ( ! [X23,X21,X22] : (s(bool,'F_0') != s(bool,X21) | p(s(bool,X22)) | ~p(s(bool,X21)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X23)) ) | $spl16),
  inference(superposition,[],[f569,f2447])).
fof(f2325,plain,(
  ( ! [X37] : (s(bool,'T_0') != s(bool,X37)) ) | $spl80),
  inference(cnf_transformation,[],[f2325_D])).
fof(f2325_D,plain,(
  ( ! [X37] : (s(bool,'T_0') != s(bool,X37)) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f4473,plain,(
  ( ! [X97,X98,X96] : (s(bool,'T_0') != s(bool,X98) | p(s(bool,X96)) | ~p(s(bool,X97)) | s(bool,X96) != s(bool,X97)) )),
  inference(forward_demodulation,[],[f4472,f1582])).
fof(f1582,plain,(
  s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)))),
  inference(resolution,[],[f831,f764])).
fof(f764,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f542])).
fof(f831,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12))))),
  inference(resolution,[],[f565,f790])).
fof(f4472,plain,(
  ( ! [X97,X98,X96] : (p(s(bool,X96)) | ~p(s(bool,X97)) | s(bool,X96) != s(bool,X97) | s(bool,X98) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)))) )),
  inference(subsumption_resolution,[],[f4208,f823])).
fof(f4208,plain,(
  ( ! [X97,X98,X96] : (~p(s(bool,'T_0')) | p(s(bool,X96)) | ~p(s(bool,X97)) | s(bool,X96) != s(bool,X97) | s(bool,X98) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)))) )),
  inference(superposition,[],[f573,f1582])).
fof(f4469,plain,(
  ( ! [X85,X86,X84] : (s(bool,'T_0') != s(bool,X84) | p(s(bool,X85)) | ~p(s(bool,X86)) | s(bool,X85) != s(bool,X86)) )),
  inference(subsumption_resolution,[],[f4204,f763])).
fof(f4204,plain,(
  ( ! [X85,X86,X84] : (s(bool,'T_0') != s(bool,X84) | p(s(bool,X85)) | ~p(s(bool,X84)) | ~p(s(bool,X86)) | s(bool,X85) != s(bool,X86)) )),
  inference(superposition,[],[f573,f1582])).
fof(f4468,plain,(
  ( ! [X83,X81,X82] : (s(bool,'T_0') != s(bool,X81) | p(s(bool,X82)) | ~p(s(bool,X83)) | s(bool,X82) != s(bool,X83)) )),
  inference(subsumption_resolution,[],[f4203,f823])).
fof(f4203,plain,(
  ( ! [X83,X81,X82] : (s(bool,'T_0') != s(bool,X81) | p(s(bool,X82)) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)))) | ~p(s(bool,X83)) | s(bool,X82) != s(bool,X83)) )),
  inference(superposition,[],[f573,f1582])).
fof(f4456,plain,(
  ( ! [X37,X38,X36] : (s(bool,'T_0') != s(bool,X38) | p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,X37)) )),
  inference(forward_demodulation,[],[f4455,f1582])).
fof(f4455,plain,(
  ( ! [X37,X38,X36] : (p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,X37) | s(bool,X38) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)))) )),
  inference(subsumption_resolution,[],[f4188,f823])).
fof(f4188,plain,(
  ( ! [X37,X38,X36] : (~p(s(bool,'T_0')) | p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,X37) | s(bool,X38) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)))) )),
  inference(superposition,[],[f569,f1582])).
fof(f4452,plain,(
  ( ! [X26,X24,X25] : (s(bool,'T_0') != s(bool,X24) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X26)) )),
  inference(subsumption_resolution,[],[f4184,f823])).
fof(f4184,plain,(
  ( ! [X26,X24,X25] : (s(bool,'T_0') != s(bool,X24) | p(s(bool,X25)) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)))) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X26)) )),
  inference(superposition,[],[f569,f1582])).
fof(f4451,plain,(
  ( ! [X23,X21,X22] : (s(bool,'T_0') != s(bool,X21) | p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X23)) )),
  inference(subsumption_resolution,[],[f4183,f763])).
fof(f4183,plain,(
  ( ! [X23,X21,X22] : (s(bool,'T_0') != s(bool,X21) | p(s(bool,X22)) | ~p(s(bool,X21)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X23)) )),
  inference(superposition,[],[f569,f1582])).
fof(f3913,plain,(
  $false | $spl4),
  inference(trivial_inequality_removal,[],[f3880])).
fof(f3880,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl4),
  inference(superposition,[],[f965,f917])).
fof(f917,plain,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0')))),
  inference(resolution,[],[f565,f714])).
fof(f714,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f511])).
fof(f511,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X0)) & (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f510])).
fof(f510,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X0)) & (s(bool,'F_0') != s(bool,X0) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f390])).
fof(f390,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> s(bool,'F_0') = s(bool,X0))),
  inference(ennf_transformation,[],[f302])).
fof(f302,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> s(bool,'F_0') = s(bool,X0))),
  inference(rectify,[],[f262])).
fof(f262,axiom,(
  ! [X22] : ((p(s(bool,X22)) => p(s(bool,'F_0'))) <=> s(bool,'F_0') = s(bool,X22))),
  file('thm.extreal.le_infty|split|1','thm.bool.IMP_F_EQ_F')).
fof(f3912,plain,(
  $false | $spl4),
  inference(equality_resolution,[],[f965])).
fof(f3678,plain,(
  ( ! [X85,X86,X84] : (s(bool,'F_0') != s(bool,X84) | p(s(bool,X85)) | ~p(s(bool,X86)) | s(bool,X85) != s(bool,X86)) )),
  inference(subsumption_resolution,[],[f3412,f571])).
fof(f3412,plain,(
  ( ! [X85,X86,X84] : (s(bool,'F_0') != s(bool,X84) | p(s(bool,X85)) | ~p(s(bool,X84)) | ~p(s(bool,X86)) | s(bool,X85) != s(bool,X86)) )),
  inference(superposition,[],[f573,f917])).
fof(f3677,plain,(
  ( ! [X83,X81,X82] : (s(bool,'F_0') != s(bool,X81) | p(s(bool,X82)) | ~p(s(bool,X83)) | s(bool,X82) != s(bool,X83)) )),
  inference(subsumption_resolution,[],[f3411,f823])).
fof(f3411,plain,(
  ( ! [X83,X81,X82] : (s(bool,'F_0') != s(bool,X81) | p(s(bool,X82)) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0')))) | ~p(s(bool,X83)) | s(bool,X82) != s(bool,X83)) )),
  inference(superposition,[],[f573,f917])).
fof(f3660,plain,(
  ( ! [X26,X24,X25] : (s(bool,'F_0') != s(bool,X24) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X26)) )),
  inference(subsumption_resolution,[],[f3392,f823])).
fof(f3392,plain,(
  ( ! [X26,X24,X25] : (s(bool,'F_0') != s(bool,X24) | p(s(bool,X25)) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0')))) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X26)) )),
  inference(superposition,[],[f569,f917])).
fof(f3659,plain,(
  ( ! [X23,X21,X22] : (s(bool,'F_0') != s(bool,X21) | p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X23)) )),
  inference(subsumption_resolution,[],[f3391,f571])).
fof(f3391,plain,(
  ( ! [X23,X21,X22] : (s(bool,'F_0') != s(bool,X21) | p(s(bool,X22)) | ~p(s(bool,X21)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X23)) )),
  inference(superposition,[],[f569,f917])).
fof(f3160,plain,(
  s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK45(sK12))))))) | $spl128),
  inference(cnf_transformation,[],[f3160_D])).
fof(f3160_D,plain,(
  s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK45(sK12))))))) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f988,plain,(
  ( ! [X0] : (s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))) ) | $spl22),
  inference(cnf_transformation,[],[f988_D])).
fof(f988_D,plain,(
  ( ! [X0] : (s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f3183,plain,(
  ( ! [X35] : (s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X35))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))))))) ) | ($spl19 | $spl24)),
  inference(forward_demodulation,[],[f3182,f990])).
fof(f3182,plain,(
  ( ! [X35] : (s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK45(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X35))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12))))) ) | ($spl19 | $spl24)),
  inference(forward_demodulation,[],[f3181,f990])).
fof(f3181,plain,(
  ( ! [X35] : (s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK45(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X35))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12))))) ) | ($spl19 | $spl24)),
  inference(subsumption_resolution,[],[f3180,f803])).
fof(f803,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))) )),
  inference(cnf_transformation,[],[f354])).
fof(f354,plain,(
  ! [X0] : s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))),
  inference(flattening,[],[f353])).
fof(f353,plain,(
  ! [X0] : ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))),
  inference(rectify,[],[f82])).
fof(f82,axiom,(
  ! [X11] : ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X11)))),
  file('thm.extreal.le_infty|split|1','thm.extreal.extreal_distinct|split|2')).
fof(f3180,plain,(
  ( ! [X35] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK45(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X35))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12))))) ) | ($spl19 | $spl24)),
  inference(subsumption_resolution,[],[f3124,f981])).
fof(f981,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',sK12) | $spl19),
  inference(cnf_transformation,[],[f981_D])).
fof(f981_D,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',sK12) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f3124,plain,(
  ( ! [X35] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK45(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X35))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12))))) ) | $spl24),
  inference(superposition,[],[f807,f990])).
fof(f807,plain,(
  ( ! [X2,X0] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(X0)))) | s('type.extreal.extreal',X0) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) )),
  inference(cnf_transformation,[],[f559])).
fof(f559,plain,(
  ! [X0] : ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0) & ! [X2] : s('type.extreal.extreal',X0) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) | s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(X0)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK45])],[f558])).
fof(f558,plain,(
  ! [X0] : ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0) & ! [X2] : s('type.extreal.extreal',X0) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) | ? [X1] : s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0))),
  inference(rectify,[],[f414])).
fof(f414,plain,(
  ! [X0] : ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0) & ! [X1] : s('type.extreal.extreal',X0) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) | ? [X2] : s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0))),
  inference(flattening,[],[f413])).
fof(f413,plain,(
  ! [X0] : ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0) & ! [X1] : s('type.extreal.extreal',X0) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) | (? [X2] : s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)))),
  inference(ennf_transformation,[],[f357])).
fof(f357,plain,(
  ! [X0] : ((s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | ? [X1] : s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) => (? [X2] : s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)))),
  inference(rectify,[],[f55])).
fof(f55,axiom,(
  ! [X11] : ((s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X11) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X11) | ? [X6] : s('type.extreal.extreal',X11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X6)))) => (? [X6] : s('type.extreal.extreal',X11) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X6))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X11) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X11)))),
  file('thm.extreal.le_infty|split|1','thm.extreal.extreal_46_prove_ltb')).
fof(f3179,plain,(
  ( ! [X34] : (s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X34))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))))))) ) | ($spl15 | $spl24)),
  inference(forward_demodulation,[],[f3178,f990])).
fof(f3178,plain,(
  ( ! [X34] : (s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK45(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X34))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12))))) ) | ($spl15 | $spl24)),
  inference(forward_demodulation,[],[f3177,f990])).
fof(f3177,plain,(
  ( ! [X34] : (s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK45(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X34))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12))))) ) | ($spl15 | $spl24)),
  inference(subsumption_resolution,[],[f3176,f821])).
fof(f821,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))) )),
  inference(cnf_transformation,[],[f369])).
fof(f369,plain,(
  ! [X0] : s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))),
  inference(flattening,[],[f368])).
fof(f368,plain,(
  ! [X0] : ~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))),
  inference(rectify,[],[f146])).
fof(f146,axiom,(
  ! [X6] : ~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X6)))),
  file('thm.extreal.le_infty|split|1','thm.extreal.extreal_not_infty|split|0')).
fof(f3176,plain,(
  ( ! [X34] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK45(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X34))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12))))) ) | ($spl15 | $spl24)),
  inference(subsumption_resolution,[],[f3123,f976])).
fof(f976,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',sK12) | $spl15),
  inference(cnf_transformation,[],[f976_D])).
fof(f976_D,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',sK12) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f3123,plain,(
  ( ! [X34] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK45(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X34))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12))))) ) | $spl24),
  inference(superposition,[],[f807,f990])).
fof(f3175,plain,(
  ( ! [X33] : (s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK45(sK12))))))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X33)))) ) | $spl24),
  inference(forward_demodulation,[],[f3174,f990])).
fof(f3174,plain,(
  ( ! [X33] : (s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X33))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))))))) ) | $spl24),
  inference(subsumption_resolution,[],[f3173,f821])).
fof(f3173,plain,(
  ( ! [X33] : (s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X33))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))))))) ) | $spl24),
  inference(subsumption_resolution,[],[f3122,f803])).
fof(f3122,plain,(
  ( ! [X33] : (s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X33))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))))))) ) | $spl24),
  inference(superposition,[],[f807,f990])).
fof(f3164,plain,(
  ( ! [X3] : (s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))))))) ) | ($spl15 | $spl24)),
  inference(forward_demodulation,[],[f3163,f990])).
fof(f3163,plain,(
  ( ! [X3] : (s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK45(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12))))) ) | ($spl15 | $spl24)),
  inference(forward_demodulation,[],[f3162,f990])).
fof(f3162,plain,(
  ( ! [X3] : (s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK45(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12))))) ) | ($spl15 | $spl24)),
  inference(subsumption_resolution,[],[f3161,f821])).
fof(f3161,plain,(
  ( ! [X3] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK45(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12))))) ) | ($spl15 | $spl24)),
  inference(subsumption_resolution,[],[f3087,f976])).
fof(f3087,plain,(
  ( ! [X3] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK45(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12))))) ) | $spl24),
  inference(superposition,[],[f990,f807])).
fof(f3158,plain,(
  ( ! [X2] : (s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))))))) ) | ($spl19 | $spl24)),
  inference(forward_demodulation,[],[f3157,f990])).
fof(f3157,plain,(
  ( ! [X2] : (s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK45(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12))))) ) | ($spl19 | $spl24)),
  inference(forward_demodulation,[],[f3156,f990])).
fof(f3156,plain,(
  ( ! [X2] : (s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK45(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12))))) ) | ($spl19 | $spl24)),
  inference(subsumption_resolution,[],[f3155,f803])).
fof(f3155,plain,(
  ( ! [X2] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK45(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12))))) ) | ($spl19 | $spl24)),
  inference(subsumption_resolution,[],[f3086,f981])).
fof(f3086,plain,(
  ( ! [X2] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK45(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12))))) ) | $spl24),
  inference(superposition,[],[f990,f807])).
fof(f2936,plain,(
  s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))))) | $spl126),
  inference(cnf_transformation,[],[f2936_D])).
fof(f2936_D,plain,(
  s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))))) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f2959,plain,(
  ( ! [X35] : (s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X35))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))))))) ) | ($spl19 | $spl20)),
  inference(forward_demodulation,[],[f2958,f985])).
fof(f2958,plain,(
  ( ! [X35] : (s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X35))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))) ) | ($spl19 | $spl20)),
  inference(forward_demodulation,[],[f2957,f985])).
fof(f2957,plain,(
  ( ! [X35] : (s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X35))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))) ) | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f2956,f803])).
fof(f2956,plain,(
  ( ! [X35] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X35))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))) ) | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f2900,f981])).
fof(f2900,plain,(
  ( ! [X35] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X35))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))) ) | $spl20),
  inference(superposition,[],[f807,f985])).
fof(f2955,plain,(
  ( ! [X34] : (s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X34))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))))))) ) | ($spl15 | $spl20)),
  inference(forward_demodulation,[],[f2954,f985])).
fof(f2954,plain,(
  ( ! [X34] : (s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X34))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))) ) | ($spl15 | $spl20)),
  inference(forward_demodulation,[],[f2953,f985])).
fof(f2953,plain,(
  ( ! [X34] : (s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X34))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))) ) | ($spl15 | $spl20)),
  inference(subsumption_resolution,[],[f2952,f821])).
fof(f2952,plain,(
  ( ! [X34] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X34))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))) ) | ($spl15 | $spl20)),
  inference(subsumption_resolution,[],[f2899,f976])).
fof(f2899,plain,(
  ( ! [X34] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X34))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))) ) | $spl20),
  inference(superposition,[],[f807,f985])).
fof(f2951,plain,(
  ( ! [X33] : (s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X33)))) ) | $spl20),
  inference(forward_demodulation,[],[f2950,f985])).
fof(f2950,plain,(
  ( ! [X33] : (s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X33))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))))))) ) | $spl20),
  inference(subsumption_resolution,[],[f2949,f821])).
fof(f2949,plain,(
  ( ! [X33] : (s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X33))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))))))) ) | $spl20),
  inference(subsumption_resolution,[],[f2898,f803])).
fof(f2898,plain,(
  ( ! [X33] : (s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X33))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))))))) ) | $spl20),
  inference(superposition,[],[f807,f985])).
fof(f2940,plain,(
  ( ! [X3] : (s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))))))) ) | ($spl15 | $spl20)),
  inference(forward_demodulation,[],[f2939,f985])).
fof(f2939,plain,(
  ( ! [X3] : (s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))) ) | ($spl15 | $spl20)),
  inference(forward_demodulation,[],[f2938,f985])).
fof(f2938,plain,(
  ( ! [X3] : (s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))) ) | ($spl15 | $spl20)),
  inference(subsumption_resolution,[],[f2937,f821])).
fof(f2937,plain,(
  ( ! [X3] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))) ) | ($spl15 | $spl20)),
  inference(subsumption_resolution,[],[f2863,f976])).
fof(f2863,plain,(
  ( ! [X3] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))) ) | $spl20),
  inference(superposition,[],[f985,f807])).
fof(f2934,plain,(
  ( ! [X2] : (s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))))))) ) | ($spl19 | $spl20)),
  inference(forward_demodulation,[],[f2933,f985])).
fof(f2933,plain,(
  ( ! [X2] : (s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))) ) | ($spl19 | $spl20)),
  inference(forward_demodulation,[],[f2932,f985])).
fof(f2932,plain,(
  ( ! [X2] : (s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))) ) | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f2931,f803])).
fof(f2931,plain,(
  ( ! [X2] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))) ) | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f2862,f981])).
fof(f2862,plain,(
  ( ! [X2] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))) ) | $spl20),
  inference(superposition,[],[f985,f807])).
fof(f2844,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl122),
  inference(cnf_transformation,[],[f2844_D])).
fof(f2844_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f2809,plain,(
  ( ! [X169,X168] : (p(s(bool,X168)) | ~p(s(bool,X169)) | s(bool,X169) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl31),
  inference(resolution,[],[f1621,f823])).
fof(f1621,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl31),
  inference(cnf_transformation,[],[f1621_D])).
fof(f1621_D,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) <=> ~$spl31),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl31])])).
fof(f2717,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s(bool,X21) != s(bool,X23)) ) | $spl31),
  inference(resolution,[],[f1621,f573])).
fof(f2715,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl31),
  inference(resolution,[],[f1621,f572])).
fof(f2845,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl124),
  inference(cnf_transformation,[],[f2845_D])).
fof(f2845_D,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f2714,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s(bool,X12) != s(bool,X13)) ) | $spl31),
  inference(resolution,[],[f1621,f571])).
fof(f2713,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl31),
  inference(resolution,[],[f1621,f571])).
fof(f2712,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s(bool,X6) != s(bool,X8)) ) | $spl31),
  inference(resolution,[],[f1621,f569])).
fof(f2711,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s(bool,X3) != s(bool,X4)) ) | $spl31),
  inference(resolution,[],[f1621,f567])).
fof(f2710,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl31),
  inference(resolution,[],[f1621,f567])).
fof(f2703,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f2572,f797])).
fof(f797,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0))))) )),
  inference(cnf_transformation,[],[f350])).
fof(f350,plain,(
  ! [X0] : p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0))))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X6] : p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X6))))),
  file('thm.extreal.le_infty|split|1','thm.extreal.le_infty|split|0')).
fof(f2572,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl30),
  inference(resolution,[],[f1622,f789])).
fof(f789,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))) )),
  inference(cnf_transformation,[],[f556])).
fof(f2571,plain,(
  $false | ($spl16 | $spl22)),
  inference(trivial_inequality_removal,[],[f2564])).
fof(f2564,plain,(
  s('type.extreal.extreal',sK12) != s('type.extreal.extreal',sK12) | ($spl16 | $spl22)),
  inference(superposition,[],[f988,f979])).
fof(f2517,plain,(
  s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))))) | $spl120),
  inference(cnf_transformation,[],[f2517_D])).
fof(f2517_D,plain,(
  s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))))) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f2540,plain,(
  ( ! [X35] : (s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X35))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))))))) ) | ($spl16 | $spl19)),
  inference(forward_demodulation,[],[f2539,f979])).
fof(f2539,plain,(
  ( ! [X35] : (s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X35))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))) ) | ($spl16 | $spl19)),
  inference(forward_demodulation,[],[f2538,f979])).
fof(f2538,plain,(
  ( ! [X35] : (s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X35))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))) ) | ($spl16 | $spl19)),
  inference(subsumption_resolution,[],[f2537,f803])).
fof(f2537,plain,(
  ( ! [X35] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X35))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))) ) | ($spl16 | $spl19)),
  inference(subsumption_resolution,[],[f2481,f981])).
fof(f2481,plain,(
  ( ! [X35] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X35))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))) ) | $spl16),
  inference(superposition,[],[f807,f979])).
fof(f2536,plain,(
  ( ! [X34] : (s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X34))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))))))) ) | ($spl15 | $spl16)),
  inference(forward_demodulation,[],[f2535,f979])).
fof(f2535,plain,(
  ( ! [X34] : (s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X34))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))) ) | ($spl15 | $spl16)),
  inference(forward_demodulation,[],[f2534,f979])).
fof(f2534,plain,(
  ( ! [X34] : (s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X34))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))) ) | ($spl15 | $spl16)),
  inference(subsumption_resolution,[],[f2533,f821])).
fof(f2533,plain,(
  ( ! [X34] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X34))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))) ) | ($spl15 | $spl16)),
  inference(subsumption_resolution,[],[f2480,f976])).
fof(f2480,plain,(
  ( ! [X34] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X34))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))) ) | $spl16),
  inference(superposition,[],[f807,f979])).
fof(f2532,plain,(
  ( ! [X33] : (s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X33)))) ) | $spl16),
  inference(forward_demodulation,[],[f2531,f979])).
fof(f2531,plain,(
  ( ! [X33] : (s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X33))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))))))) ) | $spl16),
  inference(subsumption_resolution,[],[f2530,f821])).
fof(f2530,plain,(
  ( ! [X33] : (s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X33))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))))))) ) | $spl16),
  inference(subsumption_resolution,[],[f2479,f803])).
fof(f2479,plain,(
  ( ! [X33] : (s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X33))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))))))) ) | $spl16),
  inference(superposition,[],[f807,f979])).
fof(f2521,plain,(
  ( ! [X3] : (s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))))))) ) | ($spl15 | $spl16)),
  inference(forward_demodulation,[],[f2520,f979])).
fof(f2520,plain,(
  ( ! [X3] : (s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))) ) | ($spl15 | $spl16)),
  inference(forward_demodulation,[],[f2519,f979])).
fof(f2519,plain,(
  ( ! [X3] : (s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))) ) | ($spl15 | $spl16)),
  inference(subsumption_resolution,[],[f2518,f821])).
fof(f2518,plain,(
  ( ! [X3] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))) ) | ($spl15 | $spl16)),
  inference(subsumption_resolution,[],[f2444,f976])).
fof(f2444,plain,(
  ( ! [X3] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X3))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))) ) | $spl16),
  inference(superposition,[],[f979,f807])).
fof(f2515,plain,(
  ( ! [X2] : (s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))))))) ) | ($spl16 | $spl19)),
  inference(forward_demodulation,[],[f2514,f979])).
fof(f2514,plain,(
  ( ! [X2] : (s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))) ) | ($spl16 | $spl19)),
  inference(forward_demodulation,[],[f2513,f979])).
fof(f2513,plain,(
  ( ! [X2] : (s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))) ) | ($spl16 | $spl19)),
  inference(subsumption_resolution,[],[f2512,f803])).
fof(f2512,plain,(
  ( ! [X2] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))) ) | ($spl16 | $spl19)),
  inference(subsumption_resolution,[],[f2443,f981])).
fof(f2443,plain,(
  ( ! [X2] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45('const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))))) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))) ) | $spl16),
  inference(superposition,[],[f979,f807])).
fof(f979,plain,(
  s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) | $spl16),
  inference(cnf_transformation,[],[f979_D])).
fof(f979_D,plain,(
  s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f2430,plain,(
  ( ! [X8,X9] : (sP0(X8,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X9)) ) | $spl118),
  inference(cnf_transformation,[],[f2430_D])).
fof(f2430_D,plain,(
  ( ! [X8,X9] : (sP0(X8,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X9)) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f2434,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X1) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))) ) | $spl15),
  inference(subsumption_resolution,[],[f1947,f976])).
fof(f1947,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X1) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f1504,f809])).
fof(f809,plain,(
  ( ! [X0] : (s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(X0)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0)) )),
  inference(cnf_transformation,[],[f561])).
fof(f561,plain,(
  ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK47])],[f359])).
fof(f359,plain,(
  ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | ? [X1] : s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1))))),
  inference(rectify,[],[f83])).
fof(f83,axiom,(
  ! [X51] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X51) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X51) | ? [X36] : s('type.extreal.extreal',X51) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X36))))),
  file('thm.extreal.le_infty|split|1','thm.extreal.extreal_nchotomy')).
fof(f1504,plain,(
  ( ! [X39,X38] : (sP0(X38,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)),X39)) )),
  inference(resolution,[],[f831,f583])).
fof(f583,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f441])).
fof(f441,plain,(
  ! [X0,X1,X2] : ((~sP0(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP0(X0,X1,X2)))),
  inference(rectify,[],[f440])).
fof(f440,plain,(
  ! [X2,X0,X1] : ((~sP0(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP0(X2,X0,X1)))),
  inference(flattening,[],[f439])).
fof(f439,plain,(
  ! [X2,X0,X1] : ((~sP0(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP0(X2,X0,X1)))),
  inference(nnf_transformation,[],[f415])).
fof(f415,plain,(
  ! [X2,X0,X1] : (sP0(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f2428,plain,(
  ( ! [X12,X11] : (sP0(X11,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X12)) ) | $spl116),
  inference(cnf_transformation,[],[f2428_D])).
fof(f2428_D,plain,(
  ( ! [X12,X11] : (sP0(X11,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X12)) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f2433,plain,(
  ( ! [X2,X3] : (sP0(X2,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X3) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))) ) | $spl19),
  inference(subsumption_resolution,[],[f1948,f981])).
fof(f1948,plain,(
  ( ! [X2,X3] : (sP0(X2,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X3) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f1504,f809])).
fof(f985,plain,(
  s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) | $spl20),
  inference(cnf_transformation,[],[f985_D])).
fof(f985_D,plain,(
  s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f2432,plain,(
  ( ! [X4,X5] : (sP0(X4,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X5) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))) ) | $spl15),
  inference(subsumption_resolution,[],[f1949,f976])).
fof(f1949,plain,(
  ( ! [X4,X5] : (sP0(X4,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X5) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f1504,f808])).
fof(f808,plain,(
  ( ! [X0] : (s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(X0)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0)) )),
  inference(cnf_transformation,[],[f560])).
fof(f560,plain,(
  ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK46])],[f358])).
fof(f358,plain,(
  ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | ? [X1] : s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1))))),
  inference(rectify,[],[f78])).
fof(f78,axiom,(
  ! [X6] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X6) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X6) | ? [X36] : s('type.extreal.extreal',X6) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X36))))),
  file('thm.extreal.le_infty|split|1','thm.extreal.extreal_cases')).
fof(f2431,plain,(
  ( ! [X6,X7] : (sP0(X6,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X7) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))) ) | $spl19),
  inference(subsumption_resolution,[],[f1950,f981])).
fof(f1950,plain,(
  ( ! [X6,X7] : (sP0(X6,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X7) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f1504,f808])).
fof(f990,plain,(
  s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | $spl24),
  inference(cnf_transformation,[],[f990_D])).
fof(f990_D,plain,(
  s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f2429,plain,(
  ( ! [X10,X8,X9] : (sP0(X8,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X9) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X10)))) ) | $spl15),
  inference(subsumption_resolution,[],[f1951,f976])).
fof(f1951,plain,(
  ( ! [X10,X8,X9] : (sP0(X8,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X9) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X10)))) )),
  inference(superposition,[],[f1504,f807])).
fof(f2427,plain,(
  ( ! [X12,X13,X11] : (sP0(X11,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13)))) ) | $spl19),
  inference(subsumption_resolution,[],[f1952,f981])).
fof(f1952,plain,(
  ( ! [X12,X13,X11] : (sP0(X11,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X12) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13)))) )),
  inference(superposition,[],[f1504,f807])).
fof(f2422,plain,(
  ( ! [X8,X9] : (sP2(X8,X9,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl114),
  inference(cnf_transformation,[],[f2422_D])).
fof(f2422_D,plain,(
  ( ! [X8,X9] : (sP2(X8,X9,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f2426,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))) ) | $spl15),
  inference(subsumption_resolution,[],[f1973,f976])).
fof(f1973,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f1517,f809])).
fof(f1517,plain,(
  ( ! [X64,X65] : (sP2(X64,X65,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)))) )),
  inference(resolution,[],[f831,f599])).
fof(f599,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f449])).
fof(f449,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(rectify,[],[f448])).
fof(f448,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(flattening,[],[f447])).
fof(f447,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(nnf_transformation,[],[f418])).
fof(f418,plain,(
  ! [X2,X0,X1] : (sP2(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f2420,plain,(
  ( ! [X12,X11] : (sP2(X11,X12,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl112),
  inference(cnf_transformation,[],[f2420_D])).
fof(f2420_D,plain,(
  ( ! [X12,X11] : (sP2(X11,X12,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f2425,plain,(
  ( ! [X2,X3] : (sP2(X2,X3,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))) ) | $spl19),
  inference(subsumption_resolution,[],[f1974,f981])).
fof(f1974,plain,(
  ( ! [X2,X3] : (sP2(X2,X3,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f1517,f809])).
fof(f2424,plain,(
  ( ! [X4,X5] : (sP2(X4,X5,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))) ) | $spl15),
  inference(subsumption_resolution,[],[f1975,f976])).
fof(f1975,plain,(
  ( ! [X4,X5] : (sP2(X4,X5,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f1517,f808])).
fof(f2423,plain,(
  ( ! [X6,X7] : (sP2(X6,X7,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))) ) | $spl19),
  inference(subsumption_resolution,[],[f1976,f981])).
fof(f1976,plain,(
  ( ! [X6,X7] : (sP2(X6,X7,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f1517,f808])).
fof(f2421,plain,(
  ( ! [X10,X8,X9] : (sP2(X8,X9,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X10)))) ) | $spl15),
  inference(subsumption_resolution,[],[f1977,f976])).
fof(f1977,plain,(
  ( ! [X10,X8,X9] : (sP2(X8,X9,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X10)))) )),
  inference(superposition,[],[f1517,f807])).
fof(f2419,plain,(
  ( ! [X12,X13,X11] : (sP2(X11,X12,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13)))) ) | $spl19),
  inference(subsumption_resolution,[],[f1978,f981])).
fof(f1978,plain,(
  ( ! [X12,X13,X11] : (sP2(X11,X12,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13)))) )),
  inference(superposition,[],[f1517,f807])).
fof(f2414,plain,(
  ( ! [X8,X9] : (sP2(X8,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X9)) ) | $spl110),
  inference(cnf_transformation,[],[f2414_D])).
fof(f2414_D,plain,(
  ( ! [X8,X9] : (sP2(X8,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X9)) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f2418,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X1) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))) ) | $spl15),
  inference(subsumption_resolution,[],[f1999,f976])).
fof(f1999,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X1) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f1518,f809])).
fof(f1518,plain,(
  ( ! [X66,X67] : (sP2(X66,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)),X67)) )),
  inference(resolution,[],[f831,f600])).
fof(f600,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f449])).
fof(f2412,plain,(
  ( ! [X12,X11] : (sP2(X11,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X12)) ) | $spl108),
  inference(cnf_transformation,[],[f2412_D])).
fof(f2412_D,plain,(
  ( ! [X12,X11] : (sP2(X11,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X12)) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f2417,plain,(
  ( ! [X2,X3] : (sP2(X2,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X3) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))) ) | $spl19),
  inference(subsumption_resolution,[],[f2000,f981])).
fof(f2000,plain,(
  ( ! [X2,X3] : (sP2(X2,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X3) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f1518,f809])).
fof(f2416,plain,(
  ( ! [X4,X5] : (sP2(X4,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X5) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))) ) | $spl15),
  inference(subsumption_resolution,[],[f2001,f976])).
fof(f2001,plain,(
  ( ! [X4,X5] : (sP2(X4,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X5) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f1518,f808])).
fof(f2415,plain,(
  ( ! [X6,X7] : (sP2(X6,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X7) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))) ) | $spl19),
  inference(subsumption_resolution,[],[f2002,f981])).
fof(f2002,plain,(
  ( ! [X6,X7] : (sP2(X6,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X7) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f1518,f808])).
fof(f2413,plain,(
  ( ! [X10,X8,X9] : (sP2(X8,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X9) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X10)))) ) | $spl15),
  inference(subsumption_resolution,[],[f2003,f976])).
fof(f2003,plain,(
  ( ! [X10,X8,X9] : (sP2(X8,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X9) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X10)))) )),
  inference(superposition,[],[f1518,f807])).
fof(f2411,plain,(
  ( ! [X12,X13,X11] : (sP2(X11,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13)))) ) | $spl19),
  inference(subsumption_resolution,[],[f2004,f981])).
fof(f2004,plain,(
  ( ! [X12,X13,X11] : (sP2(X11,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X12) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13)))) )),
  inference(superposition,[],[f1518,f807])).
fof(f2406,plain,(
  ( ! [X8,X9] : (sP4(X8,X9,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl106),
  inference(cnf_transformation,[],[f2406_D])).
fof(f2406_D,plain,(
  ( ! [X8,X9] : (sP4(X8,X9,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f2410,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))) ) | $spl15),
  inference(subsumption_resolution,[],[f2025,f976])).
fof(f2025,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f1534,f809])).
fof(f1534,plain,(
  ( ! [X99,X98] : (sP4(X98,X99,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)))) )),
  inference(resolution,[],[f831,f616])).
fof(f616,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f456])).
fof(f456,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(flattening,[],[f455])).
fof(f455,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(nnf_transformation,[],[f421])).
fof(f421,plain,(
  ! [X0,X1,X2] : (sP4(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f2404,plain,(
  ( ! [X12,X11] : (sP4(X11,X12,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl104),
  inference(cnf_transformation,[],[f2404_D])).
fof(f2404_D,plain,(
  ( ! [X12,X11] : (sP4(X11,X12,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f2409,plain,(
  ( ! [X2,X3] : (sP4(X2,X3,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))) ) | $spl19),
  inference(subsumption_resolution,[],[f2026,f981])).
fof(f2026,plain,(
  ( ! [X2,X3] : (sP4(X2,X3,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f1534,f809])).
fof(f2408,plain,(
  ( ! [X4,X5] : (sP4(X4,X5,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))) ) | $spl15),
  inference(subsumption_resolution,[],[f2027,f976])).
fof(f2027,plain,(
  ( ! [X4,X5] : (sP4(X4,X5,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f1534,f808])).
fof(f2407,plain,(
  ( ! [X6,X7] : (sP4(X6,X7,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))) ) | $spl19),
  inference(subsumption_resolution,[],[f2028,f981])).
fof(f2028,plain,(
  ( ! [X6,X7] : (sP4(X6,X7,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f1534,f808])).
fof(f2405,plain,(
  ( ! [X10,X8,X9] : (sP4(X8,X9,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X10)))) ) | $spl15),
  inference(subsumption_resolution,[],[f2029,f976])).
fof(f2029,plain,(
  ( ! [X10,X8,X9] : (sP4(X8,X9,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X10)))) )),
  inference(superposition,[],[f1534,f807])).
fof(f2403,plain,(
  ( ! [X12,X13,X11] : (sP4(X11,X12,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13)))) ) | $spl19),
  inference(subsumption_resolution,[],[f2030,f981])).
fof(f2030,plain,(
  ( ! [X12,X13,X11] : (sP4(X11,X12,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13)))) )),
  inference(superposition,[],[f1534,f807])).
fof(f2398,plain,(
  ( ! [X8,X9] : (sP8(X8,X9,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl102),
  inference(cnf_transformation,[],[f2398_D])).
fof(f2398_D,plain,(
  ( ! [X8,X9] : (sP8(X8,X9,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f2402,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))) ) | $spl15),
  inference(subsumption_resolution,[],[f2134,f976])).
fof(f2134,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f1547,f809])).
fof(f1547,plain,(
  ( ! [X125,X124] : (sP8(X124,X125,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)))) )),
  inference(resolution,[],[f831,f633])).
fof(f633,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f464])).
fof(f464,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f463])).
fof(f463,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f462])).
fof(f462,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f426])).
fof(f426,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f2396,plain,(
  ( ! [X12,X11] : (sP8(X11,X12,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl100),
  inference(cnf_transformation,[],[f2396_D])).
fof(f2396_D,plain,(
  ( ! [X12,X11] : (sP8(X11,X12,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f2401,plain,(
  ( ! [X2,X3] : (sP8(X2,X3,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))) ) | $spl19),
  inference(subsumption_resolution,[],[f2135,f981])).
fof(f2135,plain,(
  ( ! [X2,X3] : (sP8(X2,X3,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f1547,f809])).
fof(f2400,plain,(
  ( ! [X4,X5] : (sP8(X4,X5,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))) ) | $spl15),
  inference(subsumption_resolution,[],[f2136,f976])).
fof(f2136,plain,(
  ( ! [X4,X5] : (sP8(X4,X5,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f1547,f808])).
fof(f2399,plain,(
  ( ! [X6,X7] : (sP8(X6,X7,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))) ) | $spl19),
  inference(subsumption_resolution,[],[f2137,f981])).
fof(f2137,plain,(
  ( ! [X6,X7] : (sP8(X6,X7,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f1547,f808])).
fof(f2397,plain,(
  ( ! [X10,X8,X9] : (sP8(X8,X9,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X10)))) ) | $spl15),
  inference(subsumption_resolution,[],[f2138,f976])).
fof(f2138,plain,(
  ( ! [X10,X8,X9] : (sP8(X8,X9,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X10)))) )),
  inference(superposition,[],[f1547,f807])).
fof(f2395,plain,(
  ( ! [X12,X13,X11] : (sP8(X11,X12,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13)))) ) | $spl19),
  inference(subsumption_resolution,[],[f2139,f981])).
fof(f2139,plain,(
  ( ! [X12,X13,X11] : (sP8(X11,X12,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13)))) )),
  inference(superposition,[],[f1547,f807])).
fof(f2390,plain,(
  ( ! [X8,X9] : (sP7(X8,X9,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl98),
  inference(cnf_transformation,[],[f2390_D])).
fof(f2390_D,plain,(
  ( ! [X8,X9] : (sP7(X8,X9,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f2394,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))) ) | $spl15),
  inference(subsumption_resolution,[],[f2160,f976])).
fof(f2160,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f1550,f809])).
fof(f1550,plain,(
  ( ! [X130,X131] : (sP7(X130,X131,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)))) )),
  inference(resolution,[],[f831,f637])).
fof(f637,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f467])).
fof(f467,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f466])).
fof(f466,plain,(
  ! [X1,X0,X2] : ((~sP7(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP7(X1,X0,X2)))),
  inference(flattening,[],[f465])).
fof(f465,plain,(
  ! [X1,X0,X2] : ((~sP7(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP7(X1,X0,X2)))),
  inference(nnf_transformation,[],[f425])).
fof(f425,plain,(
  ! [X1,X0,X2] : (sP7(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f2388,plain,(
  ( ! [X12,X11] : (sP7(X11,X12,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl96),
  inference(cnf_transformation,[],[f2388_D])).
fof(f2388_D,plain,(
  ( ! [X12,X11] : (sP7(X11,X12,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f2393,plain,(
  ( ! [X2,X3] : (sP7(X2,X3,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))) ) | $spl19),
  inference(subsumption_resolution,[],[f2161,f981])).
fof(f2161,plain,(
  ( ! [X2,X3] : (sP7(X2,X3,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f1550,f809])).
fof(f2392,plain,(
  ( ! [X4,X5] : (sP7(X4,X5,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))) ) | $spl15),
  inference(subsumption_resolution,[],[f2162,f976])).
fof(f2162,plain,(
  ( ! [X4,X5] : (sP7(X4,X5,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f1550,f808])).
fof(f2391,plain,(
  ( ! [X6,X7] : (sP7(X6,X7,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))) ) | $spl19),
  inference(subsumption_resolution,[],[f2163,f981])).
fof(f2163,plain,(
  ( ! [X6,X7] : (sP7(X6,X7,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f1550,f808])).
fof(f2389,plain,(
  ( ! [X10,X8,X9] : (sP7(X8,X9,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X10)))) ) | $spl15),
  inference(subsumption_resolution,[],[f2164,f976])).
fof(f2164,plain,(
  ( ! [X10,X8,X9] : (sP7(X8,X9,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X10)))) )),
  inference(superposition,[],[f1550,f807])).
fof(f2387,plain,(
  ( ! [X12,X13,X11] : (sP7(X11,X12,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13)))) ) | $spl19),
  inference(subsumption_resolution,[],[f2165,f981])).
fof(f2165,plain,(
  ( ! [X12,X13,X11] : (sP7(X11,X12,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13)))) )),
  inference(superposition,[],[f1550,f807])).
fof(f2382,plain,(
  ( ! [X8,X9] : (sP6(X8,X9,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl94),
  inference(cnf_transformation,[],[f2382_D])).
fof(f2382_D,plain,(
  ( ! [X8,X9] : (sP6(X8,X9,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f2386,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))) ) | $spl15),
  inference(subsumption_resolution,[],[f2186,f976])).
fof(f2186,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f1551,f809])).
fof(f1551,plain,(
  ( ! [X132,X133] : (sP6(X132,X133,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)))) )),
  inference(resolution,[],[f831,f641])).
fof(f641,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f469])).
fof(f469,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(flattening,[],[f468])).
fof(f468,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(nnf_transformation,[],[f424])).
fof(f424,plain,(
  ! [X0,X1,X2] : (sP6(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f2380,plain,(
  ( ! [X12,X11] : (sP6(X11,X12,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl92),
  inference(cnf_transformation,[],[f2380_D])).
fof(f2380_D,plain,(
  ( ! [X12,X11] : (sP6(X11,X12,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f2385,plain,(
  ( ! [X2,X3] : (sP6(X2,X3,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))) ) | $spl19),
  inference(subsumption_resolution,[],[f2187,f981])).
fof(f2187,plain,(
  ( ! [X2,X3] : (sP6(X2,X3,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f1551,f809])).
fof(f2384,plain,(
  ( ! [X4,X5] : (sP6(X4,X5,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))) ) | $spl15),
  inference(subsumption_resolution,[],[f2188,f976])).
fof(f2188,plain,(
  ( ! [X4,X5] : (sP6(X4,X5,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f1551,f808])).
fof(f2383,plain,(
  ( ! [X6,X7] : (sP6(X6,X7,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))) ) | $spl19),
  inference(subsumption_resolution,[],[f2189,f981])).
fof(f2189,plain,(
  ( ! [X6,X7] : (sP6(X6,X7,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f1551,f808])).
fof(f2381,plain,(
  ( ! [X10,X8,X9] : (sP6(X8,X9,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X10)))) ) | $spl15),
  inference(subsumption_resolution,[],[f2190,f976])).
fof(f2190,plain,(
  ( ! [X10,X8,X9] : (sP6(X8,X9,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X10)))) )),
  inference(superposition,[],[f1551,f807])).
fof(f2379,plain,(
  ( ! [X12,X13,X11] : (sP6(X11,X12,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13)))) ) | $spl19),
  inference(subsumption_resolution,[],[f2191,f981])).
fof(f2191,plain,(
  ( ! [X12,X13,X11] : (sP6(X11,X12,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13)))) )),
  inference(superposition,[],[f1551,f807])).
fof(f2374,plain,(
  ( ! [X8,X9] : (sP6(X8,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X9)) ) | $spl90),
  inference(cnf_transformation,[],[f2374_D])).
fof(f2374_D,plain,(
  ( ! [X8,X9] : (sP6(X8,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X9)) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f2378,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X1) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))) ) | $spl15),
  inference(subsumption_resolution,[],[f2213,f976])).
fof(f2213,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X1) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f1552,f809])).
fof(f1552,plain,(
  ( ! [X134,X135] : (sP6(X134,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)),X135)) )),
  inference(resolution,[],[f831,f642])).
fof(f642,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f469])).
fof(f2372,plain,(
  ( ! [X12,X11] : (sP6(X11,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X12)) ) | $spl88),
  inference(cnf_transformation,[],[f2372_D])).
fof(f2372_D,plain,(
  ( ! [X12,X11] : (sP6(X11,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X12)) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f2377,plain,(
  ( ! [X2,X3] : (sP6(X2,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X3) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))) ) | $spl19),
  inference(subsumption_resolution,[],[f2214,f981])).
fof(f2214,plain,(
  ( ! [X2,X3] : (sP6(X2,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X3) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f1552,f809])).
fof(f2376,plain,(
  ( ! [X4,X5] : (sP6(X4,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X5) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))) ) | $spl15),
  inference(subsumption_resolution,[],[f2215,f976])).
fof(f2215,plain,(
  ( ! [X4,X5] : (sP6(X4,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X5) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f1552,f808])).
fof(f2375,plain,(
  ( ! [X6,X7] : (sP6(X6,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X7) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))) ) | $spl19),
  inference(subsumption_resolution,[],[f2216,f981])).
fof(f2216,plain,(
  ( ! [X6,X7] : (sP6(X6,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X7) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f1552,f808])).
fof(f2373,plain,(
  ( ! [X10,X8,X9] : (sP6(X8,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X9) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X10)))) ) | $spl15),
  inference(subsumption_resolution,[],[f2217,f976])).
fof(f2217,plain,(
  ( ! [X10,X8,X9] : (sP6(X8,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X9) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X10)))) )),
  inference(superposition,[],[f1552,f807])).
fof(f2371,plain,(
  ( ! [X12,X13,X11] : (sP6(X11,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13)))) ) | $spl19),
  inference(subsumption_resolution,[],[f2218,f981])).
fof(f2218,plain,(
  ( ! [X12,X13,X11] : (sP6(X11,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X12) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13)))) )),
  inference(superposition,[],[f1552,f807])).
fof(f2366,plain,(
  ( ! [X8,X9] : (sP6('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X8,X9)) ) | $spl86),
  inference(cnf_transformation,[],[f2366_D])).
fof(f2366_D,plain,(
  ( ! [X8,X9] : (sP6('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X8,X9)) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f2370,plain,(
  ( ! [X0,X1] : (sP6('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X0,X1) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))) ) | $spl15),
  inference(subsumption_resolution,[],[f2240,f976])).
fof(f2240,plain,(
  ( ! [X0,X1] : (sP6('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X0,X1) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f1553,f809])).
fof(f1553,plain,(
  ( ! [X136,X137] : (sP6('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)),X136,X137)) )),
  inference(resolution,[],[f831,f643])).
fof(f643,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f469])).
fof(f2364,plain,(
  ( ! [X12,X11] : (sP6('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X11,X12)) ) | $spl84),
  inference(cnf_transformation,[],[f2364_D])).
fof(f2364_D,plain,(
  ( ! [X12,X11] : (sP6('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X11,X12)) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f2369,plain,(
  ( ! [X2,X3] : (sP6('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X2,X3) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))) ) | $spl19),
  inference(subsumption_resolution,[],[f2241,f981])).
fof(f2241,plain,(
  ( ! [X2,X3] : (sP6('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X2,X3) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f1553,f809])).
fof(f2368,plain,(
  ( ! [X4,X5] : (sP6('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X4,X5) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))) ) | $spl15),
  inference(subsumption_resolution,[],[f2242,f976])).
fof(f2242,plain,(
  ( ! [X4,X5] : (sP6('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X4,X5) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f1553,f808])).
fof(f2367,plain,(
  ( ! [X6,X7] : (sP6('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X6,X7) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))) ) | $spl19),
  inference(subsumption_resolution,[],[f2243,f981])).
fof(f2243,plain,(
  ( ! [X6,X7] : (sP6('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X6,X7) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f1553,f808])).
fof(f2365,plain,(
  ( ! [X10,X8,X9] : (sP6('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X8,X9) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X10)))) ) | $spl15),
  inference(subsumption_resolution,[],[f2244,f976])).
fof(f2244,plain,(
  ( ! [X10,X8,X9] : (sP6('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X8,X9) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X10)))) )),
  inference(superposition,[],[f1553,f807])).
fof(f2363,plain,(
  ( ! [X12,X13,X11] : (sP6('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X11,X12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13)))) ) | $spl19),
  inference(subsumption_resolution,[],[f2245,f981])).
fof(f2245,plain,(
  ( ! [X12,X13,X11] : (sP6('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X11,X12) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13)))) )),
  inference(superposition,[],[f1553,f807])).
fof(f1622,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl30),
  inference(cnf_transformation,[],[f1622_D])).
fof(f1622_D,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f2362,plain,(
  ( ! [X8,X9] : (sP1(X8,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X9) | ~p(s(bool,X9))) ) | $spl82),
  inference(cnf_transformation,[],[f2362_D])).
fof(f2362_D,plain,(
  ( ! [X8,X9] : (sP1(X8,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X9) | ~p(s(bool,X9))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f2344,plain,(
  ( ! [X8,X9] : (sP1(X8,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X9) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | ~p(s(bool,X9))) ) | ($spl15 | $spl22)),
  inference(resolution,[],[f1955,f580])).
fof(f580,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | ~sP0(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f438])).
fof(f438,plain,(
  ! [X0,X1,X2] : ((~sP1(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP0(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP0(X0,X1,X2) | sP1(X0,X1,X2)))),
  inference(rectify,[],[f437])).
fof(f437,plain,(
  ! [X2,X0,X1] : ((~sP1(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP0(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP0(X2,X0,X1) | sP1(X2,X0,X1)))),
  inference(flattening,[],[f436])).
fof(f436,plain,(
  ! [X2,X0,X1] : ((~sP1(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP0(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP0(X2,X0,X1)) | sP1(X2,X0,X1)))),
  inference(nnf_transformation,[],[f416])).
fof(f416,plain,(
  ! [X2,X0,X1] : (sP1(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP0(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f1955,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X1)) ) | ($spl15 | $spl22)),
  inference(subsumption_resolution,[],[f1954,f976])).
fof(f1954,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X1) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) ) | $spl22),
  inference(subsumption_resolution,[],[f1947,f988])).
fof(f2340,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | ~p(s(bool,X1))) ) | ($spl15 | $spl22)),
  inference(resolution,[],[f1955,f581])).
fof(f581,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP0(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f441])).
fof(f2321,plain,(
  ( ! [X39,X38] : (s(bool,'T_0') != s(bool,X39) | ~p(s(bool,X38))) ) | $spl8),
  inference(duplicate_literal_removal,[],[f2270])).
fof(f2270,plain,(
  ( ! [X39,X38] : (s(bool,'T_0') != s(bool,X39) | ~p(s(bool,X38)) | ~p(s(bool,X38))) ) | $spl8),
  inference(superposition,[],[f967,f766])).
fof(f766,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f543])).
fof(f543,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f322])).
fof(f322,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X22] : (s(bool,'T_0') = s(bool,X22) <=> p(s(bool,X22)))),
  file('thm.extreal.le_infty|split|1','thm.bool.EQ_CLAUSES|split|0')).
fof(f2322,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,X37) | ~p(s(bool,X36))) ) | $spl8),
  inference(duplicate_literal_removal,[],[f2269])).
fof(f2269,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,X37) | ~p(s(bool,X36)) | ~p(s(bool,X36))) ) | $spl8),
  inference(superposition,[],[f967,f764])).
fof(f2104,plain,(
  ( ! [X35,X34] : (s(bool,'F_0') != s(bool,X35) | p(s(bool,X34))) ) | $spl2),
  inference(duplicate_literal_removal,[],[f2052])).
fof(f2052,plain,(
  ( ! [X35,X34] : (s(bool,'F_0') != s(bool,X35) | p(s(bool,X34)) | p(s(bool,X34))) ) | $spl2),
  inference(superposition,[],[f964,f762])).
fof(f762,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f541])).
fof(f541,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f320])).
fof(f320,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f319])).
fof(f319,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X22] : (s(bool,'F_0') = s(bool,X22) <=> ~p(s(bool,X22)))),
  file('thm.extreal.le_infty|split|1','thm.bool.EQ_CLAUSES|split|2')).
fof(f2105,plain,(
  ( ! [X33,X32] : (s(bool,'F_0') != s(bool,X33) | p(s(bool,X32))) ) | $spl2),
  inference(duplicate_literal_removal,[],[f2051])).
fof(f2051,plain,(
  ( ! [X33,X32] : (s(bool,'F_0') != s(bool,X33) | p(s(bool,X32)) | p(s(bool,X32))) ) | $spl2),
  inference(superposition,[],[f964,f760])).
fof(f760,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f540])).
fof(f540,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f318])).
fof(f318,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f317])).
fof(f317,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X22] : (s(bool,'F_0') = s(bool,X22) <=> ~p(s(bool,X22)))),
  file('thm.extreal.le_infty|split|1','thm.bool.EQ_CLAUSES|split|3')).
fof(f2106,plain,(
  ( ! [X28,X29] : (s(bool,'F_0') != s(bool,X29) | p(s(bool,X28))) ) | $spl2),
  inference(duplicate_literal_removal,[],[f2049])).
fof(f2049,plain,(
  ( ! [X28,X29] : (s(bool,'F_0') != s(bool,X29) | p(s(bool,X28)) | p(s(bool,X28))) ) | $spl2),
  inference(superposition,[],[f964,f714])).
fof(f1914,plain,(
  ( ! [X2,X3] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X2,X3)) ) | $spl76),
  inference(cnf_transformation,[],[f1914_D])).
fof(f1914_D,plain,(
  ( ! [X2,X3] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X2,X3)) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f1918,plain,(
  ( ! [X12,X13,X11] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X11,X12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13)))) ) | $spl19),
  inference(subsumption_resolution,[],[f1901,f981])).
fof(f1901,plain,(
  ( ! [X12,X13,X11] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X11,X12) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13)))) )),
  inference(superposition,[],[f893,f807])).
fof(f893,plain,(
  ( ! [X130,X131] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0')),X130,X131)) )),
  inference(resolution,[],[f565,f639])).
fof(f639,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f467])).
fof(f1912,plain,(
  ( ! [X0,X1] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0,X1)) ) | $spl74),
  inference(cnf_transformation,[],[f1912_D])).
fof(f1912_D,plain,(
  ( ! [X0,X1] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0,X1)) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f1917,plain,(
  ( ! [X10,X8,X9] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X8,X9) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X10)))) ) | $spl15),
  inference(subsumption_resolution,[],[f1900,f976])).
fof(f1900,plain,(
  ( ! [X10,X8,X9] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X8,X9) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X10)))) )),
  inference(superposition,[],[f893,f807])).
fof(f1916,plain,(
  ( ! [X6,X7] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X6,X7) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))) ) | $spl19),
  inference(subsumption_resolution,[],[f1899,f981])).
fof(f1899,plain,(
  ( ! [X6,X7] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X6,X7) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f893,f808])).
fof(f1915,plain,(
  ( ! [X4,X5] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X4,X5) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))) ) | $spl15),
  inference(subsumption_resolution,[],[f1898,f976])).
fof(f1898,plain,(
  ( ! [X4,X5] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X4,X5) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f893,f808])).
fof(f1913,plain,(
  ( ! [X2,X3] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X2,X3) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))) ) | $spl19),
  inference(subsumption_resolution,[],[f1897,f981])).
fof(f1897,plain,(
  ( ! [X2,X3] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X2,X3) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f893,f809])).
fof(f1911,plain,(
  ( ! [X0,X1] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0,X1) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))) ) | $spl15),
  inference(subsumption_resolution,[],[f1896,f976])).
fof(f1896,plain,(
  ( ! [X0,X1] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0,X1) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f893,f809])).
fof(f1887,plain,(
  ( ! [X2,X3] : (sP7(X2,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X3)) ) | $spl72),
  inference(cnf_transformation,[],[f1887_D])).
fof(f1887_D,plain,(
  ( ! [X2,X3] : (sP7(X2,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X3)) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f1891,plain,(
  ( ! [X12,X13,X11] : (sP7(X11,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13)))) ) | $spl19),
  inference(subsumption_resolution,[],[f1874,f981])).
fof(f1874,plain,(
  ( ! [X12,X13,X11] : (sP7(X11,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X12) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13)))) )),
  inference(superposition,[],[f892,f807])).
fof(f892,plain,(
  ( ! [X128,X129] : (sP7(X128,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0')),X129)) )),
  inference(resolution,[],[f565,f638])).
fof(f638,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f467])).
fof(f1885,plain,(
  ( ! [X0,X1] : (sP7(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X1)) ) | $spl70),
  inference(cnf_transformation,[],[f1885_D])).
fof(f1885_D,plain,(
  ( ! [X0,X1] : (sP7(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X1)) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f1890,plain,(
  ( ! [X10,X8,X9] : (sP7(X8,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X9) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X10)))) ) | $spl15),
  inference(subsumption_resolution,[],[f1873,f976])).
fof(f1873,plain,(
  ( ! [X10,X8,X9] : (sP7(X8,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X9) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X10)))) )),
  inference(superposition,[],[f892,f807])).
fof(f1889,plain,(
  ( ! [X6,X7] : (sP7(X6,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X7) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))) ) | $spl19),
  inference(subsumption_resolution,[],[f1872,f981])).
fof(f1872,plain,(
  ( ! [X6,X7] : (sP7(X6,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X7) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f892,f808])).
fof(f1888,plain,(
  ( ! [X4,X5] : (sP7(X4,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X5) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))) ) | $spl15),
  inference(subsumption_resolution,[],[f1871,f976])).
fof(f1871,plain,(
  ( ! [X4,X5] : (sP7(X4,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X5) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f892,f808])).
fof(f1886,plain,(
  ( ! [X2,X3] : (sP7(X2,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X3) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))) ) | $spl19),
  inference(subsumption_resolution,[],[f1870,f981])).
fof(f1870,plain,(
  ( ! [X2,X3] : (sP7(X2,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X3) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f892,f809])).
fof(f1884,plain,(
  ( ! [X0,X1] : (sP7(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X1) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))) ) | $spl15),
  inference(subsumption_resolution,[],[f1869,f976])).
fof(f1869,plain,(
  ( ! [X0,X1] : (sP7(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X1) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f892,f809])).
fof(f1860,plain,(
  ( ! [X2,X3] : (sP8('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X2,X3)) ) | $spl68),
  inference(cnf_transformation,[],[f1860_D])).
fof(f1860_D,plain,(
  ( ! [X2,X3] : (sP8('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X2,X3)) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f1864,plain,(
  ( ! [X12,X13,X11] : (sP8('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X11,X12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13)))) ) | $spl19),
  inference(subsumption_resolution,[],[f1847,f981])).
fof(f1847,plain,(
  ( ! [X12,X13,X11] : (sP8('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X11,X12) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13)))) )),
  inference(superposition,[],[f890,f807])).
fof(f890,plain,(
  ( ! [X125,X124] : (sP8('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0')),X124,X125)) )),
  inference(resolution,[],[f565,f635])).
fof(f635,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f464])).
fof(f1858,plain,(
  ( ! [X0,X1] : (sP8('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0,X1)) ) | $spl66),
  inference(cnf_transformation,[],[f1858_D])).
fof(f1858_D,plain,(
  ( ! [X0,X1] : (sP8('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0,X1)) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f1863,plain,(
  ( ! [X10,X8,X9] : (sP8('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X8,X9) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X10)))) ) | $spl15),
  inference(subsumption_resolution,[],[f1846,f976])).
fof(f1846,plain,(
  ( ! [X10,X8,X9] : (sP8('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X8,X9) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X10)))) )),
  inference(superposition,[],[f890,f807])).
fof(f1862,plain,(
  ( ! [X6,X7] : (sP8('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X6,X7) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))) ) | $spl19),
  inference(subsumption_resolution,[],[f1845,f981])).
fof(f1845,plain,(
  ( ! [X6,X7] : (sP8('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X6,X7) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f890,f808])).
fof(f1861,plain,(
  ( ! [X4,X5] : (sP8('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X4,X5) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))) ) | $spl15),
  inference(subsumption_resolution,[],[f1844,f976])).
fof(f1844,plain,(
  ( ! [X4,X5] : (sP8('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X4,X5) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f890,f808])).
fof(f1859,plain,(
  ( ! [X2,X3] : (sP8('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X2,X3) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))) ) | $spl19),
  inference(subsumption_resolution,[],[f1843,f981])).
fof(f1843,plain,(
  ( ! [X2,X3] : (sP8('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X2,X3) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f890,f809])).
fof(f1857,plain,(
  ( ! [X0,X1] : (sP8('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0,X1) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))) ) | $spl15),
  inference(subsumption_resolution,[],[f1842,f976])).
fof(f1842,plain,(
  ( ! [X0,X1] : (sP8('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0,X1) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f890,f809])).
fof(f1833,plain,(
  ( ! [X2,X3] : (sP8(X2,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X3)) ) | $spl64),
  inference(cnf_transformation,[],[f1833_D])).
fof(f1833_D,plain,(
  ( ! [X2,X3] : (sP8(X2,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X3)) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f1837,plain,(
  ( ! [X12,X13,X11] : (sP8(X11,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13)))) ) | $spl19),
  inference(subsumption_resolution,[],[f1821,f981])).
fof(f1821,plain,(
  ( ! [X12,X13,X11] : (sP8(X11,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X12) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13)))) )),
  inference(superposition,[],[f889,f807])).
fof(f889,plain,(
  ( ! [X123,X122] : (sP8(X122,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0')),X123)) )),
  inference(resolution,[],[f565,f634])).
fof(f634,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f464])).
fof(f1831,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X1)) ) | $spl62),
  inference(cnf_transformation,[],[f1831_D])).
fof(f1831_D,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X1)) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f1836,plain,(
  ( ! [X10,X8,X9] : (sP8(X8,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X9) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X10)))) ) | $spl15),
  inference(subsumption_resolution,[],[f1820,f976])).
fof(f1820,plain,(
  ( ! [X10,X8,X9] : (sP8(X8,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X9) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X10)))) )),
  inference(superposition,[],[f889,f807])).
fof(f1835,plain,(
  ( ! [X6,X7] : (sP8(X6,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X7) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))) ) | $spl19),
  inference(subsumption_resolution,[],[f1819,f981])).
fof(f1819,plain,(
  ( ! [X6,X7] : (sP8(X6,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X7) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f889,f808])).
fof(f1834,plain,(
  ( ! [X4,X5] : (sP8(X4,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X5) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))) ) | $spl15),
  inference(subsumption_resolution,[],[f1818,f976])).
fof(f1818,plain,(
  ( ! [X4,X5] : (sP8(X4,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X5) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f889,f808])).
fof(f1832,plain,(
  ( ! [X2,X3] : (sP8(X2,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X3) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))) ) | $spl19),
  inference(subsumption_resolution,[],[f1817,f981])).
fof(f1817,plain,(
  ( ! [X2,X3] : (sP8(X2,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X3) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f889,f809])).
fof(f1830,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X1) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))) ) | $spl15),
  inference(subsumption_resolution,[],[f1816,f976])).
fof(f1816,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X1) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f889,f809])).
fof(f1807,plain,(
  ( ! [X2,X3] : (sP4('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X2,X3)) ) | $spl60),
  inference(cnf_transformation,[],[f1807_D])).
fof(f1807_D,plain,(
  ( ! [X2,X3] : (sP4('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X2,X3)) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f1811,plain,(
  ( ! [X12,X13,X11] : (sP4('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X11,X12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13)))) ) | $spl19),
  inference(subsumption_resolution,[],[f1796,f981])).
fof(f1796,plain,(
  ( ! [X12,X13,X11] : (sP4('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X11,X12) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13)))) )),
  inference(superposition,[],[f877,f807])).
fof(f877,plain,(
  ( ! [X99,X98] : (sP4('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0')),X98,X99)) )),
  inference(resolution,[],[f565,f618])).
fof(f618,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f456])).
fof(f1805,plain,(
  ( ! [X0,X1] : (sP4('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0,X1)) ) | $spl58),
  inference(cnf_transformation,[],[f1805_D])).
fof(f1805_D,plain,(
  ( ! [X0,X1] : (sP4('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0,X1)) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f1810,plain,(
  ( ! [X10,X8,X9] : (sP4('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X8,X9) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X10)))) ) | $spl15),
  inference(subsumption_resolution,[],[f1795,f976])).
fof(f1795,plain,(
  ( ! [X10,X8,X9] : (sP4('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X8,X9) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X10)))) )),
  inference(superposition,[],[f877,f807])).
fof(f1809,plain,(
  ( ! [X6,X7] : (sP4('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X6,X7) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))) ) | $spl19),
  inference(subsumption_resolution,[],[f1794,f981])).
fof(f1794,plain,(
  ( ! [X6,X7] : (sP4('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X6,X7) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f877,f808])).
fof(f1808,plain,(
  ( ! [X4,X5] : (sP4('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X4,X5) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))) ) | $spl15),
  inference(subsumption_resolution,[],[f1793,f976])).
fof(f1793,plain,(
  ( ! [X4,X5] : (sP4('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X4,X5) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f877,f808])).
fof(f1806,plain,(
  ( ! [X2,X3] : (sP4('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X2,X3) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))) ) | $spl19),
  inference(subsumption_resolution,[],[f1792,f981])).
fof(f1792,plain,(
  ( ! [X2,X3] : (sP4('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X2,X3) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f877,f809])).
fof(f1804,plain,(
  ( ! [X0,X1] : (sP4('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0,X1) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))) ) | $spl15),
  inference(subsumption_resolution,[],[f1791,f976])).
fof(f1791,plain,(
  ( ! [X0,X1] : (sP4('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0,X1) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f877,f809])).
fof(f1781,plain,(
  ( ! [X2,X3] : (sP4(X2,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X3)) ) | $spl56),
  inference(cnf_transformation,[],[f1781_D])).
fof(f1781_D,plain,(
  ( ! [X2,X3] : (sP4(X2,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X3)) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f1785,plain,(
  ( ! [X12,X13,X11] : (sP4(X11,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13)))) ) | $spl19),
  inference(subsumption_resolution,[],[f1770,f981])).
fof(f1770,plain,(
  ( ! [X12,X13,X11] : (sP4(X11,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X12) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13)))) )),
  inference(superposition,[],[f876,f807])).
fof(f876,plain,(
  ( ! [X97,X96] : (sP4(X96,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0')),X97)) )),
  inference(resolution,[],[f565,f617])).
fof(f617,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f456])).
fof(f1779,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X1)) ) | $spl54),
  inference(cnf_transformation,[],[f1779_D])).
fof(f1779_D,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X1)) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f1784,plain,(
  ( ! [X10,X8,X9] : (sP4(X8,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X9) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X10)))) ) | $spl15),
  inference(subsumption_resolution,[],[f1769,f976])).
fof(f1769,plain,(
  ( ! [X10,X8,X9] : (sP4(X8,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X9) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X10)))) )),
  inference(superposition,[],[f876,f807])).
fof(f1783,plain,(
  ( ! [X6,X7] : (sP4(X6,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X7) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))) ) | $spl19),
  inference(subsumption_resolution,[],[f1768,f981])).
fof(f1768,plain,(
  ( ! [X6,X7] : (sP4(X6,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X7) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f876,f808])).
fof(f1782,plain,(
  ( ! [X4,X5] : (sP4(X4,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X5) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))) ) | $spl15),
  inference(subsumption_resolution,[],[f1767,f976])).
fof(f1767,plain,(
  ( ! [X4,X5] : (sP4(X4,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X5) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f876,f808])).
fof(f1780,plain,(
  ( ! [X2,X3] : (sP4(X2,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X3) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))) ) | $spl19),
  inference(subsumption_resolution,[],[f1766,f981])).
fof(f1766,plain,(
  ( ! [X2,X3] : (sP4(X2,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X3) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f876,f809])).
fof(f1778,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X1) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))) ) | $spl15),
  inference(subsumption_resolution,[],[f1765,f976])).
fof(f1765,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X1) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f876,f809])).
fof(f1755,plain,(
  ( ! [X2,X3] : (sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X2,X3)) ) | $spl52),
  inference(cnf_transformation,[],[f1755_D])).
fof(f1755_D,plain,(
  ( ! [X2,X3] : (sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X2,X3)) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f1759,plain,(
  ( ! [X12,X13,X11] : (sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X11,X12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13)))) ) | $spl19),
  inference(subsumption_resolution,[],[f1743,f981])).
fof(f1743,plain,(
  ( ! [X12,X13,X11] : (sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X11,X12) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13)))) )),
  inference(superposition,[],[f863,f807])).
fof(f863,plain,(
  ( ! [X70,X71] : (sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0')),X70,X71)) )),
  inference(resolution,[],[f565,f601])).
fof(f601,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f449])).
fof(f1753,plain,(
  ( ! [X0,X1] : (sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0,X1)) ) | $spl50),
  inference(cnf_transformation,[],[f1753_D])).
fof(f1753_D,plain,(
  ( ! [X0,X1] : (sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0,X1)) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f1758,plain,(
  ( ! [X10,X8,X9] : (sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X8,X9) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X10)))) ) | $spl15),
  inference(subsumption_resolution,[],[f1742,f976])).
fof(f1742,plain,(
  ( ! [X10,X8,X9] : (sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X8,X9) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X10)))) )),
  inference(superposition,[],[f863,f807])).
fof(f1757,plain,(
  ( ! [X6,X7] : (sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X6,X7) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))) ) | $spl19),
  inference(subsumption_resolution,[],[f1741,f981])).
fof(f1741,plain,(
  ( ! [X6,X7] : (sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X6,X7) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f863,f808])).
fof(f1756,plain,(
  ( ! [X4,X5] : (sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X4,X5) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))) ) | $spl15),
  inference(subsumption_resolution,[],[f1740,f976])).
fof(f1740,plain,(
  ( ! [X4,X5] : (sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X4,X5) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f863,f808])).
fof(f1754,plain,(
  ( ! [X2,X3] : (sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X2,X3) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))) ) | $spl19),
  inference(subsumption_resolution,[],[f1739,f981])).
fof(f1739,plain,(
  ( ! [X2,X3] : (sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X2,X3) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f863,f809])).
fof(f1752,plain,(
  ( ! [X0,X1] : (sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0,X1) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))) ) | $spl15),
  inference(subsumption_resolution,[],[f1738,f976])).
fof(f1738,plain,(
  ( ! [X0,X1] : (sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0,X1) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f863,f809])).
fof(f1728,plain,(
  ( ! [X2,X3] : (sP0('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X2,X3)) ) | $spl48),
  inference(cnf_transformation,[],[f1728_D])).
fof(f1728_D,plain,(
  ( ! [X2,X3] : (sP0('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X2,X3)) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f1732,plain,(
  ( ! [X12,X13,X11] : (sP0('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X11,X12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13)))) ) | $spl19),
  inference(subsumption_resolution,[],[f1717,f981])).
fof(f1717,plain,(
  ( ! [X12,X13,X11] : (sP0('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X11,X12) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13)))) )),
  inference(superposition,[],[f847,f807])).
fof(f847,plain,(
  ( ! [X39,X38] : (sP0('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0')),X38,X39)) )),
  inference(resolution,[],[f565,f584])).
fof(f584,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f441])).
fof(f1726,plain,(
  ( ! [X0,X1] : (sP0('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0,X1)) ) | $spl46),
  inference(cnf_transformation,[],[f1726_D])).
fof(f1726_D,plain,(
  ( ! [X0,X1] : (sP0('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0,X1)) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f1731,plain,(
  ( ! [X10,X8,X9] : (sP0('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X8,X9) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X10)))) ) | $spl15),
  inference(subsumption_resolution,[],[f1716,f976])).
fof(f1716,plain,(
  ( ! [X10,X8,X9] : (sP0('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X8,X9) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X10)))) )),
  inference(superposition,[],[f847,f807])).
fof(f1730,plain,(
  ( ! [X6,X7] : (sP0('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X6,X7) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))) ) | $spl19),
  inference(subsumption_resolution,[],[f1715,f981])).
fof(f1715,plain,(
  ( ! [X6,X7] : (sP0('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X6,X7) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f847,f808])).
fof(f1729,plain,(
  ( ! [X4,X5] : (sP0('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X4,X5) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))) ) | $spl15),
  inference(subsumption_resolution,[],[f1714,f976])).
fof(f1714,plain,(
  ( ! [X4,X5] : (sP0('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X4,X5) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f847,f808])).
fof(f1727,plain,(
  ( ! [X2,X3] : (sP0('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X2,X3) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))) ) | $spl19),
  inference(subsumption_resolution,[],[f1713,f981])).
fof(f1713,plain,(
  ( ! [X2,X3] : (sP0('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X2,X3) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f847,f809])).
fof(f1725,plain,(
  ( ! [X0,X1] : (sP0('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0,X1) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))) ) | $spl15),
  inference(subsumption_resolution,[],[f1712,f976])).
fof(f1712,plain,(
  ( ! [X0,X1] : (sP0('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0,X1) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f847,f809])).
fof(f1702,plain,(
  ( ! [X2,X3] : (sP0(X2,X3,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl44),
  inference(cnf_transformation,[],[f1702_D])).
fof(f1702_D,plain,(
  ( ! [X2,X3] : (sP0(X2,X3,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f1706,plain,(
  ( ! [X12,X13,X11] : (sP0(X11,X12,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13)))) ) | $spl19),
  inference(subsumption_resolution,[],[f1691,f981])).
fof(f1691,plain,(
  ( ! [X12,X13,X11] : (sP0(X11,X12,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X13)))) )),
  inference(superposition,[],[f846,f807])).
fof(f846,plain,(
  ( ! [X37,X36] : (sP0(X36,X37,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0')))) )),
  inference(resolution,[],[f565,f582])).
fof(f582,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f441])).
fof(f1700,plain,(
  ( ! [X0,X1] : (sP0(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl42),
  inference(cnf_transformation,[],[f1700_D])).
fof(f1700_D,plain,(
  ( ! [X0,X1] : (sP0(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f1705,plain,(
  ( ! [X10,X8,X9] : (sP0(X8,X9,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X10)))) ) | $spl15),
  inference(subsumption_resolution,[],[f1690,f976])).
fof(f1690,plain,(
  ( ! [X10,X8,X9] : (sP0(X8,X9,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X10)))) )),
  inference(superposition,[],[f846,f807])).
fof(f1704,plain,(
  ( ! [X6,X7] : (sP0(X6,X7,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))) ) | $spl19),
  inference(subsumption_resolution,[],[f1689,f981])).
fof(f1689,plain,(
  ( ! [X6,X7] : (sP0(X6,X7,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f846,f808])).
fof(f1703,plain,(
  ( ! [X4,X5] : (sP0(X4,X5,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))) ) | $spl15),
  inference(subsumption_resolution,[],[f1688,f976])).
fof(f1688,plain,(
  ( ! [X4,X5] : (sP0(X4,X5,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f846,f808])).
fof(f1701,plain,(
  ( ! [X2,X3] : (sP0(X2,X3,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))) ) | $spl19),
  inference(subsumption_resolution,[],[f1687,f981])).
fof(f1687,plain,(
  ( ! [X2,X3] : (sP0(X2,X3,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f846,f809])).
fof(f1699,plain,(
  ( ! [X0,X1] : (sP0(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))) ) | $spl15),
  inference(subsumption_resolution,[],[f1686,f976])).
fof(f1686,plain,(
  ( ! [X0,X1] : (sP0(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f846,f809])).
fof(f1676,plain,(
  ( ! [X1] : (sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X1)) ) | $spl40),
  inference(cnf_transformation,[],[f1676_D])).
fof(f1676_D,plain,(
  ( ! [X1] : (sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X1)) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f1680,plain,(
  ( ! [X6,X7] : (sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X6) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X7)))) ) | $spl19),
  inference(subsumption_resolution,[],[f1670,f981])).
fof(f1670,plain,(
  ( ! [X6,X7] : (sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X6) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X7)))) )),
  inference(superposition,[],[f1565,f807])).
fof(f1565,plain,(
  ( ! [X151] : (sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)),X151)) )),
  inference(resolution,[],[f831,f672])).
fof(f672,plain,(
  ( ! [X0,X1] : (sP11(X0,X1) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f480])).
fof(f480,plain,(
  ! [X0,X1] : ((~sP11(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP11(X0,X1)))),
  inference(flattening,[],[f479])).
fof(f479,plain,(
  ! [X0,X1] : ((~sP11(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP11(X0,X1)))),
  inference(nnf_transformation,[],[f431])).
fof(f431,plain,(
  ! [X0,X1] : (sP11(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f1674,plain,(
  ( ! [X0] : (sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X0)) ) | $spl38),
  inference(cnf_transformation,[],[f1674_D])).
fof(f1674_D,plain,(
  ( ! [X0] : (sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X0)) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f1679,plain,(
  ( ! [X4,X5] : (sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X4) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl15),
  inference(subsumption_resolution,[],[f1669,f976])).
fof(f1669,plain,(
  ( ! [X4,X5] : (sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X4) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) )),
  inference(superposition,[],[f1565,f807])).
fof(f1678,plain,(
  ( ! [X3] : (sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X3) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))) ) | $spl19),
  inference(subsumption_resolution,[],[f1668,f981])).
fof(f1668,plain,(
  ( ! [X3] : (sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X3) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f1565,f808])).
fof(f1677,plain,(
  ( ! [X2] : (sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X2) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))) ) | $spl15),
  inference(subsumption_resolution,[],[f1667,f976])).
fof(f1667,plain,(
  ( ! [X2] : (sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X2) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f1565,f808])).
fof(f1675,plain,(
  ( ! [X1] : (sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X1) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))) ) | $spl19),
  inference(subsumption_resolution,[],[f1666,f981])).
fof(f1666,plain,(
  ( ! [X1] : (sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X1) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f1565,f809])).
fof(f1673,plain,(
  ( ! [X0] : (sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X0) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))) ) | $spl15),
  inference(subsumption_resolution,[],[f1665,f976])).
fof(f1665,plain,(
  ( ! [X0] : (sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')),X0) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f1565,f809])).
fof(f1652,plain,(
  ( ! [X1] : (sP11(X1,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl36),
  inference(cnf_transformation,[],[f1652_D])).
fof(f1652_D,plain,(
  ( ! [X1] : (sP11(X1,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f1656,plain,(
  ( ! [X6,X7] : (sP11(X6,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X7)))) ) | $spl19),
  inference(subsumption_resolution,[],[f1646,f981])).
fof(f1646,plain,(
  ( ! [X6,X7] : (sP11(X6,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X7)))) )),
  inference(superposition,[],[f1564,f807])).
fof(f1564,plain,(
  ( ! [X150] : (sP11(X150,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)))) )),
  inference(resolution,[],[f831,f671])).
fof(f671,plain,(
  ( ! [X0,X1] : (sP11(X0,X1) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f480])).
fof(f1650,plain,(
  ( ! [X0] : (sP11(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) | $spl34),
  inference(cnf_transformation,[],[f1650_D])).
fof(f1650_D,plain,(
  ( ! [X0] : (sP11(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f1655,plain,(
  ( ! [X4,X5] : (sP11(X4,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) ) | $spl15),
  inference(subsumption_resolution,[],[f1645,f976])).
fof(f1645,plain,(
  ( ! [X4,X5] : (sP11(X4,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5)))) )),
  inference(superposition,[],[f1564,f807])).
fof(f1654,plain,(
  ( ! [X3] : (sP11(X3,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))) ) | $spl19),
  inference(subsumption_resolution,[],[f1644,f981])).
fof(f1644,plain,(
  ( ! [X3] : (sP11(X3,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f1564,f808])).
fof(f1653,plain,(
  ( ! [X2] : (sP11(X2,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12))))) ) | $spl15),
  inference(subsumption_resolution,[],[f1643,f976])).
fof(f1643,plain,(
  ( ! [X2] : (sP11(X2,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f1564,f808])).
fof(f1651,plain,(
  ( ! [X1] : (sP11(X1,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))) ) | $spl19),
  inference(subsumption_resolution,[],[f1642,f981])).
fof(f1642,plain,(
  ( ! [X1] : (sP11(X1,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f1564,f809])).
fof(f1649,plain,(
  ( ! [X0] : (sP11(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12))))) ) | $spl15),
  inference(subsumption_resolution,[],[f1641,f976])).
fof(f1641,plain,(
  ( ! [X0] : (sP11(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)) )),
  inference(superposition,[],[f1564,f809])).
fof(f1625,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl32),
  inference(cnf_transformation,[],[f1625_D])).
fof(f1625_D,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f1629,plain,(
  ( ! [X1] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | $spl19),
  inference(subsumption_resolution,[],[f1599,f981])).
fof(f1599,plain,(
  ( ! [X1] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) )),
  inference(superposition,[],[f831,f807])).
fof(f1628,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))) ) | $spl15),
  inference(subsumption_resolution,[],[f1598,f976])).
fof(f1598,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))) )),
  inference(superposition,[],[f831,f807])).
fof(f1627,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) | $spl19),
  inference(subsumption_resolution,[],[f1597,f981])).
fof(f1597,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)),
  inference(superposition,[],[f831,f808])).
fof(f1626,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) | $spl15),
  inference(subsumption_resolution,[],[f1596,f976])).
fof(f1596,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)),
  inference(superposition,[],[f831,f808])).
fof(f1623,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) | $spl19),
  inference(subsumption_resolution,[],[f1595,f981])).
fof(f1595,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)),
  inference(superposition,[],[f831,f809])).
fof(f1620,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) | $spl15),
  inference(subsumption_resolution,[],[f1594,f976])).
fof(f1594,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)),
  inference(superposition,[],[f831,f809])).
fof(f1616,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)))) ) | $spl26),
  inference(cnf_transformation,[],[f1616_D])).
fof(f1616_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f1585,plain,(
  ( ! [X163,X162] : (p(s(bool,X162)) | p(s(bool,X163)) | s(bool,X163) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)))) )),
  inference(resolution,[],[f831,f823])).
fof(f1496,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12))) | s(bool,X22) != s(bool,X23)) )),
  inference(resolution,[],[f831,f573])).
fof(f1617,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)))) ) | $spl28),
  inference(cnf_transformation,[],[f1617_D])).
fof(f1617_D,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f1495,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,X19) | s(bool,X20) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)))) )),
  inference(resolution,[],[f831,f573])).
fof(f1493,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)))) )),
  inference(resolution,[],[f831,f572])).
fof(f1492,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)))) )),
  inference(resolution,[],[f831,f571])).
fof(f1491,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12))) | s(bool,X7) != s(bool,X8)) )),
  inference(resolution,[],[f831,f569])).
fof(f1490,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X4) | s(bool,X5) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)))) )),
  inference(resolution,[],[f831,f569])).
fof(f1489,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK12)))) )),
  inference(resolution,[],[f831,f567])).
fof(f1484,plain,(
  ( ! [X1] : (s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | $spl19),
  inference(subsumption_resolution,[],[f1483,f981])).
fof(f1483,plain,(
  ( ! [X1] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) )),
  inference(subsumption_resolution,[],[f1444,f796])).
fof(f796,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0))))) )),
  inference(cnf_transformation,[],[f349])).
fof(f349,plain,(
  ! [X0] : p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0))))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X6] : p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X6))))),
  file('thm.extreal.le_infty|split|1','thm.extreal.le_refl')).
fof(f1444,plain,(
  ( ! [X1] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) )),
  inference(superposition,[],[f565,f807])).
fof(f1482,plain,(
  ( ! [X0] : (s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))) ) | $spl15),
  inference(subsumption_resolution,[],[f1481,f976])).
fof(f1481,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))) )),
  inference(subsumption_resolution,[],[f1443,f797])).
fof(f1443,plain,(
  ( ! [X0] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))) )),
  inference(superposition,[],[f565,f807])).
fof(f963,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl0),
  inference(cnf_transformation,[],[f963_D])).
fof(f963_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f1436,plain,(
  ( ! [X169,X168] : (p(s(bool,X168)) | ~p(s(bool,X169)) | s(bool,X169) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0')))) )),
  inference(resolution,[],[f565,f823])).
fof(f1344,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s(bool,X21) != s(bool,X23)) )),
  inference(resolution,[],[f565,f573])).
fof(f1342,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0')))) )),
  inference(resolution,[],[f565,f572])).
fof(f965,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl4),
  inference(cnf_transformation,[],[f965_D])).
fof(f965_D,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f1341,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s(bool,X12) != s(bool,X13)) )),
  inference(resolution,[],[f565,f571])).
fof(f1340,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0')))) )),
  inference(resolution,[],[f565,f571])).
fof(f1339,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s(bool,X6) != s(bool,X8)) )),
  inference(resolution,[],[f565,f569])).
fof(f1338,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s(bool,X3) != s(bool,X4)) )),
  inference(resolution,[],[f565,f567])).
fof(f1337,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0')))) )),
  inference(resolution,[],[f565,f567])).
fof(f1335,plain,(
  ( ! [X0] : (s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))) ) | ($spl15 | $spl19)),
  inference(subsumption_resolution,[],[f1330,f976])).
fof(f1330,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))) ) | $spl19),
  inference(trivial_inequality_removal,[],[f1328])).
fof(f1328,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.NegInf_0') | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))) ) | $spl19),
  inference(superposition,[],[f981,f807])).
fof(f1323,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f1322,f797])).
fof(f1322,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl18),
  inference(forward_demodulation,[],[f565,f982])).
fof(f982,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | $spl18),
  inference(cnf_transformation,[],[f982_D])).
fof(f982_D,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f1273,plain,(
  ( ! [X1] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | $spl15),
  inference(trivial_inequality_removal,[],[f1272])).
fof(f1272,plain,(
  ( ! [X1] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.PosInf_0') | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) ) | $spl15),
  inference(superposition,[],[f976,f807])).
fof(f1274,plain,(
  s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | $spl15),
  inference(trivial_inequality_removal,[],[f1270])).
fof(f1270,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.PosInf_0') | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | $spl15),
  inference(superposition,[],[f976,f808])).
fof(f1275,plain,(
  s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | $spl15),
  inference(trivial_inequality_removal,[],[f1268])).
fof(f1268,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.PosInf_0') | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | $spl15),
  inference(superposition,[],[f976,f809])).
fof(f1258,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f1220,f796])).
fof(f1220,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl14),
  inference(backward_demodulation,[],[f977,f565])).
fof(f1218,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f1154,f969])).
fof(f1154,plain,(
  ( ! [X389,X390] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X389),s('type.extreal.extreal',X390))))) ) | $spl12),
  inference(resolution,[],[f969,f789])).
fof(f1217,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f1111,f969])).
fof(f1111,plain,(
  ( ! [X273,X272] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X272),s('type.extreal.extreal',X273))))) ) | $spl12),
  inference(resolution,[],[f969,f789])).
fof(f1216,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f1110,f969])).
fof(f1110,plain,(
  ~p(s(bool,'T_0')) | $spl12),
  inference(resolution,[],[f969,f786])).
fof(f786,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f555])).
fof(f555,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f341])).
fof(f341,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f42])).
fof(f42,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.extreal.le_infty|split|1','thm.bool.NOT_CLAUSES|split|1')).
fof(f1215,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f1109,f969])).
fof(f1109,plain,(
  ~p(s(bool,'T_0')) | $spl12),
  inference(resolution,[],[f969,f784])).
fof(f784,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f554])).
fof(f554,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f340])).
fof(f340,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f43])).
fof(f43,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.extreal.le_infty|split|1','thm.bool.NOT_CLAUSES|split|2')).
fof(f1214,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f1106,f969])).
fof(f1106,plain,(
  ( ! [X270] : (~p(s(bool,X270))) ) | $spl12),
  inference(resolution,[],[f969,f773])).
fof(f773,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f399])).
fof(f399,plain,(
  ! [X0] : ((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f328])).
fof(f328,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) => ~p(s(bool,X0)))),
  inference(flattening,[],[f327])).
fof(f327,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) => ~p(s(bool,X0)))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X22] : ((p(s(bool,X22)) => p(s(bool,'F_0'))) => ~p(s(bool,X22)))),
  file('thm.extreal.le_infty|split|1','thm.bool.IMP_F')).
fof(f1213,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f1099,f969])).
fof(f1099,plain,(
  ( ! [X263] : (~p(s(bool,X263))) ) | $spl12),
  inference(resolution,[],[f969,f712])).
fof(f712,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f509])).
fof(f509,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f508])).
fof(f508,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f389])).
fof(f389,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f301])).
fof(f301,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f300])).
fof(f300,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f119])).
fof(f119,axiom,(
  ! [X22] : ((p(s(bool,X22)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X22)))),
  file('thm.extreal.le_infty|split|1','thm.bool.IMP_CLAUSES|split|4')).
fof(f1108,plain,(
  $false | $spl12),
  inference(resolution,[],[f969,f782])).
fof(f782,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f339])).
fof(f339,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.extreal.le_infty|split|1',aHL_FALSITY)).
fof(f1112,plain,(
  $false | $spl12),
  inference(resolution,[],[f969,f565])).
fof(f991,plain,(
  ( ! [X1] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) )),
  inference(subsumption_resolution,[],[f939,f796])).
fof(f939,plain,(
  ( ! [X1] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) )),
  inference(superposition,[],[f565,f807])).
fof(f977,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | $spl14),
  inference(cnf_transformation,[],[f977_D])).
fof(f977_D,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f987,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))) )),
  inference(subsumption_resolution,[],[f938,f797])).
fof(f938,plain,(
  ( ! [X0] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK45(sK12)))) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))) )),
  inference(superposition,[],[f565,f807])).
fof(f986,plain,(
  s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)),
  inference(subsumption_resolution,[],[f937,f796])).
fof(f937,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)),
  inference(superposition,[],[f565,f808])).
fof(f983,plain,(
  s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)),
  inference(subsumption_resolution,[],[f936,f797])).
fof(f936,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK46(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)),
  inference(superposition,[],[f565,f808])).
fof(f980,plain,(
  s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)),
  inference(subsumption_resolution,[],[f935,f796])).
fof(f935,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)),
  inference(superposition,[],[f565,f809])).
fof(f975,plain,(
  s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)),
  inference(subsumption_resolution,[],[f934,f797])).
fof(f934,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal',sK12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK47(sK12)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12)),
  inference(superposition,[],[f565,f809])).
fof(f931,plain,(
  ( ! [X169,X168] : (p(s(bool,X168)) | ~p(s(bool,X169)) | s(bool,X169) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0')))) )),
  inference(resolution,[],[f565,f823])).
fof(f839,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s(bool,X21) != s(bool,X23)) )),
  inference(resolution,[],[f565,f573])).
fof(f837,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0')))) )),
  inference(resolution,[],[f565,f572])).
fof(f836,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s(bool,X12) != s(bool,X13)) )),
  inference(resolution,[],[f565,f571])).
fof(f835,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0')))) )),
  inference(resolution,[],[f565,f571])).
fof(f834,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s(bool,X6) != s(bool,X8)) )),
  inference(resolution,[],[f565,f569])).
fof(f833,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s(bool,X3) != s(bool,X4)) )),
  inference(resolution,[],[f565,f567])).
fof(f832,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.PosInf_0')))) )),
  inference(resolution,[],[f565,f567])).
% SZS output end Proof for thm.extreal
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 296
% Passive clauses: 1483
% Generated clauses: 5871
% Final active clauses: 282
% Final passive clauses: 1210
% Input formulas: 261
% Initial clauses: 258
% 
% Selected by SInE selection: 85
% SInE iterations: 1
% 
% Duplicate literals: 84
% Trivial inequalities: 28
% Fw subsumption resolutions: 1104
% Fw demodulations: 315
% Bw demodulations: 4
% 
% Simple tautologies: 209
% Equational tautologies: 2
% Forward subsumptions: 2396
% 
% Binary resolution: 1290
% Forward superposition: 1688
% Backward superposition: 1246
% Equality resolution: 3
% 
% Split clauses: 226
% Split components: 504
% SAT solver clauses: 243
% SAT solver unit clauses: 14
% SAT solver binary clauses: 36
% 
% Sat splits: 226
% Sat splitting refutations: 15
% 
% Memory used [KB]: 3326
% Time elapsed: 0.096 s
% ------------------------------
% ------------------------------
% Success in time 11.849 s
