Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.nets
% (60247)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 57
% Passive clauses: 1818
% Generated clauses: 10496
% Final active clauses: 53
% Final passive clauses: 1706
% Input formulas: 11
% Initial clauses: 16
% 
% Duplicate literals: 803
% Trivial inequalities: 158
% Fw subsumption resolutions: 1274
% Fw demodulations: 432
% Bw demodulations: 3
% 
% Simple tautologies: 242
% Equational tautologies: 101
% Forward subsumptions: 4098
% 
% Binary resolution: 31
% Forward superposition: 4139
% Backward superposition: 2878
% Self superposition: 357
% Equality factoring: 290
% Equality resolution: 12
% 
% Split clauses: 540
% Split components: 1088
% SAT solver clauses: 1542
% SAT solver unit clauses: 251
% SAT solver binary clauses: 555
% 
% Sat splits: 540
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1538
% TWLsolver calls for satisfiability: 62
% 
% Memory used [KB]: 5884
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.nets
% (60271)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 56
% Passive clauses: 1048
% Generated clauses: 5376
% Final active clauses: 51
% Final passive clauses: 868
% Input formulas: 11
% Initial clauses: 16
% 
% Duplicate literals: 222
% Trivial inequalities: 93
% Fw subsumption resolutions: 154
% Fw demodulations: 342
% Bw demodulations: 64
% Condensations: 258
% 
% Simple tautologies: 78
% Equational tautologies: 17
% Forward subsumptions: 1450
% Bw demodulations to eq. taut.: 63
% 
% Binary resolution: 15
% Forward superposition: 1530
% Backward superposition: 1429
% Self superposition: 134
% Equality factoring: 44
% Equality resolution: 3
% 
% Split clauses: 73
% Split components: 152
% SAT solver clauses: 1659
% SAT solver unit clauses: 25
% SAT solver binary clauses: 58
% 
% Sat splits: 73
% Sat splitting refutations: 6
% 
% TWLsolver clauses: 1659
% TWLsolver calls for satisfiability: 135
% 
% Memory used [KB]: 15991
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.nets
% (60300)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 262
% Passive clauses: 5287
% Generated clauses: 85810
% Final active clauses: 41
% Final passive clauses: 138
% Input formulas: 11
% Initial clauses: 16
% Discarded non-redundant clauses: 1885
% 
% Duplicate literals: 6315
% Trivial inequalities: 534
% Fw subsumption resolutions: 4933
% Fw demodulations: 6219
% Bw demodulations: 3830
% Condensations: 11780
% Global subsumptions: 156
% 
% Simple tautologies: 893
% Equational tautologies: 2034
% Forward subsumptions: 46608
% Bw demodulations to eq. taut.: 26
% 
% Unit resulting resolution: 4679
% Forward superposition: 24848
% Backward superposition: 12348
% Self superposition: 551
% Equality factoring: 408
% Equality resolution: 37
% 
% Split clauses: 204
% Split components: 413
% SAT solver clauses: 6003
% SAT solver unit clauses: 107
% SAT solver binary clauses: 482
% 
% Sat splits: 204
% Sat splitting refutations: 150
% 
% Memory used [KB]: 40937
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.nets
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.nets
% SZS output start Proof for thm.nets
fof(f2741,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f64,f69,f67,f65,f73,f71,f122,f1103,f1107,f1106,f1497,f1504,f1498,f1508,f1506,f1509,f1510,f1511,f1512,f1513,f1514,f1515,f1516,f1517,f1518,f1519,f1520,f1521,f1522,f1523,f1507,f1524,f1527,f1546,f2729,f2739])).
fof(f2739,plain,(
  $false | ($spl3 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f2738,f111])).
fof(f111,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl3),
  inference(subsumption_resolution,[],[f109,f69])).
fof(f109,plain,(
  s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.nets.tends_3'(s(fun(sK0,'type.realax.real'),sK2),s('type.realax.real',sK3),s('type.pair.prod'('type.topology.topology'('type.realax.real'),fun(sK0,fun(sK0,bool))),'const.pair.,_2'(s('type.topology.topology'('type.realax.real'),'const.topology.mtop_1'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'))),s(fun(sK0,fun(sK0,bool)),sK1))))) | $spl3),
  inference(superposition,[],[f69,f40])).
fof(f40,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.nets.NET_NEG',aHL_BOOL_CASES)).
fof(f2738,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f2737,f1858])).
fof(f1858,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4(sK5(sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0),'const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))),s(sK0,sK4(sK5(sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0),'const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))) | ($spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1847,f47])).
fof(f47,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f23])).
fof(f23,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.nets.NET_NEG',aHL_FALSITY)).
fof(f1847,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4(sK5(sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0),'const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))),s(sK0,sK4(sK5(sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0),'const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))) | ($spl4 | $spl6)),
  inference(superposition,[],[f1830,f40])).
fof(f1830,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4(sK5(sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0),'const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))),s(sK0,sK4(sK5(sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0),'const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0))))) | ($spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1826,f48])).
fof(f48,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.nets.NET_NEG',aHL_TRUTH)).
fof(f1826,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4(sK5(sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0),'const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))),s(sK0,sK4(sK5(sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0),'const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0))))) | ($spl4 | $spl6)),
  inference(superposition,[],[f1670,f1589])).
fof(f1589,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK5(sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))) | $spl6),
  inference(subsumption_resolution,[],[f1588,f47])).
fof(f1588,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK5(sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))) | $spl6),
  inference(superposition,[],[f1572,f40])).
fof(f1572,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK5(sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0))))) | $spl6),
  inference(subsumption_resolution,[],[f1555,f47])).
fof(f1555,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK5(sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0))))) | $spl6),
  inference(superposition,[],[f44,f72])).
fof(f72,plain,(
  s(bool,'F_0') = s(bool,'const.nets.tends_3'(s(fun(sK0,'type.realax.real'),sK2),s('type.realax.real',sK3),s('type.pair.prod'('type.topology.topology'('type.realax.real'),fun(sK0,fun(sK0,bool))),'const.pair.,_2'(s('type.topology.topology'('type.realax.real'),'const.topology.mtop_1'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'))),s(fun(sK0,fun(sK0,bool)),sK1))))) | $spl6),
  inference(cnf_transformation,[],[f72_D])).
fof(f72_D,plain,(
  s(bool,'F_0') = s(bool,'const.nets.tends_3'(s(fun(sK0,'type.realax.real'),sK2),s('type.realax.real',sK3),s('type.pair.prod'('type.topology.topology'('type.realax.real'),fun(sK0,fun(sK0,bool))),'const.pair.,_2'(s('type.topology.topology'('type.realax.real'),'const.topology.mtop_1'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'))),s(fun(sK0,fun(sK0,bool)),sK1))))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f44,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.nets.tends_3'(s(fun(X0,X1),X4),s(X1,X5),s('type.pair.prod'('type.topology.topology'(X1),fun(X0,fun(X0,bool))),'const.pair.,_2'(s('type.topology.topology'(X1),'const.topology.mtop_1'(s('type.topology.metric'(X1),X2))),s(fun(X0,fun(X0,bool)),X3)))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK5(X5,X4,X3,X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f32])).
fof(f32,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((~p(s(bool,'const.nets.tends_3'(s(fun(X0,X1),X4),s(X1,X5),s('type.pair.prod'('type.topology.topology'(X1),fun(X0,fun(X0,bool))),'const.pair.,_2'(s('type.topology.topology'(X1),'const.topology.mtop_1'(s('type.topology.metric'(X1),X2))),s(fun(X0,fun(X0,bool)),X3)))))) | ! [X9] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X9)))) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,sK4(X9,X5,X4,X3,X2,X1,X0)))),s(X0,sK4(X9,X5,X4,X3,X2,X1,X0))))) & ! [X11] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,X11))),s(X0,sK4(X9,X5,X4,X3,X2,X1,X0))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.topology.dist_2'(s('type.topology.metric'(X1),X2),s('type.pair.prod'(X1,X1),'const.pair.,_2'(s(X1,happ(s(fun(X0,X1),X4),s(X0,X11))),s(X1,X5))))),s('type.realax.real',X9)))))))) & ((p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK5(X5,X4,X3,X2,X1,X0))))) & ! [X7] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,X7))),s(X0,X7)))) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,sK6(X7,X5,X4,X3,X2,X1,X0)))),s(X0,X7)))) & ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.topology.dist_2'(s('type.topology.metric'(X1),X2),s('type.pair.prod'(X1,X1),'const.pair.,_2'(s(X1,happ(s(fun(X0,X1),X4),s(X0,sK6(X7,X5,X4,X3,X2,X1,X0)))),s(X1,X5))))),s('type.realax.real',sK5(X5,X4,X3,X2,X1,X0)))))))) | p(s(bool,'const.nets.tends_3'(s(fun(X0,X1),X4),s(X1,X5),s('type.pair.prod'('type.topology.topology'(X1),fun(X0,fun(X0,bool))),'const.pair.,_2'(s('type.topology.topology'(X1),'const.topology.mtop_1'(s('type.topology.metric'(X1),X2))),s(fun(X0,fun(X0,bool)),X3))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4,sK5,sK6])],[f31])).
fof(f31,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((~p(s(bool,'const.nets.tends_3'(s(fun(X0,X1),X4),s(X1,X5),s('type.pair.prod'('type.topology.topology'(X1),fun(X0,fun(X0,bool))),'const.pair.,_2'(s('type.topology.topology'(X1),'const.topology.mtop_1'(s('type.topology.metric'(X1),X2))),s(fun(X0,fun(X0,bool)),X3)))))) | ! [X9] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X9)))) | ? [X10] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,X10))),s(X0,X10)))) & ! [X11] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,X11))),s(X0,X10)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.topology.dist_2'(s('type.topology.metric'(X1),X2),s('type.pair.prod'(X1,X1),'const.pair.,_2'(s(X1,happ(s(fun(X0,X1),X4),s(X0,X11))),s(X1,X5))))),s('type.realax.real',X9)))))))) & (? [X6] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X6)))) & ! [X7] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,X7))),s(X0,X7)))) | ? [X8] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,X8))),s(X0,X7)))) & ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.topology.dist_2'(s('type.topology.metric'(X1),X2),s('type.pair.prod'(X1,X1),'const.pair.,_2'(s(X1,happ(s(fun(X0,X1),X4),s(X0,X8))),s(X1,X5))))),s('type.realax.real',X6))))))) | p(s(bool,'const.nets.tends_3'(s(fun(X0,X1),X4),s(X1,X5),s('type.pair.prod'('type.topology.topology'(X1),fun(X0,fun(X0,bool))),'const.pair.,_2'(s('type.topology.topology'(X1),'const.topology.mtop_1'(s('type.topology.metric'(X1),X2))),s(fun(X0,fun(X0,bool)),X3))))))))),
  inference(rectify,[],[f30])).
fof(f30,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((~p(s(bool,'const.nets.tends_3'(s(fun(X0,X1),X4),s(X1,X5),s('type.pair.prod'('type.topology.topology'(X1),fun(X0,fun(X0,bool))),'const.pair.,_2'(s('type.topology.topology'(X1),'const.topology.mtop_1'(s('type.topology.metric'(X1),X2))),s(fun(X0,fun(X0,bool)),X3)))))) | ! [X6] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X6)))) | ? [X7] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,X7))),s(X0,X7)))) & ! [X8] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,X8))),s(X0,X7)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.topology.dist_2'(s('type.topology.metric'(X1),X2),s('type.pair.prod'(X1,X1),'const.pair.,_2'(s(X1,happ(s(fun(X0,X1),X4),s(X0,X8))),s(X1,X5))))),s('type.realax.real',X6)))))))) & (? [X6] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X6)))) & ! [X7] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,X7))),s(X0,X7)))) | ? [X8] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,X8))),s(X0,X7)))) & ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.topology.dist_2'(s('type.topology.metric'(X1),X2),s('type.pair.prod'(X1,X1),'const.pair.,_2'(s(X1,happ(s(fun(X0,X1),X4),s(X0,X8))),s(X1,X5))))),s('type.realax.real',X6))))))) | p(s(bool,'const.nets.tends_3'(s(fun(X0,X1),X4),s(X1,X5),s('type.pair.prod'('type.topology.topology'(X1),fun(X0,fun(X0,bool))),'const.pair.,_2'(s('type.topology.topology'(X1),'const.topology.mtop_1'(s('type.topology.metric'(X1),X2))),s(fun(X0,fun(X0,bool)),X3))))))))),
  inference(nnf_transformation,[],[f27])).
fof(f27,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (p(s(bool,'const.nets.tends_3'(s(fun(X0,X1),X4),s(X1,X5),s('type.pair.prod'('type.topology.topology'(X1),fun(X0,fun(X0,bool))),'const.pair.,_2'(s('type.topology.topology'(X1),'const.topology.mtop_1'(s('type.topology.metric'(X1),X2))),s(fun(X0,fun(X0,bool)),X3)))))) <=> ! [X6] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X6)))) | ? [X7] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,X7))),s(X0,X7)))) & ! [X8] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,X8))),s(X0,X7)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.topology.dist_2'(s('type.topology.metric'(X1),X2),s('type.pair.prod'(X1,X1),'const.pair.,_2'(s(X1,happ(s(fun(X0,X1),X4),s(X0,X8))),s(X1,X5))))),s('type.realax.real',X6))))))))),
  inference(ennf_transformation,[],[f22])).
fof(f22,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (p(s(bool,'const.nets.tends_3'(s(fun(X0,X1),X4),s(X1,X5),s('type.pair.prod'('type.topology.topology'(X1),fun(X0,fun(X0,bool))),'const.pair.,_2'(s('type.topology.topology'(X1),'const.topology.mtop_1'(s('type.topology.metric'(X1),X2))),s(fun(X0,fun(X0,bool)),X3)))))) <=> ! [X6] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X6)))) => ? [X7] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,X7))),s(X0,X7)))) & ! [X8] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,X8))),s(X0,X7)))) => p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.topology.dist_2'(s('type.topology.metric'(X1),X2),s('type.pair.prod'(X1,X1),'const.pair.,_2'(s(X1,happ(s(fun(X0,X1),X4),s(X0,X8))),s(X1,X5))))),s('type.realax.real',X6))))))))),
  inference(flattening,[],[f21])).
fof(f21,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5] : (p(s(bool,'const.nets.tends_3'(s(fun(X0,X1),X4),s(X1,X5),s('type.pair.prod'('type.topology.topology'(X1),fun(X0,fun(X0,bool))),'const.pair.,_2'(s('type.topology.topology'(X1),'const.topology.mtop_1'(s('type.topology.metric'(X1),X2))),s(fun(X0,fun(X0,bool)),X3)))))) <=> ! [X6] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X6)))) => ? [X7] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,X7))),s(X0,X7)))) & ! [X8] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,X8))),s(X0,X7)))) => p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.topology.dist_2'(s('type.topology.metric'(X1),X2),s('type.pair.prod'(X1,X1),'const.pair.,_2'(s(X1,happ(s(fun(X0,X1),X4),s(X0,X8))),s(X1,X5))))),s('type.realax.real',X6))))))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X10,X6] : ! [X11,X7,X8,X9] : (p(s(bool,'const.nets.tends_3'(s(fun(X10,X6),X8),s(X6,X9),s('type.pair.prod'('type.topology.topology'(X6),fun(X10,fun(X10,bool))),'const.pair.,_2'(s('type.topology.topology'(X6),'const.topology.mtop_1'(s('type.topology.metric'(X6),X11))),s(fun(X10,fun(X10,bool)),X7)))))) <=> ! [X12] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X12)))) => ? [X13] : (p(s(bool,happ(s(fun(X10,bool),happ(s(fun(X10,fun(X10,bool)),X7),s(X10,X13))),s(X10,X13)))) & ! [X14] : (p(s(bool,happ(s(fun(X10,bool),happ(s(fun(X10,fun(X10,bool)),X7),s(X10,X14))),s(X10,X13)))) => p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.topology.dist_2'(s('type.topology.metric'(X6),X11),s('type.pair.prod'(X6,X6),'const.pair.,_2'(s(X6,happ(s(fun(X10,X6),X8),s(X10,X14))),s(X6,X9))))),s('type.realax.real',X12))))))))),
  file('thm.nets.NET_NEG','thm.nets.MTOP_TENDS')).
fof(f1670,plain,(
  ( ! [X0] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4(X0,'const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))),s(sK0,sK4(X0,'const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))))) ) | $spl4),
  inference(subsumption_resolution,[],[f1651,f48])).
fof(f1651,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4(X0,'const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))),s(sK0,sK4(X0,'const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))))) ) | $spl4),
  inference(superposition,[],[f42,f71])).
fof(f42,plain,(
  ( ! [X4,X2,X0,X5,X3,X1,X9] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,sK4(X9,X5,X4,X3,X2,X1,X0)))),s(X0,sK4(X9,X5,X4,X3,X2,X1,X0))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X9)))) | ~p(s(bool,'const.nets.tends_3'(s(fun(X0,X1),X4),s(X1,X5),s('type.pair.prod'('type.topology.topology'(X1),fun(X0,fun(X0,bool))),'const.pair.,_2'(s('type.topology.topology'(X1),'const.topology.mtop_1'(s('type.topology.metric'(X1),X2))),s(fun(X0,fun(X0,bool)),X3))))))) )),
  inference(cnf_transformation,[],[f32])).
fof(f2737,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4(sK5(sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0),'const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))),s(sK0,sK4(sK5(sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0),'const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))) | ($spl4 | $spl6)),
  inference(subsumption_resolution,[],[f2736,f47])).
fof(f2736,plain,(
  p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4(sK5(sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0),'const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))),s(sK0,sK4(sK5(sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0),'const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f2735,f1832])).
fof(f1832,plain,(
  s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK3),s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,sK6(sK4(sK5(sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0),'const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0),sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))))))),s('type.realax.real',sK5(sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))) | ($spl4 | $spl6)),
  inference(resolution,[],[f1830,f1725])).
fof(f1725,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,X0)))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK3),s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,sK6(X0,sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))))))),s('type.realax.real',sK5(sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0))))) ) | $spl6),
  inference(subsumption_resolution,[],[f1723,f48])).
fof(f1723,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,X0)))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK3),s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,sK6(X0,sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))))))),s('type.realax.real',sK5(sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0))))) ) | $spl6),
  inference(superposition,[],[f1575,f40])).
fof(f1575,plain,(
  ( ! [X4] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK3),s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,sK6(X4,sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))))))),s('type.realax.real',sK5(sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X4))),s(sK0,X4))))) ) | $spl6),
  inference(forward_demodulation,[],[f1574,f37])).
fof(f37,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X1),s('type.realax.real',X0))))) = s('type.realax.real','const.topology.dist_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real',X1)))))) )),
  inference(cnf_transformation,[],[f17])).
fof(f17,plain,(
  ! [X0,X1] : s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X1),s('type.realax.real',X0))))) = s('type.realax.real','const.topology.dist_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real',X1)))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X8,X15] : s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X15),s('type.realax.real',X8))))) = s('type.realax.real','const.topology.dist_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X8),s('type.realax.real',X15)))))),
  file('thm.nets.NET_NEG','thm.topology.MR|49|_DEF')).
fof(f1574,plain,(
  ( ! [X4] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.topology.dist_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,sK6(X4,sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))),s('type.realax.real',sK3))))),s('type.realax.real',sK5(sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X4))),s(sK0,X4))))) ) | $spl6),
  inference(subsumption_resolution,[],[f1557,f47])).
fof(f1557,plain,(
  ( ! [X4] : (p(s(bool,'F_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.topology.dist_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,sK6(X4,sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))),s('type.realax.real',sK3))))),s('type.realax.real',sK5(sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X4))),s(sK0,X4))))) ) | $spl6),
  inference(superposition,[],[f46,f72])).
fof(f46,plain,(
  ( ! [X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,'const.nets.tends_3'(s(fun(X0,X1),X4),s(X1,X5),s('type.pair.prod'('type.topology.topology'(X1),fun(X0,fun(X0,bool))),'const.pair.,_2'(s('type.topology.topology'(X1),'const.topology.mtop_1'(s('type.topology.metric'(X1),X2))),s(fun(X0,fun(X0,bool)),X3)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.topology.dist_2'(s('type.topology.metric'(X1),X2),s('type.pair.prod'(X1,X1),'const.pair.,_2'(s(X1,happ(s(fun(X0,X1),X4),s(X0,sK6(X7,X5,X4,X3,X2,X1,X0)))),s(X1,X5))))),s('type.realax.real',sK5(X5,X4,X3,X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,X7))),s(X0,X7))))) )),
  inference(cnf_transformation,[],[f32])).
fof(f2735,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK3),s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,sK6(sK4(sK5(sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0),'const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0),sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))))))),s('type.realax.real',sK5(sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0))))) | s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4(sK5(sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0),'const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))),s(sK0,sK4(sK5(sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0),'const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))) | ($spl4 | $spl6)),
  inference(subsumption_resolution,[],[f2723,f48])).
fof(f2723,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK3),s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,sK6(sK4(sK5(sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0),'const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0),sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))))))),s('type.realax.real',sK5(sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0))))) | s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4(sK5(sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0),'const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))),s(sK0,sK4(sK5(sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0),'const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))) | ($spl4 | $spl6)),
  inference(superposition,[],[f2591,f1694])).
fof(f1694,plain,(
  ( ! [X3] : (s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(X3,sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))),s(sK0,X3))) | s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X3))),s(sK0,X3)))) ) | $spl6),
  inference(subsumption_resolution,[],[f1692,f48])).
fof(f1692,plain,(
  ( ! [X3] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(X3,sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))),s(sK0,X3))) | s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X3))),s(sK0,X3)))) ) | $spl6),
  inference(superposition,[],[f1629,f40])).
fof(f1629,plain,(
  ( ! [X4] : (~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X4))),s(sK0,X4)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(X4,sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))),s(sK0,X4)))) ) | $spl6),
  inference(subsumption_resolution,[],[f1627,f47])).
fof(f1627,plain,(
  ( ! [X4] : (p(s(bool,'F_0')) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X4))),s(sK0,X4)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(X4,sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))),s(sK0,X4)))) ) | $spl6),
  inference(superposition,[],[f1573,f40])).
fof(f1573,plain,(
  ( ! [X3] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(X3,sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))),s(sK0,X3)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X3))),s(sK0,X3))))) ) | $spl6),
  inference(subsumption_resolution,[],[f1556,f47])).
fof(f1556,plain,(
  ( ! [X3] : (p(s(bool,'F_0')) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(X3,sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))),s(sK0,X3)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X3))),s(sK0,X3))))) ) | $spl6),
  inference(superposition,[],[f45,f72])).
fof(f45,plain,(
  ( ! [X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,'const.nets.tends_3'(s(fun(X0,X1),X4),s(X1,X5),s('type.pair.prod'('type.topology.topology'(X1),fun(X0,fun(X0,bool))),'const.pair.,_2'(s('type.topology.topology'(X1),'const.topology.mtop_1'(s('type.topology.metric'(X1),X2))),s(fun(X0,fun(X0,bool)),X3)))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,sK6(X7,X5,X4,X3,X2,X1,X0)))),s(X0,X7)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,X7))),s(X0,X7))))) )),
  inference(cnf_transformation,[],[f32])).
fof(f2591,plain,(
  ( ! [X5] : (~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X5))),s(sK0,sK4(sK5(sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0),'const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK3),s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,X5))))))),s('type.realax.real',sK5(sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))))) ) | ($spl4 | $spl6)),
  inference(subsumption_resolution,[],[f2584,f48])).
fof(f2584,plain,(
  ( ! [X5] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X5))),s(sK0,sK4(sK5(sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0),'const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK3),s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,X5))))))),s('type.realax.real',sK5(sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))))) ) | ($spl4 | $spl6)),
  inference(superposition,[],[f1676,f1589])).
fof(f1676,plain,(
  ( ! [X2,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X1))),s(sK0,sK4(X2,'const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK3),s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,X1))))))),s('type.realax.real',X2))))) ) | $spl4),
  inference(forward_demodulation,[],[f1675,f38])).
fof(f38,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X0),s('type.realax.real',X1))))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X1),s('type.realax.real',X0)))))) )),
  inference(cnf_transformation,[],[f18])).
fof(f18,plain,(
  ! [X0,X1] : s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X0),s('type.realax.real',X1))))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X1),s('type.realax.real',X0)))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X8,X15] : s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X8),s('type.realax.real',X15))))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X15),s('type.realax.real',X8)))))),
  file('thm.nets.NET_NEG','thm.real.ABS_SUB')).
fof(f1675,plain,(
  ( ! [X2,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,X1))),s('type.realax.real',sK3))))),s('type.realax.real',X2)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X1))),s(sK0,sK4(X2,'const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2))))) ) | $spl4),
  inference(forward_demodulation,[],[f1674,f39])).
fof(f39,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X1),s('type.realax.real',X0))) = s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1)))))) )),
  inference(cnf_transformation,[],[f19])).
fof(f19,plain,(
  ! [X0,X1] : s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X1),s('type.realax.real',X0))) = s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1)))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X8,X15] : s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X15),s('type.realax.real',X8))) = s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X8))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X15)))))),
  file('thm.nets.NET_NEG','thm.real.REAL_SUB_NEG2')).
fof(f1674,plain,(
  ( ! [X2,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK3))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,X1))))))))),s('type.realax.real',X2)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X1))),s(sK0,sK4(X2,'const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2))))) ) | $spl4),
  inference(forward_demodulation,[],[f1673,f41])).
fof(f41,plain,(
  ( ! [X0] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))) = s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s('type.realax.real',X0)))) )),
  inference(cnf_transformation,[],[f20])).
fof(f20,plain,(
  ! [X0] : s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))) = s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s('type.realax.real',X0)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X16] : s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X16))) = s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s('type.realax.real',X16)))),
  file('thm.nets.NET_NEG','ARITY_0_1_const.realax.real_neg')).
fof(f1673,plain,(
  ( ! [X2,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK3))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,X1))))))))),s('type.realax.real',X2)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X1))),s(sK0,sK4(X2,'const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2))))) ) | $spl4),
  inference(forward_demodulation,[],[f1672,f36])).
fof(f36,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))) )),
  inference(cnf_transformation,[],[f16])).
fof(f16,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(flattening,[],[f15])).
fof(f15,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X17,X18,X19] : ! [X20,X21,X22] : s(X17,happ(s(fun(X18,X17),X20),s(X18,happ(s(fun(X19,X18),X21),s(X19,X22))))) = s(X17,happ(s(fun(X19,X17),'_dst_x0x1_2'(s(fun(X18,X17),X20),s(fun(X19,X18),X21))),s(X19,X22)))),
  file('thm.nets.NET_NEG','_dst_x0x1_ax')).
fof(f1672,plain,(
  ( ! [X2,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK3))),s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2))),s(sK0,X1))))))),s('type.realax.real',X2)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X1))),s(sK0,sK4(X2,'const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2))))) ) | $spl4),
  inference(forward_demodulation,[],[f1671,f37])).
fof(f1671,plain,(
  ( ! [X2,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.topology.dist_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2))),s(sK0,X1))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK3))))))),s('type.realax.real',X2)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X1))),s(sK0,sK4(X2,'const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2))))) ) | $spl4),
  inference(subsumption_resolution,[],[f1652,f48])).
fof(f1652,plain,(
  ( ! [X2,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.topology.dist_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2))),s(sK0,X1))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK3))))))),s('type.realax.real',X2)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X1))),s(sK0,sK4(X2,'const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2))))) ) | $spl4),
  inference(superposition,[],[f43,f71])).
fof(f43,plain,(
  ( ! [X4,X2,X0,X5,X3,X1,X11,X9] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.topology.dist_2'(s('type.topology.metric'(X1),X2),s('type.pair.prod'(X1,X1),'const.pair.,_2'(s(X1,happ(s(fun(X0,X1),X4),s(X0,X11))),s(X1,X5))))),s('type.realax.real',X9)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,X11))),s(X0,sK4(X9,X5,X4,X3,X2,X1,X0))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X9)))) | ~p(s(bool,'const.nets.tends_3'(s(fun(X0,X1),X4),s(X1,X5),s('type.pair.prod'('type.topology.topology'(X1),fun(X0,fun(X0,bool))),'const.pair.,_2'(s('type.topology.topology'(X1),'const.topology.mtop_1'(s('type.topology.metric'(X1),X2))),s(fun(X0,fun(X0,bool)),X3))))))) )),
  inference(cnf_transformation,[],[f32])).
fof(f2729,plain,(
  $false | ($spl4 | $spl6)),
  inference(subsumption_resolution,[],[f2728,f48])).
fof(f2728,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f2727,f1858])).
fof(f2727,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4(sK5(sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0),'const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))),s(sK0,sK4(sK5(sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0),'const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0))))) | ($spl4 | $spl6)),
  inference(subsumption_resolution,[],[f2707,f1575])).
fof(f2707,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK3),s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,sK6(sK4(sK5(sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0),'const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0),sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))))))),s('type.realax.real',sK5(sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4(sK5(sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0),'const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))),s(sK0,sK4(sK5(sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0),'const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0))))) | ($spl4 | $spl6)),
  inference(resolution,[],[f2591,f1573])).
fof(f1546,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1545,f1528])).
fof(f1528,plain,(
  s(bool,'F_0') != s(bool,'const.nets.tends_3'(s(fun(sK0,'type.realax.real'),sK2),s('type.realax.real',sK3),s('type.pair.prod'('type.topology.topology'('type.realax.real'),fun(sK0,fun(sK0,bool))),'const.pair.,_2'(s('type.topology.topology'('type.realax.real'),'const.topology.mtop_1'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'))),s(fun(sK0,fun(sK0,bool)),sK1))))) | $spl0),
  inference(forward_demodulation,[],[f50,f67])).
fof(f50,plain,(
  s(bool,'const.nets.tends_3'(s(fun(sK0,'type.realax.real'),sK2),s('type.realax.real',sK3),s('type.pair.prod'('type.topology.topology'('type.realax.real'),fun(sK0,fun(sK0,bool))),'const.pair.,_2'(s('type.topology.topology'('type.realax.real'),'const.topology.mtop_1'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'))),s(fun(sK0,fun(sK0,bool)),sK1))))) != s(bool,'const.nets.tends_3'(s(fun(sK0,'type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK3))),s('type.pair.prod'('type.topology.topology'('type.realax.real'),fun(sK0,fun(sK0,bool))),'const.pair.,_2'(s('type.topology.topology'('type.realax.real'),'const.topology.mtop_1'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'))),s(fun(sK0,fun(sK0,bool)),sK1)))))),
  inference(forward_demodulation,[],[f35,f41])).
fof(f35,plain,(
  s(bool,'const.nets.tends_3'(s(fun(sK0,'type.realax.real'),sK2),s('type.realax.real',sK3),s('type.pair.prod'('type.topology.topology'('type.realax.real'),fun(sK0,fun(sK0,bool))),'const.pair.,_2'(s('type.topology.topology'('type.realax.real'),'const.topology.mtop_1'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'))),s(fun(sK0,fun(sK0,bool)),sK1))))) != s(bool,'const.nets.tends_3'(s(fun(sK0,'type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s('type.realax.real',sK3))),s('type.pair.prod'('type.topology.topology'('type.realax.real'),fun(sK0,fun(sK0,bool))),'const.pair.,_2'(s('type.topology.topology'('type.realax.real'),'const.topology.mtop_1'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'))),s(fun(sK0,fun(sK0,bool)),sK1)))))),
  inference(cnf_transformation,[],[f29])).
fof(f29,plain,(
  p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),sK1)))) & s(bool,'const.nets.tends_3'(s(fun(sK0,'type.realax.real'),sK2),s('type.realax.real',sK3),s('type.pair.prod'('type.topology.topology'('type.realax.real'),fun(sK0,fun(sK0,bool))),'const.pair.,_2'(s('type.topology.topology'('type.realax.real'),'const.topology.mtop_1'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'))),s(fun(sK0,fun(sK0,bool)),sK1))))) != s(bool,'const.nets.tends_3'(s(fun(sK0,'type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s('type.realax.real',sK3))),s('type.pair.prod'('type.topology.topology'('type.realax.real'),fun(sK0,fun(sK0,bool))),'const.pair.,_2'(s('type.topology.topology'('type.realax.real'),'const.topology.mtop_1'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'))),s(fun(sK0,fun(sK0,bool)),sK1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f26])).
fof(f26,plain,(
  ? [X0,X1] : (p(s(bool,'const.nets.dorder_1'(s(fun(X0,fun(X0,bool)),X1)))) & ? [X2,X3] : s(bool,'const.nets.tends_3'(s(fun(X0,'type.realax.real'),X2),s('type.realax.real',X3),s('type.pair.prod'('type.topology.topology'('type.realax.real'),fun(X0,fun(X0,bool))),'const.pair.,_2'(s('type.topology.topology'('type.realax.real'),'const.topology.mtop_1'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'))),s(fun(X0,fun(X0,bool)),X1))))) != s(bool,'const.nets.tends_3'(s(fun(X0,'type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(X0,'type.realax.real'),X2))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s('type.realax.real',X3))),s('type.pair.prod'('type.topology.topology'('type.realax.real'),fun(X0,fun(X0,bool))),'const.pair.,_2'(s('type.topology.topology'('type.realax.real'),'const.topology.mtop_1'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'))),s(fun(X0,fun(X0,bool)),X1))))))),
  inference(ennf_transformation,[],[f14])).
fof(f14,plain,(
  ~! [X0,X1] : (p(s(bool,'const.nets.dorder_1'(s(fun(X0,fun(X0,bool)),X1)))) => ! [X2,X3] : s(bool,'const.nets.tends_3'(s(fun(X0,'type.realax.real'),X2),s('type.realax.real',X3),s('type.pair.prod'('type.topology.topology'('type.realax.real'),fun(X0,fun(X0,bool))),'const.pair.,_2'(s('type.topology.topology'('type.realax.real'),'const.topology.mtop_1'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'))),s(fun(X0,fun(X0,bool)),X1))))) = s(bool,'const.nets.tends_3'(s(fun(X0,'type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(X0,'type.realax.real'),X2))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s('type.realax.real',X3))),s('type.pair.prod'('type.topology.topology'('type.realax.real'),fun(X0,fun(X0,bool))),'const.pair.,_2'(s('type.topology.topology'('type.realax.real'),'const.topology.mtop_1'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'))),s(fun(X0,fun(X0,bool)),X1))))))),
  inference(flattening,[],[f13])).
fof(f13,plain,(
  ~! [X0] : ! [X1] : (p(s(bool,'const.nets.dorder_1'(s(fun(X0,fun(X0,bool)),X1)))) => ! [X2,X3] : s(bool,'const.nets.tends_3'(s(fun(X0,'type.realax.real'),X2),s('type.realax.real',X3),s('type.pair.prod'('type.topology.topology'('type.realax.real'),fun(X0,fun(X0,bool))),'const.pair.,_2'(s('type.topology.topology'('type.realax.real'),'const.topology.mtop_1'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'))),s(fun(X0,fun(X0,bool)),X1))))) = s(bool,'const.nets.tends_3'(s(fun(X0,'type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(X0,'type.realax.real'),X2))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s('type.realax.real',X3))),s('type.pair.prod'('type.topology.topology'('type.realax.real'),fun(X0,fun(X0,bool))),'const.pair.,_2'(s('type.topology.topology'('type.realax.real'),'const.topology.mtop_1'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'))),s(fun(X0,fun(X0,bool)),X1))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7] : (p(s(bool,'const.nets.dorder_1'(s(fun(X6,fun(X6,bool)),X7)))) => ! [X8,X9] : s(bool,'const.nets.tends_3'(s(fun(X6,'type.realax.real'),X8),s('type.realax.real',X9),s('type.pair.prod'('type.topology.topology'('type.realax.real'),fun(X6,fun(X6,bool))),'const.pair.,_2'(s('type.topology.topology'('type.realax.real'),'const.topology.mtop_1'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'))),s(fun(X6,fun(X6,bool)),X7))))) = s(bool,'const.nets.tends_3'(s(fun(X6,'type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(X6,'type.realax.real'),X8))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s('type.realax.real',X9))),s('type.pair.prod'('type.topology.topology'('type.realax.real'),fun(X6,fun(X6,bool))),'const.pair.,_2'(s('type.topology.topology'('type.realax.real'),'const.topology.mtop_1'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'))),s(fun(X6,fun(X6,bool)),X7))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7] : (p(s(bool,'const.nets.dorder_1'(s(fun(X6,fun(X6,bool)),X7)))) => ! [X8,X9] : s(bool,'const.nets.tends_3'(s(fun(X6,'type.realax.real'),X8),s('type.realax.real',X9),s('type.pair.prod'('type.topology.topology'('type.realax.real'),fun(X6,fun(X6,bool))),'const.pair.,_2'(s('type.topology.topology'('type.realax.real'),'const.topology.mtop_1'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'))),s(fun(X6,fun(X6,bool)),X7))))) = s(bool,'const.nets.tends_3'(s(fun(X6,'type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(X6,'type.realax.real'),X8))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s('type.realax.real',X9))),s('type.pair.prod'('type.topology.topology'('type.realax.real'),fun(X6,fun(X6,bool))),'const.pair.,_2'(s('type.topology.topology'('type.realax.real'),'const.topology.mtop_1'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'))),s(fun(X6,fun(X6,bool)),X7))))))),
  file('thm.nets.NET_NEG','thm.nets.NET_NEG')).
fof(f1545,plain,(
  s(bool,'F_0') = s(bool,'const.nets.tends_3'(s(fun(sK0,'type.realax.real'),sK2),s('type.realax.real',sK3),s('type.pair.prod'('type.topology.topology'('type.realax.real'),fun(sK0,fun(sK0,bool))),'const.pair.,_2'(s('type.topology.topology'('type.realax.real'),'const.topology.mtop_1'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'))),s(fun(sK0,fun(sK0,bool)),sK1))))) | $spl3),
  inference(trivial_inequality_removal,[],[f1543])).
fof(f1543,plain,(
  s(bool,'T_0') != s(bool,'T_0') | s(bool,'F_0') = s(bool,'const.nets.tends_3'(s(fun(sK0,'type.realax.real'),sK2),s('type.realax.real',sK3),s('type.pair.prod'('type.topology.topology'('type.realax.real'),fun(sK0,fun(sK0,bool))),'const.pair.,_2'(s('type.topology.topology'('type.realax.real'),'const.topology.mtop_1'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'))),s(fun(sK0,fun(sK0,bool)),sK1))))) | $spl3),
  inference(superposition,[],[f69,f40])).
fof(f1527,plain,(
  $false | ($spl2 | $spl4)),
  inference(subsumption_resolution,[],[f71,f128])).
fof(f128,plain,(
  s(bool,'T_0') != s(bool,'const.nets.tends_3'(s(fun(sK0,'type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK3))),s('type.pair.prod'('type.topology.topology'('type.realax.real'),fun(sK0,fun(sK0,bool))),'const.pair.,_2'(s('type.topology.topology'('type.realax.real'),'const.topology.mtop_1'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'))),s(fun(sK0,fun(sK0,bool)),sK1))))) | $spl2),
  inference(backward_demodulation,[],[f68,f50])).
fof(f68,plain,(
  s(bool,'T_0') = s(bool,'const.nets.tends_3'(s(fun(sK0,'type.realax.real'),sK2),s('type.realax.real',sK3),s('type.pair.prod'('type.topology.topology'('type.realax.real'),fun(sK0,fun(sK0,bool))),'const.pair.,_2'(s('type.topology.topology'('type.realax.real'),'const.topology.mtop_1'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'))),s(fun(sK0,fun(sK0,bool)),sK1))))) | $spl2),
  inference(cnf_transformation,[],[f68_D])).
fof(f68_D,plain,(
  s(bool,'T_0') = s(bool,'const.nets.tends_3'(s(fun(sK0,'type.realax.real'),sK2),s('type.realax.real',sK3),s('type.pair.prod'('type.topology.topology'('type.realax.real'),fun(sK0,fun(sK0,bool))),'const.pair.,_2'(s('type.topology.topology'('type.realax.real'),'const.topology.mtop_1'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'))),s(fun(sK0,fun(sK0,bool)),sK1))))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f1524,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),X2)))) | s(bool,'F_0') = s(bool,g(s(sK0,sK7(g(s(sK0,sK7(sK1,X2,fun(sK0,bool),sK0))),X3,bool,sK0)))) | s(fun(sK0,bool),X3) != s(fun(sK0,bool),f(s(sK0,sK7(sK1,X2,fun(sK0,bool),sK0))))) ) | $spl30),
  inference(cnf_transformation,[],[f1524_D])).
fof(f1524_D,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),X2)))) | s(bool,'F_0') = s(bool,g(s(sK0,sK7(g(s(sK0,sK7(sK1,X2,fun(sK0,bool),sK0))),X3,bool,sK0)))) | s(fun(sK0,bool),X3) != s(fun(sK0,bool),f(s(sK0,sK7(sK1,X2,fun(sK0,bool),sK0))))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f1507,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl15),
  inference(cnf_transformation,[],[f1507_D])).
fof(f1507_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f1523,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),X2)))) | s(fun(sK0,bool),X3) != s(fun(sK0,bool),f(s(sK0,sK7(sK1,X2,fun(sK0,bool),sK0)))) | s(bool,'F_0') = s(bool,g(s(sK0,sK7(g(s(sK0,sK7(sK1,X2,fun(sK0,bool),sK0))),X3,bool,sK0))))) )),
  inference(subsumption_resolution,[],[f1252,f595])).
fof(f595,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,f(s(sK0,sK7(g(s(sK0,sK7(sK1,X0,fun(sK0,bool),sK0))),X1,bool,sK0)))) | p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),X0)))) | s(fun(sK0,bool),X1) != s(fun(sK0,bool),f(s(sK0,sK7(sK1,X0,fun(sK0,bool),sK0)))) | s(bool,'F_0') = s(bool,g(s(sK0,sK7(g(s(sK0,sK7(sK1,X0,fun(sK0,bool),sK0))),X1,bool,sK0))))) )),
  inference(superposition,[],[f100,f40])).
fof(f100,plain,(
  ( ! [X0,X1] : (s(bool,f(s(sK0,sK7(g(s(sK0,sK7(sK1,X0,fun(sK0,bool),sK0))),X1,bool,sK0)))) != s(bool,g(s(sK0,sK7(g(s(sK0,sK7(sK1,X0,fun(sK0,bool),sK0))),X1,bool,sK0)))) | p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),X0)))) | s(fun(sK0,bool),X1) != s(fun(sK0,bool),f(s(sK0,sK7(sK1,X0,fun(sK0,bool),sK0))))) )),
  inference(superposition,[],[f51,f49])).
fof(f49,plain,(
  ( ! [X2,X0,X3,X1] : (s(fun(X0,X1),X2) = s(fun(X0,X1),X3) | s(X1,f(s(X0,sK7(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK7(X3,X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f33])).
fof(f33,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK7(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK7(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7])],[f28])).
fof(f28,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f25])).
fof(f25,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f24])).
fof(f24,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.nets.NET_NEG',aHL_EXT)).
fof(f51,plain,(
  ( ! [X0] : (s(fun(sK0,bool),f(s(sK0,sK7(sK1,X0,fun(sK0,bool),sK0)))) != s(fun(sK0,bool),g(s(sK0,sK7(sK1,X0,fun(sK0,bool),sK0)))) | p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),X0))))) )),
  inference(superposition,[],[f34,f49])).
fof(f34,plain,(
  p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),sK1))))),
  inference(cnf_transformation,[],[f29])).
fof(f1252,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),X2)))) | s(fun(sK0,bool),X3) != s(fun(sK0,bool),f(s(sK0,sK7(sK1,X2,fun(sK0,bool),sK0)))) | s(bool,'F_0') = s(bool,g(s(sK0,sK7(g(s(sK0,sK7(sK1,X2,fun(sK0,bool),sK0))),X3,bool,sK0)))) | s(bool,'T_0') = s(bool,f(s(sK0,sK7(g(s(sK0,sK7(sK1,X2,fun(sK0,bool),sK0))),X3,bool,sK0))))) )),
  inference(superposition,[],[f595,f40])).
fof(f1522,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),X0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK7(g(s(sK0,sK7(sK1,X0,fun(sK0,bool),sK0))),X1,bool,sK0)))) | s(fun(sK0,bool),X1) != s(fun(sK0,bool),f(s(sK0,sK7(sK1,X0,fun(sK0,bool),sK0))))) ) | $spl28),
  inference(cnf_transformation,[],[f1522_D])).
fof(f1522_D,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),X0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK7(g(s(sK0,sK7(sK1,X0,fun(sK0,bool),sK0))),X1,bool,sK0)))) | s(fun(sK0,bool),X1) != s(fun(sK0,bool),f(s(sK0,sK7(sK1,X0,fun(sK0,bool),sK0))))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f1521,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),X0)))) | s(fun(sK0,bool),X1) != s(fun(sK0,bool),f(s(sK0,sK7(sK1,X0,fun(sK0,bool),sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK7(g(s(sK0,sK7(sK1,X0,fun(sK0,bool),sK0))),X1,bool,sK0))))) )),
  inference(subsumption_resolution,[],[f1254,f596])).
fof(f596,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,f(s(sK0,sK7(g(s(sK0,sK7(sK1,X2,fun(sK0,bool),sK0))),X3,bool,sK0)))) | p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),X2)))) | s(fun(sK0,bool),X3) != s(fun(sK0,bool),f(s(sK0,sK7(sK1,X2,fun(sK0,bool),sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK7(g(s(sK0,sK7(sK1,X2,fun(sK0,bool),sK0))),X3,bool,sK0))))) )),
  inference(superposition,[],[f100,f40])).
fof(f1254,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),X0)))) | s(fun(sK0,bool),X1) != s(fun(sK0,bool),f(s(sK0,sK7(sK1,X0,fun(sK0,bool),sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK7(g(s(sK0,sK7(sK1,X0,fun(sK0,bool),sK0))),X1,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK7(g(s(sK0,sK7(sK1,X0,fun(sK0,bool),sK0))),X1,bool,sK0))))) )),
  inference(superposition,[],[f596,f40])).
fof(f1520,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),X3)))) | s(bool,'F_0') = s(bool,g(s(sK0,sK7(X2,g(s(sK0,sK7(sK1,X3,fun(sK0,bool),sK0))),bool,sK0)))) | s(fun(sK0,bool),X2) != s(fun(sK0,bool),f(s(sK0,sK7(sK1,X3,fun(sK0,bool),sK0))))) ) | $spl26),
  inference(cnf_transformation,[],[f1520_D])).
fof(f1520_D,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),X3)))) | s(bool,'F_0') = s(bool,g(s(sK0,sK7(X2,g(s(sK0,sK7(sK1,X3,fun(sK0,bool),sK0))),bool,sK0)))) | s(fun(sK0,bool),X2) != s(fun(sK0,bool),f(s(sK0,sK7(sK1,X3,fun(sK0,bool),sK0))))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f1519,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),X3)))) | s(fun(sK0,bool),X2) != s(fun(sK0,bool),f(s(sK0,sK7(sK1,X3,fun(sK0,bool),sK0)))) | s(bool,'F_0') = s(bool,g(s(sK0,sK7(X2,g(s(sK0,sK7(sK1,X3,fun(sK0,bool),sK0))),bool,sK0))))) )),
  inference(subsumption_resolution,[],[f1260,f597])).
fof(f597,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,f(s(sK0,sK7(X0,g(s(sK0,sK7(sK1,X1,fun(sK0,bool),sK0))),bool,sK0)))) | p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),X1)))) | s(fun(sK0,bool),X0) != s(fun(sK0,bool),f(s(sK0,sK7(sK1,X1,fun(sK0,bool),sK0)))) | s(bool,'F_0') = s(bool,g(s(sK0,sK7(X0,g(s(sK0,sK7(sK1,X1,fun(sK0,bool),sK0))),bool,sK0))))) )),
  inference(superposition,[],[f101,f40])).
fof(f101,plain,(
  ( ! [X2,X3] : (s(bool,f(s(sK0,sK7(X3,g(s(sK0,sK7(sK1,X2,fun(sK0,bool),sK0))),bool,sK0)))) != s(bool,g(s(sK0,sK7(X3,g(s(sK0,sK7(sK1,X2,fun(sK0,bool),sK0))),bool,sK0)))) | p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),X2)))) | s(fun(sK0,bool),X3) != s(fun(sK0,bool),f(s(sK0,sK7(sK1,X2,fun(sK0,bool),sK0))))) )),
  inference(superposition,[],[f51,f49])).
fof(f1260,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),X3)))) | s(fun(sK0,bool),X2) != s(fun(sK0,bool),f(s(sK0,sK7(sK1,X3,fun(sK0,bool),sK0)))) | s(bool,'F_0') = s(bool,g(s(sK0,sK7(X2,g(s(sK0,sK7(sK1,X3,fun(sK0,bool),sK0))),bool,sK0)))) | s(bool,'T_0') = s(bool,f(s(sK0,sK7(X2,g(s(sK0,sK7(sK1,X3,fun(sK0,bool),sK0))),bool,sK0))))) )),
  inference(superposition,[],[f597,f40])).
fof(f1518,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),X1)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK7(X0,g(s(sK0,sK7(sK1,X1,fun(sK0,bool),sK0))),bool,sK0)))) | s(fun(sK0,bool),X0) != s(fun(sK0,bool),f(s(sK0,sK7(sK1,X1,fun(sK0,bool),sK0))))) ) | $spl24),
  inference(cnf_transformation,[],[f1518_D])).
fof(f1518_D,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),X1)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK7(X0,g(s(sK0,sK7(sK1,X1,fun(sK0,bool),sK0))),bool,sK0)))) | s(fun(sK0,bool),X0) != s(fun(sK0,bool),f(s(sK0,sK7(sK1,X1,fun(sK0,bool),sK0))))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f1517,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),X1)))) | s(fun(sK0,bool),X0) != s(fun(sK0,bool),f(s(sK0,sK7(sK1,X1,fun(sK0,bool),sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK7(X0,g(s(sK0,sK7(sK1,X1,fun(sK0,bool),sK0))),bool,sK0))))) )),
  inference(subsumption_resolution,[],[f1262,f598])).
fof(f598,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,f(s(sK0,sK7(X2,g(s(sK0,sK7(sK1,X3,fun(sK0,bool),sK0))),bool,sK0)))) | p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),X3)))) | s(fun(sK0,bool),X2) != s(fun(sK0,bool),f(s(sK0,sK7(sK1,X3,fun(sK0,bool),sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK7(X2,g(s(sK0,sK7(sK1,X3,fun(sK0,bool),sK0))),bool,sK0))))) )),
  inference(superposition,[],[f101,f40])).
fof(f1262,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),X1)))) | s(fun(sK0,bool),X0) != s(fun(sK0,bool),f(s(sK0,sK7(sK1,X1,fun(sK0,bool),sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK7(X0,g(s(sK0,sK7(sK1,X1,fun(sK0,bool),sK0))),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK7(X0,g(s(sK0,sK7(sK1,X1,fun(sK0,bool),sK0))),bool,sK0))))) )),
  inference(superposition,[],[f598,f40])).
fof(f1516,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),X2)))) | s(bool,'F_0') = s(bool,g(s(sK0,sK7(g(s(sK0,sK7(X2,sK1,fun(sK0,bool),sK0))),X3,bool,sK0)))) | s(fun(sK0,bool),X3) != s(fun(sK0,bool),f(s(sK0,sK7(X2,sK1,fun(sK0,bool),sK0))))) ) | $spl22),
  inference(cnf_transformation,[],[f1516_D])).
fof(f1516_D,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),X2)))) | s(bool,'F_0') = s(bool,g(s(sK0,sK7(g(s(sK0,sK7(X2,sK1,fun(sK0,bool),sK0))),X3,bool,sK0)))) | s(fun(sK0,bool),X3) != s(fun(sK0,bool),f(s(sK0,sK7(X2,sK1,fun(sK0,bool),sK0))))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f1515,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),X2)))) | s(fun(sK0,bool),X3) != s(fun(sK0,bool),f(s(sK0,sK7(X2,sK1,fun(sK0,bool),sK0)))) | s(bool,'F_0') = s(bool,g(s(sK0,sK7(g(s(sK0,sK7(X2,sK1,fun(sK0,bool),sK0))),X3,bool,sK0))))) )),
  inference(subsumption_resolution,[],[f1266,f603])).
fof(f603,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,f(s(sK0,sK7(g(s(sK0,sK7(X0,sK1,fun(sK0,bool),sK0))),X1,bool,sK0)))) | p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),X0)))) | s(fun(sK0,bool),X1) != s(fun(sK0,bool),f(s(sK0,sK7(X0,sK1,fun(sK0,bool),sK0)))) | s(bool,'F_0') = s(bool,g(s(sK0,sK7(g(s(sK0,sK7(X0,sK1,fun(sK0,bool),sK0))),X1,bool,sK0))))) )),
  inference(superposition,[],[f102,f40])).
fof(f102,plain,(
  ( ! [X0,X1] : (s(bool,f(s(sK0,sK7(g(s(sK0,sK7(X0,sK1,fun(sK0,bool),sK0))),X1,bool,sK0)))) != s(bool,g(s(sK0,sK7(g(s(sK0,sK7(X0,sK1,fun(sK0,bool),sK0))),X1,bool,sK0)))) | p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),X0)))) | s(fun(sK0,bool),X1) != s(fun(sK0,bool),f(s(sK0,sK7(X0,sK1,fun(sK0,bool),sK0))))) )),
  inference(superposition,[],[f52,f49])).
fof(f52,plain,(
  ( ! [X1] : (s(fun(sK0,bool),f(s(sK0,sK7(X1,sK1,fun(sK0,bool),sK0)))) != s(fun(sK0,bool),g(s(sK0,sK7(X1,sK1,fun(sK0,bool),sK0)))) | p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),X1))))) )),
  inference(superposition,[],[f34,f49])).
fof(f1266,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),X2)))) | s(fun(sK0,bool),X3) != s(fun(sK0,bool),f(s(sK0,sK7(X2,sK1,fun(sK0,bool),sK0)))) | s(bool,'F_0') = s(bool,g(s(sK0,sK7(g(s(sK0,sK7(X2,sK1,fun(sK0,bool),sK0))),X3,bool,sK0)))) | s(bool,'T_0') = s(bool,f(s(sK0,sK7(g(s(sK0,sK7(X2,sK1,fun(sK0,bool),sK0))),X3,bool,sK0))))) )),
  inference(superposition,[],[f603,f40])).
fof(f1514,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),X0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK7(g(s(sK0,sK7(X0,sK1,fun(sK0,bool),sK0))),X1,bool,sK0)))) | s(fun(sK0,bool),X1) != s(fun(sK0,bool),f(s(sK0,sK7(X0,sK1,fun(sK0,bool),sK0))))) ) | $spl20),
  inference(cnf_transformation,[],[f1514_D])).
fof(f1514_D,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),X0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK7(g(s(sK0,sK7(X0,sK1,fun(sK0,bool),sK0))),X1,bool,sK0)))) | s(fun(sK0,bool),X1) != s(fun(sK0,bool),f(s(sK0,sK7(X0,sK1,fun(sK0,bool),sK0))))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f1513,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),X0)))) | s(fun(sK0,bool),X1) != s(fun(sK0,bool),f(s(sK0,sK7(X0,sK1,fun(sK0,bool),sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK7(g(s(sK0,sK7(X0,sK1,fun(sK0,bool),sK0))),X1,bool,sK0))))) )),
  inference(subsumption_resolution,[],[f1268,f604])).
fof(f604,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,f(s(sK0,sK7(g(s(sK0,sK7(X2,sK1,fun(sK0,bool),sK0))),X3,bool,sK0)))) | p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),X2)))) | s(fun(sK0,bool),X3) != s(fun(sK0,bool),f(s(sK0,sK7(X2,sK1,fun(sK0,bool),sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK7(g(s(sK0,sK7(X2,sK1,fun(sK0,bool),sK0))),X3,bool,sK0))))) )),
  inference(superposition,[],[f102,f40])).
fof(f1268,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),X0)))) | s(fun(sK0,bool),X1) != s(fun(sK0,bool),f(s(sK0,sK7(X0,sK1,fun(sK0,bool),sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK7(g(s(sK0,sK7(X0,sK1,fun(sK0,bool),sK0))),X1,bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK7(g(s(sK0,sK7(X0,sK1,fun(sK0,bool),sK0))),X1,bool,sK0))))) )),
  inference(superposition,[],[f604,f40])).
fof(f1512,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),X3)))) | s(bool,'F_0') = s(bool,g(s(sK0,sK7(X2,g(s(sK0,sK7(X3,sK1,fun(sK0,bool),sK0))),bool,sK0)))) | s(fun(sK0,bool),X2) != s(fun(sK0,bool),f(s(sK0,sK7(X3,sK1,fun(sK0,bool),sK0))))) ) | $spl18),
  inference(cnf_transformation,[],[f1512_D])).
fof(f1512_D,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),X3)))) | s(bool,'F_0') = s(bool,g(s(sK0,sK7(X2,g(s(sK0,sK7(X3,sK1,fun(sK0,bool),sK0))),bool,sK0)))) | s(fun(sK0,bool),X2) != s(fun(sK0,bool),f(s(sK0,sK7(X3,sK1,fun(sK0,bool),sK0))))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f1511,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),X3)))) | s(fun(sK0,bool),X2) != s(fun(sK0,bool),f(s(sK0,sK7(X3,sK1,fun(sK0,bool),sK0)))) | s(bool,'F_0') = s(bool,g(s(sK0,sK7(X2,g(s(sK0,sK7(X3,sK1,fun(sK0,bool),sK0))),bool,sK0))))) )),
  inference(subsumption_resolution,[],[f1272,f605])).
fof(f605,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,f(s(sK0,sK7(X0,g(s(sK0,sK7(X1,sK1,fun(sK0,bool),sK0))),bool,sK0)))) | p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),X1)))) | s(fun(sK0,bool),X0) != s(fun(sK0,bool),f(s(sK0,sK7(X1,sK1,fun(sK0,bool),sK0)))) | s(bool,'F_0') = s(bool,g(s(sK0,sK7(X0,g(s(sK0,sK7(X1,sK1,fun(sK0,bool),sK0))),bool,sK0))))) )),
  inference(superposition,[],[f103,f40])).
fof(f103,plain,(
  ( ! [X2,X3] : (s(bool,f(s(sK0,sK7(X3,g(s(sK0,sK7(X2,sK1,fun(sK0,bool),sK0))),bool,sK0)))) != s(bool,g(s(sK0,sK7(X3,g(s(sK0,sK7(X2,sK1,fun(sK0,bool),sK0))),bool,sK0)))) | p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),X2)))) | s(fun(sK0,bool),X3) != s(fun(sK0,bool),f(s(sK0,sK7(X2,sK1,fun(sK0,bool),sK0))))) )),
  inference(superposition,[],[f52,f49])).
fof(f1272,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),X3)))) | s(fun(sK0,bool),X2) != s(fun(sK0,bool),f(s(sK0,sK7(X3,sK1,fun(sK0,bool),sK0)))) | s(bool,'F_0') = s(bool,g(s(sK0,sK7(X2,g(s(sK0,sK7(X3,sK1,fun(sK0,bool),sK0))),bool,sK0)))) | s(bool,'T_0') = s(bool,f(s(sK0,sK7(X2,g(s(sK0,sK7(X3,sK1,fun(sK0,bool),sK0))),bool,sK0))))) )),
  inference(superposition,[],[f605,f40])).
fof(f1510,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),X1)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK7(X0,g(s(sK0,sK7(X1,sK1,fun(sK0,bool),sK0))),bool,sK0)))) | s(fun(sK0,bool),X0) != s(fun(sK0,bool),f(s(sK0,sK7(X1,sK1,fun(sK0,bool),sK0))))) ) | $spl16),
  inference(cnf_transformation,[],[f1510_D])).
fof(f1510_D,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),X1)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK7(X0,g(s(sK0,sK7(X1,sK1,fun(sK0,bool),sK0))),bool,sK0)))) | s(fun(sK0,bool),X0) != s(fun(sK0,bool),f(s(sK0,sK7(X1,sK1,fun(sK0,bool),sK0))))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f1509,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),X1)))) | s(fun(sK0,bool),X0) != s(fun(sK0,bool),f(s(sK0,sK7(X1,sK1,fun(sK0,bool),sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK7(X0,g(s(sK0,sK7(X1,sK1,fun(sK0,bool),sK0))),bool,sK0))))) )),
  inference(subsumption_resolution,[],[f1276,f606])).
fof(f606,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,f(s(sK0,sK7(X2,g(s(sK0,sK7(X3,sK1,fun(sK0,bool),sK0))),bool,sK0)))) | p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),X3)))) | s(fun(sK0,bool),X2) != s(fun(sK0,bool),f(s(sK0,sK7(X3,sK1,fun(sK0,bool),sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK7(X2,g(s(sK0,sK7(X3,sK1,fun(sK0,bool),sK0))),bool,sK0))))) )),
  inference(superposition,[],[f103,f40])).
fof(f1276,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),X1)))) | s(fun(sK0,bool),X0) != s(fun(sK0,bool),f(s(sK0,sK7(X1,sK1,fun(sK0,bool),sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK7(X0,g(s(sK0,sK7(X1,sK1,fun(sK0,bool),sK0))),bool,sK0)))) | s(bool,'F_0') = s(bool,f(s(sK0,sK7(X0,g(s(sK0,sK7(X1,sK1,fun(sK0,bool),sK0))),bool,sK0))))) )),
  inference(superposition,[],[f606,f40])).
fof(f1506,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,X0)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(X0,'const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))),s(sK0,sK4(sK5('const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0),sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0))))) ) | $spl12),
  inference(cnf_transformation,[],[f1506_D])).
fof(f1506_D,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,X0)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(X0,'const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))),s(sK0,sK4(sK5('const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0),sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0))))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f1508,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl14),
  inference(cnf_transformation,[],[f1508_D])).
fof(f1508_D,plain,(
  s(bool,'T_0') = s(bool,'F_0') <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f1498,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'F_0') | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,X0)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(X0,'const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))),s(sK0,sK4(sK5('const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0),sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0))))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f1486,f288])).
fof(f288,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK5('const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))) | $spl0),
  inference(subsumption_resolution,[],[f287,f47])).
fof(f287,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK5('const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))) | $spl0),
  inference(superposition,[],[f262,f40])).
fof(f262,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK5('const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0))))) | $spl0),
  inference(subsumption_resolution,[],[f245,f47])).
fof(f245,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK5('const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0))))) | $spl0),
  inference(superposition,[],[f44,f67])).
fof(f1486,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,X0)))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK5('const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(X0,'const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))),s(sK0,sK4(sK5('const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0),sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0))))) ) | ($spl0 | $spl2)),
  inference(resolution,[],[f269,f746])).
fof(f746,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK3),s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,X2))))))),s('type.realax.real',X3)))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X3))) | s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X2))),s(sK0,sK4(X3,sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0))))) ) | $spl2),
  inference(subsumption_resolution,[],[f743,f48])).
fof(f743,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK3),s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,X2))))))),s('type.realax.real',X3)))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X3))) | s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X2))),s(sK0,sK4(X3,sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0))))) ) | $spl2),
  inference(superposition,[],[f524,f40])).
fof(f524,plain,(
  ( ! [X2,X1] : (~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X2))),s(sK0,sK4(X1,sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK3),s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,X2))))))),s('type.realax.real',X1)))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))) ) | $spl2),
  inference(subsumption_resolution,[],[f518,f48])).
fof(f518,plain,(
  ( ! [X2,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X2))),s(sK0,sK4(X1,sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK3),s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,X2))))))),s('type.realax.real',X1)))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))) ) | $spl2),
  inference(superposition,[],[f155,f40])).
fof(f155,plain,(
  ( ! [X2,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X1))),s(sK0,sK4(X2,sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK3),s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,X1))))))),s('type.realax.real',X2))))) ) | $spl2),
  inference(forward_demodulation,[],[f154,f37])).
fof(f154,plain,(
  ( ! [X2,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.topology.dist_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,X1))),s('type.realax.real',sK3))))),s('type.realax.real',X2)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X1))),s(sK0,sK4(X2,sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2))))) ) | $spl2),
  inference(subsumption_resolution,[],[f135,f48])).
fof(f135,plain,(
  ( ! [X2,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.topology.dist_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,X1))),s('type.realax.real',sK3))))),s('type.realax.real',X2)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X1))),s(sK0,sK4(X2,sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2))))) ) | $spl2),
  inference(superposition,[],[f43,f68])).
fof(f269,plain,(
  ( ! [X4] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK3),s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,sK6(X4,'const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))))))),s('type.realax.real',sK5('const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X4))),s(sK0,X4))))) ) | $spl0),
  inference(forward_demodulation,[],[f268,f38])).
fof(f268,plain,(
  ( ! [X4] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,sK6(X4,'const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))),s('type.realax.real',sK3))))),s('type.realax.real',sK5('const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X4))),s(sK0,X4))))) ) | $spl0),
  inference(forward_demodulation,[],[f267,f39])).
fof(f267,plain,(
  ( ! [X4] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK3))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,sK6(X4,'const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))))))))),s('type.realax.real',sK5('const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X4))),s(sK0,X4))))) ) | $spl0),
  inference(forward_demodulation,[],[f266,f41])).
fof(f266,plain,(
  ( ! [X4] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK3))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,sK6(X4,'const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))))))))),s('type.realax.real',sK5('const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X4))),s(sK0,X4))))) ) | $spl0),
  inference(forward_demodulation,[],[f265,f36])).
fof(f265,plain,(
  ( ! [X4] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK3))),s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2))),s(sK0,sK6(X4,'const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))))))),s('type.realax.real',sK5('const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X4))),s(sK0,X4))))) ) | $spl0),
  inference(forward_demodulation,[],[f264,f37])).
fof(f264,plain,(
  ( ! [X4] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.topology.dist_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2))),s(sK0,sK6(X4,'const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK3))))))),s('type.realax.real',sK5('const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X4))),s(sK0,X4))))) ) | $spl0),
  inference(subsumption_resolution,[],[f247,f47])).
fof(f247,plain,(
  ( ! [X4] : (p(s(bool,'F_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.topology.dist_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2))),s(sK0,sK6(X4,'const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK3))))))),s('type.realax.real',sK5('const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X4))),s(sK0,X4))))) ) | $spl0),
  inference(superposition,[],[f46,f67])).
fof(f1504,plain,(
  $false | ($spl0 | $spl2 | $spl7)),
  inference(subsumption_resolution,[],[f1503,f123])).
fof(f123,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl7),
  inference(subsumption_resolution,[],[f119,f73])).
fof(f119,plain,(
  s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.nets.tends_3'(s(fun(sK0,'type.realax.real'),sK2),s('type.realax.real',sK3),s('type.pair.prod'('type.topology.topology'('type.realax.real'),fun(sK0,fun(sK0,bool))),'const.pair.,_2'(s('type.topology.topology'('type.realax.real'),'const.topology.mtop_1'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'))),s(fun(sK0,fun(sK0,bool)),sK1))))) | $spl7),
  inference(superposition,[],[f73,f40])).
fof(f1503,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f1502,f288])).
fof(f1502,plain,(
  s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK5('const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))) | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f1501,f48])).
fof(f1501,plain,(
  ~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK5('const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f1500,f278])).
fof(f278,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4(sK5('const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0),sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))),s(sK0,sK4(sK5('const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0),sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))) | ($spl0 | $spl2)),
  inference(resolution,[],[f262,f192])).
fof(f192,plain,(
  ( ! [X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4(X1,sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))),s(sK0,sK4(X1,sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0))))) ) | $spl2),
  inference(subsumption_resolution,[],[f187,f47])).
fof(f187,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4(X1,sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))),s(sK0,sK4(X1,sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0))))) ) | $spl2),
  inference(superposition,[],[f153,f40])).
fof(f153,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4(X0,sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))),s(sK0,sK4(X0,sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))))) ) | $spl2),
  inference(subsumption_resolution,[],[f134,f48])).
fof(f134,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4(X0,sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))),s(sK0,sK4(X0,sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))))) ) | $spl2),
  inference(superposition,[],[f42,f68])).
fof(f1500,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4(sK5('const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0),sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))),s(sK0,sK4(sK5('const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0),sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0))))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK5('const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))) | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f1491,f48])).
fof(f1491,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4(sK5('const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0),sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))),s(sK0,sK4(sK5('const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0),sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0))))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK5('const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))) | ($spl0 | $spl2)),
  inference(superposition,[],[f269,f687])).
fof(f687,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK3),s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,sK6(sK4(X1,sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0),'const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))))))),s('type.realax.real',X1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))) ) | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f683,f47])).
fof(f683,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK3),s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,sK6(sK4(X1,sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0),'const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))))))),s('type.realax.real',X1)))) ) | ($spl0 | $spl2)),
  inference(superposition,[],[f559,f40])).
fof(f559,plain,(
  ( ! [X0] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK3),s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,sK6(sK4(X0,sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0),'const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))))))),s('type.realax.real',X0)))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) ) | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f553,f48])).
fof(f553,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK3),s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,sK6(sK4(X0,sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0),'const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))))))),s('type.realax.real',X0)))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) ) | ($spl0 | $spl2)),
  inference(superposition,[],[f541,f40])).
fof(f541,plain,(
  ( ! [X0] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK3),s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,sK6(sK4(X0,sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0),'const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))))))),s('type.realax.real',X0))))) ) | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f540,f48])).
fof(f540,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK3),s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,sK6(sK4(X0,sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0),'const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))))))),s('type.realax.real',X0)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f539,f68])).
fof(f539,plain,(
  ( ! [X0] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK3),s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,sK6(sK4(X0,sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0),'const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))))))),s('type.realax.real',X0)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) | ~p(s(bool,'const.nets.tends_3'(s(fun(sK0,'type.realax.real'),sK2),s('type.realax.real',sK3),s('type.pair.prod'('type.topology.topology'('type.realax.real'),fun(sK0,fun(sK0,bool))),'const.pair.,_2'(s('type.topology.topology'('type.realax.real'),'const.topology.mtop_1'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'))),s(fun(sK0,fun(sK0,bool)),sK1))))))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f538,f37])).
fof(f538,plain,(
  ( ! [X0] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.topology.dist_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,sK6(sK4(X0,sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0),'const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))),s('type.realax.real',sK3))))),s('type.realax.real',X0)))) | ~p(s(bool,'const.nets.tends_3'(s(fun(sK0,'type.realax.real'),sK2),s('type.realax.real',sK3),s('type.pair.prod'('type.topology.topology'('type.realax.real'),fun(sK0,fun(sK0,bool))),'const.pair.,_2'(s('type.topology.topology'('type.realax.real'),'const.topology.mtop_1'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'))),s(fun(sK0,fun(sK0,bool)),sK1))))))) ) | ($spl0 | $spl2)),
  inference(duplicate_literal_removal,[],[f527])).
fof(f527,plain,(
  ( ! [X0] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.topology.dist_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),sK2),s(sK0,sK6(sK4(X0,sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0),'const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))),s('type.realax.real',sK3))))),s('type.realax.real',X0)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) | ~p(s(bool,'const.nets.tends_3'(s(fun(sK0,'type.realax.real'),sK2),s('type.realax.real',sK3),s('type.pair.prod'('type.topology.topology'('type.realax.real'),fun(sK0,fun(sK0,bool))),'const.pair.,_2'(s('type.topology.topology'('type.realax.real'),'const.topology.mtop_1'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'))),s(fun(sK0,fun(sK0,bool)),sK1))))))) ) | ($spl0 | $spl2)),
  inference(resolution,[],[f414,f43])).
fof(f414,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(sK4(X0,sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0),'const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))),s(sK0,sK4(X0,sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))))) ) | ($spl0 | $spl2)),
  inference(resolution,[],[f263,f153])).
fof(f263,plain,(
  ( ! [X3] : (~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X3))),s(sK0,X3)))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(X3,'const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))),s(sK0,X3))))) ) | $spl0),
  inference(subsumption_resolution,[],[f246,f47])).
fof(f246,plain,(
  ( ! [X3] : (p(s(bool,'F_0')) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK6(X3,'const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))),s(sK0,X3)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X3))),s(sK0,X3))))) ) | $spl0),
  inference(superposition,[],[f45,f67])).
fof(f1497,plain,(
  $false | ($spl0 | $spl2 | $spl7)),
  inference(subsumption_resolution,[],[f1496,f123])).
fof(f1496,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f1495,f288])).
fof(f1495,plain,(
  s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK5('const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))) | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f1494,f48])).
fof(f1494,plain,(
  ~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK5('const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f1485,f278])).
fof(f1485,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4(sK5('const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0),sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))),s(sK0,sK4(sK5('const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0),sK3,sK2,sK1,'const.topology.mr|49|_0','type.realax.real',sK0))))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK5('const.realax.real_neg_1'(s('type.realax.real',sK3)),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2)),sK1,'const.topology.mr|49|_0','type.realax.real',sK0)))) | ($spl0 | $spl2)),
  inference(resolution,[],[f269,f559])).
fof(f1106,plain,(
  ( ! [X12,X8,X11,X9] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK5('const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),sK1)),X9,X12,X11,bool,X8))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK5('const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),sK1)),X9,X12,X11,bool,X8))))) ) | $spl8),
  inference(cnf_transformation,[],[f1106_D])).
fof(f1106_D,plain,(
  ( ! [X12,X8,X11,X9] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK5('const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),sK1)),X9,X12,X11,bool,X8))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK5('const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),sK1)),X9,X12,X11,bool,X8))))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f1107,plain,(
  ( ! [X10] : (p(s(bool,X10)) | s(bool,'F_0') = s(bool,X10)) ) | $spl10),
  inference(cnf_transformation,[],[f1107_D])).
fof(f1107_D,plain,(
  ( ! [X10] : (p(s(bool,X10)) | s(bool,'F_0') = s(bool,X10)) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f1103,plain,(
  ( ! [X12,X10,X8,X11,X9] : (p(s(bool,X10)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK5('const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),sK1)),X9,X12,X11,bool,X8))))) | s(bool,'F_0') = s(bool,X10) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK5('const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),sK1)),X9,X12,X11,bool,X8))))) )),
  inference(duplicate_literal_removal,[],[f1080])).
fof(f1080,plain,(
  ( ! [X12,X10,X8,X11,X9] : (p(s(bool,X10)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK5('const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),sK1)),X9,X12,X11,bool,X8))))) | s(bool,'F_0') = s(bool,X10) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK5('const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),sK1)),X9,X12,X11,bool,X8)))) | s(bool,'F_0') = s(bool,X10)) )),
  inference(superposition,[],[f403,f564])).
fof(f564,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.nets.tends_3'(s(fun(X1,bool),X2),s(bool,X0),s('type.pair.prod'('type.topology.topology'(bool),fun(X1,fun(X1,bool))),'const.pair.,_2'(s('type.topology.topology'(bool),'const.topology.mtop_1'(s('type.topology.metric'(bool),X3))),s(fun(X1,fun(X1,bool)),X4))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK5('const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),sK1)),X2,X4,X3,bool,X1)))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f452,f40])).
fof(f452,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'T_0') = s(bool,'const.nets.tends_3'(s(fun(X7,bool),X4),s(bool,'T_0'),s('type.pair.prod'('type.topology.topology'(bool),fun(X7,fun(X7,bool))),'const.pair.,_2'(s('type.topology.topology'(bool),'const.topology.mtop_1'(s('type.topology.metric'(bool),X6))),s(fun(X7,fun(X7,bool)),X5))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK5('const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),sK1)),X4,X5,X6,bool,X7))))) )),
  inference(subsumption_resolution,[],[f451,f47])).
fof(f451,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.nets.tends_3'(s(fun(X7,bool),X4),s(bool,'T_0'),s('type.pair.prod'('type.topology.topology'(bool),fun(X7,fun(X7,bool))),'const.pair.,_2'(s('type.topology.topology'(bool),'const.topology.mtop_1'(s('type.topology.metric'(bool),X6))),s(fun(X7,fun(X7,bool)),X5))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK5('const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),sK1)),X4,X5,X6,bool,X7))))) )),
  inference(superposition,[],[f413,f40])).
fof(f413,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK5('const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),sK1)),X5,X7,X6,bool,X4))))) | s(bool,'T_0') = s(bool,'const.nets.tends_3'(s(fun(X4,bool),X5),s(bool,'T_0'),s('type.pair.prod'('type.topology.topology'(bool),fun(X4,fun(X4,bool))),'const.pair.,_2'(s('type.topology.topology'(bool),'const.topology.mtop_1'(s('type.topology.metric'(bool),X6))),s(fun(X4,fun(X4,bool)),X7)))))) )),
  inference(subsumption_resolution,[],[f412,f47])).
fof(f412,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,'F_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK5('const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),sK1)),X5,X7,X6,bool,X4))))) | s(bool,'T_0') = s(bool,'const.nets.tends_3'(s(fun(X4,bool),X5),s(bool,'T_0'),s('type.pair.prod'('type.topology.topology'(bool),fun(X4,fun(X4,bool))),'const.pair.,_2'(s('type.topology.topology'(bool),'const.topology.mtop_1'(s('type.topology.metric'(bool),X6))),s(fun(X4,fun(X4,bool)),X7)))))) )),
  inference(superposition,[],[f85,f40])).
fof(f85,plain,(
  ( ! [X35,X33,X36,X34] : (p(s(bool,'const.nets.tends_3'(s(fun(X33,bool),X34),s(bool,'T_0'),s('type.pair.prod'('type.topology.topology'(bool),fun(X33,fun(X33,bool))),'const.pair.,_2'(s('type.topology.topology'(bool),'const.topology.mtop_1'(s('type.topology.metric'(bool),X35))),s(fun(X33,fun(X33,bool)),X36)))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK5('const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),sK1)),X34,X36,X35,bool,X33)))))) )),
  inference(superposition,[],[f44,f55])).
fof(f55,plain,(
  s(bool,'T_0') = s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),sK1)))),
  inference(subsumption_resolution,[],[f54,f47])).
fof(f54,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),sK1)))),
  inference(superposition,[],[f34,f40])).
fof(f403,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.nets.tends_3'(s(fun(X1,bool),X2),s(bool,X0),s('type.pair.prod'('type.topology.topology'(bool),fun(X1,fun(X1,bool))),'const.pair.,_2'(s('type.topology.topology'(bool),'const.topology.mtop_1'(s('type.topology.metric'(bool),X3))),s(fun(X1,fun(X1,bool)),X4)))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK5('const.nets.dorder_1'(s(fun(sK0,fun(sK0,bool)),sK1)),X2,X4,X3,bool,X1))))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f85,f40])).
fof(f122,plain,(
  $false | ($spl3 | $spl7)),
  inference(subsumption_resolution,[],[f121,f69])).
fof(f121,plain,(
  s(bool,'T_0') = s(bool,'const.nets.tends_3'(s(fun(sK0,'type.realax.real'),sK2),s('type.realax.real',sK3),s('type.pair.prod'('type.topology.topology'('type.realax.real'),fun(sK0,fun(sK0,bool))),'const.pair.,_2'(s('type.topology.topology'('type.realax.real'),'const.topology.mtop_1'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'))),s(fun(sK0,fun(sK0,bool)),sK1))))) | $spl7),
  inference(trivial_inequality_removal,[],[f120])).
fof(f120,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.nets.tends_3'(s(fun(sK0,'type.realax.real'),sK2),s('type.realax.real',sK3),s('type.pair.prod'('type.topology.topology'('type.realax.real'),fun(sK0,fun(sK0,bool))),'const.pair.,_2'(s('type.topology.topology'('type.realax.real'),'const.topology.mtop_1'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'))),s(fun(sK0,fun(sK0,bool)),sK1))))) | $spl7),
  inference(superposition,[],[f73,f40])).
fof(f71,plain,(
  s(bool,'T_0') = s(bool,'const.nets.tends_3'(s(fun(sK0,'type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK3))),s('type.pair.prod'('type.topology.topology'('type.realax.real'),fun(sK0,fun(sK0,bool))),'const.pair.,_2'(s('type.topology.topology'('type.realax.real'),'const.topology.mtop_1'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'))),s(fun(sK0,fun(sK0,bool)),sK1))))) | $spl4),
  inference(cnf_transformation,[],[f71_D])).
fof(f71_D,plain,(
  s(bool,'T_0') = s(bool,'const.nets.tends_3'(s(fun(sK0,'type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK3))),s('type.pair.prod'('type.topology.topology'('type.realax.real'),fun(sK0,fun(sK0,bool))),'const.pair.,_2'(s('type.topology.topology'('type.realax.real'),'const.topology.mtop_1'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'))),s(fun(sK0,fun(sK0,bool)),sK1))))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f73,plain,(
  s(bool,'F_0') != s(bool,'const.nets.tends_3'(s(fun(sK0,'type.realax.real'),sK2),s('type.realax.real',sK3),s('type.pair.prod'('type.topology.topology'('type.realax.real'),fun(sK0,fun(sK0,bool))),'const.pair.,_2'(s('type.topology.topology'('type.realax.real'),'const.topology.mtop_1'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'))),s(fun(sK0,fun(sK0,bool)),sK1))))) | $spl7),
  inference(cnf_transformation,[],[f73_D])).
fof(f73_D,plain,(
  s(bool,'F_0') != s(bool,'const.nets.tends_3'(s(fun(sK0,'type.realax.real'),sK2),s('type.realax.real',sK3),s('type.pair.prod'('type.topology.topology'('type.realax.real'),fun(sK0,fun(sK0,bool))),'const.pair.,_2'(s('type.topology.topology'('type.realax.real'),'const.topology.mtop_1'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'))),s(fun(sK0,fun(sK0,bool)),sK1))))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f65,plain,(
  s(bool,'F_0') != s(bool,'const.nets.tends_3'(s(fun(sK0,'type.realax.real'),sK2),s('type.realax.real',sK3),s('type.pair.prod'('type.topology.topology'('type.realax.real'),fun(sK0,fun(sK0,bool))),'const.pair.,_2'(s('type.topology.topology'('type.realax.real'),'const.topology.mtop_1'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'))),s(fun(sK0,fun(sK0,bool)),sK1))))) | s(bool,'T_0') = s(bool,'const.nets.tends_3'(s(fun(sK0,'type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK3))),s('type.pair.prod'('type.topology.topology'('type.realax.real'),fun(sK0,fun(sK0,bool))),'const.pair.,_2'(s('type.topology.topology'('type.realax.real'),'const.topology.mtop_1'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'))),s(fun(sK0,fun(sK0,bool)),sK1)))))),
  inference(superposition,[],[f50,f40])).
fof(f67,plain,(
  s(bool,'F_0') = s(bool,'const.nets.tends_3'(s(fun(sK0,'type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK3))),s('type.pair.prod'('type.topology.topology'('type.realax.real'),fun(sK0,fun(sK0,bool))),'const.pair.,_2'(s('type.topology.topology'('type.realax.real'),'const.topology.mtop_1'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'))),s(fun(sK0,fun(sK0,bool)),sK1))))) | $spl0),
  inference(cnf_transformation,[],[f67_D])).
fof(f67_D,plain,(
  s(bool,'F_0') = s(bool,'const.nets.tends_3'(s(fun(sK0,'type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK3))),s('type.pair.prod'('type.topology.topology'('type.realax.real'),fun(sK0,fun(sK0,bool))),'const.pair.,_2'(s('type.topology.topology'('type.realax.real'),'const.topology.mtop_1'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'))),s(fun(sK0,fun(sK0,bool)),sK1))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f69,plain,(
  s(bool,'T_0') != s(bool,'const.nets.tends_3'(s(fun(sK0,'type.realax.real'),sK2),s('type.realax.real',sK3),s('type.pair.prod'('type.topology.topology'('type.realax.real'),fun(sK0,fun(sK0,bool))),'const.pair.,_2'(s('type.topology.topology'('type.realax.real'),'const.topology.mtop_1'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'))),s(fun(sK0,fun(sK0,bool)),sK1))))) | $spl3),
  inference(cnf_transformation,[],[f69_D])).
fof(f69_D,plain,(
  s(bool,'T_0') != s(bool,'const.nets.tends_3'(s(fun(sK0,'type.realax.real'),sK2),s('type.realax.real',sK3),s('type.pair.prod'('type.topology.topology'('type.realax.real'),fun(sK0,fun(sK0,bool))),'const.pair.,_2'(s('type.topology.topology'('type.realax.real'),'const.topology.mtop_1'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'))),s(fun(sK0,fun(sK0,bool)),sK1))))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f64,plain,(
  s(bool,'T_0') != s(bool,'const.nets.tends_3'(s(fun(sK0,'type.realax.real'),sK2),s('type.realax.real',sK3),s('type.pair.prod'('type.topology.topology'('type.realax.real'),fun(sK0,fun(sK0,bool))),'const.pair.,_2'(s('type.topology.topology'('type.realax.real'),'const.topology.mtop_1'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'))),s(fun(sK0,fun(sK0,bool)),sK1))))) | s(bool,'F_0') = s(bool,'const.nets.tends_3'(s(fun(sK0,'type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun(sK0,'type.realax.real'),sK2))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK3))),s('type.pair.prod'('type.topology.topology'('type.realax.real'),fun(sK0,fun(sK0,bool))),'const.pair.,_2'(s('type.topology.topology'('type.realax.real'),'const.topology.mtop_1'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'))),s(fun(sK0,fun(sK0,bool)),sK1)))))),
  inference(superposition,[],[f50,f40])).
% SZS output end Proof for thm.nets
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 234
% Passive clauses: 1550
% Generated clauses: 2713
% Final active clauses: 136
% Final passive clauses: 768
% Input formulas: 11
% Initial clauses: 16
% 
% Selected by SInE selection: 11
% SInE iterations: 1
% 
% Duplicate literals: 58
% Trivial inequalities: 22
% Fw subsumption resolutions: 166
% Fw demodulations: 307
% Bw demodulations: 16
% 
% Simple tautologies: 48
% Forward subsumptions: 561
% Bw demodulations to eq. taut.: 3
% 
% Binary resolution: 99
% Forward superposition: 1206
% Backward superposition: 780
% Self superposition: 1
% Equality resolution: 8
% 
% Split clauses: 12
% Split components: 24
% SAT solver clauses: 20
% SAT solver binary clauses: 15
% 
% Sat splits: 12
% Sat splitting refutations: 7
% 
% Memory used [KB]: 7036
% Time elapsed: 0.231 s
% ------------------------------
% ------------------------------
% Success in time 11.964 s
