Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.transc
% (45597)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 198
% Passive clauses: 3183
% Generated clauses: 30884
% Final active clauses: 145
% Final passive clauses: 2673
% Input formulas: 28
% Initial clauses: 54
% 
% Duplicate literals: 736
% Trivial inequalities: 288
% Fw subsumption resolutions: 4591
% Fw demodulations: 5969
% Bw demodulations: 49
% 
% Simple tautologies: 368
% Equational tautologies: 47
% Forward subsumptions: 14090
% Fw demodulations to eq. taut.: 2
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 1017
% Factoring: 4
% Forward superposition: 9650
% Backward superposition: 8064
% Self superposition: 221
% Equality factoring: 90
% Equality resolution: 29
% 
% Split clauses: 397
% Split components: 808
% SAT solver clauses: 1200
% SAT solver unit clauses: 21
% SAT solver binary clauses: 148
% 
% Sat splits: 397
% Sat splitting refutations: 9
% 
% TWLsolver clauses: 1193
% TWLsolver calls for satisfiability: 242
% 
% Memory used [KB]: 14072
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.transc
% (45833)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 139
% Passive clauses: 1220
% Generated clauses: 15162
% Final active clauses: 129
% Final passive clauses: 1029
% Input formulas: 28
% Initial clauses: 54
% 
% Duplicate literals: 309
% Trivial inequalities: 141
% Fw subsumption resolutions: 1731
% Fw demodulations: 3043
% Bw demodulations: 61
% Condensations: 581
% Global subsumptions: 5
% 
% Simple tautologies: 129
% Equational tautologies: 29
% Forward subsumptions: 6175
% 
% Binary resolution: 932
% Factoring: 5
% Forward superposition: 4044
% Backward superposition: 2834
% Self superposition: 124
% Equality factoring: 65
% Equality resolution: 14
% 
% Split clauses: 96
% Split components: 200
% SAT solver clauses: 1817
% SAT solver unit clauses: 72
% SAT solver binary clauses: 42
% SAT solver learnt clauses: 10
% 
% Sat splits: 96
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1797
% TWLsolver calls for satisfiability: 392
% 
% Memory used [KB]: 11897
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.transc
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.transc
% SZS output start Proof for thm.transc
fof(f3380,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f100,f160,f158,f99,f159,f157,f451,f456,f461,f776,f780,f778,f1287,f1295,f1299,f1300,f1303,f1301,f1453,f1455,f1478,f1457,f1480,f1570,f1571,f1578,f1575,f1572,f1573,f1581,f1574,f1583,f1576,f1580,f1629,f1630,f1642,f1640,f1896,f1905,f1914,f1915,f1917,f1302,f1916,f1919,f1935,f1937,f1941,f1943,f1945,f1947,f1951,f1953,f1955,f1957,f1959,f1961,f1972,f1993,f1994,f1995,f1996,f2066,f2068,f2072,f2075,f2101,f2102,f2103,f2104,f2105,f2106,f2107,f2108,f2109,f2110,f2909,f2911,f2916,f2918,f2919,f2913,f2921,f3315,f3324,f3335,f3367])).
fof(f3367,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f3366,f101])).
fof(f101,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.transc.DIVISION_EQ',aHL_TRUTH)).
fof(f3366,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl3)),
  inference(forward_demodulation,[],[f3365,f484])).
fof(f484,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))))) | $spl3),
  inference(superposition,[],[f190,f483])).
fof(f483,plain,(
  s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) | $spl3),
  inference(unit_resulting_resolution,[],[f160,f109])).
fof(f109,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK3(X0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK3(X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3])],[f35])).
fof(f35,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | ? [X1] : s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X9] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X9) | ? [X10] : s('type.num.num',X9) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))))),
  file('thm.transc.DIVISION_EQ','thm.arithmetic.num_CASES')).
fof(f190,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f107,f112])).
fof(f112,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f37])).
fof(f37,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X13] : (s(bool,'T_0') = s(bool,X13) <=> p(s(bool,X13)))),
  file('thm.transc.DIVISION_EQ','thm.bool.EQ_CLAUSES|split|0')).
fof(f107,plain,(
  ( ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))) )),
  inference(cnf_transformation,[],[f34])).
fof(f34,plain,(
  ! [X0] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X10] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X10),s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))))))),
  file('thm.transc.DIVISION_EQ','thm.prim_rec.LESS_SUC_REFL')).
fof(f3365,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f3364,f104])).
fof(f104,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f31])).
fof(f31,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.transc.DIVISION_EQ',aHL_FALSITY)).
fof(f3364,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) | ($spl0 | $spl3)),
  inference(forward_demodulation,[],[f3363,f291])).
fof(f291,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X0)))) )),
  inference(unit_resulting_resolution,[],[f277,f116])).
fof(f116,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f40])).
fof(f40,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f39])).
fof(f39,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X13] : (s(bool,'F_0') = s(bool,X13) <=> ~p(s(bool,X13)))),
  file('thm.transc.DIVISION_EQ','thm.bool.EQ_CLAUSES|split|3')).
fof(f277,plain,(
  ( ! [X0] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X0))))) )),
  inference(equality_resolution,[],[f138])).
fof(f138,plain,(
  ( ! [X0,X1] : (s('type.realax.real',X0) != s('type.realax.real',X1) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1))))) )),
  inference(cnf_transformation,[],[f68])).
fof(f68,plain,(
  ! [X0,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | s('type.realax.real',X0) != s('type.realax.real',X1))),
  inference(ennf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) => s('type.realax.real',X0) != s('type.realax.real',X1))),
  inference(flattening,[],[f54])).
fof(f54,plain,(
  ! [X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) => ~s('type.realax.real',X0) = s('type.realax.real',X1))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X15,X16] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X15),s('type.realax.real',X16)))) => ~s('type.realax.real',X15) = s('type.realax.real',X16))),
  file('thm.transc.DIVISION_EQ','thm.real.REAL_LT_IMP_NE')).
fof(f3363,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real',sK1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) | ($spl0 | $spl3)),
  inference(forward_demodulation,[],[f3343,f643])).
fof(f643,plain,(
  s('type.realax.real',sK1) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))))) | ($spl0 | $spl3)),
  inference(forward_demodulation,[],[f642,f483])).
fof(f642,plain,(
  s('type.realax.real',sK1) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))))) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f641,f101])).
fof(f641,plain,(
  ~p(s(bool,'T_0')) | s('type.realax.real',sK1) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))))) | ($spl0 | $spl3)),
  inference(forward_demodulation,[],[f634,f171])).
fof(f171,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X0)))) )),
  inference(unit_resulting_resolution,[],[f105,f112])).
fof(f105,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f32])).
fof(f32,plain,(
  ! [X0] : p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X0))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X9] : p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X9),s('type.num.num',X9))))),
  file('thm.transc.DIVISION_EQ','thm.arithmetic.LESS_EQ_REFL')).
fof(f634,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) | s('type.realax.real',sK1) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))))) | ($spl0 | $spl3)),
  inference(superposition,[],[f462,f483])).
fof(f462,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))),s('type.num.num',X0)))) | s('type.realax.real',sK1) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X0)))) ) | $spl0),
  inference(backward_demodulation,[],[f157,f438])).
fof(f438,plain,(
  ( ! [X0] : (s('type.realax.real',sK2) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X0))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))),s('type.num.num',X0))))) )),
  inference(subsumption_resolution,[],[f435,f101])).
fof(f435,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s('type.realax.real',sK2) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X0))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))),s('type.num.num',X0))))) )),
  inference(superposition,[],[f163,f170])).
fof(f170,plain,(
  s(bool,'T_0') = s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s(fun('type.num.num','type.realax.real'),sK0)))),
  inference(unit_resulting_resolution,[],[f98,f112])).
fof(f98,plain,(
  p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s(fun('type.num.num','type.realax.real'),sK0))))),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK1),s('type.realax.real',sK2))),s(fun('type.num.num','type.realax.real'),sK0)))) & (s('type.realax.real',sK1) = s('type.realax.real',sK2) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))) & (s('type.realax.real',sK1) != s('type.realax.real',sK2) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f75])).
fof(f75,plain,(
  ? [X0,X1,X2] : (p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s(fun('type.num.num','type.realax.real'),X0)))) & (s('type.realax.real',X1) = s('type.realax.real',X2) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X0)))) & (s('type.realax.real',X1) != s('type.realax.real',X2) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X0)))))),
  inference(flattening,[],[f74])).
fof(f74,plain,(
  ? [X0,X1,X2] : (p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s(fun('type.num.num','type.realax.real'),X0)))) & ((s('type.realax.real',X1) = s('type.realax.real',X2) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X0)))) & (s('type.realax.real',X1) != s('type.realax.real',X2) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X0))))))),
  inference(nnf_transformation,[],[f61])).
fof(f61,plain,(
  ? [X0,X1,X2] : (p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s(fun('type.num.num','type.realax.real'),X0)))) & (s('type.realax.real',X1) = s('type.realax.real',X2) <~> s('type.num.num','const.num.|48|_0') = s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X0)))))),
  inference(ennf_transformation,[],[f30])).
fof(f30,plain,(
  ~! [X0,X1,X2] : (p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s(fun('type.num.num','type.realax.real'),X0)))) => (s('type.realax.real',X1) = s('type.realax.real',X2) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X0)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : (p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X7),s('type.realax.real',X8))),s(fun('type.num.num','type.realax.real'),X6)))) => (s('type.realax.real',X7) = s('type.realax.real',X8) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X6)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : (p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X7),s('type.realax.real',X8))),s(fun('type.num.num','type.realax.real'),X6)))) => (s('type.realax.real',X7) = s('type.realax.real',X8) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X6)))))),
  file('thm.transc.DIVISION_EQ','thm.transc.DIVISION_EQ')).
fof(f163,plain,(
  ( ! [X2,X0,X5,X1] : (~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s(fun('type.num.num','type.realax.real'),X0)))) | s('type.realax.real',X2) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X5))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X0))),s('type.num.num',X5))))) )),
  inference(forward_demodulation,[],[f147,f137])).
fof(f137,plain,(
  ( ! [X0,X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0,X1] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num',X1)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X10,X9] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X9),s('type.num.num',X10))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X10),s('type.num.num',X9)))),
  file('thm.transc.DIVISION_EQ','thm.arithmetic.GREATER_EQ')).
fof(f147,plain,(
  ( ! [X2,X0,X5,X1] : (s('type.realax.real',X2) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X5))) | ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X5),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X0)))))) | ~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s(fun('type.num.num','type.realax.real'),X0))))) )),
  inference(cnf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s(fun('type.num.num','type.realax.real'),X0)))) | (s('type.realax.real',X1) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num','const.num.|48|_0'))) & ! [X6] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X0)))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X6))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))))))))) & ! [X5] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X5),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X0)))))) | s('type.realax.real',X2) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X5)))))) & (s('type.realax.real',X1) != s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num','const.num.|48|_0'))) | (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK5(X2,X1,X0)),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X0)))))) & ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',sK5(X2,X1,X0)))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK5(X2,X1,X0)))))))))) | (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK6(X2,X1,X0)),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X0)))))) & s('type.realax.real',X2) != s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',sK6(X2,X1,X0))))) | p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s(fun('type.num.num','type.realax.real'),X0))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5,sK6])],[f96])).
fof(f96,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s(fun('type.num.num','type.realax.real'),X0)))) | (s('type.realax.real',X1) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num','const.num.|48|_0'))) & ! [X6] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X0)))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X6))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))))))))) & ! [X5] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X5),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X0)))))) | s('type.realax.real',X2) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X5)))))) & (s('type.realax.real',X1) != s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num','const.num.|48|_0'))) | ? [X4] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X0)))))) & ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X4))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))))))))) | ? [X3] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X3),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X0)))))) & s('type.realax.real',X2) != s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X3)))) | p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s(fun('type.num.num','type.realax.real'),X0))))))),
  inference(rectify,[],[f95])).
fof(f95,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s(fun('type.num.num','type.realax.real'),X0)))) | (s('type.realax.real',X1) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num','const.num.|48|_0'))) & ! [X4] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X0)))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X4))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))))))))) & ! [X3] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X3),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X0)))))) | s('type.realax.real',X2) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X3)))))) & (s('type.realax.real',X1) != s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num','const.num.|48|_0'))) | ? [X4] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X0)))))) & ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X4))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))))))))) | ? [X3] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X3),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X0)))))) & s('type.realax.real',X2) != s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X3)))) | p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s(fun('type.num.num','type.realax.real'),X0))))))),
  inference(flattening,[],[f94])).
fof(f94,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s(fun('type.num.num','type.realax.real'),X0)))) | (s('type.realax.real',X1) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num','const.num.|48|_0'))) & ! [X4] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X0)))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X4))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))))))))) & ! [X3] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X3),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X0)))))) | s('type.realax.real',X2) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X3)))))) & ((s('type.realax.real',X1) != s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num','const.num.|48|_0'))) | ? [X4] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X0)))))) & ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X4))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))))))))) | ? [X3] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X3),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X0)))))) & s('type.realax.real',X2) != s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X3))))) | p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s(fun('type.num.num','type.realax.real'),X0))))))),
  inference(nnf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s(fun('type.num.num','type.realax.real'),X0)))) <=> (s('type.realax.real',X1) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num','const.num.|48|_0'))) & ! [X4] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X0)))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X4))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))))))))) & ! [X3] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X3),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X0)))))) | s('type.realax.real',X2) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X3))))))),
  inference(ennf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s(fun('type.num.num','type.realax.real'),X0)))) <=> (s('type.realax.real',X1) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num','const.num.|48|_0'))) & ! [X4] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X0)))))) => p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X4))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))))))))) & ! [X3] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X3),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X0)))))) => s('type.realax.real',X2) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X3))))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X6,X7,X8] : (p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X7),s('type.realax.real',X8))),s(fun('type.num.num','type.realax.real'),X6)))) <=> (s('type.realax.real',X7) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X6),s('type.num.num','const.num.|48|_0'))) & ! [X10] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X10),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X6)))))) => p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X6),s('type.num.num',X10))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X6),s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))))))))) & ! [X10] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X10),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X6)))))) => s('type.realax.real',X8) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X6),s('type.num.num',X10))))))),
  file('thm.transc.DIVISION_EQ','thm.transc.DIVISION_THM')).
fof(f3343,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) | ($spl0 | $spl3)),
  inference(superposition,[],[f523,f483])).
fof(f523,plain,(
  ( ! [X0] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))))))) ) | $spl0),
  inference(forward_demodulation,[],[f522,f395])).
fof(f395,plain,(
  s('type.realax.real',sK1) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num','const.num.|48|_0')))),
  inference(subsumption_resolution,[],[f393,f101])).
fof(f393,plain,(
  ~p(s(bool,'T_0')) | s('type.realax.real',sK1) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f145,f170])).
fof(f145,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s(fun('type.num.num','type.realax.real'),X0)))) | s('type.realax.real',X1) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f97])).
fof(f522,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))))) ) | $spl0),
  inference(subsumption_resolution,[],[f509,f101])).
fof(f509,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))))) ) | $spl0),
  inference(superposition,[],[f144,f464])).
fof(f464,plain,(
  s(bool,'T_0') = s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK1),s('type.realax.real',sK1))),s(fun('type.num.num','type.realax.real'),sK0))) | $spl0),
  inference(backward_demodulation,[],[f157,f170])).
fof(f144,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s(fun('type.num.num','type.realax.real'),X0)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X0)))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))))))))) )),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s(fun('type.num.num','type.realax.real'),X0)))) | ! [X3] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X0)))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))))))))))),
  inference(ennf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s(fun('type.num.num','type.realax.real'),X0)))) => ! [X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X0)))))) => p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))))))))))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X6,X7,X8] : (p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X7),s('type.realax.real',X8))),s(fun('type.num.num','type.realax.real'),X6)))) => ! [X10] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X10),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X6)))))) => p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X6),s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X6),s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))))))))))),
  file('thm.transc.DIVISION_EQ','thm.transc.DIVISION_LT')).
fof(f3335,plain,(
  $false | ($spl0 | $spl3 | $spl7 | $spl15 | $spl38)),
  inference(subsumption_resolution,[],[f3334,f101])).
fof(f3334,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl3 | $spl7 | $spl15 | $spl38)),
  inference(forward_demodulation,[],[f3333,f2937])).
fof(f2937,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))))) | ($spl3 | $spl7 | $spl15 | $spl38)),
  inference(backward_demodulation,[],[f2929,f484])).
fof(f2929,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num',sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))) | ($spl7 | $spl15 | $spl38)),
  inference(backward_demodulation,[],[f2927,f1306])).
fof(f1306,plain,(
  s('type.num.num',sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK3(sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))))))) | $spl7),
  inference(unit_resulting_resolution,[],[f779,f109])).
fof(f779,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))) | $spl7),
  inference(cnf_transformation,[],[f779_D])).
fof(f779_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f2927,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num',sK3(sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))))) | ($spl15 | $spl38)),
  inference(backward_demodulation,[],[f2913,f2271])).
fof(f2271,plain,(
  s('type.num.num',sK3(sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK3(sK3(sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))))) | $spl15),
  inference(unit_resulting_resolution,[],[f1456,f109])).
fof(f1456,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK3(sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))))) | $spl15),
  inference(cnf_transformation,[],[f1456_D])).
fof(f1456_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK3(sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))))) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f3333,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) | ($spl0 | $spl3 | $spl7 | $spl15 | $spl38)),
  inference(subsumption_resolution,[],[f3332,f104])).
fof(f3332,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) | ($spl0 | $spl3 | $spl7 | $spl15 | $spl38)),
  inference(forward_demodulation,[],[f3331,f291])).
fof(f3331,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real',sK1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) | ($spl0 | $spl3 | $spl7 | $spl15 | $spl38)),
  inference(forward_demodulation,[],[f3302,f643])).
fof(f3302,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) | ($spl0 | $spl3 | $spl7 | $spl15 | $spl38)),
  inference(superposition,[],[f523,f2936])).
fof(f2936,plain,(
  s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | ($spl3 | $spl7 | $spl15 | $spl38)),
  inference(backward_demodulation,[],[f2929,f483])).
fof(f3324,plain,(
  $false | ($spl0 | $spl3 | $spl7 | $spl15 | $spl38)),
  inference(subsumption_resolution,[],[f3323,f101])).
fof(f3323,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl3 | $spl7 | $spl15 | $spl38)),
  inference(forward_demodulation,[],[f3322,f2937])).
fof(f3322,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) | ($spl0 | $spl3 | $spl7 | $spl15 | $spl38)),
  inference(forward_demodulation,[],[f3321,f2929])).
fof(f3321,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) | ($spl0 | $spl3 | $spl7 | $spl15 | $spl38)),
  inference(subsumption_resolution,[],[f3320,f104])).
fof(f3320,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) | ($spl0 | $spl3 | $spl7 | $spl15 | $spl38)),
  inference(forward_demodulation,[],[f3319,f291])).
fof(f3319,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real',sK1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) | ($spl0 | $spl3 | $spl7 | $spl15 | $spl38)),
  inference(forward_demodulation,[],[f3318,f643])).
fof(f3318,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) | ($spl0 | $spl3 | $spl7 | $spl15 | $spl38)),
  inference(forward_demodulation,[],[f3294,f2936])).
fof(f3294,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) | ($spl0 | $spl7 | $spl15 | $spl38)),
  inference(superposition,[],[f523,f2929])).
fof(f3315,plain,(
  $false | ($spl0 | $spl3 | $spl7 | $spl15 | $spl38)),
  inference(subsumption_resolution,[],[f3314,f101])).
fof(f3314,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl3 | $spl7 | $spl15 | $spl38)),
  inference(forward_demodulation,[],[f3313,f2937])).
fof(f3313,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) | ($spl0 | $spl3 | $spl7 | $spl15 | $spl38)),
  inference(forward_demodulation,[],[f3312,f2938])).
fof(f2938,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))))) | ($spl3 | $spl7 | $spl15 | $spl38)),
  inference(backward_demodulation,[],[f2929,f485])).
fof(f485,plain,(
  s('type.num.num',sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))))) | $spl3),
  inference(superposition,[],[f106,f483])).
fof(f106,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f33])).
fof(f33,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X9] : s('type.num.num',X9) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X9)))))),
  file('thm.transc.DIVISION_EQ','thm.prim_rec.PRE|split|1')).
fof(f3312,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) | ($spl0 | $spl3 | $spl7 | $spl15 | $spl38)),
  inference(subsumption_resolution,[],[f3311,f104])).
fof(f3311,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) | ($spl0 | $spl3 | $spl7 | $spl15 | $spl38)),
  inference(forward_demodulation,[],[f3310,f291])).
fof(f3310,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real',sK1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) | ($spl0 | $spl3 | $spl7 | $spl15 | $spl38)),
  inference(forward_demodulation,[],[f3309,f643])).
fof(f3309,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) | ($spl0 | $spl3 | $spl7 | $spl15 | $spl38)),
  inference(forward_demodulation,[],[f3291,f2936])).
fof(f3291,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) | ($spl0 | $spl3 | $spl7 | $spl15 | $spl38)),
  inference(superposition,[],[f523,f2938])).
fof(f2921,plain,(
  s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',sK3(sK3(sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))))) = s('type.num.num',sK3(sK3(sK3('const.prim_rec.PRE_1'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))))) | $spl42),
  inference(cnf_transformation,[],[f2921_D])).
fof(f2921_D,plain,(
  s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',sK3(sK3(sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))))) = s('type.num.num',sK3(sK3(sK3('const.prim_rec.PRE_1'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))))) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f2913,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK3(sK3(sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) | $spl38),
  inference(cnf_transformation,[],[f2913_D])).
fof(f2913_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK3(sK3(sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f2919,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK3(sK3(sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) | s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',sK3(sK3(sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))))) = s('type.num.num',sK3(sK3(sK3('const.prim_rec.PRE_1'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))))) | ($spl12 | $spl15)),
  inference(forward_demodulation,[],[f2880,f2291])).
fof(f2291,plain,(
  s('type.num.num',sK3(sK3(sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) = s('type.num.num',sK3(sK3('const.prim_rec.PRE_1'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))))))) | ($spl12 | $spl15)),
  inference(backward_demodulation,[],[f2283,f1455])).
fof(f2283,plain,(
  s('type.num.num',sK3(sK3(sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',sK3(sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))))))) | $spl15),
  inference(superposition,[],[f106,f2271])).
fof(f2880,plain,(
  s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',sK3(sK3(sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))))) = s('type.num.num',sK3(sK3(sK3('const.prim_rec.PRE_1'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK3(sK3('const.prim_rec.PRE_1'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))))))) | ($spl12 | $spl15)),
  inference(superposition,[],[f281,f2291])).
fof(f281,plain,(
  ( ! [X1] : (s('type.num.num',sK3(X1)) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',X1))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) )),
  inference(superposition,[],[f106,f109])).
fof(f2918,plain,(
  s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',sK3(sK3(sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))))) = s('type.num.num',sK3(sK3('const.prim_rec.PRE_1'(s('type.num.num',sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))))) | $spl40),
  inference(cnf_transformation,[],[f2918_D])).
fof(f2918_D,plain,(
  s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',sK3(sK3(sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))))) = s('type.num.num',sK3(sK3('const.prim_rec.PRE_1'(s('type.num.num',sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))))) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f2916,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK3(sK3(sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) | s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',sK3(sK3(sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))))) = s('type.num.num',sK3(sK3('const.prim_rec.PRE_1'(s('type.num.num',sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))))) | ($spl15 | $spl16)),
  inference(forward_demodulation,[],[f2879,f2292])).
fof(f2292,plain,(
  s('type.num.num',sK3(sK3(sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) = s('type.num.num',sK3('const.prim_rec.PRE_1'(s('type.num.num',sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))))))) | ($spl15 | $spl16)),
  inference(backward_demodulation,[],[f2283,f1480])).
fof(f2879,plain,(
  s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',sK3(sK3(sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))))) = s('type.num.num',sK3(sK3('const.prim_rec.PRE_1'(s('type.num.num',sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK3('const.prim_rec.PRE_1'(s('type.num.num',sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))))))) | ($spl15 | $spl16)),
  inference(superposition,[],[f281,f2292])).
fof(f2911,plain,(
  s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',sK3(sK3(sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))))) = s('type.num.num',sK3('const.prim_rec.PRE_1'(s('type.num.num',sK3(sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))))) | $spl36),
  inference(cnf_transformation,[],[f2911_D])).
fof(f2911_D,plain,(
  s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',sK3(sK3(sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))))) = s('type.num.num',sK3('const.prim_rec.PRE_1'(s('type.num.num',sK3(sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))))) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f2909,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK3(sK3(sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) | s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',sK3(sK3(sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))))) = s('type.num.num',sK3('const.prim_rec.PRE_1'(s('type.num.num',sK3(sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))))) | $spl15),
  inference(forward_demodulation,[],[f2876,f2283])).
fof(f2876,plain,(
  s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',sK3(sK3(sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))))) = s('type.num.num',sK3('const.prim_rec.PRE_1'(s('type.num.num',sK3(sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',sK3(sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))))))) | $spl15),
  inference(superposition,[],[f281,f2283])).
fof(f2110,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f1936,f1949])).
fof(f1949,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1)) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1302,f1302,f142])).
fof(f142,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f70])).
fof(f70,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f69])).
fof(f69,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f56])).
fof(f56,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X11,X12] : ((p(s(bool,X11)) => p(s(bool,X12))) => ((p(s(bool,X12)) => p(s(bool,X11))) => s(bool,X11) = s(bool,X12)))),
  file('thm.transc.DIVISION_EQ','thm.bool.IMP_ANTISYM_AX')).
fof(f1936,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X1)))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f292,f1302,f204])).
fof(f204,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f115,f108])).
fof(f108,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.transc.DIVISION_EQ',aHL_BOOL_CASES)).
fof(f115,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f80])).
fof(f292,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X0)))) )),
  inference(unit_resulting_resolution,[],[f277,f111])).
fof(f111,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f78])).
fof(f2109,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f1940,f101])).
fof(f1940,plain,(
  ~p(s(bool,'T_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f168,f1302,f142])).
fof(f168,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f104,f111])).
fof(f2108,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f1942,f101])).
fof(f1942,plain,(
  ~p(s(bool,'T_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f292,f1302,f142])).
fof(f2107,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f1944,f1302])).
fof(f1944,plain,(
  ~p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f168,f1302,f142])).
fof(f2106,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f1946,f1302])).
fof(f1946,plain,(
  ( ! [X0] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X0))))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f292,f1302,f142])).
fof(f2105,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f1950,f101])).
fof(f1950,plain,(
  ~p(s(bool,'T_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f168,f1302,f142])).
fof(f2104,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f1952,f101])).
fof(f1952,plain,(
  ~p(s(bool,'T_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f292,f1302,f142])).
fof(f2103,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f1954,f1302])).
fof(f1954,plain,(
  ~p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f168,f1302,f142])).
fof(f2102,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f1956,f1302])).
fof(f1956,plain,(
  ( ! [X0] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X0))))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f292,f1302,f142])).
fof(f2101,plain,(
  $false | ($spl0 | $spl10)),
  inference(subsumption_resolution,[],[f2100,f1971])).
fof(f1971,plain,(
  ( ! [X0,X1] : (s('type.realax.real',X0) != s('type.realax.real',X1)) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1302,f138])).
fof(f2100,plain,(
  ( ! [X0] : (s('type.realax.real',sK1) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X0)))) ) | ($spl0 | $spl10)),
  inference(subsumption_resolution,[],[f2045,f101])).
fof(f2045,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s('type.realax.real',sK1) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X0)))) ) | ($spl0 | $spl10)),
  inference(backward_demodulation,[],[f1960,f462])).
fof(f1960,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1302,f112])).
fof(f2075,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f2074,f1302])).
fof(f2074,plain,(
  ( ! [X0,X5] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X0))),s('type.num.num',X5))))) ) | $spl10),
  inference(subsumption_resolution,[],[f2073,f1971])).
fof(f2073,plain,(
  ( ! [X2,X0,X5] : (s('type.realax.real',X2) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X5))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X0))),s('type.num.num',X5))))) ) | $spl10),
  inference(subsumption_resolution,[],[f2020,f101])).
fof(f2020,plain,(
  ( ! [X2,X0,X5] : (~p(s(bool,'T_0')) | s('type.realax.real',X2) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X5))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X0))),s('type.num.num',X5))))) ) | $spl10),
  inference(backward_demodulation,[],[f1960,f163])).
fof(f2072,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f2071,f1971])).
fof(f2071,plain,(
  ( ! [X0,X1] : (s('type.realax.real',X1) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num','const.num.|48|_0')))) ) | $spl10),
  inference(subsumption_resolution,[],[f2014,f101])).
fof(f2014,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s('type.realax.real',X1) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num','const.num.|48|_0')))) ) | $spl10),
  inference(backward_demodulation,[],[f1960,f145])).
fof(f2068,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f2067,f1302])).
fof(f2067,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl10),
  inference(subsumption_resolution,[],[f1998,f1960])).
fof(f1998,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl10),
  inference(backward_demodulation,[],[f1960,f115])).
fof(f2066,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f1997,f101])).
fof(f1997,plain,(
  ~p(s(bool,'T_0')) | $spl10),
  inference(backward_demodulation,[],[f1960,f104])).
fof(f1996,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f1964,f1971])).
fof(f1964,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real',X0) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X1),s('type.num.num',X2)))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1302,f1302,f163])).
fof(f1995,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f1965,f1971])).
fof(f1965,plain,(
  ( ! [X0,X1] : (s('type.realax.real',X0) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X1),s('type.num.num','const.num.|48|_0')))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1302,f145])).
fof(f1994,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f1966,f1971])).
fof(f1966,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real',X0) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X1),s('type.num.num',X2)))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1302,f1302,f163])).
fof(f1993,plain,(
  $false | ($spl0 | $spl10)),
  inference(subsumption_resolution,[],[f1967,f1971])).
fof(f1967,plain,(
  ( ! [X0] : (s('type.realax.real',sK1) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X0)))) ) | ($spl0 | $spl10)),
  inference(unit_resulting_resolution,[],[f1302,f462])).
fof(f1972,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f395,f1302,f138])).
fof(f1961,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f168,f1302,f112])).
fof(f1959,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f291,f1302,f115])).
fof(f1957,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1302,f292,f1302,f142])).
fof(f1955,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1302,f168,f1302,f142])).
fof(f1953,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f101,f292,f1302,f142])).
fof(f1951,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f101,f168,f1302,f142])).
fof(f1947,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1302,f292,f1302,f142])).
fof(f1945,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1302,f168,f1302,f142])).
fof(f1943,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f101,f292,f1302,f142])).
fof(f1941,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f101,f168,f1302,f142])).
fof(f1937,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f291,f292,f1302,f204])).
fof(f1935,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f291,f168,f1302,f204])).
fof(f1919,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f104,f1302])).
fof(f1916,plain,(
  ( ! [X4] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X4)))))))) ) | $spl34),
  inference(cnf_transformation,[],[f1916_D])).
fof(f1916_D,plain,(
  ( ! [X4] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X4)))))))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f1302,plain,(
  ( ! [X5] : (p(s(bool,X5))) ) | $spl10),
  inference(cnf_transformation,[],[f1302_D])).
fof(f1302_D,plain,(
  ( ! [X5] : (p(s(bool,X5))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f1917,plain,(
  ( ! [X6,X7] : (p(s(bool,X7)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))))) ) | $spl0),
  inference(subsumption_resolution,[],[f1887,f111])).
fof(f1887,plain,(
  ( ! [X6,X7] : (p(s(bool,X7)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) | s(bool,'T_0') = s(bool,X7) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))))) ) | $spl0),
  inference(superposition,[],[f523,f201])).
fof(f201,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f108,f108])).
fof(f1915,plain,(
  ( ! [X4,X5] : (p(s(bool,X5)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X4)))))))) ) | $spl0),
  inference(subsumption_resolution,[],[f1886,f111])).
fof(f1886,plain,(
  ( ! [X4,X5] : (p(s(bool,X5)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))))))) | s(bool,'T_0') = s(bool,X5)) ) | $spl0),
  inference(superposition,[],[f523,f201])).
fof(f1914,plain,(
  $false | ($spl0 | $spl3 | $spl7 | $spl14)),
  inference(subsumption_resolution,[],[f1913,f101])).
fof(f1913,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl3 | $spl7 | $spl14)),
  inference(forward_demodulation,[],[f1912,f1493])).
fof(f1493,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))))) | ($spl3 | $spl7 | $spl14)),
  inference(backward_demodulation,[],[f1490,f484])).
fof(f1490,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num',sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))) | ($spl7 | $spl14)),
  inference(backward_demodulation,[],[f1457,f1306])).
fof(f1912,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) | ($spl0 | $spl3 | $spl7 | $spl14)),
  inference(subsumption_resolution,[],[f1911,f104])).
fof(f1911,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) | ($spl0 | $spl3 | $spl7 | $spl14)),
  inference(forward_demodulation,[],[f1910,f291])).
fof(f1910,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real',sK1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) | ($spl0 | $spl3 | $spl7 | $spl14)),
  inference(forward_demodulation,[],[f1883,f643])).
fof(f1883,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) | ($spl0 | $spl3 | $spl7 | $spl14)),
  inference(superposition,[],[f523,f1492])).
fof(f1492,plain,(
  s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl3 | $spl7 | $spl14)),
  inference(backward_demodulation,[],[f1490,f483])).
fof(f1905,plain,(
  $false | ($spl0 | $spl3 | $spl7 | $spl14)),
  inference(subsumption_resolution,[],[f1904,f101])).
fof(f1904,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl3 | $spl7 | $spl14)),
  inference(forward_demodulation,[],[f1903,f1493])).
fof(f1903,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) | ($spl0 | $spl3 | $spl7 | $spl14)),
  inference(forward_demodulation,[],[f1902,f1490])).
fof(f1902,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) | ($spl0 | $spl3 | $spl7 | $spl14)),
  inference(subsumption_resolution,[],[f1901,f104])).
fof(f1901,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) | ($spl0 | $spl3 | $spl7 | $spl14)),
  inference(forward_demodulation,[],[f1900,f291])).
fof(f1900,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real',sK1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) | ($spl0 | $spl3 | $spl7 | $spl14)),
  inference(forward_demodulation,[],[f1899,f643])).
fof(f1899,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) | ($spl0 | $spl3 | $spl7 | $spl14)),
  inference(forward_demodulation,[],[f1877,f1492])).
fof(f1877,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) | ($spl0 | $spl7 | $spl14)),
  inference(superposition,[],[f523,f1490])).
fof(f1896,plain,(
  $false | ($spl0 | $spl3 | $spl7 | $spl14)),
  inference(subsumption_resolution,[],[f1895,f101])).
fof(f1895,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl3 | $spl7 | $spl14)),
  inference(forward_demodulation,[],[f1894,f1493])).
fof(f1894,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) | ($spl0 | $spl3 | $spl7 | $spl14)),
  inference(forward_demodulation,[],[f1893,f1494])).
fof(f1494,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))))) | ($spl3 | $spl7 | $spl14)),
  inference(backward_demodulation,[],[f1490,f485])).
fof(f1893,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) | ($spl0 | $spl3 | $spl7 | $spl14)),
  inference(subsumption_resolution,[],[f1892,f104])).
fof(f1892,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) | ($spl0 | $spl3 | $spl7 | $spl14)),
  inference(forward_demodulation,[],[f1891,f291])).
fof(f1891,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real',sK1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) | ($spl0 | $spl3 | $spl7 | $spl14)),
  inference(forward_demodulation,[],[f1890,f643])).
fof(f1890,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) | ($spl0 | $spl3 | $spl7 | $spl14)),
  inference(forward_demodulation,[],[f1874,f1492])).
fof(f1874,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) | ($spl0 | $spl3 | $spl7 | $spl14)),
  inference(superposition,[],[f523,f1494])).
fof(f1640,plain,(
  ( ! [X2,X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X2) = s(bool,sK4(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK4(X0,X1,bool,bool))))) ) | $spl30),
  inference(cnf_transformation,[],[f1640_D])).
fof(f1640_D,plain,(
  ( ! [X2,X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X2) = s(bool,sK4(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK4(X0,X1,bool,bool))))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f1642,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | $spl33),
  inference(cnf_transformation,[],[f1642_D])).
fof(f1642_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f1630,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,g(s(bool,sK4(X3,X4,bool,bool)))) | s(bool,X5) = s(bool,sK4(X3,X4,bool,bool)) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f362,f201])).
fof(f362,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK4(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK4(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f143,f108])).
fof(f143,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK4(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK4(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f93])).
fof(f93,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK4(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK4(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f71])).
fof(f71,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f57])).
fof(f57,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.transc.DIVISION_EQ',aHL_EXT)).
fof(f1629,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,g(s(bool,sK4(X0,X1,bool,bool)))) | s(bool,X2) = s(bool,sK4(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f362,f201])).
fof(f1580,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl24),
  inference(cnf_transformation,[],[f1580_D])).
fof(f1580_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f1576,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK4(X0,X1,bool,bool))) ) | $spl20),
  inference(cnf_transformation,[],[f1576_D])).
fof(f1576_D,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK4(X0,X1,bool,bool))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f1583,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl29),
  inference(cnf_transformation,[],[f1583_D])).
fof(f1583_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f1574,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X12) = s(fun(bool,bool),X13) | s(bool,'T_0') = s(bool,sK4(X12,X13,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f361,f108])).
fof(f361,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK4(X0,X1,X2,bool))) )),
  inference(superposition,[],[f143,f108])).
fof(f1581,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl26),
  inference(cnf_transformation,[],[f1581_D])).
fof(f1581_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f1573,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X10) = s(fun(bool,bool),X11) | s(bool,'T_0') = s(bool,sK4(X10,X11,bool,bool)) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f361,f201])).
fof(f1572,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X7) = s(fun(bool,bool),X8) | s(bool,'T_0') = s(bool,sK4(X7,X8,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f361,f201])).
fof(f1575,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl18),
  inference(cnf_transformation,[],[f1575_D])).
fof(f1575_D,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f1578,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | $spl23),
  inference(cnf_transformation,[],[f1578_D])).
fof(f1578_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f1571,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,sK4(X3,X4,bool,bool)) | s(bool,X5) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f361,f201])).
fof(f1570,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK4(X0,X1,bool,bool)) | s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f361,f201])).
fof(f1480,plain,(
  s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',sK3(sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))))))) = s('type.num.num',sK3('const.prim_rec.PRE_1'(s('type.num.num',sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))))))) | $spl16),
  inference(cnf_transformation,[],[f1480_D])).
fof(f1480_D,plain,(
  s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',sK3(sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))))))) = s('type.num.num',sK3('const.prim_rec.PRE_1'(s('type.num.num',sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))))))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f1457,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK3(sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))))) | $spl14),
  inference(cnf_transformation,[],[f1457_D])).
fof(f1457_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK3(sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f1478,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK3(sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))))) | s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',sK3(sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))))))) = s('type.num.num',sK3('const.prim_rec.PRE_1'(s('type.num.num',sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))))))) | $spl7),
  inference(forward_demodulation,[],[f1470,f1305])).
fof(f1305,plain,(
  s('type.num.num',sK3(sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))))) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) | $spl7),
  inference(unit_resulting_resolution,[],[f779,f281])).
fof(f1470,plain,(
  s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',sK3(sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))))))) = s('type.num.num',sK3('const.prim_rec.PRE_1'(s('type.num.num',sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) | $spl7),
  inference(superposition,[],[f281,f1305])).
fof(f1455,plain,(
  s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',sK3(sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))))))) = s('type.num.num',sK3(sK3('const.prim_rec.PRE_1'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))))))) | $spl12),
  inference(cnf_transformation,[],[f1455_D])).
fof(f1455_D,plain,(
  s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',sK3(sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))))))) = s('type.num.num',sK3(sK3('const.prim_rec.PRE_1'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))))))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f1453,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK3(sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))))) | s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',sK3(sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))))))) = s('type.num.num',sK3(sK3('const.prim_rec.PRE_1'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))))))) | ($spl4 | $spl7)),
  inference(forward_demodulation,[],[f1444,f1307])).
fof(f1307,plain,(
  s('type.num.num',sK3(sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))))) = s('type.num.num',sK3('const.prim_rec.PRE_1'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) | ($spl4 | $spl7)),
  inference(backward_demodulation,[],[f1305,f778])).
fof(f1444,plain,(
  s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',sK3(sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))))))) = s('type.num.num',sK3(sK3('const.prim_rec.PRE_1'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK3('const.prim_rec.PRE_1'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) | ($spl4 | $spl7)),
  inference(superposition,[],[f281,f1307])).
fof(f1301,plain,(
  ( ! [X4] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X4)))))))) ) | $spl8),
  inference(cnf_transformation,[],[f1301_D])).
fof(f1301_D,plain,(
  ( ! [X4] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X4)))))))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f1303,plain,(
  ( ! [X6,X7] : (p(s(bool,X7)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))))) ) | ($spl0 | $spl3 | $spl6)),
  inference(subsumption_resolution,[],[f1279,f111])).
fof(f1279,plain,(
  ( ! [X6,X7] : (p(s(bool,X7)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,'T_0') = s(bool,X7) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))))) ) | ($spl0 | $spl3 | $spl6)),
  inference(superposition,[],[f788,f201])).
fof(f788,plain,(
  ( ! [X0] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl0 | $spl3 | $spl6)),
  inference(backward_demodulation,[],[f786,f523])).
fof(f786,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))) | ($spl3 | $spl6)),
  inference(backward_demodulation,[],[f780,f483])).
fof(f1300,plain,(
  ( ! [X4,X5] : (p(s(bool,X5)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X4)))))))) ) | ($spl0 | $spl3 | $spl6)),
  inference(subsumption_resolution,[],[f1278,f111])).
fof(f1278,plain,(
  ( ! [X4,X5] : (p(s(bool,X5)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))))))) | s(bool,'T_0') = s(bool,X5)) ) | ($spl0 | $spl3 | $spl6)),
  inference(superposition,[],[f788,f201])).
fof(f1299,plain,(
  $false | ($spl0 | $spl3 | $spl6)),
  inference(subsumption_resolution,[],[f1298,f101])).
fof(f1298,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl3 | $spl6)),
  inference(forward_demodulation,[],[f1297,f190])).
fof(f1297,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl0 | $spl3 | $spl6)),
  inference(subsumption_resolution,[],[f1296,f104])).
fof(f1296,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl0 | $spl3 | $spl6)),
  inference(forward_demodulation,[],[f1275,f291])).
fof(f1275,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real',sK1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl0 | $spl3 | $spl6)),
  inference(superposition,[],[f788,f791])).
fof(f791,plain,(
  s('type.realax.real',sK1) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl3 | $spl6)),
  inference(backward_demodulation,[],[f786,f643])).
fof(f1295,plain,(
  $false | ($spl0 | $spl3 | $spl6)),
  inference(subsumption_resolution,[],[f1294,f101])).
fof(f1294,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl3 | $spl6)),
  inference(forward_demodulation,[],[f1293,f190])).
fof(f1293,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl0 | $spl3 | $spl6)),
  inference(forward_demodulation,[],[f1292,f780])).
fof(f1292,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl0 | $spl3 | $spl6)),
  inference(subsumption_resolution,[],[f1291,f104])).
fof(f1291,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl0 | $spl3 | $spl6)),
  inference(forward_demodulation,[],[f1290,f291])).
fof(f1290,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real',sK1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl0 | $spl3 | $spl6)),
  inference(forward_demodulation,[],[f1273,f791])).
fof(f1273,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl0 | $spl3 | $spl6)),
  inference(superposition,[],[f788,f780])).
fof(f1287,plain,(
  $false | ($spl0 | $spl3 | $spl6)),
  inference(subsumption_resolution,[],[f1286,f101])).
fof(f1286,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl3 | $spl6)),
  inference(forward_demodulation,[],[f1285,f190])).
fof(f1285,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl0 | $spl3 | $spl6)),
  inference(forward_demodulation,[],[f1284,f103])).
fof(f103,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0')))),
  inference(cnf_transformation,[],[f24])).
fof(f24,axiom,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0')))),
  file('thm.transc.DIVISION_EQ','thm.prim_rec.PRE|split|0')).
fof(f1284,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl0 | $spl3 | $spl6)),
  inference(subsumption_resolution,[],[f1283,f104])).
fof(f1283,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl0 | $spl3 | $spl6)),
  inference(forward_demodulation,[],[f1282,f291])).
fof(f1282,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real',sK1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl0 | $spl3 | $spl6)),
  inference(forward_demodulation,[],[f1270,f791])).
fof(f1270,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl0 | $spl3 | $spl6)),
  inference(superposition,[],[f788,f103])).
fof(f778,plain,(
  s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) = s('type.num.num',sK3('const.prim_rec.PRE_1'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) | $spl4),
  inference(cnf_transformation,[],[f778_D])).
fof(f778_D,plain,(
  s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) = s('type.num.num',sK3('const.prim_rec.PRE_1'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f780,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))) | $spl6),
  inference(cnf_transformation,[],[f780_D])).
fof(f780_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f776,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))) | s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) = s('type.num.num',sK3('const.prim_rec.PRE_1'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) | $spl3),
  inference(forward_demodulation,[],[f751,f485])).
fof(f751,plain,(
  s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',sK3('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) = s('type.num.num',sK3('const.prim_rec.PRE_1'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))))) | $spl3),
  inference(superposition,[],[f281,f485])).
fof(f461,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f460,f158])).
fof(f460,plain,(
  s('type.realax.real',sK1) = s('type.realax.real',sK2) | $spl2),
  inference(forward_demodulation,[],[f459,f395])).
fof(f459,plain,(
  s('type.realax.real',sK2) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num','const.num.|48|_0'))) | $spl2),
  inference(subsumption_resolution,[],[f445,f101])).
fof(f445,plain,(
  ~p(s(bool,'T_0')) | s('type.realax.real',sK2) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num','const.num.|48|_0'))) | $spl2),
  inference(superposition,[],[f439,f171])).
fof(f439,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | s('type.realax.real',sK2) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X0)))) ) | $spl2),
  inference(forward_demodulation,[],[f438,f159])).
fof(f456,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f455,f158])).
fof(f455,plain,(
  s('type.realax.real',sK1) = s('type.realax.real',sK2) | $spl2),
  inference(forward_demodulation,[],[f454,f395])).
fof(f454,plain,(
  s('type.realax.real',sK2) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num','const.num.|48|_0'))) | $spl2),
  inference(forward_demodulation,[],[f453,f103])).
fof(f453,plain,(
  s('type.realax.real',sK2) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))))) | $spl2),
  inference(subsumption_resolution,[],[f452,f101])).
fof(f452,plain,(
  ~p(s(bool,'T_0')) | s('type.realax.real',sK2) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))))) | $spl2),
  inference(forward_demodulation,[],[f442,f171])).
fof(f442,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | s('type.realax.real',sK2) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))))) | $spl2),
  inference(superposition,[],[f439,f103])).
fof(f451,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f450,f158])).
fof(f450,plain,(
  s('type.realax.real',sK1) = s('type.realax.real',sK2) | $spl2),
  inference(forward_demodulation,[],[f449,f395])).
fof(f449,plain,(
  s('type.realax.real',sK2) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num','const.num.|48|_0'))) | $spl2),
  inference(forward_demodulation,[],[f448,f159])).
fof(f448,plain,(
  s('type.realax.real',sK2) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))))) | $spl2),
  inference(subsumption_resolution,[],[f447,f101])).
fof(f447,plain,(
  ~p(s(bool,'T_0')) | s('type.realax.real',sK2) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))))) | $spl2),
  inference(forward_demodulation,[],[f440,f171])).
fof(f440,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | s('type.realax.real',sK2) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))))) | $spl2),
  inference(superposition,[],[f439,f159])).
fof(f157,plain,(
  s('type.realax.real',sK1) = s('type.realax.real',sK2) | $spl0),
  inference(cnf_transformation,[],[f157_D])).
fof(f157_D,plain,(
  s('type.realax.real',sK1) = s('type.realax.real',sK2) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f159,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))) | $spl2),
  inference(cnf_transformation,[],[f159_D])).
fof(f159_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f99,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))) | s('type.realax.real',sK1) = s('type.realax.real',sK2)),
  inference(cnf_transformation,[],[f76])).
fof(f158,plain,(
  s('type.realax.real',sK1) != s('type.realax.real',sK2) | $spl1),
  inference(cnf_transformation,[],[f158_D])).
fof(f158_D,plain,(
  s('type.realax.real',sK1) != s('type.realax.real',sK2) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f160,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))) | $spl3),
  inference(cnf_transformation,[],[f160_D])).
fof(f160_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f100,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK0))) | s('type.realax.real',sK1) != s('type.realax.real',sK2)),
  inference(cnf_transformation,[],[f76])).
% SZS output end Proof for thm.transc
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 140
% Passive clauses: 321
% Generated clauses: 3358
% Final active clauses: 53
% Final passive clauses: 78
% Input formulas: 28
% Initial clauses: 54
% 
% Duplicate literals: 127
% Trivial inequalities: 36
% Fw subsumption resolutions: 190
% Fw demodulations: 555
% Bw demodulations: 139
% Condensations: 173
% 
% Simple tautologies: 67
% Equational tautologies: 50
% Forward subsumptions: 1774
% Fw demodulations to eq. taut.: 1
% Bw demodulations to eq. taut.: 5
% 
% Unit resulting resolution: 216
% Factoring: 1
% Forward superposition: 762
% Backward superposition: 894
% Self superposition: 28
% Equality factoring: 26
% Equality resolution: 18
% 
% Split clauses: 19
% Split components: 43
% SAT solver clauses: 406
% SAT solver unit clauses: 56
% SAT solver binary clauses: 106
% 
% Sat splits: 19
% Sat splitting refutations: 45
% 
% Memory used [KB]: 2430
% Time elapsed: 0.082 s
% ------------------------------
% ------------------------------
% Success in time 1.397 s
