Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.prelim
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.prelim
% SZS output start Proof for thm.prelim
fof(f16588,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f16554,f7277,f4550,f4548,f156,f229])).
fof(f229,plain,(
  ~p(s(bool,'F_0')) | $spl1),
  inference(cnf_transformation,[],[f229_D])).
fof(f229_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f156,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f67])).
fof(f67,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.prelim.ordering_case_def|split|1',aHL_FALSITY)).
fof(f4548,plain,(
  ( ! [X14,X17,X15,X16] : (s(bool,'T_0') = s(bool,sK4(X14,X15,X16,'const.prelim.EQUAL_0',X17))) ) | $spl60),
  inference(cnf_transformation,[],[f4548_D])).
fof(f4548_D,plain,(
  ( ! [X14,X17,X15,X16] : (s(bool,'T_0') = s(bool,sK4(X14,X15,X16,'const.prelim.EQUAL_0',X17))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f4550,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,'T_0') = s(bool,sK4(X24,X25,X26,'const.prelim.EQUAL_0',X27))) ) | $spl1),
  inference(subsumption_resolution,[],[f4543,f229])).
fof(f4543,plain,(
  ( ! [X26,X24,X27,X25] : (p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,sK4(X24,X25,X26,'const.prelim.EQUAL_0',X27))) )),
  inference(superposition,[],[f4535,f181])).
fof(f181,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.prelim.ordering_case_def|split|1',aHL_BOOL_CASES)).
fof(f4535,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,sK4(X4,X5,X6,'const.prelim.EQUAL_0',X7)))) )),
  inference(trivial_inequality_removal,[],[f4534])).
fof(f4534,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,sK4(X4,X5,X6,'const.prelim.EQUAL_0',X7)))) )),
  inference(superposition,[],[f3909,f435])).
fof(f435,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.EQUAL_0')))),
  inference(backward_demodulation,[],[f155,f154])).
fof(f154,plain,(
  s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.EQUAL_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(cnf_transformation,[],[f62])).
fof(f62,axiom,(
  s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.EQUAL_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  file('thm.prelim.ordering_case_def|split|1','thm.prelim.ordering2num_thm|split|1')).
fof(f155,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(cnf_transformation,[],[f23])).
fof(f23,axiom,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  file('thm.prelim.ordering_case_def|split|1','thm.numeral.numeral_distrib|split|13')).
fof(f3909,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X1))) | p(s(bool,sK4(X4,X3,X2,X1,X0)))) )),
  inference(forward_demodulation,[],[f201,f155])).
fof(f201,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,sK4(X4,X3,X2,X1,X0))) | s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X1))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) )),
  inference(cnf_transformation,[],[f140])).
fof(f140,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,sK4(X4,X3,X2,X1,X0))) | s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X1))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) & (s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X1))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | p(s(bool,sK4(X4,X3,X2,X1,X0)))) & s(X0,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X1),s(X0,X2),s(X0,X3),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s(X0,X2),s(X0,'const.bool.COND_3'(s(bool,sK4(X4,X3,X2,X1,X0)),s(X0,X3),s(X0,X4))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f139])).
fof(f139,plain,(
  ! [X0,X1,X2,X3,X4] : ? [X5] : ((~p(s(bool,X5)) | s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X1))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) & (s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X1))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | p(s(bool,X5))) & s(X0,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X1),s(X0,X2),s(X0,X3),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s(X0,X2),s(X0,'const.bool.COND_3'(s(bool,X5),s(X0,X3),s(X0,X4))))))),
  inference(flattening,[],[f138])).
fof(f138,plain,(
  ! [X0,X1,X2,X3,X4] : ? [X5] : (((~p(s(bool,X5)) | s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X1))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) & (s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X1))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | p(s(bool,X5)))) & s(X0,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X1),s(X0,X2),s(X0,X3),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s(X0,X2),s(X0,'const.bool.COND_3'(s(bool,X5),s(X0,X3),s(X0,X4))))))),
  inference(nnf_transformation,[],[f106])).
fof(f106,plain,(
  ! [X0,X1,X2,X3,X4] : ? [X5] : ((p(s(bool,X5)) <=> s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X1))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) & s(X0,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X1),s(X0,X2),s(X0,X3),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s(X0,X2),s(X0,'const.bool.COND_3'(s(bool,X5),s(X0,X3),s(X0,X4))))))),
  inference(flattening,[],[f105])).
fof(f105,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : ? [X5] : ((p(s(bool,X5)) <=> s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X1))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) & s(X0,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X1),s(X0,X2),s(X0,X3),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s(X0,X2),s(X0,'const.bool.COND_3'(s(bool,X5),s(X0,X3),s(X0,X4))))))),
  inference(rectify,[],[f64])).
fof(f64,axiom,(
  ! [X6] : ! [X10,X7,X8,X9] : ? [X15] : ((p(s(bool,X15)) <=> s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X10))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) & s(X6,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X10),s(X6,X7),s(X6,X8),s(X6,X9))) = s(X6,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X10))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s(X6,X7),s(X6,'const.bool.COND_3'(s(bool,X15),s(X6,X8),s(X6,X9))))))),
  file('thm.prelim.ordering_case_def|split|1','thm.prelim.ordering_CASE')).
fof(f7277,plain,(
  $false | $spl29),
  inference(subsumption_resolution,[],[f7272,f998])).
fof(f998,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.ZERO_0') | $spl29),
  inference(cnf_transformation,[],[f998_D])).
fof(f998_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.ZERO_0') <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f7272,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.ZERO_0')),
  inference(trivial_inequality_removal,[],[f7197])).
fof(f7197,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.|48|_0') | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.ZERO_0')),
  inference(superposition,[],[f190,f7126])).
fof(f7126,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0')))),
  inference(forward_demodulation,[],[f7125,f714])).
fof(f714,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(equality_resolution,[],[f191])).
fof(f191,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f134])).
fof(f134,plain,(
  ! [X0] : ((s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0)) & (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))),
  inference(nnf_transformation,[],[f96])).
fof(f96,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) <=> s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X13] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X13))) <=> s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X13))),
  file('thm.prelim.ordering_case_def|split|1','thm.numeral.numeral_distrib|split|18')).
fof(f7125,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(subsumption_resolution,[],[f7124,f149])).
fof(f149,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.prelim.ordering_case_def|split|1',aHL_TRUTH)).
fof(f7124,plain,(
  ~p(s(bool,'T_0')) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(forward_demodulation,[],[f7081,f164])).
fof(f164,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0] : s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X13] : s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13)))),
  file('thm.prelim.ordering_case_def|split|1','thm.numeral.numeral_distrib|split|26')).
fof(f7081,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(superposition,[],[f6876,f714])).
fof(f6876,plain,(
  ( ! [X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.ZERO_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))) )),
  inference(superposition,[],[f194,f6121])).
fof(f6121,plain,(
  ( ! [X2] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.ZERO_0'))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2)))))) )),
  inference(superposition,[],[f203,f714])).
fof(f203,plain,(
  ( ! [X0,X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f107])).
fof(f107,plain,(
  ! [X0,X1] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X13,X14] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X14),s('type.num.num',X13))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X13))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X14)))))),
  file('thm.prelim.ordering_case_def|split|1','thm.numeral.numeral_distrib|split|31')).
fof(f194,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f136])).
fof(f136,plain,(
  ! [X0] : ((~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) & (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))))),
  inference(nnf_transformation,[],[f98])).
fof(f98,plain,(
  ! [X0] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X0))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X13] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13)))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X13))),
  file('thm.prelim.ordering_case_def|split|1','thm.numeral.numeral_distrib|split|30')).
fof(f190,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f134])).
fof(f16554,plain,(
  $false | ($spl28 | $spl60)),
  inference(trivial_inequality_removal,[],[f16553])).
fof(f16553,plain,(
  s(sK0,sK2) != s(sK0,sK2) | ($spl28 | $spl60)),
  inference(superposition,[],[f148,f16096])).
fof(f16096,plain,(
  ( ! [X10,X8,X11,X9] : (s(X8,X10) = s(X8,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.EQUAL_0'),s(X8,X9),s(X8,X10),s(X8,X11)))) ) | ($spl28 | $spl60)),
  inference(forward_demodulation,[],[f16095,f199])).
fof(f199,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f104])).
fof(f104,plain,(
  ! [X0,X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f103])).
fof(f103,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6] : ! [X11,X12] : s(X6,X12) = s(X6,'const.bool.COND_3'(s(bool,'F_0'),s(X6,X11),s(X6,X12)))),
  file('thm.prelim.ordering_case_def|split|1','thm.bool.COND_CLAUSES|split|1')).
fof(f16095,plain,(
  ( ! [X10,X8,X11,X9] : (s(X8,'const.bool.COND_3'(s(bool,'F_0'),s(X8,X9),s(X8,X10))) = s(X8,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.EQUAL_0'),s(X8,X9),s(X8,X10),s(X8,X11)))) ) | ($spl28 | $spl60)),
  inference(forward_demodulation,[],[f16094,f165])).
fof(f165,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0] : s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X13] : s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X13),s('type.num.num','const.num.|48|_0')))),
  file('thm.prelim.ordering_case_def|split|1','thm.numeral.numeral_distrib|split|20')).
fof(f16094,plain,(
  ( ! [X10,X8,X11,X9] : (s(X8,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.EQUAL_0'),s(X8,X9),s(X8,X10),s(X8,X11))) = s(X8,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))),s(X8,X9),s(X8,X10)))) ) | ($spl28 | $spl60)),
  inference(forward_demodulation,[],[f16093,f207])).
fof(f207,plain,(
  ( ! [X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f111])).
fof(f111,plain,(
  ! [X0,X1] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X13,X14] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X13),s('type.num.num',X14))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X13))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X14)))))),
  file('thm.prelim.ordering_case_def|split|1','thm.numeral.numeral_lt|split|3')).
fof(f16093,plain,(
  ( ! [X10,X8,X11,X9] : (s(X8,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.EQUAL_0'),s(X8,X9),s(X8,X10),s(X8,X11))) = s(X8,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(X8,X9),s(X8,X10)))) ) | ($spl28 | $spl60)),
  inference(forward_demodulation,[],[f16092,f198])).
fof(f198,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f101])).
fof(f101,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : ! [X11,X12] : s(X6,X11) = s(X6,'const.bool.COND_3'(s(bool,'T_0'),s(X6,X11),s(X6,X12)))),
  file('thm.prelim.ordering_case_def|split|1','thm.bool.COND_CLAUSES|split|0')).
fof(f16092,plain,(
  ( ! [X10,X8,X11,X9] : (s(X8,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.EQUAL_0'),s(X8,X9),s(X8,X10),s(X8,X11))) = s(X8,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(X8,X9),s(X8,'const.bool.COND_3'(s(bool,'T_0'),s(X8,X10),s(X8,X11)))))) ) | ($spl28 | $spl60)),
  inference(forward_demodulation,[],[f15678,f4548])).
fof(f15678,plain,(
  ( ! [X10,X8,X11,X9] : (s(X8,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.EQUAL_0'),s(X8,X9),s(X8,X10),s(X8,X11))) = s(X8,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(X8,X9),s(X8,'const.bool.COND_3'(s(bool,sK4(X11,X10,X9,'const.prelim.EQUAL_0',X8)),s(X8,X10),s(X8,X11)))))) ) | $spl28),
  inference(superposition,[],[f15008,f14851])).
fof(f14851,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.EQUAL_0'))) | $spl28),
  inference(subsumption_resolution,[],[f14816,f7336])).
fof(f7336,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl28),
  inference(backward_demodulation,[],[f997,f890])).
fof(f890,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(superposition,[],[f173,f714])).
fof(f173,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X13] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X13))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X13)))))),
  file('thm.prelim.ordering_case_def|split|1','thm.numeral.numeral_distrib|split|14')).
fof(f997,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.ZERO_0') | $spl28),
  inference(cnf_transformation,[],[f997_D])).
fof(f997_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.ZERO_0') <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f14816,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.EQUAL_0'))) | $spl28),
  inference(superposition,[],[f7561,f435])).
fof(f7561,plain,(
  ( ! [X4] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X4))) | s('type.num.num',X4) = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl28),
  inference(superposition,[],[f222,f7320])).
fof(f7320,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | $spl28),
  inference(backward_demodulation,[],[f997,f155])).
fof(f222,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0,X1] : ((s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) & (s('type.num.num',X0) != s('type.num.num',X1) | s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))),
  inference(nnf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0,X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))) <=> s('type.num.num',X0) = s('type.num.num',X1))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X13,X14] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X13))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X14))) <=> s('type.num.num',X13) = s('type.num.num',X14))),
  file('thm.prelim.ordering_case_def|split|1','thm.numeral.numeral_distrib|split|19')).
fof(f15008,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X1),s(X0,X2),s(X0,X3),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X1))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(X0,X2),s(X0,'const.bool.COND_3'(s(bool,sK4(X4,X3,X2,X1,X0)),s(X0,X3),s(X0,X4)))))) ) | $spl28),
  inference(backward_demodulation,[],[f14853,f12913])).
fof(f12913,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X1),s(X0,X2),s(X0,X3),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X1))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s(X0,X2),s(X0,'const.bool.COND_3'(s(bool,sK4(X4,X3,X2,X1,X0)),s(X0,X3),s(X0,X4)))))) ) | $spl28),
  inference(forward_demodulation,[],[f12912,f7320])).
fof(f12912,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X1),s(X0,X2),s(X0,X3),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))),s(X0,X2),s(X0,'const.bool.COND_3'(s(bool,sK4(X4,X3,X2,X1,X0)),s(X0,X3),s(X0,X4)))))) ) | $spl28),
  inference(forward_demodulation,[],[f202,f997])).
fof(f202,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X1),s(X0,X2),s(X0,X3),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s(X0,X2),s(X0,'const.bool.COND_3'(s(bool,sK4(X4,X3,X2,X1,X0)),s(X0,X3),s(X0,X4)))))) )),
  inference(cnf_transformation,[],[f140])).
fof(f14853,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | $spl28),
  inference(backward_demodulation,[],[f14851,f435])).
fof(f148,plain,(
  s(sK0,sK2) != s(sK0,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.EQUAL_0'),s(sK0,sK1),s(sK0,sK2),s(sK0,sK3)))),
  inference(cnf_transformation,[],[f129])).
fof(f129,plain,(
  s(sK0,sK2) != s(sK0,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.EQUAL_0'),s(sK0,sK1),s(sK0,sK2),s(sK0,sK3)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f127])).
fof(f127,plain,(
  ? [X0,X1,X2,X3] : s(X0,X2) != s(X0,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.EQUAL_0'),s(X0,X1),s(X0,X2),s(X0,X3)))),
  inference(ennf_transformation,[],[f66])).
fof(f66,plain,(
  ~! [X0,X1,X2,X3] : s(X0,X2) = s(X0,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.EQUAL_0'),s(X0,X1),s(X0,X2),s(X0,X3)))),
  inference(flattening,[],[f65])).
fof(f65,plain,(
  ~! [X0] : ! [X1,X2,X3] : s(X0,X2) = s(X0,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.EQUAL_0'),s(X0,X1),s(X0,X2),s(X0,X3)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : s(X6,X8) = s(X6,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.EQUAL_0'),s(X6,X7),s(X6,X8),s(X6,X9)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : s(X6,X8) = s(X6,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.EQUAL_0'),s(X6,X7),s(X6,X8),s(X6,X9)))),
  file('thm.prelim.ordering_case_def|split|1','thm.prelim.ordering_case_def|split|1')).
% SZS output end Proof for thm.prelim
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 311
% Passive clauses: 2654
% Generated clauses: 16436
% Final active clauses: 129
% Final passive clauses: 1959
% Input formulas: 63
% Initial clauses: 79
% 
% Duplicate literals: 291
% Trivial inequalities: 145
% Fw subsumption resolutions: 364
% Fw demodulations: 2801
% Bw demodulations: 146
% 
% Simple tautologies: 42
% Equational tautologies: 26
% Forward subsumptions: 10202
% Fw demodulations to eq. taut.: 1
% 
% Binary resolution: 28
% Forward superposition: 5710
% Backward superposition: 6586
% Self superposition: 91
% Equality factoring: 36
% Equality resolution: 12
% 
% Split clauses: 55
% Split components: 110
% SAT solver clauses: 99
% SAT solver unit clauses: 8
% SAT solver binary clauses: 40
% SAT solver learnt clauses: 1
% 
% Sat splits: 55
% Sat splitting refutations: 15
% 
% TWLsolver clauses: 98
% TWLsolver calls for satisfiability: 46
% 
% Memory used [KB]: 11257
% Time elapsed: 0.261 s
% ------------------------------
% ------------------------------
% Success in time 0.262 s
