Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 44
dis+11_7_41 on thm.extreal
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.extreal
% SZS output start Proof for thm.extreal
fof(f207102,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f140670,f140735,f140733,f1740,f2584,f146874,f146864,f1966,f2585,f2583,f199322,f198760,f198764,f198762,f206961,f199324])).
fof(f199324,plain,(
  $false | ($spl1 | $spl870 | $spl1018)),
  inference(subsumption_resolution,[],[f199323,f2584])).
fof(f199323,plain,(
  p(s(bool,'F_0')) | ($spl870 | $spl1018)),
  inference(forward_demodulation,[],[f198953,f2875])).
fof(f2875,plain,(
  ( ! [X6] : (s(bool,'F_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal','const.extreal.NegInf_0')))) )),
  inference(resolution,[],[f1897,f1991])).
fof(f1991,plain,(
  ( ! [X0] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0'))))) )),
  inference(cnf_transformation,[],[f712])).
fof(f712,plain,(
  ! [X0] : ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0'))))),
  inference(flattening,[],[f711])).
fof(f711,plain,(
  ! [X0] : ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0'))))),
  inference(rectify,[],[f96])).
fof(f96,axiom,(
  ! [X6] : ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal','const.extreal.NegInf_0'))))),
  file('thm.extreal.sub_not_infty|split|0','thm.extreal.lt_infty|split|3')).
fof(f1897,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f1437])).
fof(f1437,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f656])).
fof(f656,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f655])).
fof(f655,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X21] : (s(bool,'F_0') = s(bool,X21) <=> ~p(s(bool,X21)))),
  file('thm.extreal.sub_not_infty|split|0','thm.bool.EQ_CLAUSES|split|3')).
fof(f198953,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK19),s('type.extreal.extreal','const.extreal.NegInf_0')))) | ($spl870 | $spl1018)),
  inference(backward_demodulation,[],[f198762,f140735])).
fof(f206961,plain,(
  $false | $spl1020),
  inference(subsumption_resolution,[],[f206960,f1697])).
fof(f1697,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',sK20)),
  inference(cnf_transformation,[],[f1411])).
fof(f1411,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',sK19) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',sK20) & s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK19),s('type.extreal.extreal',sK20)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK19,sK20])],[f1143])).
fof(f1143,plain,(
  ? [X0,X1] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X1) & s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))),
  inference(flattening,[],[f1142])).
fof(f1142,plain,(
  ? [X0,X1] : ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X1)) & s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))),
  inference(ennf_transformation,[],[f520])).
fof(f520,plain,(
  ~! [X0,X1] : ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X1)) => s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))),
  inference(flattening,[],[f519])).
fof(f519,plain,(
  ~! [X0,X1] : ((~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) & ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X1)) => ~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ((~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X6) & ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X7)) => ~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X7))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ((~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X6) & ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X7)) => ~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X7))))),
  file('thm.extreal.sub_not_infty|split|0','thm.extreal.sub_not_infty|split|0')).
fof(f206960,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK20) | $spl1020),
  inference(subsumption_resolution,[],[f206597,f1696])).
fof(f1696,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',sK19)),
  inference(cnf_transformation,[],[f1411])).
fof(f206597,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK19) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK20) | $spl1020),
  inference(superposition,[],[f198764,f1890])).
fof(f1890,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)) )),
  inference(cnf_transformation,[],[f1173])).
fof(f1173,plain,(
  ! [X0] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0'))) & s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)))))),
  inference(ennf_transformation,[],[f650])).
fof(f650,plain,(
  ! [X0] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0) => (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0'))) & s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)))))),
  inference(flattening,[],[f649])).
fof(f649,plain,(
  ! [X0] : (~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) => (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0'))) & s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)))))),
  inference(rectify,[],[f119])).
fof(f119,axiom,(
  ! [X6] : (~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X6) => (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal','const.extreal.NegInf_0'))) & s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X6)))))),
  file('thm.extreal.sub_not_infty|split|0','thm.extreal.add_infty|split|2')).
fof(f198762,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK20) | $spl1018),
  inference(cnf_transformation,[],[f198762_D])).
fof(f198762_D,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK20) <=> ~$spl1018),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1018])])).
fof(f198764,plain,(
  s('type.extreal.extreal',sK19) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK20))) | $spl1020),
  inference(cnf_transformation,[],[f198764_D])).
fof(f198764_D,plain,(
  s('type.extreal.extreal',sK19) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK20))) <=> ~$spl1020),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1020])])).
fof(f198760,plain,(
  s('type.extreal.extreal',sK19) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK20))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK20)),
  inference(subsumption_resolution,[],[f198313,f1697])).
fof(f198313,plain,(
  s('type.extreal.extreal',sK19) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK20))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK20) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK20)),
  inference(superposition,[],[f2186,f1698])).
fof(f1698,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',sK19),s('type.extreal.extreal',sK20)))),
  inference(cnf_transformation,[],[f1411])).
fof(f2186,plain,(
  ( ! [X0,X1] : (s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))),s('type.extreal.extreal',X1))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X1) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X1)) )),
  inference(cnf_transformation,[],[f1245])).
fof(f1245,plain,(
  ! [X0,X1] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X1) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X1) | s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))),s('type.extreal.extreal',X1))))),
  inference(flattening,[],[f1244])).
fof(f1244,plain,(
  ! [X0,X1] : ((s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X1) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X1)) | s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))),s('type.extreal.extreal',X1))))),
  inference(ennf_transformation,[],[f901])).
fof(f901,plain,(
  ! [X0,X1] : ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X1) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X1)) => s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))),s('type.extreal.extreal',X1))))),
  inference(flattening,[],[f900])).
fof(f900,plain,(
  ! [X0,X1] : ((~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X1) & ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X1)) => s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))),s('type.extreal.extreal',X1))))),
  inference(rectify,[],[f73])).
fof(f73,axiom,(
  ! [X6,X7] : ((~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X7) & ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X7)) => s('type.extreal.extreal',X6) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X7))),s('type.extreal.extreal',X7))))),
  file('thm.extreal.sub_not_infty|split|0','thm.extreal.sub_add')).
fof(f199322,plain,(
  $false | ($spl1 | $spl884 | $spl1018)),
  inference(subsumption_resolution,[],[f198945,f146864])).
fof(f198945,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK19)))) | ($spl1 | $spl1018)),
  inference(backward_demodulation,[],[f198762,f140668])).
fof(f140668,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK20),s('type.extreal.extreal',sK19)))) | $spl1),
  inference(subsumption_resolution,[],[f140667,f2584])).
fof(f140667,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK20),s('type.extreal.extreal',sK19))))),
  inference(forward_demodulation,[],[f140491,f2875])).
fof(f140491,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK20),s('type.extreal.extreal',sK19))))),
  inference(superposition,[],[f2158,f1698])).
fof(f2158,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))) )),
  inference(cnf_transformation,[],[f1194])).
fof(f1194,plain,(
  ! [X0,X1] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))))))),
  inference(ennf_transformation,[],[f869])).
fof(f869,plain,(
  ! [X0,X1] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) => p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))))))),
  inference(rectify,[],[f107])).
fof(f107,axiom,(
  ! [X6,X7] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X7)))) => p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_sub_2'(s('type.extreal.extreal',X7),s('type.extreal.extreal',X6)))))))),
  file('thm.extreal.sub_not_infty|split|0','thm.extreal.sub_zero_lt')).
fof(f2583,plain,(
  p(s(bool,'F_0')) | $spl0),
  inference(cnf_transformation,[],[f2583_D])).
fof(f2583_D,plain,(
  p(s(bool,'F_0')) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f2585,plain,(
  p(s(bool,'T_0')) | $spl2),
  inference(cnf_transformation,[],[f2585_D])).
fof(f2585_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f1966,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f1479])).
fof(f1479,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f1478])).
fof(f1478,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f1177])).
fof(f1177,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f686])).
fof(f686,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X21] : ((p(s(bool,'F_0')) => p(s(bool,X21))) <=> p(s(bool,'T_0')))),
  file('thm.extreal.sub_not_infty|split|0','thm.bool.IMP_CLAUSES|split|2')).
fof(f146864,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK19)))) | $spl884),
  inference(cnf_transformation,[],[f146864_D])).
fof(f146864_D,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK19)))) <=> ~$spl884),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl884])])).
fof(f146874,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK19)))) | ($spl1 | $spl2 | $spl868)),
  inference(subsumption_resolution,[],[f146873,f2584])).
fof(f146873,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK19)))) | p(s(bool,'F_0')) | ($spl1 | $spl2 | $spl868)),
  inference(resolution,[],[f146369,f2314])).
fof(f2314,plain,(
  ( ! [X0,X1] : (~sP2(X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f1614])).
fof(f1614,plain,(
  ! [X0,X1] : ((~sP2(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP2(X0,X1)))),
  inference(rectify,[],[f1613])).
fof(f1613,plain,(
  ! [X1,X0] : ((~sP2(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP2(X1,X0)))),
  inference(flattening,[],[f1612])).
fof(f1612,plain,(
  ! [X1,X0] : ((~sP2(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP2(X1,X0)))),
  inference(nnf_transformation,[],[f1386])).
fof(f1386,plain,(
  ! [X1,X0] : (sP2(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f146369,plain,(
  sP2('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK19)),'F_0') | ($spl1 | $spl2 | $spl868)),
  inference(superposition,[],[f8467,f141037])).
fof(f141037,plain,(
  s('type.extreal.extreal',sK19) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.extreal.real_1'(s('type.extreal.extreal',sK19))))) | $spl868),
  inference(subsumption_resolution,[],[f141036,f1696])).
fof(f141036,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK19) | s('type.extreal.extreal',sK19) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.extreal.real_1'(s('type.extreal.extreal',sK19))))) | $spl868),
  inference(forward_demodulation,[],[f141035,f140733])).
fof(f141035,plain,(
  s('type.extreal.extreal',sK19) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.extreal.real_1'(s('type.extreal.extreal',sK19))))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK20) | $spl868),
  inference(subsumption_resolution,[],[f140752,f1697])).
fof(f140752,plain,(
  s('type.extreal.extreal',sK19) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.extreal.real_1'(s('type.extreal.extreal',sK19))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK20) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK20) | $spl868),
  inference(superposition,[],[f1851,f140733])).
fof(f1851,plain,(
  ( ! [X0] : (s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.extreal.real_1'(s('type.extreal.extreal',X0))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0)) )),
  inference(cnf_transformation,[],[f1153])).
fof(f1153,plain,(
  ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.extreal.real_1'(s('type.extreal.extreal',X0))))))),
  inference(flattening,[],[f1152])).
fof(f1152,plain,(
  ! [X0] : ((s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)) | s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.extreal.real_1'(s('type.extreal.extreal',X0))))))),
  inference(ennf_transformation,[],[f635])).
fof(f635,plain,(
  ! [X0] : ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0)) => s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.extreal.real_1'(s('type.extreal.extreal',X0))))))),
  inference(flattening,[],[f634])).
fof(f634,plain,(
  ! [X0] : ((~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) & ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)) => s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.extreal.real_1'(s('type.extreal.extreal',X0))))))),
  inference(rectify,[],[f117])).
fof(f117,axiom,(
  ! [X6] : ((~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X6) & ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X6)) => s('type.extreal.extreal',X6) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.extreal.real_1'(s('type.extreal.extreal',X6))))))),
  file('thm.extreal.sub_not_infty|split|0','thm.extreal.normal_real')).
fof(f8467,plain,(
  ( ! [X9] : (sP2('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X9)))),'F_0')) ) | ($spl1 | $spl2)),
  inference(resolution,[],[f1778,f3132])).
fof(f3132,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | sP2(X0,'F_0')) ) | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f3128,f2585])).
fof(f3128,plain,(
  ( ! [X0] : (sP2(X0,'F_0') | ~p(s(bool,X0)) | ~p(s(bool,'T_0'))) ) | ($spl1 | $spl2)),
  inference(resolution,[],[f3113,f2315])).
fof(f2315,plain,(
  ( ! [X0,X1] : (~sP2(X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f1614])).
fof(f3113,plain,(
  ( ! [X14] : (sP2('T_0',X14) | sP2(X14,'F_0')) ) | ($spl1 | $spl2)),
  inference(resolution,[],[f2916,f2584])).
fof(f2916,plain,(
  ( ! [X0,X1] : (sP2('T_0',X0) | p(s(bool,X1)) | sP2(X0,X1)) ) | $spl2),
  inference(resolution,[],[f2901,f2316])).
fof(f2316,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP2(X0,X1)) )),
  inference(cnf_transformation,[],[f1614])).
fof(f2901,plain,(
  ( ! [X2] : (p(s(bool,X2)) | sP2('T_0',X2)) ) | $spl2),
  inference(resolution,[],[f2316,f2585])).
fof(f1778,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))))))) )),
  inference(cnf_transformation,[],[f561])).
fof(f561,plain,(
  ! [X0] : p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))))))),
  inference(rectify,[],[f95])).
fof(f95,axiom,(
  ! [X7] : p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X7))))))),
  file('thm.extreal.sub_not_infty|split|0','thm.extreal.lt_infty|split|0')).
fof(f2584,plain,(
  ~p(s(bool,'F_0')) | $spl1),
  inference(cnf_transformation,[],[f2584_D])).
fof(f2584_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f1740,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f523])).
fof(f523,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.extreal.sub_not_infty|split|0',aHL_FALSITY)).
fof(f140733,plain,(
  s('type.extreal.extreal',sK19) = s('type.extreal.extreal',sK20) | $spl868),
  inference(cnf_transformation,[],[f140733_D])).
fof(f140733_D,plain,(
  s('type.extreal.extreal',sK19) = s('type.extreal.extreal',sK20) <=> ~$spl868),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl868])])).
fof(f140735,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK19),s('type.extreal.extreal',sK20)))) | $spl870),
  inference(cnf_transformation,[],[f140735_D])).
fof(f140735_D,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK19),s('type.extreal.extreal',sK20)))) <=> ~$spl870),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl870])])).
fof(f140670,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK19),s('type.extreal.extreal',sK20)))) | s('type.extreal.extreal',sK19) = s('type.extreal.extreal',sK20) | $spl1),
  inference(resolution,[],[f140668,f2151])).
fof(f2151,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) | s('type.extreal.extreal',X0) = s('type.extreal.extreal',X1)) )),
  inference(cnf_transformation,[],[f862])).
fof(f862,plain,(
  ! [X0,X1] : (s('type.extreal.extreal',X0) = s('type.extreal.extreal',X1) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))))),
  inference(rectify,[],[f319])).
fof(f319,axiom,(
  ! [X6,X7] : (s('type.extreal.extreal',X6) = s('type.extreal.extreal',X7) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X7)))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X7),s('type.extreal.extreal',X6)))))),
  file('thm.extreal.sub_not_infty|split|0','thm.extreal.lt_total')).
% SZS output end Proof for thm.extreal
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 3253
% Passive clauses: 36594
% Generated clauses: 205430
% Final active clauses: 2810
% Final passive clauses: 30275
% Input formulas: 517
% Initial clauses: 877
% 
% Duplicate literals: 378
% Trivial inequalities: 389
% Fw subsumption resolutions: 7134
% Fw demodulations: 31223
% Bw demodulations: 328
% 
% Simple tautologies: 646
% Equational tautologies: 46
% Forward subsumptions: 128625
% Fw demodulations to eq. taut.: 85
% 
% Binary resolution: 25353
% Factoring: 1230
% Forward superposition: 59021
% Backward superposition: 78504
% Self superposition: 187
% Equality factoring: 2
% Equality resolution: 55
% 
% Split clauses: 1021
% Split components: 2045
% SAT solver clauses: 1628
% SAT solver unit clauses: 95
% SAT solver binary clauses: 307
% SAT solver learnt clauses: 3
% SAT solver learnt literals: 3
% 
% Sat splits: 1021
% Sat splitting refutations: 47
% 
% TWLsolver clauses: 1619
% TWLsolver calls for satisfiability: 1201
% 
% Memory used [KB]: 109124
% Time elapsed: 4.238 s
% ------------------------------
% ------------------------------
% Success in time 4.28 s
