Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.gcd
% (66404)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 157
% Passive clauses: 3292
% Generated clauses: 26105
% Final active clauses: 141
% Final passive clauses: 2812
% Input formulas: 32
% Initial clauses: 71
% 
% Duplicate literals: 880
% Trivial inequalities: 479
% Fw subsumption resolutions: 3877
% Fw demodulations: 4473
% Bw demodulations: 15
% 
% Simple tautologies: 465
% Equational tautologies: 59
% Forward subsumptions: 11977
% 
% Binary resolution: 257
% Factoring: 2
% Forward superposition: 8388
% Backward superposition: 7499
% Self superposition: 239
% Equality factoring: 146
% Equality resolution: 47
% 
% Split clauses: 408
% Split components: 844
% SAT solver clauses: 891
% SAT solver unit clauses: 20
% SAT solver binary clauses: 145
% 
% Sat splits: 408
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 891
% TWLsolver calls for satisfiability: 182
% 
% Memory used [KB]: 12920
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 294 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.gcd
% (66423)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 132
% Passive clauses: 793
% Generated clauses: 14367
% Final active clauses: 110
% Final passive clauses: 642
% Input formulas: 32
% Initial clauses: 71
% 
% Duplicate literals: 881
% Trivial inequalities: 137
% Fw subsumption resolutions: 667
% Fw demodulations: 1021
% Bw demodulations: 35
% Condensations: 1616
% Global subsumptions: 9
% 
% Simple tautologies: 208
% Equational tautologies: 35
% Forward subsumptions: 7642
% 
% Binary resolution: 222
% Factoring: 1
% Forward superposition: 5629
% Backward superposition: 3011
% Self superposition: 135
% Equality factoring: 60
% Equality resolution: 24
% 
% Split clauses: 71
% Split components: 154
% SAT solver clauses: 1453
% SAT solver unit clauses: 74
% SAT solver binary clauses: 50
% SAT solver learnt clauses: 18
% 
% Sat splits: 71
% Sat splitting refutations: 6
% 
% TWLsolver clauses: 1417
% TWLsolver calls for satisfiability: 361
% 
% Memory used [KB]: 12153
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 287 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.gcd
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.gcd
% SZS output start Proof for thm.gcd
fof(f427340,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f2151,f2153,f2152,f34157,f34172,f34160,f34236,f34237,f34238,f34243,f34246,f34247,f34248,f34249,f34251,f34253,f34255,f34257,f34260,f34263,f34265,f34269,f34271,f34273,f34275,f34277,f34279,f34377,f34689,f34693,f34695,f34699,f34701,f34703,f34723,f34724,f87024,f87036,f87025,f87675,f87719,f87717,f103057,f103121,f151757,f151759,f151761,f152941,f152943,f152945,f152947,f152951,f152953,f152955,f152957,f152987,f152989,f152995,f152998,f153183,f153185,f153186,f153322,f153323,f153343,f153344,f153345,f153346,f153347,f153348,f153349,f153350,f186802,f186854,f225564,f34159,f225569,f225568,f225566,f242770,f242790,f242791,f242793,f352822,f352882,f389627,f389692,f34158,f352645,f352644,f389737,f389754,f389758,f87026,f389756,f389763,f389765,f427325])).
fof(f427325,plain,(
  $false | ($spl16 | $spl20 | $spl23 | $spl25)),
  inference(subsumption_resolution,[],[f427324,f389766])).
fof(f389766,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.gcd.lcm_2'(s('type.num.num',sK1),s('type.num.num',sK0))) | $spl16),
  inference(backward_demodulation,[],[f87718,f116])).
fof(f116,plain,(
  s('type.num.num','const.gcd.lcm_2'(s('type.num.num',sK0),s('type.num.num',sK1))) != s('type.num.num','const.gcd.lcm_2'(s('type.num.num',sK1),s('type.num.num',sK0)))),
  inference(cnf_transformation,[],[f86])).
fof(f86,plain,(
  s('type.num.num','const.gcd.lcm_2'(s('type.num.num',sK0),s('type.num.num',sK1))) != s('type.num.num','const.gcd.lcm_2'(s('type.num.num',sK1),s('type.num.num',sK0)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f75])).
fof(f75,plain,(
  ? [X0,X1] : s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X0),s('type.num.num',X1))) != s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X1),s('type.num.num',X0)))),
  inference(ennf_transformation,[],[f34])).
fof(f34,plain,(
  ~! [X0,X1] : s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X1),s('type.num.num',X0)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X6),s('type.num.num',X7))) = s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X7),s('type.num.num',X6)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X6),s('type.num.num',X7))) = s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X7),s('type.num.num',X6)))),
  file('thm.gcd.LCM_COMM','thm.gcd.LCM_COMM')).
fof(f87718,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.gcd.lcm_2'(s('type.num.num',sK0),s('type.num.num',sK1))) | $spl16),
  inference(cnf_transformation,[],[f87718_D])).
fof(f87718_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.gcd.lcm_2'(s('type.num.num',sK0),s('type.num.num',sK1))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f427324,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.gcd.lcm_2'(s('type.num.num',sK1),s('type.num.num',sK0))) | ($spl16 | $spl20 | $spl23 | $spl25)),
  inference(backward_demodulation,[],[f427323,f411100])).
fof(f411100,plain,(
  s('type.num.num','const.gcd.lcm_2'(s('type.num.num',sK1),s('type.num.num',sK0))) = s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num',sK1))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',sK0),s('type.num.num',sK1))))))) | $spl20),
  inference(forward_demodulation,[],[f411099,f162])).
fof(f162,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f64])).
fof(f64,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',X0)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X8,X9] : s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X8),s('type.num.num',X9))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X9),s('type.num.num',X8)))),
  file('thm.gcd.LCM_COMM','thm.arithmetic.MULT_SYM')).
fof(f411099,plain,(
  s('type.num.num','const.gcd.lcm_2'(s('type.num.num',sK1),s('type.num.num',sK0))) = s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK1),s('type.num.num',sK0))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',sK0),s('type.num.num',sK1))))))) | $spl20),
  inference(forward_demodulation,[],[f401252,f163])).
fof(f163,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X1),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f65])).
fof(f65,plain,(
  ! [X0,X1] : s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X1),s('type.num.num',X0)))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X7,X6] : s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X6),s('type.num.num',X7))) = s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X7),s('type.num.num',X6)))),
  file('thm.gcd.LCM_COMM','thm.gcd.GCD_SYM')).
fof(f401252,plain,(
  s('type.num.num','const.gcd.lcm_2'(s('type.num.num',sK1),s('type.num.num',sK0))) = s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK1),s('type.num.num',sK0))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',sK1),s('type.num.num',sK0))))))) | $spl20),
  inference(superposition,[],[f180,f352631])).
fof(f352631,plain,(
  s(bool,'F_0') = s(bool,sK3(sK0,sK1)) | $spl20),
  inference(equality_resolution,[],[f186854])).
fof(f180,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3(X1,X0)),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X0),s('type.num.num',X1)))))))) )),
  inference(cnf_transformation,[],[f113])).
fof(f113,plain,(
  ! [X0,X1] : ((~p(s(bool,sK3(X1,X0))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) & ((s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)) | p(s(bool,sK3(X1,X0)))) & s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3(X1,X0)),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X0),s('type.num.num',X1))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3])],[f112])).
fof(f112,plain,(
  ! [X0,X1] : ? [X2] : ((~p(s(bool,X2)) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) & ((s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)) | p(s(bool,X2))) & s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.bool.COND_3'(s(bool,X2),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X0),s('type.num.num',X1))))))))),
  inference(flattening,[],[f111])).
fof(f111,plain,(
  ! [X0,X1] : ? [X2] : (((~p(s(bool,X2)) | (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1))) & ((s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)) | p(s(bool,X2)))) & s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.bool.COND_3'(s(bool,X2),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X0),s('type.num.num',X1))))))))),
  inference(nnf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0,X1] : ? [X2] : ((p(s(bool,X2)) <=> (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1))) & s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.bool.COND_3'(s(bool,X2),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X0),s('type.num.num',X1))))))))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X8,X9] : ? [X22] : ((p(s(bool,X22)) <=> (s('type.num.num','const.num.|48|_0') = s('type.num.num',X8) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X9))) & s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X8),s('type.num.num',X9))) = s('type.num.num','const.bool.COND_3'(s(bool,X22),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X8),s('type.num.num',X9))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X8),s('type.num.num',X9))))))))),
  file('thm.gcd.LCM_COMM','thm.gcd.lcm_def')).
fof(f427323,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num',sK1))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',sK0),s('type.num.num',sK1))))))) | ($spl16 | $spl23 | $spl25)),
  inference(forward_demodulation,[],[f415867,f87718])).
fof(f415867,plain,(
  s('type.num.num','const.gcd.lcm_2'(s('type.num.num',sK0),s('type.num.num',sK1))) = s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num',sK1))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',sK0),s('type.num.num',sK1))))))) | ($spl23 | $spl25)),
  inference(superposition,[],[f180,f389853])).
fof(f389853,plain,(
  s(bool,'F_0') = s(bool,sK3(sK1,sK0)) | ($spl23 | $spl25)),
  inference(unit_resulting_resolution,[],[f189,f389771,f164])).
fof(f164,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f79])).
fof(f79,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f66])).
fof(f66,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X10,X11] : ((p(s(bool,X10)) => p(s(bool,X11))) => ((p(s(bool,X11)) => p(s(bool,X10))) => s(bool,X10) = s(bool,X11)))),
  file('thm.gcd.LCM_COMM','thm.bool.IMP_ANTISYM_AX')).
fof(f389771,plain,(
  ~p(s(bool,sK3(sK1,sK0))) | ($spl23 | $spl25)),
  inference(unit_resulting_resolution,[],[f225565,f225567,f177])).
fof(f177,plain,(
  ( ! [X0,X1] : (~p(s(bool,sK3(X1,X0))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f113])).
fof(f225567,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK1) | $spl25),
  inference(cnf_transformation,[],[f225567_D])).
fof(f225567_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK1) <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f225565,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK0) | $spl23),
  inference(cnf_transformation,[],[f225565_D])).
fof(f225565_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK0) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f189,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f120,f117])).
fof(f117,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.gcd.LCM_COMM',aHL_TRUTH)).
fof(f120,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f87])).
fof(f87,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f35])).
fof(f35,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f18])).
fof(f18,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.gcd.LCM_COMM','thm.bool.NOT_CLAUSES|split|1')).
fof(f389765,plain,(
  $false | ($spl17 | $spl24)),
  inference(subsumption_resolution,[],[f389764,f683])).
fof(f683,plain,(
  ( ! [X3] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X3),s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f645,f154])).
fof(f154,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f56])).
fof(f56,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X13] : ! [X10,X11] : s(X13,X10) = s(X13,'const.bool.COND_3'(s(bool,'T_0'),s(X13,X10),s(X13,X11)))),
  file('thm.gcd.LCM_COMM','thm.bool.COND_CLAUSES|split|0')).
fof(f645,plain,(
  ( ! [X3] : (s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X3),s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X3),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X3),s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f180,f204])).
fof(f204,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,sK3('const.num.|48|_0',X0))) )),
  inference(unit_resulting_resolution,[],[f202,f128])).
fof(f128,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f89])).
fof(f89,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f41])).
fof(f41,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X12] : (s(bool,'T_0') = s(bool,X12) <=> p(s(bool,X12)))),
  file('thm.gcd.LCM_COMM','thm.bool.EQ_CLAUSES|split|1')).
fof(f202,plain,(
  ( ! [X0] : (p(s(bool,sK3('const.num.|48|_0',X0)))) )),
  inference(equality_resolution,[],[f179])).
fof(f179,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X1) | p(s(bool,sK3(X1,X0)))) )),
  inference(cnf_transformation,[],[f113])).
fof(f389764,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.gcd.lcm_2'(s('type.num.num',sK0),s('type.num.num','const.num.|48|_0'))) | ($spl17 | $spl24)),
  inference(forward_demodulation,[],[f87719,f225568])).
fof(f389763,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f389762,f680])).
fof(f680,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.gcd.lcm_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(forward_demodulation,[],[f643,f154])).
fof(f643,plain,(
  ( ! [X0] : (s('type.num.num','const.gcd.lcm_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))))))) )),
  inference(superposition,[],[f180,f199])).
fof(f199,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,sK3(X0,'const.num.|48|_0'))) )),
  inference(unit_resulting_resolution,[],[f197,f128])).
fof(f197,plain,(
  ( ! [X0] : (p(s(bool,sK3(X0,'const.num.|48|_0')))) )),
  inference(equality_resolution,[],[f178])).
fof(f178,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | p(s(bool,sK3(X1,X0)))) )),
  inference(cnf_transformation,[],[f113])).
fof(f389762,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.gcd.lcm_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0))) | $spl24),
  inference(forward_demodulation,[],[f389761,f683])).
fof(f389761,plain,(
  s('type.num.num','const.gcd.lcm_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0))) != s('type.num.num','const.gcd.lcm_2'(s('type.num.num',sK0),s('type.num.num','const.num.|48|_0'))) | $spl24),
  inference(backward_demodulation,[],[f225568,f116])).
fof(f389756,plain,(
  s(bool,'T_0') != s(bool,sK3(sK0,sK1)) | $spl33),
  inference(cnf_transformation,[],[f389756_D])).
fof(f389756_D,plain,(
  s(bool,'T_0') != s(bool,sK3(sK0,sK1)) <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f87026,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl12),
  inference(cnf_transformation,[],[f87026_D])).
fof(f87026_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f389758,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') != s(bool,sK3(sK0,sK1))) ) | $spl20),
  inference(condensation,[],[f389757])).
fof(f389757,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') != s(bool,sK3(sK0,sK1)) | s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,X2)) ) | $spl20),
  inference(subsumption_resolution,[],[f352591,f86651])).
fof(f86651,plain,(
  ( ! [X14,X13] : (s(bool,'T_0') != s(bool,X13) | s(bool,X13) = s(bool,X14) | s(bool,'F_0') = s(bool,X14)) )),
  inference(superposition,[],[f190,f85854])).
fof(f85854,plain,(
  ( ! [X2,X1] : (s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,'F_0') = s(bool,X1)) )),
  inference(superposition,[],[f85821,f85821])).
fof(f85821,plain,(
  ( ! [X26,X25] : (s(bool,X26) = s(bool,sK3(X25,'const.num.|48|_0')) | s(bool,'F_0') = s(bool,X26)) )),
  inference(subsumption_resolution,[],[f85802,f117])).
fof(f85802,plain,(
  ( ! [X26,X25] : (~p(s(bool,'T_0')) | s(bool,X26) = s(bool,sK3(X25,'const.num.|48|_0')) | s(bool,'F_0') = s(bool,X26)) )),
  inference(superposition,[],[f85788,f199])).
fof(f85788,plain,(
  ( ! [X50,X48] : (~p(s(bool,X50)) | s(bool,X48) = s(bool,X50) | s(bool,'F_0') = s(bool,X48)) )),
  inference(forward_demodulation,[],[f85787,f154])).
fof(f85787,plain,(
  ( ! [X50,X48,X49] : (~p(s(bool,X50)) | s(bool,X50) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X48),s(bool,X49))) | s(bool,'F_0') = s(bool,X48)) )),
  inference(subsumption_resolution,[],[f85751,f117])).
fof(f85751,plain,(
  ( ! [X50,X48,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,X50)) | s(bool,X50) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X48),s(bool,X49))) | s(bool,'F_0') = s(bool,X48)) )),
  inference(superposition,[],[f167,f286])).
fof(f286,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f154,f125])).
fof(f125,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.gcd.LCM_COMM',aHL_BOOL_CASES)).
fof(f167,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f80])).
fof(f190,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f189,f127])).
fof(f127,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f89])).
fof(f352591,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') != s(bool,sK3(sK0,sK1)) | s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl20),
  inference(superposition,[],[f186854,f215])).
fof(f215,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f125,f125])).
fof(f389754,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') != s(bool,sK3(sK0,sK1))) ) | $spl20),
  inference(condensation,[],[f389753])).
fof(f389753,plain,(
  ( ! [X8,X7] : (s(bool,'T_0') != s(bool,sK3(sK0,sK1)) | s(bool,'F_0') = s(bool,X7) | s(bool,X7) = s(bool,X8)) ) | $spl20),
  inference(subsumption_resolution,[],[f352594,f86651])).
fof(f352594,plain,(
  ( ! [X8,X7] : (s(bool,'T_0') != s(bool,sK3(sK0,sK1)) | s(bool,'F_0') = s(bool,X7) | s(bool,X7) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) ) | $spl20),
  inference(superposition,[],[f186854,f215])).
fof(f389737,plain,(
  ( ! [X11] : (s(bool,X11) = s(bool,sK3(sK0,sK1)) | s(bool,'T_0') = s(bool,X11)) ) | $spl30),
  inference(cnf_transformation,[],[f389737_D])).
fof(f389737_D,plain,(
  ( ! [X11] : (s(bool,X11) = s(bool,sK3(sK0,sK1)) | s(bool,'T_0') = s(bool,X11)) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f352644,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') = s(bool,X4) | s(bool,X5) = s(bool,sK3(sK0,sK1)) | s(bool,'T_0') = s(bool,X5)) ) | $spl20),
  inference(subsumption_resolution,[],[f352622,f125])).
fof(f352622,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | s(bool,'F_0') = s(bool,X4) | s(bool,X5) = s(bool,sK3(sK0,sK1)) | s(bool,'T_0') = s(bool,X5)) ) | $spl20),
  inference(superposition,[],[f186854,f215])).
fof(f352645,plain,(
  ( ! [X10,X11] : (s(bool,'F_0') = s(bool,X10) | s(bool,X11) = s(bool,sK3(sK0,sK1)) | s(bool,'T_0') = s(bool,X11)) ) | $spl20),
  inference(subsumption_resolution,[],[f352625,f125])).
fof(f352625,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') != s(bool,X10) | s(bool,'F_0') = s(bool,X10) | s(bool,X11) = s(bool,sK3(sK0,sK1)) | s(bool,'T_0') = s(bool,X11)) ) | $spl20),
  inference(superposition,[],[f186854,f215])).
fof(f34158,plain,(
  ( ! [X30] : (s(bool,'T_0') = s(bool,X30)) ) | $spl4),
  inference(cnf_transformation,[],[f34158_D])).
fof(f34158_D,plain,(
  ( ! [X30] : (s(bool,'T_0') = s(bool,X30)) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f389692,plain,(
  ( ! [X78] : (s(bool,X78) != s(bool,sK3('const.gcd.lcm_2'(s('type.num.num',sK1),s('type.num.num',sK0)),sK0)) | s(bool,'F_0') = s(bool,X78)) ) | $spl28),
  inference(cnf_transformation,[],[f389692_D])).
fof(f389692_D,plain,(
  ( ! [X78] : (s(bool,X78) != s(bool,sK3('const.gcd.lcm_2'(s('type.num.num',sK1),s('type.num.num',sK0)),sK0)) | s(bool,'F_0') = s(bool,X78)) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f389627,plain,(
  ( ! [X78,X77] : (s(bool,X78) != s(bool,sK3('const.gcd.lcm_2'(s('type.num.num',sK1),s('type.num.num',sK0)),sK0)) | s(bool,'T_0') = s(bool,X77) | s(bool,'F_0') = s(bool,X78)) ) | ($spl16 | $spl23)),
  inference(duplicate_literal_removal,[],[f389520])).
fof(f389520,plain,(
  ( ! [X78,X77] : (s(bool,X78) != s(bool,sK3('const.gcd.lcm_2'(s('type.num.num',sK1),s('type.num.num',sK0)),sK0)) | s(bool,'T_0') = s(bool,X77) | s(bool,'F_0') = s(bool,X78) | s(bool,'T_0') = s(bool,X77)) ) | ($spl16 | $spl23)),
  inference(superposition,[],[f85974,f267978])).
fof(f267978,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,sK3('const.gcd.lcm_2'(s('type.num.num',sK1),s('type.num.num',sK0)),sK0)) | s(bool,'T_0') = s(bool,X0)) ) | ($spl16 | $spl23)),
  inference(subsumption_resolution,[],[f267703,f117])).
fof(f267703,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,sK3('const.gcd.lcm_2'(s('type.num.num',sK1),s('type.num.num',sK0)),sK0)) | s(bool,'T_0') = s(bool,X0)) ) | ($spl16 | $spl23)),
  inference(superposition,[],[f243109,f215])).
fof(f243109,plain,(
  ~p(s(bool,sK3('const.gcd.lcm_2'(s('type.num.num',sK1),s('type.num.num',sK0)),sK0))) | ($spl16 | $spl23)),
  inference(unit_resulting_resolution,[],[f225565,f242794,f177])).
fof(f242794,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.gcd.lcm_2'(s('type.num.num',sK1),s('type.num.num',sK0))) | $spl16),
  inference(backward_demodulation,[],[f87718,f116])).
fof(f85974,plain,(
  ( ! [X65,X63] : (s(bool,X63) != s(bool,X65) | s(bool,'T_0') = s(bool,X65) | s(bool,'F_0') = s(bool,X63)) )),
  inference(subsumption_resolution,[],[f85973,f117])).
fof(f85973,plain,(
  ( ! [X65,X63] : (~p(s(bool,'T_0')) | s(bool,X63) != s(bool,X65) | s(bool,'T_0') = s(bool,X65) | s(bool,'F_0') = s(bool,X63)) )),
  inference(forward_demodulation,[],[f85883,f199])).
fof(f85883,plain,(
  ( ! [X64,X65,X63] : (~p(s(bool,sK3(X64,'const.num.|48|_0'))) | s(bool,X63) != s(bool,X65) | s(bool,'T_0') = s(bool,X65) | s(bool,'F_0') = s(bool,X63)) )),
  inference(superposition,[],[f218,f85821])).
fof(f218,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f131,f125])).
fof(f131,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f91])).
fof(f91,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f44])).
fof(f44,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f43])).
fof(f43,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X12] : (s(bool,'F_0') = s(bool,X12) <=> ~p(s(bool,X12)))),
  file('thm.gcd.LCM_COMM','thm.bool.EQ_CLAUSES|split|3')).
fof(f352882,plain,(
  ( ! [X78] : (s(bool,X78) != s(bool,sK3(sK0,sK0)) | s(bool,'F_0') = s(bool,X78)) ) | $spl26),
  inference(cnf_transformation,[],[f352882_D])).
fof(f352882_D,plain,(
  ( ! [X78] : (s(bool,X78) != s(bool,sK3(sK0,sK0)) | s(bool,'F_0') = s(bool,X78)) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f352822,plain,(
  ( ! [X78,X77] : (s(bool,X78) != s(bool,sK3(sK0,sK0)) | s(bool,'T_0') = s(bool,X77) | s(bool,'F_0') = s(bool,X78)) ) | $spl23),
  inference(duplicate_literal_removal,[],[f352729])).
fof(f352729,plain,(
  ( ! [X78,X77] : (s(bool,X78) != s(bool,sK3(sK0,sK0)) | s(bool,'T_0') = s(bool,X77) | s(bool,'F_0') = s(bool,X78) | s(bool,'T_0') = s(bool,X77)) ) | $spl23),
  inference(superposition,[],[f85974,f243108])).
fof(f243108,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,sK3(sK0,sK0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl23),
  inference(subsumption_resolution,[],[f242867,f117])).
fof(f242867,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,sK3(sK0,sK0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl23),
  inference(superposition,[],[f242796,f215])).
fof(f242796,plain,(
  ~p(s(bool,sK3(sK0,sK0))) | $spl23),
  inference(unit_resulting_resolution,[],[f225565,f225565,f177])).
fof(f242793,plain,(
  $false | $spl22),
  inference(subsumption_resolution,[],[f242792,f680])).
fof(f242792,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.gcd.lcm_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK1))) | $spl22),
  inference(forward_demodulation,[],[f242786,f683])).
fof(f242786,plain,(
  s('type.num.num','const.gcd.lcm_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK1))) != s('type.num.num','const.gcd.lcm_2'(s('type.num.num',sK1),s('type.num.num','const.num.|48|_0'))) | $spl22),
  inference(backward_demodulation,[],[f225566,f116])).
fof(f242791,plain,(
  $false | ($spl17 | $spl22)),
  inference(subsumption_resolution,[],[f242785,f680])).
fof(f242785,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.gcd.lcm_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK1))) | ($spl17 | $spl22)),
  inference(backward_demodulation,[],[f225566,f87719])).
fof(f242790,plain,(
  $false | ($spl17 | $spl22)),
  inference(subsumption_resolution,[],[f242789,f117])).
fof(f242789,plain,(
  ~p(s(bool,'T_0')) | ($spl17 | $spl22)),
  inference(forward_demodulation,[],[f242784,f1487])).
fof(f1487,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK3(X0,'const.gcd.lcm_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1))))) )),
  inference(unit_resulting_resolution,[],[f689,f128])).
fof(f689,plain,(
  ( ! [X0,X1] : (p(s(bool,sK3(X0,'const.gcd.lcm_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f680,f178])).
fof(f242784,plain,(
  ~p(s(bool,sK3('const.gcd.lcm_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK1)),'const.gcd.lcm_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK1))))) | ($spl17 | $spl22)),
  inference(backward_demodulation,[],[f225566,f225534])).
fof(f225534,plain,(
  ~p(s(bool,sK3('const.gcd.lcm_2'(s('type.num.num',sK0),s('type.num.num',sK1)),'const.gcd.lcm_2'(s('type.num.num',sK0),s('type.num.num',sK1))))) | $spl17),
  inference(unit_resulting_resolution,[],[f87719,f87719,f177])).
fof(f242770,plain,(
  $false | ($spl14 | $spl18)),
  inference(subsumption_resolution,[],[f242769,f116])).
fof(f242769,plain,(
  s('type.num.num','const.gcd.lcm_2'(s('type.num.num',sK0),s('type.num.num',sK1))) = s('type.num.num','const.gcd.lcm_2'(s('type.num.num',sK1),s('type.num.num',sK0))) | ($spl14 | $spl18)),
  inference(backward_demodulation,[],[f239684,f242768])).
fof(f242768,plain,(
  s('type.num.num','const.gcd.lcm_2'(s('type.num.num',sK1),s('type.num.num',sK0))) = s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num',sK1))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',sK0),s('type.num.num',sK1))))))) | $spl14),
  inference(forward_demodulation,[],[f242767,f162])).
fof(f242767,plain,(
  s('type.num.num','const.gcd.lcm_2'(s('type.num.num',sK1),s('type.num.num',sK0))) = s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK1),s('type.num.num',sK0))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',sK0),s('type.num.num',sK1))))))) | $spl14),
  inference(forward_demodulation,[],[f239682,f163])).
fof(f239682,plain,(
  s('type.num.num','const.gcd.lcm_2'(s('type.num.num',sK1),s('type.num.num',sK0))) = s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK1),s('type.num.num',sK0))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',sK1),s('type.num.num',sK0))))))) | $spl14),
  inference(superposition,[],[f180,f87717])).
fof(f239684,plain,(
  s('type.num.num','const.gcd.lcm_2'(s('type.num.num',sK0),s('type.num.num',sK1))) = s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num',sK1))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',sK0),s('type.num.num',sK1))))))) | $spl18),
  inference(superposition,[],[f180,f195468])).
fof(f195468,plain,(
  s(bool,'F_0') = s(bool,sK3(sK1,sK0)) | $spl18),
  inference(condensation,[],[f195467])).
fof(f195467,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,sK3(sK1,sK0)) | s(bool,'F_0') = s(bool,X0)) ) | $spl18),
  inference(condensation,[],[f195466])).
fof(f195466,plain,(
  ( ! [X10,X9] : (s(bool,'F_0') = s(bool,X10) | s(bool,'F_0') = s(bool,sK3(sK1,sK0)) | s(bool,'F_0') = s(bool,X9)) ) | $spl18),
  inference(subsumption_resolution,[],[f195451,f85854])).
fof(f195451,plain,(
  ( ! [X10,X9] : (s(bool,X9) != s(bool,X10) | s(bool,'F_0') = s(bool,X10) | s(bool,'F_0') = s(bool,sK3(sK1,sK0)) | s(bool,'F_0') = s(bool,X9)) ) | $spl18),
  inference(superposition,[],[f103121,f85854])).
fof(f225566,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK0) | $spl22),
  inference(cnf_transformation,[],[f225566_D])).
fof(f225566_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK0) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f225568,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK1) | $spl24),
  inference(cnf_transformation,[],[f225568_D])).
fof(f225568_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK1) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f225569,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK1) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK0) | $spl15),
  inference(subsumption_resolution,[],[f225544,f117])).
fof(f225544,plain,(
  ~p(s(bool,'T_0')) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK1) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK0) | $spl15),
  inference(superposition,[],[f177,f186911])).
fof(f186911,plain,(
  s(bool,'T_0') = s(bool,sK3(sK0,sK1)) | $spl15),
  inference(unit_resulting_resolution,[],[f189,f87716,f258])).
fof(f258,plain,(
  ( ! [X6,X7] : (p(s(bool,X7)) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f253,f189])).
fof(f253,plain,(
  ( ! [X6,X7] : (p(s(bool,'F_0')) | p(s(bool,X7)) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f164,f125])).
fof(f87716,plain,(
  s(bool,'F_0') != s(bool,sK3(sK0,sK1)) | $spl15),
  inference(cnf_transformation,[],[f87716_D])).
fof(f87716_D,plain,(
  s(bool,'F_0') != s(bool,sK3(sK0,sK1)) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f34159,plain,(
  ( ! [X23] : (~p(s(bool,X23))) ) | $spl6),
  inference(cnf_transformation,[],[f34159_D])).
fof(f34159_D,plain,(
  ( ! [X23] : (~p(s(bool,X23))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f225564,plain,(
  ( ! [X14] : (~p(s(bool,X14)) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK1) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK0)) ) | $spl15),
  inference(subsumption_resolution,[],[f225543,f131])).
fof(f225543,plain,(
  ( ! [X14] : (~p(s(bool,X14)) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK1) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK0) | s(bool,'F_0') = s(bool,X14)) ) | $spl15),
  inference(superposition,[],[f177,f186923])).
fof(f186923,plain,(
  ( ! [X6] : (s(bool,X6) = s(bool,sK3(sK0,sK1)) | s(bool,'F_0') = s(bool,X6)) ) | $spl15),
  inference(trivial_inequality_removal,[],[f186920])).
fof(f186920,plain,(
  ( ! [X6] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,X6) = s(bool,sK3(sK0,sK1)) | s(bool,'F_0') = s(bool,X6)) ) | $spl15),
  inference(superposition,[],[f87716,f85854])).
fof(f186854,plain,(
  ( ! [X70] : (s(bool,X70) != s(bool,sK3(sK0,sK1)) | s(bool,'F_0') = s(bool,X70)) ) | $spl20),
  inference(cnf_transformation,[],[f186854_D])).
fof(f186854_D,plain,(
  ( ! [X70] : (s(bool,X70) != s(bool,sK3(sK0,sK1)) | s(bool,'F_0') = s(bool,X70)) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f186802,plain,(
  ( ! [X70,X69] : (s(bool,X70) != s(bool,sK3(sK0,sK1)) | s(bool,'T_0') = s(bool,X69) | s(bool,'F_0') = s(bool,X70)) ) | $spl14),
  inference(duplicate_literal_removal,[],[f186736])).
fof(f186736,plain,(
  ( ! [X70,X69] : (s(bool,X70) != s(bool,sK3(sK0,sK1)) | s(bool,'T_0') = s(bool,X69) | s(bool,'F_0') = s(bool,X70) | s(bool,'T_0') = s(bool,X69)) ) | $spl14),
  inference(superposition,[],[f85974,f118163])).
fof(f118163,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,sK3(sK0,sK1)) | s(bool,'T_0') = s(bool,X0)) ) | $spl14),
  inference(subsumption_resolution,[],[f107640,f190])).
fof(f107640,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X0) = s(bool,sK3(sK0,sK1)) | s(bool,'T_0') = s(bool,X0)) ) | $spl14),
  inference(superposition,[],[f87717,f215])).
fof(f153350,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f152940,f117])).
fof(f152940,plain,(
  ~p(s(bool,'T_0')) | $spl0),
  inference(unit_resulting_resolution,[],[f190,f2152,f167])).
fof(f153349,plain,(
  $false | ($spl0 | $spl14)),
  inference(subsumption_resolution,[],[f152942,f117])).
fof(f152942,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl14)),
  inference(unit_resulting_resolution,[],[f118166,f2152,f167])).
fof(f118166,plain,(
  s(bool,'T_0') != s(bool,sK3(sK0,sK1)) | $spl14),
  inference(subsumption_resolution,[],[f107650,f189])).
fof(f107650,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,sK3(sK0,sK1)) | $spl14),
  inference(superposition,[],[f127,f87717])).
fof(f153348,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f152944,f2152])).
fof(f152944,plain,(
  ~p(s(bool,'F_0')) | $spl0),
  inference(unit_resulting_resolution,[],[f190,f2152,f167])).
fof(f153347,plain,(
  $false | ($spl0 | $spl14)),
  inference(subsumption_resolution,[],[f152946,f2152])).
fof(f152946,plain,(
  ~p(s(bool,sK3(sK0,sK1))) | ($spl0 | $spl14)),
  inference(unit_resulting_resolution,[],[f118166,f2152,f167])).
fof(f153346,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f152950,f117])).
fof(f152950,plain,(
  ~p(s(bool,'T_0')) | $spl0),
  inference(unit_resulting_resolution,[],[f190,f2152,f167])).
fof(f153345,plain,(
  $false | ($spl0 | $spl14)),
  inference(subsumption_resolution,[],[f152952,f117])).
fof(f152952,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl14)),
  inference(unit_resulting_resolution,[],[f118166,f2152,f167])).
fof(f153344,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f152954,f2152])).
fof(f152954,plain,(
  ~p(s(bool,'F_0')) | $spl0),
  inference(unit_resulting_resolution,[],[f190,f2152,f167])).
fof(f153343,plain,(
  $false | ($spl0 | $spl14)),
  inference(subsumption_resolution,[],[f152956,f2152])).
fof(f152956,plain,(
  ~p(s(bool,sK3(sK0,sK1))) | ($spl0 | $spl14)),
  inference(unit_resulting_resolution,[],[f118166,f2152,f167])).
fof(f153323,plain,(
  $false | ($spl0 | $spl16)),
  inference(subsumption_resolution,[],[f153191,f117])).
fof(f153191,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl16)),
  inference(backward_demodulation,[],[f153187,f189])).
fof(f153187,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl0 | $spl16)),
  inference(subsumption_resolution,[],[f153151,f47840])).
fof(f47840,plain,(
  ( ! [X72,X71,X73] : (s(bool,X73) != s(bool,sK3(X71,'const.gcd.lcm_2'(s('type.num.num',X72),s('type.num.num','const.num.|48|_0')))) | s(bool,'T_0') = s(bool,X73)) )),
  inference(subsumption_resolution,[],[f42979,f117])).
fof(f42979,plain,(
  ( ! [X72,X71,X73] : (~p(s(bool,'T_0')) | s(bool,X73) != s(bool,sK3(X71,'const.gcd.lcm_2'(s('type.num.num',X72),s('type.num.num','const.num.|48|_0')))) | s(bool,'T_0') = s(bool,X73)) )),
  inference(superposition,[],[f218,f2746])).
fof(f2746,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK3(X0,'const.gcd.lcm_2'(s('type.num.num',X1),s('type.num.num','const.num.|48|_0'))))) )),
  inference(unit_resulting_resolution,[],[f1871,f128])).
fof(f1871,plain,(
  ( ! [X0,X1] : (p(s(bool,sK3(X0,'const.gcd.lcm_2'(s('type.num.num',X1),s('type.num.num','const.num.|48|_0')))))) )),
  inference(unit_resulting_resolution,[],[f683,f178])).
fof(f153151,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,sK3('const.gcd.lcm_2'(s('type.num.num',sK1),s('type.num.num','const.num.|48|_0')),'const.gcd.lcm_2'(s('type.num.num',sK1),s('type.num.num','const.num.|48|_0')))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl0 | $spl16)),
  inference(backward_demodulation,[],[f152997,f124367])).
fof(f124367,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,sK3('const.gcd.lcm_2'(s('type.num.num',sK1),s('type.num.num',sK0)),'const.gcd.lcm_2'(s('type.num.num',sK1),s('type.num.num',sK0)))) | s(bool,'T_0') = s(bool,X0)) ) | $spl16),
  inference(subsumption_resolution,[],[f124059,f117])).
fof(f124059,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,sK3('const.gcd.lcm_2'(s('type.num.num',sK1),s('type.num.num',sK0)),'const.gcd.lcm_2'(s('type.num.num',sK1),s('type.num.num',sK0)))) | s(bool,'T_0') = s(bool,X0)) ) | $spl16),
  inference(superposition,[],[f118187,f215])).
fof(f118187,plain,(
  ~p(s(bool,sK3('const.gcd.lcm_2'(s('type.num.num',sK1),s('type.num.num',sK0)),'const.gcd.lcm_2'(s('type.num.num',sK1),s('type.num.num',sK0))))) | $spl16),
  inference(unit_resulting_resolution,[],[f103154,f103154,f177])).
fof(f103154,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.gcd.lcm_2'(s('type.num.num',sK1),s('type.num.num',sK0))) | $spl16),
  inference(backward_demodulation,[],[f87718,f116])).
fof(f152997,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) ) | ($spl0 | $spl16)),
  inference(unit_resulting_resolution,[],[f103154,f2152,f177])).
fof(f153322,plain,(
  $false | ($spl0 | $spl16)),
  inference(subsumption_resolution,[],[f153321,f2152])).
fof(f153321,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl16)),
  inference(subsumption_resolution,[],[f153188,f153187])).
fof(f153188,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | ($spl0 | $spl16)),
  inference(backward_demodulation,[],[f153187,f131])).
fof(f153186,plain,(
  $false | ($spl0 | $spl16)),
  inference(subsumption_resolution,[],[f153150,f2746])).
fof(f153150,plain,(
  s(bool,'T_0') != s(bool,sK3('const.gcd.lcm_2'(s('type.num.num',sK1),s('type.num.num','const.num.|48|_0')),'const.gcd.lcm_2'(s('type.num.num',sK1),s('type.num.num','const.num.|48|_0')))) | ($spl0 | $spl16)),
  inference(backward_demodulation,[],[f152997,f124057])).
fof(f124057,plain,(
  s(bool,'T_0') != s(bool,sK3('const.gcd.lcm_2'(s('type.num.num',sK1),s('type.num.num',sK0)),'const.gcd.lcm_2'(s('type.num.num',sK1),s('type.num.num',sK0)))) | $spl16),
  inference(unit_resulting_resolution,[],[f118187,f127])).
fof(f153185,plain,(
  $false | ($spl0 | $spl16)),
  inference(subsumption_resolution,[],[f153184,f190])).
fof(f153184,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl0 | $spl16)),
  inference(forward_demodulation,[],[f153149,f2746])).
fof(f153149,plain,(
  s(bool,'F_0') = s(bool,sK3('const.gcd.lcm_2'(s('type.num.num',sK1),s('type.num.num','const.num.|48|_0')),'const.gcd.lcm_2'(s('type.num.num',sK1),s('type.num.num','const.num.|48|_0')))) | ($spl0 | $spl16)),
  inference(backward_demodulation,[],[f152997,f124055])).
fof(f124055,plain,(
  s(bool,'F_0') = s(bool,sK3('const.gcd.lcm_2'(s('type.num.num',sK1),s('type.num.num',sK0)),'const.gcd.lcm_2'(s('type.num.num',sK1),s('type.num.num',sK0)))) | $spl16),
  inference(unit_resulting_resolution,[],[f189,f118187,f164])).
fof(f153183,plain,(
  $false | ($spl0 | $spl16)),
  inference(subsumption_resolution,[],[f153148,f152997])).
fof(f153148,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.gcd.lcm_2'(s('type.num.num',sK1),s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl16)),
  inference(backward_demodulation,[],[f152997,f103154])).
fof(f152998,plain,(
  $false | ($spl0 | $spl16)),
  inference(unit_resulting_resolution,[],[f103154,f103154,f2152,f177])).
fof(f152995,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f190,f190,f2152,f85788])).
fof(f152989,plain,(
  $false | ($spl0 | $spl14)),
  inference(unit_resulting_resolution,[],[f87717,f118166,f2152,f218])).
fof(f152987,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f154,f190,f2152,f218])).
fof(f152957,plain,(
  $false | ($spl0 | $spl14)),
  inference(unit_resulting_resolution,[],[f2152,f118166,f2152,f167])).
fof(f152955,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f2152,f190,f2152,f167])).
fof(f152953,plain,(
  $false | ($spl0 | $spl14)),
  inference(unit_resulting_resolution,[],[f117,f118166,f2152,f167])).
fof(f152951,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f117,f190,f2152,f167])).
fof(f152947,plain,(
  $false | ($spl0 | $spl14)),
  inference(unit_resulting_resolution,[],[f2152,f118166,f2152,f167])).
fof(f152945,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f2152,f190,f2152,f167])).
fof(f152943,plain,(
  $false | ($spl0 | $spl14)),
  inference(unit_resulting_resolution,[],[f117,f118166,f2152,f167])).
fof(f152941,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f117,f190,f2152,f167])).
fof(f151761,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f154,f2152,f131])).
fof(f151759,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f190,f2152,f128])).
fof(f151757,plain,(
  $false | $spl0),
  inference(unit_resulting_resolution,[],[f189,f2152])).
fof(f103121,plain,(
  ( ! [X90] : (s(bool,X90) != s(bool,sK3(sK1,sK0)) | s(bool,'F_0') = s(bool,X90)) ) | $spl18),
  inference(cnf_transformation,[],[f103121_D])).
fof(f103121_D,plain,(
  ( ! [X90] : (s(bool,X90) != s(bool,sK3(sK1,sK0)) | s(bool,'F_0') = s(bool,X90)) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f103057,plain,(
  ( ! [X90,X89] : (s(bool,X90) != s(bool,sK3(sK1,sK0)) | s(bool,'T_0') = s(bool,X89) | s(bool,'F_0') = s(bool,X90)) ) | $spl17),
  inference(duplicate_literal_removal,[],[f102967])).
fof(f102967,plain,(
  ( ! [X90,X89] : (s(bool,X90) != s(bool,sK3(sK1,sK0)) | s(bool,'T_0') = s(bool,X89) | s(bool,'F_0') = s(bool,X90) | s(bool,'T_0') = s(bool,X89)) ) | $spl17),
  inference(superposition,[],[f85974,f102769])).
fof(f102769,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,sK3(sK1,sK0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl17),
  inference(subsumption_resolution,[],[f92246,f190])).
fof(f92246,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X0) = s(bool,sK3(sK1,sK0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl17),
  inference(superposition,[],[f87722,f215])).
fof(f87722,plain,(
  s(bool,'F_0') = s(bool,sK3(sK1,sK0)) | $spl17),
  inference(unit_resulting_resolution,[],[f87719,f85979])).
fof(f85979,plain,(
  ( ! [X103,X102] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X103),s('type.num.num',X102))) | s(bool,'F_0') = s(bool,sK3(X102,X103))) )),
  inference(forward_demodulation,[],[f85978,f154])).
fof(f85978,plain,(
  ( ! [X103,X102] : (s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X103),s('type.num.num',X102))) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X103),s('type.num.num',X102))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X103),s('type.num.num',X102))))))) | s(bool,'F_0') = s(bool,sK3(X102,X103))) )),
  inference(forward_demodulation,[],[f85891,f199])).
fof(f85891,plain,(
  ( ! [X103,X102,X104] : (s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X103),s('type.num.num',X102))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3(X104,'const.num.|48|_0')),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X103),s('type.num.num',X102))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X103),s('type.num.num',X102))))))) | s(bool,'F_0') = s(bool,sK3(X102,X103))) )),
  inference(superposition,[],[f180,f85821])).
fof(f87717,plain,(
  s(bool,'F_0') = s(bool,sK3(sK0,sK1)) | $spl14),
  inference(cnf_transformation,[],[f87717_D])).
fof(f87717_D,plain,(
  s(bool,'F_0') = s(bool,sK3(sK0,sK1)) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f87719,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.gcd.lcm_2'(s('type.num.num',sK0),s('type.num.num',sK1))) | $spl17),
  inference(cnf_transformation,[],[f87719_D])).
fof(f87719_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.gcd.lcm_2'(s('type.num.num',sK0),s('type.num.num',sK1))) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f87675,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.gcd.lcm_2'(s('type.num.num',sK0),s('type.num.num',sK1))) | s(bool,'F_0') = s(bool,sK3(sK0,sK1))),
  inference(superposition,[],[f116,f85979])).
fof(f87025,plain,(
  ( ! [X6,X4,X2,X5,X3,X1] : (s(X1,X2) != s(X1,X3) | s(bool,'F_0') != s(bool,X5) | s(X1,'const.bool.COND_3'(s(bool,'F_0'),s(X1,X4),s(X1,X3))) = s(X1,'const.bool.COND_3'(s(bool,X5),s(X1,X6),s(X1,X2)))) ) | $spl10),
  inference(cnf_transformation,[],[f87025_D])).
fof(f87025_D,plain,(
  ( ! [X6,X4,X2,X5,X3,X1] : (s(X1,X2) != s(X1,X3) | s(bool,'F_0') != s(bool,X5) | s(X1,'const.bool.COND_3'(s(bool,'F_0'),s(X1,X4),s(X1,X3))) = s(X1,'const.bool.COND_3'(s(bool,X5),s(X1,X6),s(X1,X2)))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f87036,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,'F_0') = s(bool,X0) | s(X1,X2) != s(X1,X3) | s(X1,'const.bool.COND_3'(s(bool,'F_0'),s(X1,X4),s(X1,X3))) = s(X1,'const.bool.COND_3'(s(bool,X5),s(X1,X6),s(X1,X2))) | s(bool,'F_0') != s(bool,X5)) )),
  inference(condensation,[],[f87035])).
fof(f87035,plain,(
  ( ! [X39,X45,X43,X41,X46,X44,X42,X40] : (s(X40,X41) != s(X40,X42) | s(X40,'const.bool.COND_3'(s(bool,'F_0'),s(X40,X45),s(X40,X42))) = s(X40,'const.bool.COND_3'(s(bool,X43),s(X40,X44),s(X40,X41))) | s(bool,'F_0') != s(bool,X43) | s(bool,X39) = s(bool,X46) | s(bool,'F_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f86753,f85788])).
fof(f86753,plain,(
  ( ! [X39,X45,X43,X41,X46,X44,X42,X40] : (p(s(bool,X39)) | s(X40,X41) != s(X40,X42) | s(X40,'const.bool.COND_3'(s(bool,'F_0'),s(X40,X45),s(X40,X42))) = s(X40,'const.bool.COND_3'(s(bool,X43),s(X40,X44),s(X40,X41))) | s(bool,'F_0') != s(bool,X43) | s(bool,X39) = s(bool,X46) | s(bool,'F_0') = s(bool,X46)) )),
  inference(superposition,[],[f160,f85854])).
fof(f160,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,X2)) | s(X0,X5) != s(X0,X6) | s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))) | s(bool,X1) != s(bool,X2)) )),
  inference(cnf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : (s(bool,X1) != s(bool,X2) | (p(s(bool,X2)) & s(X0,X3) != s(X0,X4)) | (~p(s(bool,X2)) & s(X0,X5) != s(X0,X6)) | s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))))),
  inference(flattening,[],[f77])).
fof(f77,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ((s(bool,X1) != s(bool,X2) | (p(s(bool,X2)) & s(X0,X3) != s(X0,X4)) | (~p(s(bool,X2)) & s(X0,X5) != s(X0,X6))) | s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))))),
  inference(ennf_transformation,[],[f63])).
fof(f63,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ((s(bool,X1) = s(bool,X2) & (p(s(bool,X2)) => s(X0,X3) = s(X0,X4)) & (~p(s(bool,X2)) => s(X0,X5) = s(X0,X6))) => s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))))),
  inference(flattening,[],[f62])).
fof(f62,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6] : ((s(bool,X1) = s(bool,X2) & (p(s(bool,X2)) => s(X0,X3) = s(X0,X4)) & (~p(s(bool,X2)) => s(X0,X5) = s(X0,X6))) => s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X13] : ! [X18,X19,X14,X20,X15,X21] : ((s(bool,X18) = s(bool,X19) & (p(s(bool,X19)) => s(X13,X14) = s(X13,X20)) & (~p(s(bool,X19)) => s(X13,X15) = s(X13,X21))) => s(X13,'const.bool.COND_3'(s(bool,X18),s(X13,X14),s(X13,X15))) = s(X13,'const.bool.COND_3'(s(bool,X19),s(X13,X20),s(X13,X21))))),
  file('thm.gcd.LCM_COMM','thm.bool.COND_CONG')).
fof(f87024,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,'F_0') = s(bool,X0) | s(X1,X2) != s(X1,X3) | s(X1,'const.bool.COND_3'(s(bool,'F_0'),s(X1,X4),s(X1,X3))) = s(X1,'const.bool.COND_3'(s(bool,X5),s(X1,X6),s(X1,X2))) | s(bool,'F_0') != s(bool,X5)) )),
  inference(condensation,[],[f87023])).
fof(f87023,plain,(
  ( ! [X39,X45,X43,X41,X46,X44,X42,X40] : (s(X40,X41) != s(X40,X42) | s(X40,'const.bool.COND_3'(s(bool,'F_0'),s(X40,X45),s(X40,X42))) = s(X40,'const.bool.COND_3'(s(bool,X43),s(X40,X44),s(X40,X41))) | s(bool,'F_0') != s(bool,X43) | s(bool,X39) = s(bool,X46) | s(bool,'F_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f86659,f85788])).
fof(f86659,plain,(
  ( ! [X39,X45,X43,X41,X46,X44,X42,X40] : (p(s(bool,X39)) | s(X40,X41) != s(X40,X42) | s(X40,'const.bool.COND_3'(s(bool,'F_0'),s(X40,X45),s(X40,X42))) = s(X40,'const.bool.COND_3'(s(bool,X43),s(X40,X44),s(X40,X41))) | s(bool,'F_0') != s(bool,X43) | s(bool,X39) = s(bool,X46) | s(bool,'F_0') = s(bool,X46)) )),
  inference(superposition,[],[f160,f85854])).
fof(f34724,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f34694,f34159])).
fof(f34694,plain,(
  p(s(bool,'F_0')) | $spl6),
  inference(unit_resulting_resolution,[],[f190,f34159,f164])).
fof(f34723,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f34700,f34159])).
fof(f34700,plain,(
  p(s(bool,'F_0')) | $spl6),
  inference(unit_resulting_resolution,[],[f190,f34159,f164])).
fof(f34703,plain,(
  $false | $spl6),
  inference(unit_resulting_resolution,[],[f154,f34159,f127])).
fof(f34701,plain,(
  $false | $spl6),
  inference(unit_resulting_resolution,[],[f189,f190,f34159,f164])).
fof(f34699,plain,(
  $false | $spl6),
  inference(unit_resulting_resolution,[],[f34159,f190,f34159,f164])).
fof(f34695,plain,(
  $false | $spl6),
  inference(unit_resulting_resolution,[],[f189,f190,f34159,f164])).
fof(f34693,plain,(
  $false | $spl6),
  inference(unit_resulting_resolution,[],[f34159,f190,f34159,f164])).
fof(f34689,plain,(
  $false | $spl6),
  inference(unit_resulting_resolution,[],[f190,f190,f34159,f258])).
fof(f34377,plain,(
  $false | $spl6),
  inference(unit_resulting_resolution,[],[f117,f34159])).
fof(f34279,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f116,f117,f189,f125,f128,f127,f132,f131,f154,f155,f161,f160,f159,f162,f163,f167,f164,f176,f180,f179,f178,f177,f190,f197,f199,f202,f204,f215,f218,f258,f275,f274,f286,f295,f296,f311,f323,f322,f324,f338,f347,f348,f365,f381,f382,f402,f403,f419,f420,f451,f452,f469,f470,f491,f492,f512,f513,f572,f573,f574,f607,f611,f642,f680,f683,f648,f649,f651,f653,f655,f656,f657,f658,f660,f662,f663,f664,f665,f671,f672,f689,f688,f715,f716,f727,f934,f945,f1353,f1355,f1357,f1359,f1361,f1487,f1871,f1870,f1955,f1956,f2030,f2041,f2151,f2154,f2265,f2746,f12998,f13000,f13005,f13023,f13031,f13061,f13069,f13146,f13148,f13153,f8434,f8435,f19837,f8447,f26086,f26091,f33665,f33670,f33689,f33690,f33708,f33718,f33719,f33739,f33758,f33773,f33776,f33823,f33824,f33825,f34157,f33831,f34165,f33849,f34166,f34167,f34169,f34170,f33876,f33879,f34172,f34158,f34235,f34181,f34242,f34245,f34197,f34198,f34199,f34252,f34254,f34256,f34259,f34262,f34264,f34268,f34270,f34272,f34274,f34276,f34278])).
fof(f34278,plain,(
  ( ! [X6,X8,X7] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X8),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X7),s('type.num.num',X6)))))) ) | $spl4),
  inference(forward_demodulation,[],[f34230,f154])).
fof(f34230,plain,(
  ( ! [X6,X8,X7] : (s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X8),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X7),s('type.num.num',X6))))) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X8),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X7),s('type.num.num',X6))))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X8),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X7),s('type.num.num',X6)))))))))) ) | $spl4),
  inference(backward_demodulation,[],[f34158,f660])).
fof(f34276,plain,(
  ( ! [X17,X15] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X17),s('type.num.num',X15)))) ) | $spl4),
  inference(forward_demodulation,[],[f34228,f154])).
fof(f34228,plain,(
  ( ! [X17,X15] : (s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X17),s('type.num.num',X15))) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X17),s('type.num.num',X15))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X17),s('type.num.num',X15)))))))) ) | $spl4),
  inference(backward_demodulation,[],[f34158,f663])).
fof(f34274,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',X0)))))) ) | $spl4),
  inference(forward_demodulation,[],[f34225,f154])).
fof(f34225,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',X0))))) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',X0))))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',X0)))))))))) ) | $spl4),
  inference(backward_demodulation,[],[f34158,f658])).
fof(f34272,plain,(
  ( ! [X6,X8,X7] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.gcd.lcm_2'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X7),s('type.num.num',X6))),s('type.num.num',X8)))) ) | $spl4),
  inference(forward_demodulation,[],[f34223,f154])).
fof(f34223,plain,(
  ( ! [X6,X8,X7] : (s('type.num.num','const.gcd.lcm_2'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X7),s('type.num.num',X6))),s('type.num.num',X8))) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X7),s('type.num.num',X6))),s('type.num.num',X8))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X7),s('type.num.num',X6))),s('type.num.num',X8)))))))) ) | $spl4),
  inference(backward_demodulation,[],[f34158,f653])).
fof(f34270,plain,(
  ( ! [X17,X15] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X15),s('type.num.num',X17)))) ) | $spl4),
  inference(forward_demodulation,[],[f34221,f154])).
fof(f34221,plain,(
  ( ! [X17,X15] : (s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X15),s('type.num.num',X17))) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X15),s('type.num.num',X17))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X15),s('type.num.num',X17)))))))) ) | $spl4),
  inference(backward_demodulation,[],[f34158,f656])).
fof(f34268,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.gcd.lcm_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.num.num',X2)))) ) | $spl4),
  inference(forward_demodulation,[],[f34218,f154])).
fof(f34218,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.gcd.lcm_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.num.num',X2))) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.num.num',X2))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.num.num',X2)))))))) ) | $spl4),
  inference(backward_demodulation,[],[f34158,f651])).
fof(f34264,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X0),s('type.num.num',X1)))) ) | $spl4),
  inference(forward_demodulation,[],[f34214,f154])).
fof(f34214,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X1),s('type.num.num',X0)))))))) ) | $spl4),
  inference(backward_demodulation,[],[f34158,f665])).
fof(f34262,plain,(
  ( ! [X14,X12,X13] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X14),s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X13),s('type.num.num',X12)))))) ) | $spl4),
  inference(forward_demodulation,[],[f34261,f154])).
fof(f34261,plain,(
  ( ! [X14,X12,X13] : (s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X14),s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X13),s('type.num.num',X12))))) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X14),s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X13),s('type.num.num',X12))))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X14),s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X13),s('type.num.num',X12)))))))))) ) | $spl4),
  inference(forward_demodulation,[],[f34213,f311])).
fof(f34213,plain,(
  ( ! [X14,X12,X13] : (s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X14),s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X13),s('type.num.num',X12))))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3('const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X13),s('type.num.num',X12))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X13),s('type.num.num',X12)))))),X14)),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X14),s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X13),s('type.num.num',X12))))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X14),s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X13),s('type.num.num',X12)))))))))) ) | $spl4),
  inference(backward_demodulation,[],[f34158,f662])).
fof(f34259,plain,(
  ( ! [X14,X12,X13] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.gcd.lcm_2'(s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X13),s('type.num.num',X12))),s('type.num.num',X14)))) ) | $spl4),
  inference(forward_demodulation,[],[f34258,f154])).
fof(f34258,plain,(
  ( ! [X14,X12,X13] : (s('type.num.num','const.gcd.lcm_2'(s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X13),s('type.num.num',X12))),s('type.num.num',X14))) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X13),s('type.num.num',X12))),s('type.num.num',X14))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X13),s('type.num.num',X12))),s('type.num.num',X14)))))))) ) | $spl4),
  inference(forward_demodulation,[],[f34212,f365])).
fof(f34212,plain,(
  ( ! [X14,X12,X13] : (s('type.num.num','const.gcd.lcm_2'(s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X13),s('type.num.num',X12))),s('type.num.num',X14))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3(X14,'const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X13),s('type.num.num',X12))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X13),s('type.num.num',X12)))))))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X13),s('type.num.num',X12))),s('type.num.num',X14))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X13),s('type.num.num',X12))),s('type.num.num',X14)))))))) ) | $spl4),
  inference(backward_demodulation,[],[f34158,f655])).
fof(f34256,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X0),s('type.num.num',X1)))) ) | $spl4),
  inference(forward_demodulation,[],[f34211,f154])).
fof(f34211,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X0),s('type.num.num',X1)))))))) ) | $spl4),
  inference(backward_demodulation,[],[f34158,f649])).
fof(f34254,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X0),s('type.num.num',X1)))) ) | $spl4),
  inference(forward_demodulation,[],[f34210,f154])).
fof(f34210,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X0),s('type.num.num',X1)))))))) ) | $spl4),
  inference(backward_demodulation,[],[f34158,f180])).
fof(f34252,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) ) | $spl4),
  inference(subsumption_resolution,[],[f34233,f117])).
fof(f34233,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | ~p(s(bool,'T_0'))) ) | $spl4),
  inference(condensation,[],[f34207])).
fof(f34207,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) ) | $spl4),
  inference(backward_demodulation,[],[f34158,f177])).
fof(f34199,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num',X1))) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',X0),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',X2)))))) ) | $spl4),
  inference(backward_demodulation,[],[f34158,f2041])).
fof(f34198,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X2),s('type.num.num',X1))) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',X0),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X1),s('type.num.num',X2)))))) ) | $spl4),
  inference(backward_demodulation,[],[f34158,f2030])).
fof(f34197,plain,(
  ( ! [X30,X37,X35,X33,X31,X36,X34,X32] : (s(bool,X33) != s(bool,sK3('const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',X30),s('type.num.num','const.num.|48|_0')),X31)) | s(X32,'const.bool.COND_3'(s(bool,'T_0'),s(X32,X36),s(X32,X37))) = s(X32,'const.bool.COND_3'(s(bool,X33),s(X32,X34),s(X32,X35))) | s(X32,X34) != s(X32,X36)) ) | $spl4),
  inference(backward_demodulation,[],[f34158,f1361])).
fof(f34245,plain,(
  ( ! [X19,X20] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X20),s('type.num.num',X19)))) ) | $spl4),
  inference(forward_demodulation,[],[f34244,f154])).
fof(f34244,plain,(
  ( ! [X19,X20] : (s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X20),s('type.num.num',X19))) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X20),s('type.num.num',X19))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X20),s('type.num.num',X19)))))))) ) | $spl4),
  inference(forward_demodulation,[],[f34196,f34158])).
fof(f34196,plain,(
  ( ! [X19,X20,X18] : (s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X20),s('type.num.num',X19))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3('const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',X18),s('type.num.num',X19)),X20)),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X20),s('type.num.num',X19))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X20),s('type.num.num',X19)))))))) ) | $spl4),
  inference(backward_demodulation,[],[f34158,f664])).
fof(f34242,plain,(
  ( ! [X19,X20] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X19),s('type.num.num',X20)))) ) | $spl4),
  inference(forward_demodulation,[],[f34241,f154])).
fof(f34241,plain,(
  ( ! [X19,X20] : (s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X19),s('type.num.num',X20))) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X19),s('type.num.num',X20))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X19),s('type.num.num',X20)))))))) ) | $spl4),
  inference(forward_demodulation,[],[f34195,f34158])).
fof(f34195,plain,(
  ( ! [X19,X20,X18] : (s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X19),s('type.num.num',X20))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3(X20,'const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',X18),s('type.num.num',X19)))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X19),s('type.num.num',X20))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X19),s('type.num.num',X20)))))))) ) | $spl4),
  inference(backward_demodulation,[],[f34158,f657])).
fof(f34181,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) ) | $spl4),
  inference(backward_demodulation,[],[f34158,f155])).
fof(f34235,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl4),
  inference(subsumption_resolution,[],[f34179,f34158])).
fof(f34179,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl4),
  inference(backward_demodulation,[],[f34158,f131])).
fof(f33879,plain,(
  ( ! [X138,X136,X139,X137] : (s(X138,f(s(bool,'T_0'))) != s(X138,g(s(bool,'T_0'))) | s(fun(bool,X138),X136) = s(fun(bool,X138),X137) | s(bool,X139) = s(bool,sK2(X136,X137,X138,bool)) | s(bool,'T_0') = s(bool,X139)) )),
  inference(superposition,[],[f176,f215])).
fof(f33876,plain,(
  ( ! [X127,X128,X126,X129] : (s(bool,'T_0') != s(bool,f(s(X126,sK2(X127,X128,bool,X126)))) | s(fun(X126,bool),X127) = s(fun(X126,bool),X128) | s(bool,X129) = s(bool,g(s(X126,sK2(X127,X128,bool,X126)))) | s(bool,'T_0') = s(bool,X129)) )),
  inference(superposition,[],[f176,f215])).
fof(f34170,plain,(
  ( ! [X92,X90,X91] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X91) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X90) | s(bool,X92) = s(bool,sK3(X90,X91)) | s(bool,'T_0') = s(bool,X92)) )),
  inference(subsumption_resolution,[],[f33869,f117])).
fof(f33869,plain,(
  ( ! [X92,X90,X91] : (~p(s(bool,'T_0')) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X91) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X90) | s(bool,X92) = s(bool,sK3(X90,X91)) | s(bool,'T_0') = s(bool,X92)) )),
  inference(superposition,[],[f177,f215])).
fof(f34169,plain,(
  ( ! [X88,X87,X89] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X88),s('type.num.num',X87))) | s(bool,X89) = s(bool,sK3(X87,X88)) | s(bool,'T_0') = s(bool,X89)) )),
  inference(forward_demodulation,[],[f33868,f154])).
fof(f33868,plain,(
  ( ! [X88,X87,X89] : (s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X88),s('type.num.num',X87))) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X88),s('type.num.num',X87))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X88),s('type.num.num',X87))))))) | s(bool,X89) = s(bool,sK3(X87,X88)) | s(bool,'T_0') = s(bool,X89)) )),
  inference(superposition,[],[f180,f215])).
fof(f34167,plain,(
  ( ! [X52,X53,X51] : (s(bool,X51) != s(bool,X52) | s(bool,'T_0') = s(bool,X52) | s(bool,X51) = s(bool,X53) | s(bool,'T_0') = s(bool,X53)) )),
  inference(subsumption_resolution,[],[f33854,f117])).
fof(f33854,plain,(
  ( ! [X52,X53,X51] : (~p(s(bool,'T_0')) | s(bool,X51) != s(bool,X52) | s(bool,'T_0') = s(bool,X52) | s(bool,X51) = s(bool,X53) | s(bool,'T_0') = s(bool,X53)) )),
  inference(superposition,[],[f218,f215])).
fof(f34166,plain,(
  ( ! [X50,X48,X49] : (~p(s(bool,X49)) | s(bool,X48) = s(bool,X49) | s(bool,X48) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(subsumption_resolution,[],[f33853,f117])).
fof(f33853,plain,(
  ( ! [X50,X48,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,X49)) | s(bool,X48) = s(bool,X49) | s(bool,X48) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f167,f215])).
fof(f33849,plain,(
  ( ! [X30,X28,X26,X33,X31,X29,X27,X32] : (s(bool,'T_0') != s(bool,X27) | s(X28,X29) != s(X28,X30) | s(X28,X31) != s(X28,X32) | s(X28,'const.bool.COND_3'(s(bool,X26),s(X28,X32),s(X28,X30))) = s(X28,'const.bool.COND_3'(s(bool,X27),s(X28,X31),s(X28,X29))) | s(bool,X26) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f161,f215])).
fof(f34165,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X16] : (s(X11,'const.bool.COND_3'(s(bool,X10),s(X11,X15),s(X11,X16))) = s(X11,'const.bool.COND_3'(s(bool,X12),s(X11,X13),s(X11,X14))) | s(X11,X13) != s(X11,X15) | s(bool,X10) != s(bool,X12) | s(bool,X10) = s(bool,X17) | s(bool,'T_0') = s(bool,X17)) )),
  inference(subsumption_resolution,[],[f33847,f117])).
fof(f33847,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X16] : (~p(s(bool,'T_0')) | s(X11,'const.bool.COND_3'(s(bool,X10),s(X11,X15),s(X11,X16))) = s(X11,'const.bool.COND_3'(s(bool,X12),s(X11,X13),s(X11,X14))) | s(X11,X13) != s(X11,X15) | s(bool,X10) != s(bool,X12) | s(bool,X10) = s(bool,X17) | s(bool,'T_0') = s(bool,X17)) )),
  inference(superposition,[],[f159,f215])).
fof(f33831,plain,(
  ( ! [X39,X45,X43,X41,X46,X44,X42,X40] : (s(bool,X39) != s(bool,X40) | s(X41,X42) != s(X41,X43) | s(X41,X44) != s(X41,X45) | s(X41,'const.bool.COND_3'(s(bool,'T_0'),s(X41,X45),s(X41,X43))) = s(X41,'const.bool.COND_3'(s(bool,X40),s(X41,X44),s(X41,X42))) | s(bool,X39) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f161,f215])).
fof(f33825,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,sK3(X13,'const.bool.COND_3'(s(bool,X12),s('type.num.num','const.num.|48|_0'),s('type.num.num',X14)))) | s(bool,X12) = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f365,f215])).
fof(f33824,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,sK3('const.bool.COND_3'(s(bool,X8),s('type.num.num','const.num.|48|_0'),s('type.num.num',X9)),X10)) | s(bool,X8) = s(bool,X11) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f311,f215])).
fof(f33823,plain,(
  ( ! [X6,X4,X2,X5,X3] : (s(X3,X4) = s(X3,'const.bool.COND_3'(s(bool,X2),s(X3,X4),s(X3,X5))) | s(bool,X2) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f154,f215])).
fof(f33776,plain,(
  ( ! [X146,X147,X145,X148] : (s(X147,f(s(bool,X148))) != s(X147,g(s(bool,X148))) | s(fun(bool,X147),X145) = s(fun(bool,X147),X146) | s(bool,'T_0') = s(bool,X148) | s(bool,'T_0') = s(bool,sK2(X145,X146,X147,bool))) )),
  inference(superposition,[],[f176,f215])).
fof(f33773,plain,(
  ( ! [X138,X136,X135,X137] : (s(bool,X138) != s(bool,f(s(X135,sK2(X136,X137,bool,X135)))) | s(fun(X135,bool),X136) = s(fun(X135,bool),X137) | s(bool,'T_0') = s(bool,X138) | s(bool,'T_0') = s(bool,g(s(X135,sK2(X136,X137,bool,X135))))) )),
  inference(superposition,[],[f176,f215])).
fof(f33758,plain,(
  ( ! [X90,X88,X89] : (s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X89),s('type.num.num',X88))) = s('type.num.num','const.bool.COND_3'(s(bool,X90),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X89),s('type.num.num',X88))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X89),s('type.num.num',X88))))))) | s(bool,'T_0') = s(bool,X90) | s(bool,'T_0') = s(bool,sK3(X88,X89))) )),
  inference(superposition,[],[f180,f215])).
fof(f33739,plain,(
  ( ! [X30,X28,X26,X33,X31,X29,X27,X32] : (s(bool,X27) != s(bool,X28) | s(X29,X30) != s(X29,X31) | s(X29,X32) != s(X29,X33) | s(X29,'const.bool.COND_3'(s(bool,X26),s(X29,X33),s(X29,X31))) = s(X29,'const.bool.COND_3'(s(bool,X28),s(X29,X32),s(X29,X30))) | s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,X26)) )),
  inference(superposition,[],[f161,f215])).
fof(f33719,plain,(
  ( ! [X30,X31,X29] : (s(bool,'T_0') = s(bool,X30) | s(bool,X31) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X29),s(bool,X30))) | s(bool,'T_0') = s(bool,X31)) )),
  inference(superposition,[],[f215,f155])).
fof(f33718,plain,(
  ( ! [X28,X26,X27] : (s(bool,'T_0') = s(bool,X26) | s(bool,X28) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X26),s(bool,X27))) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f215,f154])).
fof(f33708,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f215,f215])).
fof(f33690,plain,(
  ( ! [X43,X44,X42] : (s(bool,X43) = s(bool,X44) | s(bool,'T_0') = s(bool,X44) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X42),s(bool,X43)))) )),
  inference(superposition,[],[f215,f155])).
fof(f33689,plain,(
  ( ! [X39,X41,X40] : (s(bool,X39) = s(bool,X41) | s(bool,'T_0') = s(bool,X41) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X39),s(bool,X40)))) )),
  inference(superposition,[],[f215,f154])).
fof(f33670,plain,(
  ( ! [X66,X64,X65] : (s(bool,X66) != s(bool,sK3(X64,'const.gcd.lcm_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X65)))) | s(bool,'T_0') = s(bool,X66)) )),
  inference(subsumption_resolution,[],[f28474,f117])).
fof(f28474,plain,(
  ( ! [X66,X64,X65] : (~p(s(bool,'T_0')) | s(bool,X66) != s(bool,sK3(X64,'const.gcd.lcm_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X65)))) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f218,f1487])).
fof(f33665,plain,(
  ( ! [X30,X28,X26,X24,X23,X29,X27,X25] : (s(X25,'const.bool.COND_3'(s(bool,'T_0'),s(X25,X29),s(X25,X30))) = s(X25,'const.bool.COND_3'(s(bool,X26),s(X25,X27),s(X25,X28))) | s(X25,X27) != s(X25,X29) | s(bool,X26) != s(bool,sK3(X23,'const.gcd.lcm_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X24))))) )),
  inference(forward_demodulation,[],[f33664,f1487])).
fof(f33664,plain,(
  ( ! [X30,X28,X26,X24,X23,X29,X27,X25] : (s(X25,'const.bool.COND_3'(s(bool,X26),s(X25,X27),s(X25,X28))) = s(X25,'const.bool.COND_3'(s(bool,sK3(X23,'const.gcd.lcm_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X24)))),s(X25,X29),s(X25,X30))) | s(X25,X27) != s(X25,X29) | s(bool,X26) != s(bool,sK3(X23,'const.gcd.lcm_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X24))))) )),
  inference(subsumption_resolution,[],[f28467,f117])).
fof(f28467,plain,(
  ( ! [X30,X28,X26,X24,X23,X29,X27,X25] : (~p(s(bool,'T_0')) | s(X25,'const.bool.COND_3'(s(bool,X26),s(X25,X27),s(X25,X28))) = s(X25,'const.bool.COND_3'(s(bool,sK3(X23,'const.gcd.lcm_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X24)))),s(X25,X29),s(X25,X30))) | s(X25,X27) != s(X25,X29) | s(bool,X26) != s(bool,sK3(X23,'const.gcd.lcm_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X24))))) )),
  inference(superposition,[],[f159,f1487])).
fof(f26091,plain,(
  ( ! [X61,X62,X60] : (s(bool,X62) != s(bool,sK3('const.gcd.lcm_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X60)),X61)) | s(bool,'T_0') = s(bool,X62)) )),
  inference(subsumption_resolution,[],[f21817,f117])).
fof(f21817,plain,(
  ( ! [X61,X62,X60] : (~p(s(bool,'T_0')) | s(bool,X62) != s(bool,sK3('const.gcd.lcm_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X60)),X61)) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f218,f727])).
fof(f26086,plain,(
  ( ! [X26,X24,X23,X21,X19,X25,X22,X20] : (s(X21,'const.bool.COND_3'(s(bool,'T_0'),s(X21,X25),s(X21,X26))) = s(X21,'const.bool.COND_3'(s(bool,X22),s(X21,X23),s(X21,X24))) | s(X21,X23) != s(X21,X25) | s(bool,X22) != s(bool,sK3('const.gcd.lcm_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X19)),X20))) )),
  inference(forward_demodulation,[],[f26085,f727])).
fof(f26085,plain,(
  ( ! [X26,X24,X23,X21,X19,X25,X22,X20] : (s(X21,'const.bool.COND_3'(s(bool,X22),s(X21,X23),s(X21,X24))) = s(X21,'const.bool.COND_3'(s(bool,sK3('const.gcd.lcm_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X19)),X20)),s(X21,X25),s(X21,X26))) | s(X21,X23) != s(X21,X25) | s(bool,X22) != s(bool,sK3('const.gcd.lcm_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X19)),X20))) )),
  inference(subsumption_resolution,[],[f21810,f117])).
fof(f21810,plain,(
  ( ! [X26,X24,X23,X21,X19,X25,X22,X20] : (~p(s(bool,'T_0')) | s(X21,'const.bool.COND_3'(s(bool,X22),s(X21,X23),s(X21,X24))) = s(X21,'const.bool.COND_3'(s(bool,sK3('const.gcd.lcm_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X19)),X20)),s(X21,X25),s(X21,X26))) | s(X21,X23) != s(X21,X25) | s(bool,X22) != s(bool,sK3('const.gcd.lcm_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X19)),X20))) )),
  inference(superposition,[],[f159,f727])).
fof(f8447,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,X1) != s(X0,X2) | s(X0,X3) != s(X0,X4) | s(X0,'const.bool.COND_3'(s(bool,X5),s(X0,X3),s(X0,X1))) = s(X0,'const.bool.COND_3'(s(bool,X5),s(X0,X4),s(X0,X2)))) )),
  inference(equality_resolution,[],[f161])).
fof(f19837,plain,(
  ( ! [X47,X45,X50,X48,X46,X51,X49] : (s(bool,'T_0') = s(bool,X45) | s(bool,'F_0') != s(bool,X46) | s(X47,X48) != s(X47,X49) | s(X47,'const.bool.COND_3'(s(bool,X45),s(X47,X50),s(X47,X48))) = s(X47,'const.bool.COND_3'(s(bool,X46),s(X47,X51),s(X47,X49)))) )),
  inference(global_subsumption,[],[f116,f117,f189,f125,f128,f127,f132,f131,f154,f155,f161,f160,f159,f162,f163,f167,f164,f176,f180,f179,f178,f177,f190,f197,f199,f202,f204,f215,f218,f258,f275,f274,f286,f295,f296,f311,f323,f322,f324,f338,f347,f348,f365,f381,f382,f402,f403,f419,f420,f451,f452,f469,f470,f491,f492,f512,f513,f572,f573,f574,f607,f611,f642,f680,f683,f648,f649,f651,f653,f655,f656,f657,f658,f660,f662,f663,f664,f665,f671,f672,f689,f688,f715,f716,f727,f934,f945,f1353,f1355,f1357,f1359,f1361,f1487,f1871,f1870,f1955,f1956,f2030,f2041,f2151,f2154,f2265,f2746,f12998,f13000,f13005,f13023,f13031,f13061,f13069,f13146,f13148,f13153,f8434,f8435])).
fof(f8435,plain,(
  ( ! [X47,X45,X50,X48,X46,X51,X49] : (s(bool,'F_0') != s(bool,X46) | s(X47,X48) != s(X47,X49) | s(X47,X50) != s(X47,X51) | s(X47,'const.bool.COND_3'(s(bool,X45),s(X47,X50),s(X47,X48))) = s(X47,'const.bool.COND_3'(s(bool,X46),s(X47,X51),s(X47,X49))) | s(bool,'T_0') = s(bool,X45)) )),
  inference(superposition,[],[f161,f125])).
fof(f8434,plain,(
  ( ! [X39,X43,X41,X38,X44,X42,X40] : (s(bool,X38) != s(bool,X39) | s(X40,X41) != s(X40,X42) | s(X40,X43) != s(X40,X44) | s(X40,'const.bool.COND_3'(s(bool,'F_0'),s(X40,X43),s(X40,X41))) = s(X40,'const.bool.COND_3'(s(bool,X39),s(X40,X44),s(X40,X42))) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f161,f125])).
fof(f13153,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X3),s('type.num.num',X4))))) = s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num',X1))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X4),s('type.num.num',X3)))))) )),
  inference(forward_demodulation,[],[f6398,f155])).
fof(f6398,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X3),s('type.num.num',X4))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X5),s(bool,X0))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num',X1))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X4),s('type.num.num',X3)))))) )),
  inference(unit_resulting_resolution,[],[f162,f162,f155,f161])).
fof(f13148,plain,(
  ( ! [X4,X2,X0,X1] : (s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.num.num',X4))) = s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num',X1))),s('type.num.num',X4)))) )),
  inference(forward_demodulation,[],[f13147,f155])).
fof(f13147,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num',X1))),s('type.num.num',X4))) = s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X3),s('type.num.num',X4)))))) )),
  inference(forward_demodulation,[],[f6401,f155])).
fof(f6401,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X3),s('type.num.num',X4))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X5),s(bool,X0))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num',X1))),s('type.num.num',X4)))) )),
  inference(unit_resulting_resolution,[],[f155,f162,f155,f161])).
fof(f13146,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X3),s('type.num.num',X4))))) = s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num',X1))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X4),s('type.num.num',X3)))))) )),
  inference(forward_demodulation,[],[f6402,f155])).
fof(f6402,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X3),s('type.num.num',X4))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X5),s(bool,X0))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num',X1))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X4),s('type.num.num',X3)))))) )),
  inference(unit_resulting_resolution,[],[f163,f162,f155,f161])).
fof(f13069,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num',X1),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num',X3))))) = s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num',X1),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X3),s('type.num.num',X2)))))) )),
  inference(forward_demodulation,[],[f13068,f155])).
fof(f13068,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num',X1),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num',X3))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X4),s(bool,X0))),s('type.num.num',X1),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X3),s('type.num.num',X2)))))) )),
  inference(forward_demodulation,[],[f6441,f155])).
fof(f6441,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num',X1),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num',X3))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X4),s(bool,X0))),s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X5),s('type.num.num',X1))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X3),s('type.num.num',X2)))))) )),
  inference(unit_resulting_resolution,[],[f162,f155,f155,f161])).
fof(f13061,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num',X1),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X2),s('type.num.num',X3))))) = s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num',X1),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X3),s('type.num.num',X2)))))) )),
  inference(forward_demodulation,[],[f13060,f155])).
fof(f13060,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num',X1),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X2),s('type.num.num',X3))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X4),s(bool,X0))),s('type.num.num',X1),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X3),s('type.num.num',X2)))))) )),
  inference(forward_demodulation,[],[f6445,f155])).
fof(f6445,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num',X1),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X2),s('type.num.num',X3))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X4),s(bool,X0))),s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X5),s('type.num.num',X1))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X3),s('type.num.num',X2)))))) )),
  inference(unit_resulting_resolution,[],[f163,f155,f155,f161])).
fof(f13031,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X3),s('type.num.num',X4))))) = s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X4),s('type.num.num',X3)))))) )),
  inference(forward_demodulation,[],[f13030,f155])).
fof(f13030,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X3),s('type.num.num',X4))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X5),s(bool,X0))),s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X4),s('type.num.num',X3)))))) )),
  inference(forward_demodulation,[],[f6460,f180])).
fof(f6460,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X3),s('type.num.num',X4))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X5),s(bool,X0))),s('type.num.num','const.bool.COND_3'(s(bool,sK3(X2,X1)),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X1),s('type.num.num',X2))))))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X4),s('type.num.num',X3)))))) )),
  inference(unit_resulting_resolution,[],[f162,f180,f155,f161])).
fof(f13023,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X3),s('type.num.num',X4))))) = s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X4),s('type.num.num',X3)))))) )),
  inference(forward_demodulation,[],[f13022,f155])).
fof(f13022,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X3),s('type.num.num',X4))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X5),s(bool,X0))),s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X4),s('type.num.num',X3)))))) )),
  inference(forward_demodulation,[],[f6464,f180])).
fof(f6464,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X3),s('type.num.num',X4))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X5),s(bool,X0))),s('type.num.num','const.bool.COND_3'(s(bool,sK3(X2,X1)),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X1),s('type.num.num',X2))))))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X4),s('type.num.num',X3)))))) )),
  inference(unit_resulting_resolution,[],[f163,f180,f155,f161])).
fof(f13005,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X3),s('type.num.num',X4))))) = s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X2),s('type.num.num',X1))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X4),s('type.num.num',X3)))))) )),
  inference(forward_demodulation,[],[f6474,f155])).
fof(f6474,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X3),s('type.num.num',X4))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X5),s(bool,X0))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X2),s('type.num.num',X1))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X4),s('type.num.num',X3)))))) )),
  inference(unit_resulting_resolution,[],[f162,f163,f155,f161])).
fof(f13000,plain,(
  ( ! [X4,X2,X0,X1] : (s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.num.num',X4))) = s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X2),s('type.num.num',X1))),s('type.num.num',X4)))) )),
  inference(forward_demodulation,[],[f12999,f155])).
fof(f12999,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X2),s('type.num.num',X1))),s('type.num.num',X4))) = s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X3),s('type.num.num',X4)))))) )),
  inference(forward_demodulation,[],[f6477,f155])).
fof(f6477,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X3),s('type.num.num',X4))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X5),s(bool,X0))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X2),s('type.num.num',X1))),s('type.num.num',X4)))) )),
  inference(unit_resulting_resolution,[],[f155,f163,f155,f161])).
fof(f12998,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X3),s('type.num.num',X4))))) = s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X2),s('type.num.num',X1))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X4),s('type.num.num',X3)))))) )),
  inference(forward_demodulation,[],[f6478,f155])).
fof(f6478,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.num.num','const.bool.COND_3'(s(bool,X0),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X3),s('type.num.num',X4))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X5),s(bool,X0))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X2),s('type.num.num',X1))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X4),s('type.num.num',X3)))))) )),
  inference(unit_resulting_resolution,[],[f163,f163,f155,f161])).
fof(f2265,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK3('const.gcd.lcm_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),X1))) )),
  inference(unit_resulting_resolution,[],[f1870,f128])).
fof(f2154,plain,(
  ( ! [X47,X45,X50,X48,X46,X51,X49] : (s(X46,X47) != s(X46,X48) | s(X46,'const.bool.COND_3'(s(bool,X45),s(X46,X51),s(X46,X48))) = s(X46,'const.bool.COND_3'(s(bool,X49),s(X46,X50),s(X46,X47))) | s(bool,X45) != s(bool,X49) | s(bool,'T_0') = s(bool,X45)) )),
  inference(subsumption_resolution,[],[f2011,f189])).
fof(f2011,plain,(
  ( ! [X47,X45,X50,X48,X46,X51,X49] : (p(s(bool,'F_0')) | s(X46,X47) != s(X46,X48) | s(X46,'const.bool.COND_3'(s(bool,X45),s(X46,X51),s(X46,X48))) = s(X46,'const.bool.COND_3'(s(bool,X49),s(X46,X50),s(X46,X47))) | s(bool,X45) != s(bool,X49) | s(bool,'T_0') = s(bool,X45)) )),
  inference(superposition,[],[f160,f125])).
fof(f2041,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num',X1))) = s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X0),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',X2)))))) )),
  inference(forward_demodulation,[],[f2040,f155])).
fof(f2040,plain,(
  ( ! [X4,X2,X0,X1] : (s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X0),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',X2))))) = s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X4),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num',X1)))))) )),
  inference(forward_demodulation,[],[f1995,f155])).
fof(f1995,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X0),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',X2))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,'F_0'))),s('type.num.num',X4),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f189,f162,f155,f160])).
fof(f2030,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X2),s('type.num.num',X1))) = s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X0),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X1),s('type.num.num',X2)))))) )),
  inference(forward_demodulation,[],[f2029,f155])).
fof(f2029,plain,(
  ( ! [X4,X2,X0,X1] : (s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X0),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X1),s('type.num.num',X2))))) = s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X4),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X2),s('type.num.num',X1)))))) )),
  inference(forward_demodulation,[],[f1999,f155])).
fof(f1999,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X0),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X1),s('type.num.num',X2))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,'F_0'))),s('type.num.num',X4),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X2),s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f189,f163,f155,f160])).
fof(f1956,plain,(
  ( ! [X14,X15] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.bool.COND_3'(s(bool,sK3('const.gcd.lcm_2'(s('type.num.num',X14),s('type.num.num','const.num.|48|_0')),X15)),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X15),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X15),s('type.num.num','const.num.|48|_0')))))))) )),
  inference(forward_demodulation,[],[f1902,f683])).
fof(f1902,plain,(
  ( ! [X14,X15] : (s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X15),s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3('const.gcd.lcm_2'(s('type.num.num',X14),s('type.num.num','const.num.|48|_0')),X15)),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X15),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X15),s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f180,f683])).
fof(f1955,plain,(
  ( ! [X12,X13] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.bool.COND_3'(s(bool,sK3(X13,'const.gcd.lcm_2'(s('type.num.num',X12),s('type.num.num','const.num.|48|_0')))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13)))))))) )),
  inference(forward_demodulation,[],[f1901,f680])).
fof(f1901,plain,(
  ( ! [X12,X13] : (s('type.num.num','const.gcd.lcm_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3(X13,'const.gcd.lcm_2'(s('type.num.num',X12),s('type.num.num','const.num.|48|_0')))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13)))))))) )),
  inference(superposition,[],[f180,f683])).
fof(f1870,plain,(
  ( ! [X0,X1] : (p(s(bool,sK3('const.gcd.lcm_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),X1)))) )),
  inference(unit_resulting_resolution,[],[f683,f179])).
fof(f1361,plain,(
  ( ! [X30,X37,X35,X33,X31,X36,X34,X32] : (s(X32,'const.bool.COND_3'(s(bool,'T_0'),s(X32,X36),s(X32,X37))) = s(X32,'const.bool.COND_3'(s(bool,X33),s(X32,X34),s(X32,X35))) | s(X32,X34) != s(X32,X36) | s(bool,X33) != s(bool,sK3('const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X30),s('type.num.num','const.num.|48|_0')),X31))) )),
  inference(forward_demodulation,[],[f1360,f381])).
fof(f1360,plain,(
  ( ! [X30,X37,X35,X33,X31,X36,X34,X32] : (s(X32,'const.bool.COND_3'(s(bool,X33),s(X32,X34),s(X32,X35))) = s(X32,'const.bool.COND_3'(s(bool,sK3('const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X30),s('type.num.num','const.num.|48|_0')),X31)),s(X32,X36),s(X32,X37))) | s(X32,X34) != s(X32,X36) | s(bool,X33) != s(bool,sK3('const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X30),s('type.num.num','const.num.|48|_0')),X31))) )),
  inference(subsumption_resolution,[],[f913,f117])).
fof(f913,plain,(
  ( ! [X30,X37,X35,X33,X31,X36,X34,X32] : (~p(s(bool,'T_0')) | s(X32,'const.bool.COND_3'(s(bool,X33),s(X32,X34),s(X32,X35))) = s(X32,'const.bool.COND_3'(s(bool,sK3('const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X30),s('type.num.num','const.num.|48|_0')),X31)),s(X32,X36),s(X32,X37))) | s(X32,X34) != s(X32,X36) | s(bool,X33) != s(bool,sK3('const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X30),s('type.num.num','const.num.|48|_0')),X31))) )),
  inference(superposition,[],[f159,f381])).
fof(f1359,plain,(
  ( ! [X28,X26,X24,X23,X29,X27,X25,X22] : (s(X24,'const.bool.COND_3'(s(bool,'T_0'),s(X24,X28),s(X24,X29))) = s(X24,'const.bool.COND_3'(s(bool,X25),s(X24,X26),s(X24,X27))) | s(X24,X26) != s(X24,X28) | s(bool,X25) != s(bool,sK3('const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',X22)),X23))) )),
  inference(forward_demodulation,[],[f1358,f311])).
fof(f1358,plain,(
  ( ! [X28,X26,X24,X23,X29,X27,X25,X22] : (s(X24,'const.bool.COND_3'(s(bool,X25),s(X24,X26),s(X24,X27))) = s(X24,'const.bool.COND_3'(s(bool,sK3('const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',X22)),X23)),s(X24,X28),s(X24,X29))) | s(X24,X26) != s(X24,X28) | s(bool,X25) != s(bool,sK3('const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',X22)),X23))) )),
  inference(subsumption_resolution,[],[f912,f117])).
fof(f912,plain,(
  ( ! [X28,X26,X24,X23,X29,X27,X25,X22] : (~p(s(bool,'T_0')) | s(X24,'const.bool.COND_3'(s(bool,X25),s(X24,X26),s(X24,X27))) = s(X24,'const.bool.COND_3'(s(bool,sK3('const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',X22)),X23)),s(X24,X28),s(X24,X29))) | s(X24,X26) != s(X24,X28) | s(bool,X25) != s(bool,sK3('const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',X22)),X23))) )),
  inference(superposition,[],[f159,f311])).
fof(f1357,plain,(
  ( ! [X21,X19,X17,X15,X20,X18,X16] : (s(X16,'const.bool.COND_3'(s(bool,'T_0'),s(X16,X20),s(X16,X21))) = s(X16,'const.bool.COND_3'(s(bool,X17),s(X16,X18),s(X16,X19))) | s(X16,X18) != s(X16,X20) | s(bool,X17) != s(bool,sK3('const.num.|48|_0',X15))) )),
  inference(forward_demodulation,[],[f1356,f204])).
fof(f1356,plain,(
  ( ! [X21,X19,X17,X15,X20,X18,X16] : (s(X16,'const.bool.COND_3'(s(bool,X17),s(X16,X18),s(X16,X19))) = s(X16,'const.bool.COND_3'(s(bool,sK3('const.num.|48|_0',X15)),s(X16,X20),s(X16,X21))) | s(X16,X18) != s(X16,X20) | s(bool,X17) != s(bool,sK3('const.num.|48|_0',X15))) )),
  inference(subsumption_resolution,[],[f911,f117])).
fof(f911,plain,(
  ( ! [X21,X19,X17,X15,X20,X18,X16] : (~p(s(bool,'T_0')) | s(X16,'const.bool.COND_3'(s(bool,X17),s(X16,X18),s(X16,X19))) = s(X16,'const.bool.COND_3'(s(bool,sK3('const.num.|48|_0',X15)),s(X16,X20),s(X16,X21))) | s(X16,X18) != s(X16,X20) | s(bool,X17) != s(bool,sK3('const.num.|48|_0',X15))) )),
  inference(superposition,[],[f159,f204])).
fof(f1355,plain,(
  ( ! [X14,X12,X10,X8,X7,X13,X11,X9] : (s(X9,'const.bool.COND_3'(s(bool,'T_0'),s(X9,X13),s(X9,X14))) = s(X9,'const.bool.COND_3'(s(bool,X10),s(X9,X11),s(X9,X12))) | s(X9,X11) != s(X9,X13) | s(bool,X10) != s(bool,sK3(X7,'const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',X8))))) )),
  inference(forward_demodulation,[],[f1354,f365])).
fof(f1354,plain,(
  ( ! [X14,X12,X10,X8,X7,X13,X11,X9] : (s(X9,'const.bool.COND_3'(s(bool,X10),s(X9,X11),s(X9,X12))) = s(X9,'const.bool.COND_3'(s(bool,sK3(X7,'const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',X8)))),s(X9,X13),s(X9,X14))) | s(X9,X11) != s(X9,X13) | s(bool,X10) != s(bool,sK3(X7,'const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',X8))))) )),
  inference(subsumption_resolution,[],[f910,f117])).
fof(f910,plain,(
  ( ! [X14,X12,X10,X8,X7,X13,X11,X9] : (~p(s(bool,'T_0')) | s(X9,'const.bool.COND_3'(s(bool,X10),s(X9,X11),s(X9,X12))) = s(X9,'const.bool.COND_3'(s(bool,sK3(X7,'const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',X8)))),s(X9,X13),s(X9,X14))) | s(X9,X11) != s(X9,X13) | s(bool,X10) != s(bool,sK3(X7,'const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',X8))))) )),
  inference(superposition,[],[f159,f365])).
fof(f1353,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X1,'const.bool.COND_3'(s(bool,'T_0'),s(X1,X5),s(X1,X6))) = s(X1,'const.bool.COND_3'(s(bool,X2),s(X1,X3),s(X1,X4))) | s(X1,X3) != s(X1,X5) | s(bool,X2) != s(bool,sK3(X0,'const.num.|48|_0'))) )),
  inference(forward_demodulation,[],[f1352,f199])).
fof(f1352,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X1,'const.bool.COND_3'(s(bool,X2),s(X1,X3),s(X1,X4))) = s(X1,'const.bool.COND_3'(s(bool,sK3(X0,'const.num.|48|_0')),s(X1,X5),s(X1,X6))) | s(X1,X3) != s(X1,X5) | s(bool,X2) != s(bool,sK3(X0,'const.num.|48|_0'))) )),
  inference(subsumption_resolution,[],[f909,f117])).
fof(f909,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'T_0')) | s(X1,'const.bool.COND_3'(s(bool,X2),s(X1,X3),s(X1,X4))) = s(X1,'const.bool.COND_3'(s(bool,sK3(X0,'const.num.|48|_0')),s(X1,X5),s(X1,X6))) | s(X1,X3) != s(X1,X5) | s(bool,X2) != s(bool,sK3(X0,'const.num.|48|_0'))) )),
  inference(superposition,[],[f159,f199])).
fof(f945,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',X0))) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X2)))) )),
  inference(forward_demodulation,[],[f944,f154])).
fof(f944,plain,(
  ( ! [X2,X0,X5,X1] : (s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X2))) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.num.num',X5)))) )),
  inference(forward_demodulation,[],[f899,f381])).
fof(f899,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X2))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3('const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X3),s('type.num.num','const.num.|48|_0')),X4)),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.num.num',X5)))) )),
  inference(unit_resulting_resolution,[],[f117,f162,f381,f159])).
fof(f934,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X1),s('type.num.num',X0))) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X2)))) )),
  inference(forward_demodulation,[],[f933,f154])).
fof(f933,plain,(
  ( ! [X2,X0,X5,X1] : (s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X2))) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.num.num',X5)))) )),
  inference(forward_demodulation,[],[f903,f381])).
fof(f903,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X2))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3('const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X3),s('type.num.num','const.num.|48|_0')),X4)),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.num.num',X5)))) )),
  inference(unit_resulting_resolution,[],[f117,f163,f381,f159])).
fof(f727,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK3('const.gcd.lcm_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)),X1))) )),
  inference(unit_resulting_resolution,[],[f688,f128])).
fof(f716,plain,(
  ( ! [X14,X15] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.bool.COND_3'(s(bool,sK3('const.gcd.lcm_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X14)),X15)),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X15),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X15),s('type.num.num','const.num.|48|_0')))))))) )),
  inference(forward_demodulation,[],[f705,f683])).
fof(f705,plain,(
  ( ! [X14,X15] : (s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X15),s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3('const.gcd.lcm_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X14)),X15)),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X15),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X15),s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f180,f680])).
fof(f715,plain,(
  ( ! [X12,X13] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.bool.COND_3'(s(bool,sK3(X13,'const.gcd.lcm_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X12)))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13)))))))) )),
  inference(forward_demodulation,[],[f704,f680])).
fof(f704,plain,(
  ( ! [X12,X13] : (s('type.num.num','const.gcd.lcm_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3(X13,'const.gcd.lcm_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X12)))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13)))))))) )),
  inference(superposition,[],[f180,f680])).
fof(f688,plain,(
  ( ! [X0,X1] : (p(s(bool,sK3('const.gcd.lcm_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)),X1)))) )),
  inference(unit_resulting_resolution,[],[f680,f179])).
fof(f672,plain,(
  ( ! [X17,X15,X16] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X16),s('type.num.num',X15))) | p(s(bool,sK3('const.bool.COND_3'(s(bool,sK3(X15,X16)),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X16),s('type.num.num',X15))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X16),s('type.num.num',X15)))))),X17)))) )),
  inference(superposition,[],[f179,f180])).
fof(f671,plain,(
  ( ! [X14,X12,X13] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X13),s('type.num.num',X12))) | p(s(bool,sK3(X14,'const.bool.COND_3'(s(bool,sK3(X12,X13)),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X13),s('type.num.num',X12))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X13),s('type.num.num',X12)))))))))) )),
  inference(superposition,[],[f178,f180])).
fof(f665,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3(X1,X0)),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X1),s('type.num.num',X0)))))))) )),
  inference(superposition,[],[f180,f163])).
fof(f664,plain,(
  ( ! [X19,X20,X18] : (s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X20),s('type.num.num',X19))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3('const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X18),s('type.num.num',X19)),X20)),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X20),s('type.num.num',X19))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X20),s('type.num.num',X19)))))))) )),
  inference(superposition,[],[f180,f155])).
fof(f663,plain,(
  ( ! [X17,X15,X16] : (s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X17),s('type.num.num',X15))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3('const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',X15),s('type.num.num',X16)),X17)),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X17),s('type.num.num',X15))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X17),s('type.num.num',X15)))))))) )),
  inference(superposition,[],[f180,f154])).
fof(f662,plain,(
  ( ! [X14,X12,X13] : (s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X14),s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X13),s('type.num.num',X12))))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3('const.bool.COND_3'(s(bool,sK3(X12,X13)),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X13),s('type.num.num',X12))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X13),s('type.num.num',X12)))))),X14)),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X14),s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X13),s('type.num.num',X12))))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X14),s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X13),s('type.num.num',X12)))))))))) )),
  inference(superposition,[],[f180,f180])).
fof(f660,plain,(
  ( ! [X6,X8,X7] : (s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X8),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X7),s('type.num.num',X6))))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3('const.gcd.gcd_2'(s('type.num.num',X6),s('type.num.num',X7)),X8)),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X8),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X7),s('type.num.num',X6))))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X8),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X7),s('type.num.num',X6)))))))))) )),
  inference(superposition,[],[f180,f163])).
fof(f658,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',X0))))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3('const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1)),X2)),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',X0))))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',X0)))))))))) )),
  inference(superposition,[],[f180,f162])).
fof(f657,plain,(
  ( ! [X19,X20,X18] : (s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X19),s('type.num.num',X20))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3(X20,'const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X18),s('type.num.num',X19)))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X19),s('type.num.num',X20))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X19),s('type.num.num',X20)))))))) )),
  inference(superposition,[],[f180,f155])).
fof(f656,plain,(
  ( ! [X17,X15,X16] : (s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X15),s('type.num.num',X17))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3(X17,'const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',X15),s('type.num.num',X16)))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X15),s('type.num.num',X17))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X15),s('type.num.num',X17)))))))) )),
  inference(superposition,[],[f180,f154])).
fof(f655,plain,(
  ( ! [X14,X12,X13] : (s('type.num.num','const.gcd.lcm_2'(s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X13),s('type.num.num',X12))),s('type.num.num',X14))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3(X14,'const.bool.COND_3'(s(bool,sK3(X12,X13)),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X13),s('type.num.num',X12))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X13),s('type.num.num',X12)))))))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X13),s('type.num.num',X12))),s('type.num.num',X14))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X13),s('type.num.num',X12))),s('type.num.num',X14)))))))) )),
  inference(superposition,[],[f180,f180])).
fof(f653,plain,(
  ( ! [X6,X8,X7] : (s('type.num.num','const.gcd.lcm_2'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X7),s('type.num.num',X6))),s('type.num.num',X8))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3(X8,'const.gcd.gcd_2'(s('type.num.num',X6),s('type.num.num',X7)))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X7),s('type.num.num',X6))),s('type.num.num',X8))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X7),s('type.num.num',X6))),s('type.num.num',X8)))))))) )),
  inference(superposition,[],[f180,f163])).
fof(f651,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.gcd.lcm_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.num.num',X2))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3(X2,'const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1)))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.num.num',X2))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.num.num',X2)))))))) )),
  inference(superposition,[],[f180,f162])).
fof(f649,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3(X1,X0)),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X0),s('type.num.num',X1)))))))) )),
  inference(superposition,[],[f180,f162])).
fof(f648,plain,(
  ( ! [X8,X9] : (s('type.num.num','const.gcd.lcm_2'(s('type.num.num',X9),s('type.num.num',X8))) = s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.DIV_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X9),s('type.num.num',X8))),s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X9),s('type.num.num',X8))))))) | s(bool,'T_0') = s(bool,sK3(X8,X9))) )),
  inference(superposition,[],[f180,f125])).
fof(f642,plain,(
  ( ! [X28,X26,X27] : (s(bool,X28) != s(bool,sK3('const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X26),s('type.num.num','const.num.|48|_0')),X27)) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f630,f117])).
fof(f630,plain,(
  ( ! [X28,X26,X27] : (~p(s(bool,'T_0')) | s(bool,X28) != s(bool,sK3('const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X26),s('type.num.num','const.num.|48|_0')),X27)) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f218,f381])).
fof(f611,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,sK3('const.bool.COND_3'(s(bool,X0),s('type.num.num',X1),s('type.num.num','const.num.|48|_0')),X2)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f381,f125])).
fof(f607,plain,(
  ( ! [X28,X26,X27] : (s(bool,X28) != s(bool,sK3(X26,'const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',X27)))) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f595,f117])).
fof(f595,plain,(
  ( ! [X28,X26,X27] : (~p(s(bool,'T_0')) | s(bool,X28) != s(bool,sK3(X26,'const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',X27)))) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f218,f365])).
fof(f574,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK2(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK2(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f176,f125])).
fof(f573,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK2(X0,X1,X2,bool))) )),
  inference(superposition,[],[f176,f125])).
fof(f572,plain,(
  ( ! [X26,X24,X25] : (s(bool,X26) != s(bool,sK3('const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',X24)),X25)) | s(bool,'T_0') = s(bool,X26)) )),
  inference(subsumption_resolution,[],[f560,f117])).
fof(f560,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,'T_0')) | s(bool,X26) != s(bool,sK3('const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',X24)),X25)) | s(bool,'T_0') = s(bool,X26)) )),
  inference(superposition,[],[f218,f311])).
fof(f513,plain,(
  ( ! [X12,X13,X11] : (s(bool,X12) != s(bool,sK3('const.num.|48|_0',X13)) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X11),s(bool,X12)))) )),
  inference(superposition,[],[f452,f155])).
fof(f512,plain,(
  ( ! [X10,X8,X9] : (s(bool,X8) != s(bool,sK3('const.num.|48|_0',X10)) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X8),s(bool,X9)))) )),
  inference(superposition,[],[f452,f154])).
fof(f492,plain,(
  ( ! [X12,X13,X11] : (s(bool,X12) != s(bool,sK3(X13,'const.num.|48|_0')) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X11),s(bool,X12)))) )),
  inference(superposition,[],[f451,f155])).
fof(f491,plain,(
  ( ! [X10,X8,X9] : (s(bool,X8) != s(bool,sK3(X10,'const.num.|48|_0')) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X8),s(bool,X9)))) )),
  inference(superposition,[],[f451,f154])).
fof(f470,plain,(
  ( ! [X10,X11,X9] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X10),s('type.num.num',X9))) | p(s(bool,sK3(X11,'const.gcd.gcd_2'(s('type.num.num',X9),s('type.num.num',X10)))))) )),
  inference(superposition,[],[f178,f163])).
fof(f469,plain,(
  ( ! [X6,X8,X7] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.gcd.gcd_2'(s('type.num.num',X7),s('type.num.num',X6))) | p(s(bool,sK3('const.gcd.gcd_2'(s('type.num.num',X6),s('type.num.num',X7)),X8)))) )),
  inference(superposition,[],[f179,f163])).
fof(f452,plain,(
  ( ! [X2,X3] : (s(bool,X3) != s(bool,sK3('const.num.|48|_0',X2)) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f444,f117])).
fof(f444,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | s(bool,X3) != s(bool,sK3('const.num.|48|_0',X2)) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f218,f204])).
fof(f451,plain,(
  ( ! [X0,X1] : (s(bool,X1) != s(bool,sK3(X0,'const.num.|48|_0')) | s(bool,'T_0') = s(bool,X1)) )),
  inference(subsumption_resolution,[],[f443,f117])).
fof(f443,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s(bool,X1) != s(bool,sK3(X0,'const.num.|48|_0')) | s(bool,'T_0') = s(bool,X1)) )),
  inference(superposition,[],[f218,f199])).
fof(f420,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,sK3(X1,'const.bool.COND_3'(s(bool,X0),s('type.num.num',X2),s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f323,f125])).
fof(f419,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK3(X0,'const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X1),s('type.num.num','const.num.|48|_0'))))) )),
  inference(unit_resulting_resolution,[],[f323,f128])).
fof(f403,plain,(
  ( ! [X4,X5,X3] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X4),s('type.num.num',X3))) | p(s(bool,sK3(X5,'const.arithmetic.*_2'(s('type.num.num',X3),s('type.num.num',X4)))))) )),
  inference(superposition,[],[f178,f162])).
fof(f402,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',X0))) | p(s(bool,sK3('const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1)),X2)))) )),
  inference(superposition,[],[f179,f162])).
fof(f382,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,sK3('const.bool.COND_3'(s(bool,X0),s('type.num.num',X1),s('type.num.num','const.num.|48|_0')),X2))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f322,f125])).
fof(f381,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK3('const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),X1))) )),
  inference(unit_resulting_resolution,[],[f322,f128])).
fof(f365,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK3(X0,'const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',X1))))) )),
  inference(unit_resulting_resolution,[],[f275,f128])).
fof(f348,plain,(
  ( ! [X24,X23,X22] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X23) | p(s(bool,sK3(X24,'const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X22),s('type.num.num',X23)))))) )),
  inference(superposition,[],[f178,f155])).
fof(f347,plain,(
  ( ! [X21,X19,X20] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X20) | p(s(bool,sK3('const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X19),s('type.num.num',X20)),X21)))) )),
  inference(superposition,[],[f179,f155])).
fof(f338,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,X1))) | s(bool,'F_0') = s(bool,X1)) )),
  inference(superposition,[],[f155,f125])).
fof(f324,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.bool.COND_3'(s(bool,X0),s(X1,X3),s(X1,X2))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f155,f125])).
fof(f322,plain,(
  ( ! [X0,X1] : (p(s(bool,sK3('const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),X1)))) )),
  inference(unit_resulting_resolution,[],[f155,f179])).
fof(f323,plain,(
  ( ! [X0,X1] : (p(s(bool,sK3(X0,'const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X1),s('type.num.num','const.num.|48|_0')))))) )),
  inference(unit_resulting_resolution,[],[f155,f178])).
fof(f311,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK3('const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)),X1))) )),
  inference(unit_resulting_resolution,[],[f274,f128])).
fof(f296,plain,(
  ( ! [X24,X23,X22] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X22) | p(s(bool,sK3(X24,'const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',X22),s('type.num.num',X23)))))) )),
  inference(superposition,[],[f178,f154])).
fof(f295,plain,(
  ( ! [X21,X19,X20] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X19) | p(s(bool,sK3('const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',X19),s('type.num.num',X20)),X21)))) )),
  inference(superposition,[],[f179,f154])).
fof(f274,plain,(
  ( ! [X0,X1] : (p(s(bool,sK3('const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)),X1)))) )),
  inference(unit_resulting_resolution,[],[f154,f179])).
fof(f275,plain,(
  ( ! [X0,X1] : (p(s(bool,sK3(X0,'const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f154,f178])).
fof(f176,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK2(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK2(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f110])).
fof(f110,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK2(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK2(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2])],[f83])).
fof(f83,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.gcd.LCM_COMM',aHL_EXT)).
fof(f159,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,X2)) | s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))) | s(X0,X3) != s(X0,X4) | s(bool,X1) != s(bool,X2)) )),
  inference(cnf_transformation,[],[f78])).
fof(f161,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,X1) != s(bool,X2) | s(X0,X5) != s(X0,X6) | s(X0,X3) != s(X0,X4) | s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6)))) )),
  inference(cnf_transformation,[],[f78])).
fof(f155,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0,X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f58])).
fof(f58,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X13] : ! [X10,X11] : s(X13,X11) = s(X13,'const.bool.COND_3'(s(bool,'F_0'),s(X13,X10),s(X13,X11)))),
  file('thm.gcd.LCM_COMM','thm.bool.COND_CLAUSES|split|1')).
fof(f132,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f91])).
fof(f34277,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f116,f117,f189,f125,f128,f127,f132,f131,f154,f155,f161,f160,f159,f162,f163,f167,f164,f176,f180,f179,f178,f177,f190,f197,f199,f202,f204,f215,f218,f258,f275,f274,f286,f295,f296,f311,f323,f322,f324,f338,f347,f348,f365,f381,f382,f402,f403,f419,f420,f451,f452,f469,f470,f491,f492,f512,f513,f572,f573,f574,f607,f611,f642,f680,f683,f648,f649,f651,f653,f655,f656,f657,f658,f660,f662,f663,f664,f665,f671,f672,f689,f688,f715,f716,f727,f934,f945,f1353,f1355,f1357,f1359,f1361,f1487,f1871,f1870,f1955,f1956,f2030,f2041,f2151,f2154,f2265,f2746,f12998,f13000,f13005,f13023,f13031,f13061,f13069,f13146,f13148,f13153,f8434,f8435,f19837,f8447,f26086,f26091,f33665,f33670,f33689,f33690,f33708,f33718,f33719,f33739,f33758,f33773,f33776,f33823,f33824,f33825,f34157,f33831,f34165,f33849,f34166,f34167,f34169,f34170,f33876,f33879,f34172,f34158,f34235,f34181,f34242,f34245,f34197,f34198,f34199,f34252,f34254,f34256,f34259,f34262,f34264,f34268,f34270,f34272,f34274,f34276])).
fof(f34275,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f116,f117,f189,f125,f128,f127,f132,f131,f154,f155,f161,f160,f159,f162,f163,f167,f164,f176,f180,f179,f178,f177,f190,f197,f199,f202,f204,f215,f218,f258,f275,f274,f286,f295,f296,f311,f323,f322,f324,f338,f347,f348,f365,f381,f382,f402,f403,f419,f420,f451,f452,f469,f470,f491,f492,f512,f513,f572,f573,f574,f607,f611,f642,f680,f683,f648,f649,f651,f653,f655,f656,f657,f658,f660,f662,f663,f664,f665,f671,f672,f689,f688,f715,f716,f727,f934,f945,f1353,f1355,f1357,f1359,f1361,f1487,f1871,f1870,f1955,f1956,f2030,f2041,f2151,f2154,f2265,f2746,f12998,f13000,f13005,f13023,f13031,f13061,f13069,f13146,f13148,f13153,f8434,f8435,f19837,f8447,f26086,f26091,f33665,f33670,f33689,f33690,f33708,f33718,f33719,f33739,f33758,f33773,f33776,f33823,f33824,f33825,f34157,f33831,f34165,f33849,f34166,f34167,f34169,f34170,f33876,f33879,f34172,f34158,f34235,f34181,f34242,f34245,f34197,f34198,f34199,f34252,f34254,f34256,f34259,f34262,f34264,f34268,f34270,f34272,f34274])).
fof(f34273,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f116,f117,f189,f125,f128,f127,f132,f131,f154,f155,f161,f160,f159,f162,f163,f167,f164,f176,f180,f179,f178,f177,f190,f197,f199,f202,f204,f215,f218,f258,f275,f274,f286,f295,f296,f311,f323,f322,f324,f338,f347,f348,f365,f381,f382,f402,f403,f419,f420,f451,f452,f469,f470,f491,f492,f512,f513,f572,f573,f574,f607,f611,f642,f680,f683,f648,f649,f651,f653,f655,f656,f657,f658,f660,f662,f663,f664,f665,f671,f672,f689,f688,f715,f716,f727,f934,f945,f1353,f1355,f1357,f1359,f1361,f1487,f1871,f1870,f1955,f1956,f2030,f2041,f2151,f2154,f2265,f2746,f12998,f13000,f13005,f13023,f13031,f13061,f13069,f13146,f13148,f13153,f8434,f8435,f19837,f8447,f26086,f26091,f33665,f33670,f33689,f33690,f33708,f33718,f33719,f33739,f33758,f33773,f33776,f33823,f33824,f33825,f34157,f33831,f34165,f33849,f34166,f34167,f34169,f34170,f33876,f33879,f34172,f34158,f34235,f34181,f34242,f34245,f34197,f34198,f34199,f34252,f34254,f34256,f34259,f34262,f34264,f34268,f34270,f34272])).
fof(f34271,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f116,f117,f189,f125,f128,f127,f132,f131,f154,f155,f161,f160,f159,f162,f163,f167,f164,f176,f180,f179,f178,f177,f190,f197,f199,f202,f204,f215,f218,f258,f275,f274,f286,f295,f296,f311,f323,f322,f324,f338,f347,f348,f365,f381,f382,f402,f403,f419,f420,f451,f452,f469,f470,f491,f492,f512,f513,f572,f573,f574,f607,f611,f642,f680,f683,f648,f649,f651,f653,f655,f656,f657,f658,f660,f662,f663,f664,f665,f671,f672,f689,f688,f715,f716,f727,f934,f945,f1353,f1355,f1357,f1359,f1361,f1487,f1871,f1870,f1955,f1956,f2030,f2041,f2151,f2154,f2265,f2746,f12998,f13000,f13005,f13023,f13031,f13061,f13069,f13146,f13148,f13153,f8434,f8435,f19837,f8447,f26086,f26091,f33665,f33670,f33689,f33690,f33708,f33718,f33719,f33739,f33758,f33773,f33776,f33823,f33824,f33825,f34157,f33831,f34165,f33849,f34166,f34167,f34169,f34170,f33876,f33879,f34172,f34158,f34235,f34181,f34242,f34245,f34197,f34198,f34199,f34252,f34254,f34256,f34259,f34262,f34264,f34268,f34270])).
fof(f34269,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f116,f117,f189,f125,f128,f127,f132,f131,f154,f155,f161,f160,f159,f162,f163,f167,f164,f176,f180,f179,f178,f177,f190,f197,f199,f202,f204,f215,f218,f258,f275,f274,f286,f295,f296,f311,f323,f322,f324,f338,f347,f348,f365,f381,f382,f402,f403,f419,f420,f451,f452,f469,f470,f491,f492,f512,f513,f572,f573,f574,f607,f611,f642,f680,f683,f648,f649,f651,f653,f655,f656,f657,f658,f660,f662,f663,f664,f665,f671,f672,f689,f688,f715,f716,f727,f934,f945,f1353,f1355,f1357,f1359,f1361,f1487,f1871,f1870,f1955,f1956,f2030,f2041,f2151,f2154,f2265,f2746,f12998,f13000,f13005,f13023,f13031,f13061,f13069,f13146,f13148,f13153,f8434,f8435,f19837,f8447,f26086,f26091,f33665,f33670,f33689,f33690,f33708,f33718,f33719,f33739,f33758,f33773,f33776,f33823,f33824,f33825,f34157,f33831,f34165,f33849,f34166,f34167,f34169,f34170,f33876,f33879,f34172,f34158,f34235,f34181,f34242,f34245,f34197,f34198,f34199,f34252,f34254,f34256,f34259,f34262,f34264,f34268])).
fof(f34265,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f116,f117,f189,f125,f128,f127,f132,f131,f154,f155,f161,f160,f159,f162,f163,f167,f164,f176,f180,f179,f178,f177,f190,f197,f199,f202,f204,f215,f218,f258,f275,f274,f286,f295,f296,f311,f323,f322,f324,f338,f347,f348,f365,f381,f382,f402,f403,f419,f420,f451,f452,f469,f470,f491,f492,f512,f513,f572,f573,f574,f607,f611,f642,f680,f683,f648,f649,f651,f653,f655,f656,f657,f658,f660,f662,f663,f664,f665,f671,f672,f689,f688,f715,f716,f727,f934,f945,f1353,f1355,f1357,f1359,f1361,f1487,f1871,f1870,f1955,f1956,f2030,f2041,f2151,f2154,f2265,f2746,f12998,f13000,f13005,f13023,f13031,f13061,f13069,f13146,f13148,f13153,f8434,f8435,f19837,f8447,f26086,f26091,f33665,f33670,f33689,f33690,f33708,f33718,f33719,f33739,f33758,f33773,f33776,f33823,f33824,f33825,f34157,f33831,f34165,f33849,f34166,f34167,f34169,f34170,f33876,f33879,f34172,f34158,f34235,f34181,f34242,f34245,f34197,f34198,f34199,f34252,f34254,f34256,f34259,f34262,f34264])).
fof(f34263,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f116,f117,f189,f125,f128,f127,f132,f131,f154,f155,f161,f160,f159,f162,f163,f167,f164,f176,f180,f179,f178,f177,f190,f197,f199,f202,f204,f215,f218,f258,f275,f274,f286,f295,f296,f311,f323,f322,f324,f338,f347,f348,f365,f381,f382,f402,f403,f419,f420,f451,f452,f469,f470,f491,f492,f512,f513,f572,f573,f574,f607,f611,f642,f680,f683,f648,f649,f651,f653,f655,f656,f657,f658,f660,f662,f663,f664,f665,f671,f672,f689,f688,f715,f716,f727,f934,f945,f1353,f1355,f1357,f1359,f1361,f1487,f1871,f1870,f1955,f1956,f2030,f2041,f2151,f2154,f2265,f2746,f12998,f13000,f13005,f13023,f13031,f13061,f13069,f13146,f13148,f13153,f8434,f8435,f19837,f8447,f26086,f26091,f33665,f33670,f33689,f33690,f33708,f33718,f33719,f33739,f33758,f33773,f33776,f33823,f33824,f33825,f34157,f33831,f34165,f33849,f34166,f34167,f34169,f34170,f33876,f33879,f34172,f34158,f34235,f34181,f34242,f34245,f34197,f34198,f34199,f34252,f34254,f34256,f34259,f34262])).
fof(f34260,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f116,f117,f189,f125,f128,f127,f132,f131,f154,f155,f161,f160,f159,f162,f163,f167,f164,f176,f180,f179,f178,f177,f190,f197,f199,f202,f204,f215,f218,f258,f275,f274,f286,f295,f296,f311,f323,f322,f324,f338,f347,f348,f365,f381,f382,f402,f403,f419,f420,f451,f452,f469,f470,f491,f492,f512,f513,f572,f573,f574,f607,f611,f642,f680,f683,f648,f649,f651,f653,f655,f656,f657,f658,f660,f662,f663,f664,f665,f671,f672,f689,f688,f715,f716,f727,f934,f945,f1353,f1355,f1357,f1359,f1361,f1487,f1871,f1870,f1955,f1956,f2030,f2041,f2151,f2154,f2265,f2746,f12998,f13000,f13005,f13023,f13031,f13061,f13069,f13146,f13148,f13153,f8434,f8435,f19837,f8447,f26086,f26091,f33665,f33670,f33689,f33690,f33708,f33718,f33719,f33739,f33758,f33773,f33776,f33823,f33824,f33825,f34157,f33831,f34165,f33849,f34166,f34167,f34169,f34170,f33876,f33879,f34172,f34158,f34235,f34181,f34242,f34245,f34197,f34198,f34199,f34252,f34254,f34256,f34259])).
fof(f34257,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f116,f117,f189,f125,f128,f127,f132,f131,f154,f155,f161,f160,f159,f162,f163,f167,f164,f176,f180,f179,f178,f177,f190,f197,f199,f202,f204,f215,f218,f258,f275,f274,f286,f295,f296,f311,f323,f322,f324,f338,f347,f348,f365,f381,f382,f402,f403,f419,f420,f451,f452,f469,f470,f491,f492,f512,f513,f572,f573,f574,f607,f611,f642,f680,f683,f648,f649,f651,f653,f655,f656,f657,f658,f660,f662,f663,f664,f665,f671,f672,f689,f688,f715,f716,f727,f934,f945,f1353,f1355,f1357,f1359,f1361,f1487,f1871,f1870,f1955,f1956,f2030,f2041,f2151,f2154,f2265,f2746,f12998,f13000,f13005,f13023,f13031,f13061,f13069,f13146,f13148,f13153,f8434,f8435,f19837,f8447,f26086,f26091,f33665,f33670,f33689,f33690,f33708,f33718,f33719,f33739,f33758,f33773,f33776,f33823,f33824,f33825,f34157,f33831,f34165,f33849,f34166,f34167,f34169,f34170,f33876,f33879,f34172,f34158,f34235,f34181,f34242,f34245,f34197,f34198,f34199,f34252,f34254,f34256])).
fof(f34255,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f116,f117,f189,f125,f128,f127,f132,f131,f154,f155,f161,f160,f159,f162,f163,f167,f164,f176,f180,f179,f178,f177,f190,f197,f199,f202,f204,f215,f218,f258,f275,f274,f286,f295,f296,f311,f323,f322,f324,f338,f347,f348,f365,f381,f382,f402,f403,f419,f420,f451,f452,f469,f470,f491,f492,f512,f513,f572,f573,f574,f607,f611,f642,f680,f683,f648,f649,f651,f653,f655,f656,f657,f658,f660,f662,f663,f664,f665,f671,f672,f689,f688,f715,f716,f727,f934,f945,f1353,f1355,f1357,f1359,f1361,f1487,f1871,f1870,f1955,f1956,f2030,f2041,f2151,f2154,f2265,f2746,f12998,f13000,f13005,f13023,f13031,f13061,f13069,f13146,f13148,f13153,f8434,f8435,f19837,f8447,f26086,f26091,f33665,f33670,f33689,f33690,f33708,f33718,f33719,f33739,f33758,f33773,f33776,f33823,f33824,f33825,f34157,f33831,f34165,f33849,f34166,f34167,f34169,f34170,f33876,f33879,f34172,f34158,f34235,f34181,f34242,f34245,f34197,f34198,f34199,f34252,f34254])).
fof(f34253,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f116,f117,f189,f125,f128,f127,f132,f131,f154,f155,f161,f160,f159,f162,f163,f167,f164,f176,f180,f179,f178,f177,f190,f197,f199,f202,f204,f215,f218,f258,f275,f274,f286,f295,f296,f311,f323,f322,f324,f338,f347,f348,f365,f381,f382,f402,f403,f419,f420,f451,f452,f469,f470,f491,f492,f512,f513,f572,f573,f574,f607,f611,f642,f680,f683,f648,f649,f651,f653,f655,f656,f657,f658,f660,f662,f663,f664,f665,f671,f672,f689,f688,f715,f716,f727,f934,f945,f1353,f1355,f1357,f1359,f1361,f1487,f1871,f1870,f1955,f1956,f2030,f2041,f2151,f2154,f2265,f2746,f12998,f13000,f13005,f13023,f13031,f13061,f13069,f13146,f13148,f13153,f8434,f8435,f19837,f8447,f26086,f26091,f33665,f33670,f33689,f33690,f33708,f33718,f33719,f33739,f33758,f33773,f33776,f33823,f33824,f33825,f34157,f33831,f34165,f33849,f34166,f34167,f34169,f34170,f33876,f33879,f34172,f34158,f34235,f34181,f34242,f34245,f34197,f34198,f34199,f34252])).
fof(f34251,plain,(
  $false | ($spl2 | $spl4)),
  inference(global_subsumption,[],[f34250,f116,f117,f189,f125,f128,f127,f132,f131,f154,f155,f161,f160,f159,f162,f163,f167,f164,f176,f180,f179,f178,f177,f190,f197,f199,f202,f204,f215,f218,f258,f275,f274,f286,f295,f296,f311,f323,f322,f324,f338,f347,f348,f365,f381,f382,f402,f403,f419,f420,f451,f452,f469,f470,f491,f492,f512,f513,f572,f573,f574,f607,f611,f642,f680,f683,f648,f649,f651,f653,f655,f656,f657,f658,f660,f662,f663,f664,f665,f671,f672,f689,f688,f715,f716,f727,f934,f945,f1353,f1355,f1357,f1359,f1361,f1487,f1871,f1870,f1955,f1956,f2030,f2041,f2151,f2154,f2265,f2746,f12998,f13000,f13005,f13023,f13031,f13061,f13069,f13146,f13148,f13153,f8434,f8435,f19837,f8447,f26086,f26091,f33665,f33670,f33689,f33690,f33708,f33718,f33719,f33739,f33758,f33773,f33776,f33823,f33824,f33825,f34157,f33831,f34165,f33849,f34166,f34167,f34169,f34170,f33876,f33879,f34172,f34158,f34235,f34181,f34242,f34245,f34197,f34198,f34199])).
fof(f34250,plain,(
  ( ! [X39,X43,X41,X42,X40] : (s(X39,X41) = s(X39,'const.bool.COND_3'(s(bool,X42),s(X39,X43),s(X39,X40))) | s(X39,X40) != s(X39,X41)) ) | ($spl2 | $spl4)),
  inference(subsumption_resolution,[],[f34200,f34158])).
fof(f34200,plain,(
  ( ! [X39,X43,X41,X42,X40] : (s(bool,'T_0') != s(bool,X42) | s(X39,X41) = s(X39,'const.bool.COND_3'(s(bool,X42),s(X39,X43),s(X39,X40))) | s(X39,X40) != s(X39,X41)) ) | ($spl2 | $spl4)),
  inference(backward_demodulation,[],[f34158,f2153])).
fof(f34249,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f116,f117,f189,f125,f128,f127,f132,f131,f154,f155,f161,f160,f159,f162,f163,f167,f164,f176,f180,f179,f178,f177,f190,f197,f199,f202,f204,f215,f218,f258,f275,f274,f286,f295,f296,f311,f323,f322,f324,f338,f347,f348,f365,f381,f382,f402,f403,f419,f420,f451,f452,f469,f470,f491,f492,f512,f513,f572,f573,f574,f607,f611,f642,f680,f683,f648,f649,f651,f653,f655,f656,f657,f658,f660,f662,f663,f664,f665,f671,f672,f689,f688,f715,f716,f727,f934,f945,f1353,f1355,f1357,f1359,f1361,f1487,f1871,f1870,f1955,f1956,f2030,f2041,f2151,f2154,f2265,f2746,f12998,f13000,f13005,f13023,f13031,f13061,f13069,f13146,f13148,f13153,f8434,f8435,f19837,f8447,f26086,f26091,f33665,f33670,f33689,f33690,f33708,f33718,f33719,f33739,f33758,f33773,f33776,f33823,f33824,f33825,f34157,f33831,f34165,f33849,f34166,f34167,f34169,f34170,f33876,f33879,f34172,f34158,f34235,f34181,f34242,f34245,f34197,f34198,f34199])).
fof(f34248,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f116,f117,f189,f125,f128,f127,f132,f131,f154,f155,f161,f160,f159,f162,f163,f167,f164,f176,f180,f179,f178,f177,f190,f197,f199,f202,f204,f215,f218,f258,f275,f274,f286,f295,f296,f311,f323,f322,f324,f338,f347,f348,f365,f381,f382,f402,f403,f419,f420,f451,f452,f469,f470,f491,f492,f512,f513,f572,f573,f574,f607,f611,f642,f680,f683,f648,f649,f651,f653,f655,f656,f657,f658,f660,f662,f663,f664,f665,f671,f672,f689,f688,f715,f716,f727,f934,f945,f1353,f1355,f1357,f1359,f1361,f1487,f1871,f1870,f1955,f1956,f2030,f2041,f2151,f2154,f2265,f2746,f12998,f13000,f13005,f13023,f13031,f13061,f13069,f13146,f13148,f13153,f8434,f8435,f19837,f8447,f26086,f26091,f33665,f33670,f33689,f33690,f33708,f33718,f33719,f33739,f33758,f33773,f33776,f33823,f33824,f33825,f34157,f33831,f34165,f33849,f34166,f34167,f34169,f34170,f33876,f33879,f34172,f34158,f34235,f34181,f34242,f34245,f34197,f34198])).
fof(f34247,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f116,f117,f189,f125,f128,f127,f132,f131,f154,f155,f161,f160,f159,f162,f163,f167,f164,f176,f180,f179,f178,f177,f190,f197,f199,f202,f204,f215,f218,f258,f275,f274,f286,f295,f296,f311,f323,f322,f324,f338,f347,f348,f365,f381,f382,f402,f403,f419,f420,f451,f452,f469,f470,f491,f492,f512,f513,f572,f573,f574,f607,f611,f642,f680,f683,f648,f649,f651,f653,f655,f656,f657,f658,f660,f662,f663,f664,f665,f671,f672,f689,f688,f715,f716,f727,f934,f945,f1353,f1355,f1357,f1359,f1361,f1487,f1871,f1870,f1955,f1956,f2030,f2041,f2151,f2154,f2265,f2746,f12998,f13000,f13005,f13023,f13031,f13061,f13069,f13146,f13148,f13153,f8434,f8435,f19837,f8447,f26086,f26091,f33665,f33670,f33689,f33690,f33708,f33718,f33719,f33739,f33758,f33773,f33776,f33823,f33824,f33825,f34157,f33831,f34165,f33849,f34166,f34167,f34169,f34170,f33876,f33879,f34172,f34158,f34235,f34181,f34242,f34245,f34197])).
fof(f34246,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f116,f117,f189,f125,f128,f127,f132,f131,f154,f155,f161,f160,f159,f162,f163,f167,f164,f176,f180,f179,f178,f177,f190,f197,f199,f202,f204,f215,f218,f258,f275,f274,f286,f295,f296,f311,f323,f322,f324,f338,f347,f348,f365,f381,f382,f402,f403,f419,f420,f451,f452,f469,f470,f491,f492,f512,f513,f572,f573,f574,f607,f611,f642,f680,f683,f648,f649,f651,f653,f655,f656,f657,f658,f660,f662,f663,f664,f665,f671,f672,f689,f688,f715,f716,f727,f934,f945,f1353,f1355,f1357,f1359,f1361,f1487,f1871,f1870,f1955,f1956,f2030,f2041,f2151,f2154,f2265,f2746,f12998,f13000,f13005,f13023,f13031,f13061,f13069,f13146,f13148,f13153,f8434,f8435,f19837,f8447,f26086,f26091,f33665,f33670,f33689,f33690,f33708,f33718,f33719,f33739,f33758,f33773,f33776,f33823,f33824,f33825,f34157,f33831,f34165,f33849,f34166,f34167,f34169,f34170,f33876,f33879,f34172,f34158,f34235,f34181,f34242,f34245])).
fof(f34243,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f116,f117,f189,f125,f128,f127,f132,f131,f154,f155,f161,f160,f159,f162,f163,f167,f164,f176,f180,f179,f178,f177,f190,f197,f199,f202,f204,f215,f218,f258,f275,f274,f286,f295,f296,f311,f323,f322,f324,f338,f347,f348,f365,f381,f382,f402,f403,f419,f420,f451,f452,f469,f470,f491,f492,f512,f513,f572,f573,f574,f607,f611,f642,f680,f683,f648,f649,f651,f653,f655,f656,f657,f658,f660,f662,f663,f664,f665,f671,f672,f689,f688,f715,f716,f727,f934,f945,f1353,f1355,f1357,f1359,f1361,f1487,f1871,f1870,f1955,f1956,f2030,f2041,f2151,f2154,f2265,f2746,f12998,f13000,f13005,f13023,f13031,f13061,f13069,f13146,f13148,f13153,f8434,f8435,f19837,f8447,f26086,f26091,f33665,f33670,f33689,f33690,f33708,f33718,f33719,f33739,f33758,f33773,f33776,f33823,f33824,f33825,f34157,f33831,f34165,f33849,f34166,f34167,f34169,f34170,f33876,f33879,f34172,f34158,f34235,f34181,f34242])).
fof(f34238,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f34182,f117])).
fof(f34182,plain,(
  ~p(s(bool,'T_0')) | $spl4),
  inference(backward_demodulation,[],[f34158,f189])).
fof(f34237,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f116,f117,f189,f125,f128,f127,f132,f131,f154,f155,f161,f160,f159,f162,f163,f167,f164,f176,f180,f179,f178,f177,f190,f197,f199,f202,f204,f215,f218,f258,f275,f274,f286,f295,f296,f311,f323,f322,f324,f338,f347,f348,f365,f381,f382,f402,f403,f419,f420,f451,f452,f469,f470,f491,f492,f512,f513,f572,f573,f574,f607,f611,f642,f680,f683,f648,f649,f651,f653,f655,f656,f657,f658,f660,f662,f663,f664,f665,f671,f672,f689,f688,f715,f716,f727,f934,f945,f1353,f1355,f1357,f1359,f1361,f1487,f1871,f1870,f1955,f1956,f2030,f2041,f2151,f2154,f2265,f2746,f12998,f13000,f13005,f13023,f13031,f13061,f13069,f13146,f13148,f13153,f8434,f8435,f19837,f8447,f26086,f26091,f33665,f33670,f33689,f33690,f33708,f33718,f33719,f33739,f33758,f33773,f33776,f33823,f33824,f33825,f34157,f33831,f34165,f33849,f34166,f34167,f34169,f34170,f33876,f33879,f34172,f34158,f34235,f34181])).
fof(f34236,plain,(
  $false | $spl4),
  inference(global_subsumption,[],[f116,f117,f189,f125,f128,f127,f132,f131,f154,f155,f161,f160,f159,f162,f163,f167,f164,f176,f180,f179,f178,f177,f190,f197,f199,f202,f204,f215,f218,f258,f275,f274,f286,f295,f296,f311,f323,f322,f324,f338,f347,f348,f365,f381,f382,f402,f403,f419,f420,f451,f452,f469,f470,f491,f492,f512,f513,f572,f573,f574,f607,f611,f642,f680,f683,f648,f649,f651,f653,f655,f656,f657,f658,f660,f662,f663,f664,f665,f671,f672,f689,f688,f715,f716,f727,f934,f945,f1353,f1355,f1357,f1359,f1361,f1487,f1871,f1870,f1955,f1956,f2030,f2041,f2151,f2154,f2265,f2746,f12998,f13000,f13005,f13023,f13031,f13061,f13069,f13146,f13148,f13153,f8434,f8435,f19837,f8447,f26086,f26091,f33665,f33670,f33689,f33690,f33708,f33718,f33719,f33739,f33758,f33773,f33776,f33823,f33824,f33825,f34157,f33831,f34165,f33849,f34166,f34167,f34169,f34170,f33876,f33879,f34172,f34158,f34235])).
fof(f34160,plain,(
  ( ! [X28,X26,X24,X27,X25] : (s(X24,X28) = s(X24,'const.bool.COND_3'(s(bool,X25),s(X24,X26),s(X24,X27))) | s(bool,'T_0') != s(bool,X25) | s(X24,X26) != s(X24,X28)) ) | $spl8),
  inference(cnf_transformation,[],[f34160_D])).
fof(f34160_D,plain,(
  ( ! [X28,X26,X24,X27,X25] : (s(X24,X28) = s(X24,'const.bool.COND_3'(s(bool,X25),s(X24,X26),s(X24,X27))) | s(bool,'T_0') != s(bool,X25) | s(X24,X26) != s(X24,X28)) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f34172,plain,(
  ( ! [X30,X28,X26,X24,X23,X27,X25] : (s(X24,X28) = s(X24,'const.bool.COND_3'(s(bool,X25),s(X24,X26),s(X24,X27))) | ~p(s(bool,X23)) | s(X24,X26) != s(X24,X28) | s(bool,'T_0') != s(bool,X25) | s(bool,'T_0') = s(bool,X30)) )),
  inference(forward_demodulation,[],[f34171,f154])).
fof(f34171,plain,(
  ( ! [X30,X28,X26,X24,X23,X29,X27,X25] : (~p(s(bool,X23)) | s(X24,'const.bool.COND_3'(s(bool,'T_0'),s(X24,X28),s(X24,X29))) = s(X24,'const.bool.COND_3'(s(bool,X25),s(X24,X26),s(X24,X27))) | s(X24,X26) != s(X24,X28) | s(bool,'T_0') != s(bool,X25) | s(bool,'T_0') = s(bool,X30)) )),
  inference(subsumption_resolution,[],[f33887,f218])).
fof(f33887,plain,(
  ( ! [X30,X28,X26,X24,X23,X29,X27,X25] : (~p(s(bool,X23)) | s(X24,'const.bool.COND_3'(s(bool,'T_0'),s(X24,X28),s(X24,X29))) = s(X24,'const.bool.COND_3'(s(bool,X25),s(X24,X26),s(X24,X27))) | s(X24,X26) != s(X24,X28) | s(bool,'T_0') != s(bool,X25) | s(bool,X23) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f159,f215])).
fof(f34157,plain,(
  ( ! [X30,X28,X26,X24,X23,X27,X25] : (s(X24,X28) = s(X24,'const.bool.COND_3'(s(bool,X25),s(X24,X26),s(X24,X27))) | ~p(s(bool,X23)) | s(X24,X26) != s(X24,X28) | s(bool,'T_0') != s(bool,X25) | s(bool,'T_0') = s(bool,X30)) )),
  inference(forward_demodulation,[],[f34156,f154])).
fof(f34156,plain,(
  ( ! [X30,X28,X26,X24,X23,X29,X27,X25] : (~p(s(bool,X23)) | s(X24,'const.bool.COND_3'(s(bool,'T_0'),s(X24,X28),s(X24,X29))) = s(X24,'const.bool.COND_3'(s(bool,X25),s(X24,X26),s(X24,X27))) | s(X24,X26) != s(X24,X28) | s(bool,'T_0') != s(bool,X25) | s(bool,'T_0') = s(bool,X30)) )),
  inference(subsumption_resolution,[],[f33829,f218])).
fof(f33829,plain,(
  ( ! [X30,X28,X26,X24,X23,X29,X27,X25] : (~p(s(bool,X23)) | s(X24,'const.bool.COND_3'(s(bool,'T_0'),s(X24,X28),s(X24,X29))) = s(X24,'const.bool.COND_3'(s(bool,X25),s(X24,X26),s(X24,X27))) | s(X24,X26) != s(X24,X28) | s(bool,'T_0') != s(bool,X25) | s(bool,X23) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f159,f215])).
fof(f2152,plain,(
  ( ! [X38] : (p(s(bool,X38))) ) | $spl0),
  inference(cnf_transformation,[],[f2152_D])).
fof(f2152_D,plain,(
  ( ! [X38] : (p(s(bool,X38))) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f2153,plain,(
  ( ! [X39,X43,X41,X42,X40] : (s(X39,X41) = s(X39,'const.bool.COND_3'(s(bool,X42),s(X39,X43),s(X39,X40))) | s(bool,'F_0') != s(bool,X42) | s(X39,X40) != s(X39,X41)) ) | $spl2),
  inference(cnf_transformation,[],[f2153_D])).
fof(f2153_D,plain,(
  ( ! [X39,X43,X41,X42,X40] : (s(X39,X41) = s(X39,'const.bool.COND_3'(s(bool,X42),s(X39,X43),s(X39,X40))) | s(bool,'F_0') != s(bool,X42) | s(X39,X40) != s(X39,X41)) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f2151,plain,(
  ( ! [X39,X43,X41,X38,X42,X40] : (s(X39,X41) = s(X39,'const.bool.COND_3'(s(bool,X42),s(X39,X43),s(X39,X40))) | p(s(bool,X38)) | s(X39,X40) != s(X39,X41) | s(bool,'F_0') != s(bool,X42)) )),
  inference(forward_demodulation,[],[f2150,f155])).
fof(f2150,plain,(
  ( ! [X39,X43,X41,X38,X44,X42,X40] : (p(s(bool,X38)) | s(X39,X40) != s(X39,X41) | s(X39,'const.bool.COND_3'(s(bool,'F_0'),s(X39,X44),s(X39,X41))) = s(X39,'const.bool.COND_3'(s(bool,X42),s(X39,X43),s(X39,X40))) | s(bool,'F_0') != s(bool,X42)) )),
  inference(subsumption_resolution,[],[f2010,f127])).
fof(f2010,plain,(
  ( ! [X39,X43,X41,X38,X44,X42,X40] : (p(s(bool,X38)) | s(X39,X40) != s(X39,X41) | s(X39,'const.bool.COND_3'(s(bool,'F_0'),s(X39,X44),s(X39,X41))) = s(X39,'const.bool.COND_3'(s(bool,X42),s(X39,X43),s(X39,X40))) | s(bool,'F_0') != s(bool,X42) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f160,f125])).
% SZS output end Proof for thm.gcd
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 175
% Passive clauses: 834
% Generated clauses: 427447
% Final active clauses: 63
% Final passive clauses: 285
% Input formulas: 32
% Initial clauses: 71
% 
% Duplicate literals: 599
% Trivial inequalities: 171
% Fw subsumption resolutions: 594
% Fw demodulations: 291418
% Bw demodulations: 344
% Condensations: 849
% Global subsumptions: 23
% 
% Simple tautologies: 224
% Equational tautologies: 245
% Forward subsumptions: 65041
% Fw demodulations to eq. taut.: 67368
% Bw demodulations to eq. taut.: 12
% 
% Unit resulting resolution: 126338
% Factoring: 1
% Forward superposition: 3180
% Backward superposition: 3261
% Self superposition: 130
% Equality factoring: 72
% Equality resolution: 29
% 
% Split clauses: 16
% Split components: 35
% SAT solver clauses: 956
% SAT solver unit clauses: 165
% SAT solver binary clauses: 361
% 
% Sat splits: 16
% Sat splitting refutations: 65
% 
% Memory used [KB]: 58719
% Time elapsed: 6.141 s
% ------------------------------
% ------------------------------
% Success in time 7.431 s
