Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.ring
% (54416)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 619
% Passive clauses: 6762
% Generated clauses: 24915
% Final active clauses: 613
% Final passive clauses: 5896
% Input formulas: 260
% Initial clauses: 380
% 
% Duplicate literals: 82
% Trivial inequalities: 81
% Fw subsumption resolutions: 773
% Fw demodulations: 4288
% Bw demodulations: 6
% 
% Simple tautologies: 201
% Equational tautologies: 26
% Forward subsumptions: 12667
% 
% Binary resolution: 3607
% Factoring: 350
% Forward superposition: 6356
% Backward superposition: 8801
% Self superposition: 77
% Equality factoring: 2
% Equality resolution: 20
% 
% Split clauses: 163
% Split components: 327
% SAT solver clauses: 281
% SAT solver unit clauses: 19
% SAT solver binary clauses: 61
% 
% Sat splits: 163
% 
% TWLsolver clauses: 281
% TWLsolver calls for satisfiability: 231
% 
% Memory used [KB]: 20340
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.ring
% (54431)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 359
% Passive clauses: 2808
% Generated clauses: 20181
% Final active clauses: 356
% Final passive clauses: 2314
% Input formulas: 260
% Initial clauses: 380
% 
% Duplicate literals: 61
% Trivial inequalities: 8
% Fw subsumption resolutions: 408
% Fw demodulations: 2256
% Bw demodulations: 138
% Condensations: 105
% Global subsumptions: 18
% 
% Simple tautologies: 160
% Equational tautologies: 11
% Forward subsumptions: 10967
% Fw demodulations to eq. taut.: 4
% 
% Binary resolution: 1695
% Factoring: 268
% Forward superposition: 6045
% Backward superposition: 5851
% Self superposition: 113
% Equality factoring: 2
% Equality resolution: 5
% 
% Split clauses: 18
% Split components: 36
% SAT solver clauses: 3549
% SAT solver unit clauses: 579
% SAT solver binary clauses: 79
% SAT solver learnt clauses: 36
% SAT solver learnt literals: 10
% 
% Sat splits: 18
% 
% TWLsolver clauses: 3477
% TWLsolver calls for satisfiability: 1379
% 
% Memory used [KB]: 29551
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.ring
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.ring
% SZS output start Proof for thm.ring
fof(f83137,plain,(
  $false),
  inference(subsumption_resolution,[],[f83136,f1038])).
fof(f1038,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),X3),s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),X1))))) = s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),X3),s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f470])).
fof(f470,plain,(
  ! [X0,X1,X2,X3] : s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),X3),s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),X1))))) = s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),X3),s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f469])).
fof(f469,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),X3),s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),X1))))) = s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),X3),s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X7] : ! [X19,X9,X10] : s('type.ring.ring'(X7),'const.ring.ring_R|49|_fupd_2'(s(fun(X7,X7),X10),s('type.ring.ring'(X7),'const.ring.ring_R|48|_fupd_2'(s(fun(X7,X7),X9),s('type.ring.ring'(X7),X19))))) = s('type.ring.ring'(X7),'const.ring.ring_R|48|_fupd_2'(s(fun(X7,X7),X9),s('type.ring.ring'(X7),'const.ring.ring_R|49|_fupd_2'(s(fun(X7,X7),X10),s('type.ring.ring'(X7),X19)))))),
  file('thm.ring.ring_fupdcanon_comp|split|1','thm.ring.ring_5|49|_prove_ltb')).
fof(f83136,plain,(
  s('type.ring.ring'(sK7),'const.ring.ring_R|49|_fupd_2'(s(fun(sK7,sK7),sK10),s('type.ring.ring'(sK7),'const.ring.ring_R|48|_fupd_2'(s(fun(sK7,sK7),sK9),s('type.ring.ring'(sK7),happ(s(fun(sK6,'type.ring.ring'(sK7)),sK8),s(sK6,sK25('const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK7,sK7),sK10))),s(fun(sK6,'type.ring.ring'(sK7)),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK7,sK7),sK9))),s(fun(sK6,'type.ring.ring'(sK7)),sK8)))),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK7,sK7),sK9))),s(fun(sK6,'type.ring.ring'(sK7)),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK7,sK7),sK10))),s(fun(sK6,'type.ring.ring'(sK7)),sK8)))),sK6,'type.ring.ring'(sK7))))))))) != s('type.ring.ring'(sK7),'const.ring.ring_R|48|_fupd_2'(s(fun(sK7,sK7),sK9),s('type.ring.ring'(sK7),'const.ring.ring_R|49|_fupd_2'(s(fun(sK7,sK7),sK10),s('type.ring.ring'(sK7),happ(s(fun(sK6,'type.ring.ring'(sK7)),sK8),s(sK6,sK25('const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK7,sK7),sK10))),s(fun(sK6,'type.ring.ring'(sK7)),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK7,sK7),sK9))),s(fun(sK6,'type.ring.ring'(sK7)),sK8)))),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK7,sK7),sK9))),s(fun(sK6,'type.ring.ring'(sK7)),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK7,sK7),sK10))),s(fun(sK6,'type.ring.ring'(sK7)),sK8)))),sK6,'type.ring.ring'(sK7)))))))))),
  inference(forward_demodulation,[],[f83135,f17424])).
fof(f17424,plain,(
  ( ! [X70,X68,X66,X69,X67] : (s('type.ring.ring'(X66),'const.ring.ring_R|48|_fupd_2'(s(fun(X66,X66),X67),s('type.ring.ring'(X66),happ(s(fun(X68,'type.ring.ring'(X66)),X69),s(X68,X70))))) = s('type.ring.ring'(X66),'const.combin.o_3'(s(fun('type.ring.ring'(X66),'type.ring.ring'(X66)),'const.ring.ring_R|48|_fupd_1'(s(fun(X66,X66),X67))),s(fun(X68,'type.ring.ring'(X66)),X69),s(X68,X70)))) )),
  inference(superposition,[],[f1213,f988])).
fof(f988,plain,(
  ( ! [X2,X0,X1] : (s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),X1),s('type.ring.ring'(X0),X2))) = s('type.ring.ring'(X0),happ(s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_R|48|_fupd_1'(s(fun(X0,X0),X1))),s('type.ring.ring'(X0),X2)))) )),
  inference(cnf_transformation,[],[f378])).
fof(f378,plain,(
  ! [X0,X1,X2] : s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),X1),s('type.ring.ring'(X0),X2))) = s('type.ring.ring'(X0),happ(s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_R|48|_fupd_1'(s(fun(X0,X0),X1))),s('type.ring.ring'(X0),X2)))),
  inference(flattening,[],[f377])).
fof(f377,plain,(
  ! [X0] : ! [X1,X2] : s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),X1),s('type.ring.ring'(X0),X2))) = s('type.ring.ring'(X0),happ(s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_R|48|_fupd_1'(s(fun(X0,X0),X1))),s('type.ring.ring'(X0),X2)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X7] : ! [X12,X13] : s('type.ring.ring'(X7),'const.ring.ring_R|48|_fupd_2'(s(fun(X7,X7),X12),s('type.ring.ring'(X7),X13))) = s('type.ring.ring'(X7),happ(s(fun('type.ring.ring'(X7),'type.ring.ring'(X7)),'const.ring.ring_R|48|_fupd_1'(s(fun(X7,X7),X12))),s('type.ring.ring'(X7),X13)))),
  file('thm.ring.ring_fupdcanon_comp|split|1','ARITY_1_2_const.ring.ring_R|48|_fupd')).
fof(f1213,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5)))))) )),
  inference(cnf_transformation,[],[f681])).
fof(f681,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5)))))),
  inference(flattening,[],[f680])).
fof(f680,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5)))))),
  inference(rectify,[],[f63])).
fof(f63,axiom,(
  ! [X6,X11,X7] : ! [X10,X9,X5] : s(X6,'const.combin.o_3'(s(fun(X11,X6),X10),s(fun(X7,X11),X9),s(X7,X5))) = s(X6,happ(s(fun(X11,X6),X10),s(X11,happ(s(fun(X7,X11),X9),s(X7,X5)))))),
  file('thm.ring.ring_fupdcanon_comp|split|1','thm.combin.o_DEF')).
fof(f83135,plain,(
  s('type.ring.ring'(sK7),'const.ring.ring_R|48|_fupd_2'(s(fun(sK7,sK7),sK9),s('type.ring.ring'(sK7),'const.ring.ring_R|49|_fupd_2'(s(fun(sK7,sK7),sK10),s('type.ring.ring'(sK7),happ(s(fun(sK6,'type.ring.ring'(sK7)),sK8),s(sK6,sK25('const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK7,sK7),sK10))),s(fun(sK6,'type.ring.ring'(sK7)),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK7,sK7),sK9))),s(fun(sK6,'type.ring.ring'(sK7)),sK8)))),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK7,sK7),sK9))),s(fun(sK6,'type.ring.ring'(sK7)),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK7,sK7),sK10))),s(fun(sK6,'type.ring.ring'(sK7)),sK8)))),sK6,'type.ring.ring'(sK7))))))))) != s('type.ring.ring'(sK7),'const.ring.ring_R|49|_fupd_2'(s(fun(sK7,sK7),sK10),s('type.ring.ring'(sK7),'const.combin.o_3'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK7,sK7),sK9))),s(fun(sK6,'type.ring.ring'(sK7)),sK8),s(sK6,sK25('const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK7,sK7),sK10))),s(fun(sK6,'type.ring.ring'(sK7)),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK7,sK7),sK9))),s(fun(sK6,'type.ring.ring'(sK7)),sK8)))),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK7,sK7),sK9))),s(fun(sK6,'type.ring.ring'(sK7)),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK7,sK7),sK10))),s(fun(sK6,'type.ring.ring'(sK7)),sK8)))),sK6,'type.ring.ring'(sK7)))))))),
  inference(forward_demodulation,[],[f83134,f17423])).
fof(f17423,plain,(
  ( ! [X61,X64,X62,X65,X63] : (s('type.ring.ring'(X61),'const.ring.ring_R|49|_fupd_2'(s(fun(X61,X61),X62),s('type.ring.ring'(X61),happ(s(fun(X63,'type.ring.ring'(X61)),X64),s(X63,X65))))) = s('type.ring.ring'(X61),'const.combin.o_3'(s(fun('type.ring.ring'(X61),'type.ring.ring'(X61)),'const.ring.ring_R|49|_fupd_1'(s(fun(X61,X61),X62))),s(fun(X63,'type.ring.ring'(X61)),X64),s(X63,X65)))) )),
  inference(superposition,[],[f1213,f987])).
fof(f987,plain,(
  ( ! [X2,X0,X1] : (s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),X1),s('type.ring.ring'(X0),X2))) = s('type.ring.ring'(X0),happ(s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_R|49|_fupd_1'(s(fun(X0,X0),X1))),s('type.ring.ring'(X0),X2)))) )),
  inference(cnf_transformation,[],[f376])).
fof(f376,plain,(
  ! [X0,X1,X2] : s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),X1),s('type.ring.ring'(X0),X2))) = s('type.ring.ring'(X0),happ(s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_R|49|_fupd_1'(s(fun(X0,X0),X1))),s('type.ring.ring'(X0),X2)))),
  inference(flattening,[],[f375])).
fof(f375,plain,(
  ! [X0] : ! [X1,X2] : s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),X1),s('type.ring.ring'(X0),X2))) = s('type.ring.ring'(X0),happ(s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_R|49|_fupd_1'(s(fun(X0,X0),X1))),s('type.ring.ring'(X0),X2)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X7] : ! [X12,X13] : s('type.ring.ring'(X7),'const.ring.ring_R|49|_fupd_2'(s(fun(X7,X7),X12),s('type.ring.ring'(X7),X13))) = s('type.ring.ring'(X7),happ(s(fun('type.ring.ring'(X7),'type.ring.ring'(X7)),'const.ring.ring_R|49|_fupd_1'(s(fun(X7,X7),X12))),s('type.ring.ring'(X7),X13)))),
  file('thm.ring.ring_fupdcanon_comp|split|1','ARITY_1_2_const.ring.ring_R|49|_fupd')).
fof(f83134,plain,(
  s('type.ring.ring'(sK7),'const.ring.ring_R|49|_fupd_2'(s(fun(sK7,sK7),sK10),s('type.ring.ring'(sK7),'const.combin.o_3'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK7,sK7),sK9))),s(fun(sK6,'type.ring.ring'(sK7)),sK8),s(sK6,sK25('const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK7,sK7),sK10))),s(fun(sK6,'type.ring.ring'(sK7)),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK7,sK7),sK9))),s(fun(sK6,'type.ring.ring'(sK7)),sK8)))),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK7,sK7),sK9))),s(fun(sK6,'type.ring.ring'(sK7)),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK7,sK7),sK10))),s(fun(sK6,'type.ring.ring'(sK7)),sK8)))),sK6,'type.ring.ring'(sK7))))))) != s('type.ring.ring'(sK7),'const.ring.ring_R|48|_fupd_2'(s(fun(sK7,sK7),sK9),s('type.ring.ring'(sK7),'const.combin.o_3'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK7,sK7),sK10))),s(fun(sK6,'type.ring.ring'(sK7)),sK8),s(sK6,sK25('const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK7,sK7),sK10))),s(fun(sK6,'type.ring.ring'(sK7)),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK7,sK7),sK9))),s(fun(sK6,'type.ring.ring'(sK7)),sK8)))),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK7,sK7),sK9))),s(fun(sK6,'type.ring.ring'(sK7)),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK7,sK7),sK10))),s(fun(sK6,'type.ring.ring'(sK7)),sK8)))),sK6,'type.ring.ring'(sK7)))))))),
  inference(forward_demodulation,[],[f83133,f1217])).
fof(f1217,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))) )),
  inference(cnf_transformation,[],[f689])).
fof(f689,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(flattening,[],[f688])).
fof(f688,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6,X11,X7] : ! [X12,X13,X14] : s(X6,'const.combin.o_3'(s(fun(X11,X6),X12),s(fun(X7,X11),X13),s(X7,X14))) = s(X6,happ(s(fun(X7,X6),'const.combin.o_2'(s(fun(X11,X6),X12),s(fun(X7,X11),X13))),s(X7,X14)))),
  file('thm.ring.ring_fupdcanon_comp|split|1','ARITY_2_3_const.combin.o')).
fof(f83133,plain,(
  s('type.ring.ring'(sK7),'const.ring.ring_R|48|_fupd_2'(s(fun(sK7,sK7),sK9),s('type.ring.ring'(sK7),'const.combin.o_3'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK7,sK7),sK10))),s(fun(sK6,'type.ring.ring'(sK7)),sK8),s(sK6,sK25('const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK7,sK7),sK10))),s(fun(sK6,'type.ring.ring'(sK7)),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK7,sK7),sK9))),s(fun(sK6,'type.ring.ring'(sK7)),sK8)))),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK7,sK7),sK9))),s(fun(sK6,'type.ring.ring'(sK7)),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK7,sK7),sK10))),s(fun(sK6,'type.ring.ring'(sK7)),sK8)))),sK6,'type.ring.ring'(sK7))))))) != s('type.ring.ring'(sK7),'const.ring.ring_R|49|_fupd_2'(s(fun(sK7,sK7),sK10),s('type.ring.ring'(sK7),happ(s(fun(sK6,'type.ring.ring'(sK7)),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK7,sK7),sK9))),s(fun(sK6,'type.ring.ring'(sK7)),sK8))),s(sK6,sK25('const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK7,sK7),sK10))),s(fun(sK6,'type.ring.ring'(sK7)),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK7,sK7),sK9))),s(fun(sK6,'type.ring.ring'(sK7)),sK8)))),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK7,sK7),sK9))),s(fun(sK6,'type.ring.ring'(sK7)),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK7,sK7),sK10))),s(fun(sK6,'type.ring.ring'(sK7)),sK8)))),sK6,'type.ring.ring'(sK7)))))))),
  inference(forward_demodulation,[],[f83132,f1217])).
fof(f83132,plain,(
  s('type.ring.ring'(sK7),'const.ring.ring_R|49|_fupd_2'(s(fun(sK7,sK7),sK10),s('type.ring.ring'(sK7),happ(s(fun(sK6,'type.ring.ring'(sK7)),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK7,sK7),sK9))),s(fun(sK6,'type.ring.ring'(sK7)),sK8))),s(sK6,sK25('const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK7,sK7),sK10))),s(fun(sK6,'type.ring.ring'(sK7)),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK7,sK7),sK9))),s(fun(sK6,'type.ring.ring'(sK7)),sK8)))),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK7,sK7),sK9))),s(fun(sK6,'type.ring.ring'(sK7)),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK7,sK7),sK10))),s(fun(sK6,'type.ring.ring'(sK7)),sK8)))),sK6,'type.ring.ring'(sK7))))))) != s('type.ring.ring'(sK7),'const.ring.ring_R|48|_fupd_2'(s(fun(sK7,sK7),sK9),s('type.ring.ring'(sK7),happ(s(fun(sK6,'type.ring.ring'(sK7)),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK7,sK7),sK10))),s(fun(sK6,'type.ring.ring'(sK7)),sK8))),s(sK6,sK25('const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK7,sK7),sK10))),s(fun(sK6,'type.ring.ring'(sK7)),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK7,sK7),sK9))),s(fun(sK6,'type.ring.ring'(sK7)),sK8)))),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK7,sK7),sK9))),s(fun(sK6,'type.ring.ring'(sK7)),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK7,sK7),sK10))),s(fun(sK6,'type.ring.ring'(sK7)),sK8)))),sK6,'type.ring.ring'(sK7)))))))),
  inference(forward_demodulation,[],[f83131,f17423])).
fof(f83131,plain,(
  s('type.ring.ring'(sK7),'const.ring.ring_R|48|_fupd_2'(s(fun(sK7,sK7),sK9),s('type.ring.ring'(sK7),happ(s(fun(sK6,'type.ring.ring'(sK7)),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK7,sK7),sK10))),s(fun(sK6,'type.ring.ring'(sK7)),sK8))),s(sK6,sK25('const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK7,sK7),sK10))),s(fun(sK6,'type.ring.ring'(sK7)),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK7,sK7),sK9))),s(fun(sK6,'type.ring.ring'(sK7)),sK8)))),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK7,sK7),sK9))),s(fun(sK6,'type.ring.ring'(sK7)),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK7,sK7),sK10))),s(fun(sK6,'type.ring.ring'(sK7)),sK8)))),sK6,'type.ring.ring'(sK7))))))) != s('type.ring.ring'(sK7),'const.combin.o_3'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK7,sK7),sK10))),s(fun(sK6,'type.ring.ring'(sK7)),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK7,sK7),sK9))),s(fun(sK6,'type.ring.ring'(sK7)),sK8))),s(sK6,sK25('const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK7,sK7),sK10))),s(fun(sK6,'type.ring.ring'(sK7)),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK7,sK7),sK9))),s(fun(sK6,'type.ring.ring'(sK7)),sK8)))),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK7,sK7),sK9))),s(fun(sK6,'type.ring.ring'(sK7)),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK7,sK7),sK10))),s(fun(sK6,'type.ring.ring'(sK7)),sK8)))),sK6,'type.ring.ring'(sK7)))))),
  inference(forward_demodulation,[],[f83130,f17424])).
fof(f83130,plain,(
  s('type.ring.ring'(sK7),'const.combin.o_3'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK7,sK7),sK10))),s(fun(sK6,'type.ring.ring'(sK7)),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK7,sK7),sK9))),s(fun(sK6,'type.ring.ring'(sK7)),sK8))),s(sK6,sK25('const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK7,sK7),sK10))),s(fun(sK6,'type.ring.ring'(sK7)),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK7,sK7),sK9))),s(fun(sK6,'type.ring.ring'(sK7)),sK8)))),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK7,sK7),sK9))),s(fun(sK6,'type.ring.ring'(sK7)),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK7,sK7),sK10))),s(fun(sK6,'type.ring.ring'(sK7)),sK8)))),sK6,'type.ring.ring'(sK7))))) != s('type.ring.ring'(sK7),'const.combin.o_3'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK7,sK7),sK9))),s(fun(sK6,'type.ring.ring'(sK7)),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK7,sK7),sK10))),s(fun(sK6,'type.ring.ring'(sK7)),sK8))),s(sK6,sK25('const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK7,sK7),sK10))),s(fun(sK6,'type.ring.ring'(sK7)),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK7,sK7),sK9))),s(fun(sK6,'type.ring.ring'(sK7)),sK8)))),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK7,sK7),sK9))),s(fun(sK6,'type.ring.ring'(sK7)),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK7,sK7),sK10))),s(fun(sK6,'type.ring.ring'(sK7)),sK8)))),sK6,'type.ring.ring'(sK7)))))),
  inference(forward_demodulation,[],[f83129,f1217])).
fof(f83129,plain,(
  s('type.ring.ring'(sK7),'const.combin.o_3'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK7,sK7),sK9))),s(fun(sK6,'type.ring.ring'(sK7)),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK7,sK7),sK10))),s(fun(sK6,'type.ring.ring'(sK7)),sK8))),s(sK6,sK25('const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK7,sK7),sK10))),s(fun(sK6,'type.ring.ring'(sK7)),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK7,sK7),sK9))),s(fun(sK6,'type.ring.ring'(sK7)),sK8)))),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK7,sK7),sK9))),s(fun(sK6,'type.ring.ring'(sK7)),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK7,sK7),sK10))),s(fun(sK6,'type.ring.ring'(sK7)),sK8)))),sK6,'type.ring.ring'(sK7))))) != s('type.ring.ring'(sK7),happ(s(fun(sK6,'type.ring.ring'(sK7)),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK7,sK7),sK10))),s(fun(sK6,'type.ring.ring'(sK7)),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK7,sK7),sK9))),s(fun(sK6,'type.ring.ring'(sK7)),sK8))))),s(sK6,sK25('const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK7,sK7),sK10))),s(fun(sK6,'type.ring.ring'(sK7)),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK7,sK7),sK9))),s(fun(sK6,'type.ring.ring'(sK7)),sK8)))),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK7,sK7),sK9))),s(fun(sK6,'type.ring.ring'(sK7)),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK7,sK7),sK10))),s(fun(sK6,'type.ring.ring'(sK7)),sK8)))),sK6,'type.ring.ring'(sK7)))))),
  inference(forward_demodulation,[],[f83124,f1217])).
fof(f83124,plain,(
  s('type.ring.ring'(sK7),happ(s(fun(sK6,'type.ring.ring'(sK7)),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK7,sK7),sK10))),s(fun(sK6,'type.ring.ring'(sK7)),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK7,sK7),sK9))),s(fun(sK6,'type.ring.ring'(sK7)),sK8))))),s(sK6,sK25('const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK7,sK7),sK10))),s(fun(sK6,'type.ring.ring'(sK7)),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK7,sK7),sK9))),s(fun(sK6,'type.ring.ring'(sK7)),sK8)))),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK7,sK7),sK9))),s(fun(sK6,'type.ring.ring'(sK7)),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK7,sK7),sK10))),s(fun(sK6,'type.ring.ring'(sK7)),sK8)))),sK6,'type.ring.ring'(sK7))))) != s('type.ring.ring'(sK7),happ(s(fun(sK6,'type.ring.ring'(sK7)),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK7,sK7),sK9))),s(fun(sK6,'type.ring.ring'(sK7)),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK7,sK7),sK10))),s(fun(sK6,'type.ring.ring'(sK7)),sK8))))),s(sK6,sK25('const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK7,sK7),sK10))),s(fun(sK6,'type.ring.ring'(sK7)),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK7,sK7),sK9))),s(fun(sK6,'type.ring.ring'(sK7)),sK8)))),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK7,sK7),sK9))),s(fun(sK6,'type.ring.ring'(sK7)),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK7,sK7),sK10))),s(fun(sK6,'type.ring.ring'(sK7)),sK8)))),sK6,'type.ring.ring'(sK7)))))),
  inference(unit_resulting_resolution,[],[f881,f1144])).
fof(f1144,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,sK25(X3,X2,X1,X0)))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,sK25(X3,X2,X1,X0)))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)) )),
  inference(cnf_transformation,[],[f857])).
fof(f857,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X5] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))) & (s(X0,happ(s(fun(X1,X0),X2),s(X1,sK25(X3,X2,X1,X0)))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,sK25(X3,X2,X1,X0)))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK25])],[f856])).
fof(f856,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X5] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))) & (? [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(rectify,[],[f855])).
fof(f855,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4)))) & (? [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(nnf_transformation,[],[f611])).
fof(f611,plain,(
  ! [X0,X1,X2,X3] : (s(fun(X1,X0),X2) = s(fun(X1,X0),X3) <=> ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))))),
  inference(flattening,[],[f610])).
fof(f610,plain,(
  ! [X0,X1] : ! [X2,X3] : (s(fun(X1,X0),X2) = s(fun(X1,X0),X3) <=> ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))))),
  inference(rectify,[],[f69])).
fof(f69,axiom,(
  ! [X6,X7] : ! [X10,X9] : (s(fun(X7,X6),X9) = s(fun(X7,X6),X10) <=> ! [X15] : s(X6,happ(s(fun(X7,X6),X9),s(X7,X15))) = s(X6,happ(s(fun(X7,X6),X10),s(X7,X15))))),
  file('thm.ring.ring_fupdcanon_comp|split|1','thm.bool.FUN_EQ_THM')).
fof(f881,plain,(
  s(fun(sK6,'type.ring.ring'(sK7)),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK7,sK7),sK10))),s(fun(sK6,'type.ring.ring'(sK7)),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK7,sK7),sK9))),s(fun(sK6,'type.ring.ring'(sK7)),sK8))))) != s(fun(sK6,'type.ring.ring'(sK7)),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK7,sK7),sK9))),s(fun(sK6,'type.ring.ring'(sK7)),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK7,sK7),sK10))),s(fun(sK6,'type.ring.ring'(sK7)),sK8)))))),
  inference(cnf_transformation,[],[f789])).
fof(f789,plain,(
  s(fun(sK6,'type.ring.ring'(sK7)),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK7,sK7),sK10))),s(fun(sK6,'type.ring.ring'(sK7)),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK7,sK7),sK9))),s(fun(sK6,'type.ring.ring'(sK7)),sK8))))) != s(fun(sK6,'type.ring.ring'(sK7)),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK7,sK7),sK9))),s(fun(sK6,'type.ring.ring'(sK7)),'const.combin.o_2'(s(fun('type.ring.ring'(sK7),'type.ring.ring'(sK7)),'const.ring.ring_R|49|_fupd_1'(s(fun(sK7,sK7),sK10))),s(fun(sK6,'type.ring.ring'(sK7)),sK8)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK6,sK7,sK8,sK9,sK10])],[f746])).
fof(f746,plain,(
  ? [X0,X1,X2,X3,X4] : s(fun(X0,'type.ring.ring'(X1)),'const.combin.o_2'(s(fun('type.ring.ring'(X1),'type.ring.ring'(X1)),'const.ring.ring_R|49|_fupd_1'(s(fun(X1,X1),X4))),s(fun(X0,'type.ring.ring'(X1)),'const.combin.o_2'(s(fun('type.ring.ring'(X1),'type.ring.ring'(X1)),'const.ring.ring_R|48|_fupd_1'(s(fun(X1,X1),X3))),s(fun(X0,'type.ring.ring'(X1)),X2))))) != s(fun(X0,'type.ring.ring'(X1)),'const.combin.o_2'(s(fun('type.ring.ring'(X1),'type.ring.ring'(X1)),'const.ring.ring_R|48|_fupd_1'(s(fun(X1,X1),X3))),s(fun(X0,'type.ring.ring'(X1)),'const.combin.o_2'(s(fun('type.ring.ring'(X1),'type.ring.ring'(X1)),'const.ring.ring_R|49|_fupd_1'(s(fun(X1,X1),X4))),s(fun(X0,'type.ring.ring'(X1)),X2)))))),
  inference(ennf_transformation,[],[f263])).
fof(f263,plain,(
  ~! [X0,X1,X2,X3,X4] : s(fun(X0,'type.ring.ring'(X1)),'const.combin.o_2'(s(fun('type.ring.ring'(X1),'type.ring.ring'(X1)),'const.ring.ring_R|49|_fupd_1'(s(fun(X1,X1),X4))),s(fun(X0,'type.ring.ring'(X1)),'const.combin.o_2'(s(fun('type.ring.ring'(X1),'type.ring.ring'(X1)),'const.ring.ring_R|48|_fupd_1'(s(fun(X1,X1),X3))),s(fun(X0,'type.ring.ring'(X1)),X2))))) = s(fun(X0,'type.ring.ring'(X1)),'const.combin.o_2'(s(fun('type.ring.ring'(X1),'type.ring.ring'(X1)),'const.ring.ring_R|48|_fupd_1'(s(fun(X1,X1),X3))),s(fun(X0,'type.ring.ring'(X1)),'const.combin.o_2'(s(fun('type.ring.ring'(X1),'type.ring.ring'(X1)),'const.ring.ring_R|49|_fupd_1'(s(fun(X1,X1),X4))),s(fun(X0,'type.ring.ring'(X1)),X2)))))),
  inference(flattening,[],[f262])).
fof(f262,plain,(
  ~! [X0,X1] : ! [X2,X3,X4] : s(fun(X0,'type.ring.ring'(X1)),'const.combin.o_2'(s(fun('type.ring.ring'(X1),'type.ring.ring'(X1)),'const.ring.ring_R|49|_fupd_1'(s(fun(X1,X1),X4))),s(fun(X0,'type.ring.ring'(X1)),'const.combin.o_2'(s(fun('type.ring.ring'(X1),'type.ring.ring'(X1)),'const.ring.ring_R|48|_fupd_1'(s(fun(X1,X1),X3))),s(fun(X0,'type.ring.ring'(X1)),X2))))) = s(fun(X0,'type.ring.ring'(X1)),'const.combin.o_2'(s(fun('type.ring.ring'(X1),'type.ring.ring'(X1)),'const.ring.ring_R|48|_fupd_1'(s(fun(X1,X1),X3))),s(fun(X0,'type.ring.ring'(X1)),'const.combin.o_2'(s(fun('type.ring.ring'(X1),'type.ring.ring'(X1)),'const.ring.ring_R|49|_fupd_1'(s(fun(X1,X1),X4))),s(fun(X0,'type.ring.ring'(X1)),X2)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10] : s(fun(X6,'type.ring.ring'(X7)),'const.combin.o_2'(s(fun('type.ring.ring'(X7),'type.ring.ring'(X7)),'const.ring.ring_R|49|_fupd_1'(s(fun(X7,X7),X10))),s(fun(X6,'type.ring.ring'(X7)),'const.combin.o_2'(s(fun('type.ring.ring'(X7),'type.ring.ring'(X7)),'const.ring.ring_R|48|_fupd_1'(s(fun(X7,X7),X9))),s(fun(X6,'type.ring.ring'(X7)),X8))))) = s(fun(X6,'type.ring.ring'(X7)),'const.combin.o_2'(s(fun('type.ring.ring'(X7),'type.ring.ring'(X7)),'const.ring.ring_R|48|_fupd_1'(s(fun(X7,X7),X9))),s(fun(X6,'type.ring.ring'(X7)),'const.combin.o_2'(s(fun('type.ring.ring'(X7),'type.ring.ring'(X7)),'const.ring.ring_R|49|_fupd_1'(s(fun(X7,X7),X10))),s(fun(X6,'type.ring.ring'(X7)),X8)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10] : s(fun(X6,'type.ring.ring'(X7)),'const.combin.o_2'(s(fun('type.ring.ring'(X7),'type.ring.ring'(X7)),'const.ring.ring_R|49|_fupd_1'(s(fun(X7,X7),X10))),s(fun(X6,'type.ring.ring'(X7)),'const.combin.o_2'(s(fun('type.ring.ring'(X7),'type.ring.ring'(X7)),'const.ring.ring_R|48|_fupd_1'(s(fun(X7,X7),X9))),s(fun(X6,'type.ring.ring'(X7)),X8))))) = s(fun(X6,'type.ring.ring'(X7)),'const.combin.o_2'(s(fun('type.ring.ring'(X7),'type.ring.ring'(X7)),'const.ring.ring_R|48|_fupd_1'(s(fun(X7,X7),X9))),s(fun(X6,'type.ring.ring'(X7)),'const.combin.o_2'(s(fun('type.ring.ring'(X7),'type.ring.ring'(X7)),'const.ring.ring_R|49|_fupd_1'(s(fun(X7,X7),X10))),s(fun(X6,'type.ring.ring'(X7)),X8)))))),
  file('thm.ring.ring_fupdcanon_comp|split|1','thm.ring.ring_fupdcanon_comp|split|1')).
% SZS output end Proof for thm.ring
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 431
% Passive clauses: 4684
% Generated clauses: 82238
% Final active clauses: 424
% Final passive clauses: 4059
% Input formulas: 260
% Initial clauses: 380
% 
% Duplicate literals: 3
% Trivial inequalities: 11
% Fw subsumption resolutions: 100
% Fw demodulations: 32281
% Bw demodulations: 137
% Condensations: 8
% 
% Simple tautologies: 92
% Equational tautologies: 40
% Forward subsumptions: 16989
% Fw demodulations to eq. taut.: 27852
% 
% Unit resulting resolution: 31396
% Factoring: 1
% Forward superposition: 8783
% Backward superposition: 8983
% Self superposition: 133
% Equality factoring: 2
% Equality resolution: 7
% 
% Split clauses: 12
% Split components: 24
% SAT solver clauses: 4870
% SAT solver unit clauses: 2280
% SAT solver binary clauses: 2238
% 
% Sat splits: 12
% Sat splitting refutations: 160
% 
% Memory used [KB]: 72408
% Time elapsed: 2.666 s
% ------------------------------
% ------------------------------
% Success in time 4 s
