Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.divides
% (49650)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 848
% Passive clauses: 6874
% Generated clauses: 21825
% Final active clauses: 787
% Final passive clauses: 5586
% Input formulas: 45
% Initial clauses: 170
% 
% Duplicate literals: 645
% Trivial inequalities: 30
% Fw subsumption resolutions: 2544
% Fw demodulations: 1085
% Bw demodulations: 59
% 
% Simple tautologies: 245
% Equational tautologies: 60
% Forward subsumptions: 9477
% 
% Binary resolution: 8880
% Factoring: 1464
% Forward superposition: 3792
% Backward superposition: 2760
% Self superposition: 38
% Equality factoring: 34
% Equality resolution: 10
% 
% Split clauses: 644
% Split components: 1298
% SAT solver clauses: 1246
% SAT solver unit clauses: 29
% SAT solver binary clauses: 186
% 
% Sat splits: 644
% 
% TWLsolver clauses: 1245
% TWLsolver calls for satisfiability: 580
% 
% Memory used [KB]: 8827
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.divides
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.divides
% SZS output start Proof for thm.divides
fof(f8456,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f8168,f248,f436,f2135,f227,f421,f3812,f3816,f3814,f3065,f3066,f524,f529,f224,f419,f245,f433,f933,f955,f249,f435,f293,f447,f3779,f3780,f1134,f1151,f405,f516,f7497,f5323,f3332,f3333,f422,f424,f543,f547,f286,f445,f972,f989,f290,f448,f6960,f6962,f544,f549,f226,f415,f6781,f6782,f313,f458,f5869,f5872,f237,f431,f5702,f5706,f289,f446,f5659,f5660,f599,f618,f312,f459,f5533,f2111,f5534,f1663,f1682,f322,f462,f3069,f3070,f523,f527,f225,f417,f7971])).
fof(f7971,plain,(
  $false | ($spl6 | $spl9 | $spl1472)),
  inference(subsumption_resolution,[],[f7638,f421])).
fof(f7638,plain,(
  ~p(s(bool,'T_0')) | ($spl9 | $spl1472)),
  inference(backward_demodulation,[],[f5323,f424])).
fof(f417,plain,(
  p(s(bool,'const.divides.divides_2'(s('type.num.num',sK10),s('type.num.num',sK11)))) | $spl2),
  inference(cnf_transformation,[],[f417_D])).
fof(f417_D,plain,(
  p(s(bool,'const.divides.divides_2'(s('type.num.num',sK10),s('type.num.num',sK11)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f225,plain,(
  p(s(bool,'const.divides.divides_2'(s('type.num.num',sK10),s('type.num.num',sK11))))),
  inference(cnf_transformation,[],[f150])).
fof(f150,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) & p(s(bool,'const.divides.divides_2'(s('type.num.num',sK10),s('type.num.num',sK11)))) & ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',sK11),s('type.num.num',sK12))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK10,sK11,sK12])],[f107])).
fof(f107,plain,(
  ? [X0,X1,X2] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X2)))) & p(s(bool,'const.divides.divides_2'(s('type.num.num',X0),s('type.num.num',X1)))) & ~p(s(bool,'const.divides.divides_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X1),s('type.num.num',X2)))))))),
  inference(flattening,[],[f106])).
fof(f106,plain,(
  ? [X0,X1,X2] : ((p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X2)))) & p(s(bool,'const.divides.divides_2'(s('type.num.num',X0),s('type.num.num',X1))))) & ~p(s(bool,'const.divides.divides_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X1),s('type.num.num',X2)))))))),
  inference(ennf_transformation,[],[f47])).
fof(f47,plain,(
  ~! [X0,X1,X2] : ((p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X2)))) & p(s(bool,'const.divides.divides_2'(s('type.num.num',X0),s('type.num.num',X1))))) => p(s(bool,'const.divides.divides_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X1),s('type.num.num',X2)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : ((p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X8)))) & p(s(bool,'const.divides.divides_2'(s('type.num.num',X6),s('type.num.num',X7))))) => p(s(bool,'const.divides.divides_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X7),s('type.num.num',X8)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : ((p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X8)))) & p(s(bool,'const.divides.divides_2'(s('type.num.num',X6),s('type.num.num',X7))))) => p(s(bool,'const.divides.divides_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X7),s('type.num.num',X8)))))))),
  file('thm.divides.divides_|49|6_prove_ltb','thm.divides.divides_|49|6_prove_ltb')).
fof(f527,plain,(
  s(bool,'T_0') = s(bool,'const.divides.divides_2'(s('type.num.num',sK10),s('type.num.num',sK11))) | $spl166),
  inference(cnf_transformation,[],[f527_D])).
fof(f527_D,plain,(
  s(bool,'T_0') = s(bool,'const.divides.divides_2'(s('type.num.num',sK10),s('type.num.num',sK11))) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f523,plain,(
  s(bool,'T_0') = s(bool,'const.divides.divides_2'(s('type.num.num',sK10),s('type.num.num',sK11))) | ($spl2 | $spl18)),
  inference(resolution,[],[f433,f417])).
fof(f3070,plain,(
  ( ! [X9] : (s(bool,X9) = s(bool,'const.divides.divides_2'(s('type.num.num',sK10),s('type.num.num',sK11))) | s(bool,'F_0') = s(bool,X9)) ) | $spl1270),
  inference(cnf_transformation,[],[f3070_D])).
fof(f3070_D,plain,(
  ( ! [X9] : (s(bool,X9) = s(bool,'const.divides.divides_2'(s('type.num.num',sK10),s('type.num.num',sK11))) | s(bool,'F_0') = s(bool,X9)) ) <=> ~$spl1270),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1270])])).
fof(f3069,plain,(
  ( ! [X9] : (s(bool,X9) = s(bool,'const.divides.divides_2'(s('type.num.num',sK10),s('type.num.num',sK11))) | s(bool,'F_0') = s(bool,X9)) ) | ($spl6 | $spl166 | $spl330)),
  inference(subsumption_resolution,[],[f3038,f421])).
fof(f3038,plain,(
  ( ! [X9] : (~p(s(bool,'T_0')) | s(bool,X9) = s(bool,'const.divides.divides_2'(s('type.num.num',sK10),s('type.num.num',sK11))) | s(bool,'F_0') = s(bool,X9)) ) | ($spl166 | $spl330)),
  inference(superposition,[],[f955,f527])).
fof(f462,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.divides.divides_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',X2)))))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',X0),s('type.num.num',X1))))) ) | $spl50),
  inference(cnf_transformation,[],[f462_D])).
fof(f462_D,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.divides.divides_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',X2)))))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',X0),s('type.num.num',X1))))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f322,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.divides.divides_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',X2)))))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f129])).
fof(f129,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.divides.divides_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.divides.divides_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',X2)))))))),
  inference(ennf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.divides.divides_2'(s('type.num.num',X0),s('type.num.num',X1)))) => p(s(bool,'const.divides.divides_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',X2)))))))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X6,X7,X19] : (p(s(bool,'const.divides.divides_2'(s('type.num.num',X6),s('type.num.num',X7)))) => p(s(bool,'const.divides.divides_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X7),s('type.num.num',X19)))))))),
  file('thm.divides.divides_|49|6_prove_ltb','thm.divides.DIVIDES_MULT')).
fof(f1682,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.divides.divides_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s(bool,'T_0') = s(bool,'const.divides.divides_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',X2)))))) ) | $spl576),
  inference(cnf_transformation,[],[f1682_D])).
fof(f1682_D,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.divides.divides_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s(bool,'T_0') = s(bool,'const.divides.divides_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',X2)))))) ) <=> ~$spl576),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl576])])).
fof(f1663,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.divides.divides_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s(bool,'T_0') = s(bool,'const.divides.divides_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',X2)))))) ) | ($spl18 | $spl50)),
  inference(resolution,[],[f462,f433])).
fof(f5534,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,'const.divides.divides_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK11),s('type.num.num',X1)))))) ) | $spl1544),
  inference(cnf_transformation,[],[f5534_D])).
fof(f5534_D,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,'const.divides.divides_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK11),s('type.num.num',X1)))))) ) <=> ~$spl1544),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1544])])).
fof(f2111,plain,(
  ( ! [X2] : (~p(s(bool,X2))) ) | $spl738),
  inference(cnf_transformation,[],[f2111_D])).
fof(f2111_D,plain,(
  ( ! [X2] : (~p(s(bool,X2))) ) <=> ~$spl738),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl738])])).
fof(f5533,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,'const.divides.divides_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK11),s('type.num.num',X1)))))) ) | ($spl24 | $spl576 | $spl1270)),
  inference(subsumption_resolution,[],[f5487,f436])).
fof(f5487,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,'const.divides.divides_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK11),s('type.num.num',X1))))) | s(bool,'F_0') = s(bool,X0)) ) | ($spl576 | $spl1270)),
  inference(superposition,[],[f1682,f3070])).
fof(f459,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | ~sP1(X0,X1)) ) | $spl44),
  inference(cnf_transformation,[],[f459_D])).
fof(f459_D,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | ~sP1(X0,X1)) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f312,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | ~sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f189])).
fof(f189,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1)))),
  inference(rectify,[],[f188])).
fof(f188,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0)))),
  inference(flattening,[],[f187])).
fof(f187,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP1(X1,X0)))),
  inference(nnf_transformation,[],[f137])).
fof(f137,plain,(
  ! [X1,X0] : (sP1(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f618,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~sP1('T_0',X0)) ) | $spl192),
  inference(cnf_transformation,[],[f618_D])).
fof(f618_D,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~sP1('T_0',X0)) ) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f599,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~sP1('T_0',X0)) ) | ($spl6 | $spl44)),
  inference(resolution,[],[f459,f421])).
fof(f5660,plain,(
  ( ! [X29] : (~sP1('T_0','const.divides.divides_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK11),s('type.num.num',X29)))))) ) | $spl1556),
  inference(cnf_transformation,[],[f5660_D])).
fof(f5660_D,plain,(
  ( ! [X29] : (~sP1('T_0','const.divides.divides_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK11),s('type.num.num',X29)))))) ) <=> ~$spl1556),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1556])])).
fof(f5659,plain,(
  ( ! [X29] : (~sP1('T_0','const.divides.divides_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK11),s('type.num.num',X29)))))) ) | ($spl6 | $spl192 | $spl1544)),
  inference(subsumption_resolution,[],[f5571,f421])).
fof(f5571,plain,(
  ( ! [X29] : (~p(s(bool,'T_0')) | ~sP1('T_0','const.divides.divides_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK11),s('type.num.num',X29)))))) ) | ($spl192 | $spl1544)),
  inference(superposition,[],[f618,f5534])).
fof(f446,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num',X1)))))) ) | $spl28),
  inference(cnf_transformation,[],[f446_D])).
fof(f446_D,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num',X1)))))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f289,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num',X1)))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X9,X10] : s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X9),s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X9),s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X9),s('type.num.num',X10)))))),
  file('thm.divides.divides_|49|6_prove_ltb','thm.arithmetic.EXP|split|1')).
fof(f5706,plain,(
  ( ! [X0] : (~sP1('T_0','const.divides.divides_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',sK11),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))))) ) | $spl1582),
  inference(cnf_transformation,[],[f5706_D])).
fof(f5706_D,plain,(
  ( ! [X0] : (~sP1('T_0','const.divides.divides_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',sK11),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))))) ) <=> ~$spl1582),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1582])])).
fof(f5702,plain,(
  ( ! [X0] : (~sP1('T_0','const.divides.divides_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',sK11),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))))) ) | ($spl28 | $spl1556)),
  inference(superposition,[],[f5660,f446])).
fof(f431,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK13(X0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) ) | $spl16),
  inference(cnf_transformation,[],[f431_D])).
fof(f431_D,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK13(X0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f237,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK13(X0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f153])).
fof(f153,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK13(X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13])],[f53])).
fof(f53,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | ? [X1] : s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X9] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X9) | ? [X10] : s('type.num.num',X9) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))))),
  file('thm.divides.divides_|49|6_prove_ltb','thm.arithmetic.num_CASES')).
fof(f5872,plain,(
  ( ! [X0] : (~sP1('T_0','const.divides.divides_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',sK11),s('type.num.num',X0))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) ) | $spl1630),
  inference(cnf_transformation,[],[f5872_D])).
fof(f5872_D,plain,(
  ( ! [X0] : (~sP1('T_0','const.divides.divides_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',sK11),s('type.num.num',X0))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) ) <=> ~$spl1630),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1630])])).
fof(f5869,plain,(
  ( ! [X0] : (~sP1('T_0','const.divides.divides_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',sK11),s('type.num.num',X0))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) ) | ($spl16 | $spl1582)),
  inference(superposition,[],[f5706,f431])).
fof(f458,plain,(
  ( ! [X0,X1] : (sP1(X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0))) ) | $spl42),
  inference(cnf_transformation,[],[f458_D])).
fof(f458_D,plain,(
  ( ! [X0,X1] : (sP1(X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f313,plain,(
  ( ! [X0,X1] : (sP1(X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f189])).
fof(f6782,plain,(
  ( ! [X0] : (p(s(bool,'const.divides.divides_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',sK11),s('type.num.num',X0)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) ) | $spl1956),
  inference(cnf_transformation,[],[f6782_D])).
fof(f6782_D,plain,(
  ( ! [X0] : (p(s(bool,'const.divides.divides_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',sK11),s('type.num.num',X0)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) ) <=> ~$spl1956),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1956])])).
fof(f6781,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | p(s(bool,'const.divides.divides_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',sK11),s('type.num.num',X0))))))) ) | ($spl6 | $spl42 | $spl1630)),
  inference(subsumption_resolution,[],[f6773,f421])).
fof(f6773,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | p(s(bool,'const.divides.divides_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',sK11),s('type.num.num',X0)))))) | ~p(s(bool,'T_0'))) ) | ($spl42 | $spl1630)),
  inference(resolution,[],[f5872,f458])).
fof(f415,plain,(
  ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',sK11),s('type.num.num',sK12)))))) | $spl1),
  inference(cnf_transformation,[],[f415_D])).
fof(f415_D,plain,(
  ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',sK11),s('type.num.num',sK12)))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f226,plain,(
  ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',sK11),s('type.num.num',sK12))))))),
  inference(cnf_transformation,[],[f150])).
fof(f549,plain,(
  s(bool,'F_0') = s(bool,'const.divides.divides_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',sK11),s('type.num.num',sK12))))) | $spl172),
  inference(cnf_transformation,[],[f549_D])).
fof(f549_D,plain,(
  s(bool,'F_0') = s(bool,'const.divides.divides_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',sK11),s('type.num.num',sK12))))) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f544,plain,(
  s(bool,'F_0') = s(bool,'const.divides.divides_2'(s('type.num.num',sK10),s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',sK11),s('type.num.num',sK12))))) | ($spl1 | $spl22)),
  inference(resolution,[],[f435,f415])).
fof(f6962,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK12) | $spl2046),
  inference(cnf_transformation,[],[f6962_D])).
fof(f6962_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK12) <=> ~$spl2046),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2046])])).
fof(f6960,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK12) | ($spl9 | $spl172 | $spl1956)),
  inference(subsumption_resolution,[],[f6935,f424])).
fof(f6935,plain,(
  p(s(bool,'F_0')) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK12) | ($spl172 | $spl1956)),
  inference(superposition,[],[f6782,f549])).
fof(f448,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) = s(bool,X1)) ) | $spl32),
  inference(cnf_transformation,[],[f448_D])).
fof(f448_D,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) = s(bool,X1)) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f290,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f118])).
fof(f118,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f85])).
fof(f85,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X11,X12] : ((p(s(bool,X11)) => p(s(bool,X12))) => ((p(s(bool,X12)) => p(s(bool,X11))) => s(bool,X11) = s(bool,X12)))),
  file('thm.divides.divides_|49|6_prove_ltb','thm.bool.IMP_ANTISYM_AX')).
fof(f989,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | $spl332),
  inference(cnf_transformation,[],[f989_D])).
fof(f989_D,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) <=> ~$spl332),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl332])])).
fof(f972,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | ($spl18 | $spl32)),
  inference(resolution,[],[f448,f433])).
fof(f445,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))) ) | $spl26),
  inference(cnf_transformation,[],[f445_D])).
fof(f445_D,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f286,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))) )),
  inference(cnf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))),
  inference(flattening,[],[f80])).
fof(f80,plain,(
  ! [X0] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X10] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X10),s('type.num.num','const.num.|48|_0'))))),
  file('thm.divides.divides_|49|6_prove_ltb','thm.prim_rec.NOT_LESS_|48|')).
fof(f547,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.num.|48|_0')))) ) | $spl170),
  inference(cnf_transformation,[],[f547_D])).
fof(f547_D,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f543,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.num.|48|_0')))) ) | ($spl22 | $spl26)),
  inference(resolution,[],[f435,f445])).
fof(f424,plain,(
  ~p(s(bool,'F_0')) | $spl9),
  inference(cnf_transformation,[],[f424_D])).
fof(f424_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f422,plain,(
  ~p(s(bool,'F_0'))),
  inference(global_subsumption,[],[f230,f227])).
fof(f230,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f151])).
fof(f151,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f48])).
fof(f48,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f25])).
fof(f25,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.divides.divides_|49|6_prove_ltb','thm.bool.NOT_CLAUSES|split|2')).
fof(f3333,plain,(
  ( ! [X12,X11] : (s(bool,X12) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X12)) ) | $spl1314),
  inference(cnf_transformation,[],[f3333_D])).
fof(f3333_D,plain,(
  ( ! [X12,X11] : (s(bool,X12) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X12)) ) <=> ~$spl1314),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1314])])).
fof(f3332,plain,(
  ( ! [X12,X11] : (s(bool,X12) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X12)) ) | ($spl9 | $spl170 | $spl332)),
  inference(subsumption_resolution,[],[f3310,f424])).
fof(f3310,plain,(
  ( ! [X12,X11] : (p(s(bool,'F_0')) | s(bool,X12) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X12)) ) | ($spl170 | $spl332)),
  inference(superposition,[],[f989,f547])).
fof(f5323,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,X2)) ) | $spl1472),
  inference(cnf_transformation,[],[f5323_D])).
fof(f5323_D,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl1472),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1472])])).
fof(f7497,plain,(
  ( ! [X3] : (s(bool,'T_0') = s(bool,X3)) ) | ($spl1314 | $spl1400 | $spl2046)),
  inference(subsumption_resolution,[],[f7123,f3333])).
fof(f7123,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X3)) ) | ($spl1400 | $spl2046)),
  inference(backward_demodulation,[],[f6962,f3816])).
fof(f516,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) != s(bool,X1)) ) | $spl158),
  inference(cnf_transformation,[],[f516_D])).
fof(f516_D,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) != s(bool,X1)) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f405,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | ~p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(condensation,[],[f404])).
fof(f404,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2)) )),
  inference(condensation,[],[f393])).
fof(f393,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f223])).
fof(f223,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f222])).
fof(f222,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f134])).
fof(f134,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f133])).
fof(f133,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f105])).
fof(f105,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X8,X17,X15,X18] : ((s(bool,X8) = s(bool,X17) & (p(s(bool,X17)) => s(bool,X15) = s(bool,X18))) => ((p(s(bool,X8)) => p(s(bool,X15))) <=> (p(s(bool,X17)) => p(s(bool,X18)))))),
  file('thm.divides.divides_|49|6_prove_ltb','thm.bool.IMP_CONG')).
fof(f1151,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | ~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X1)) ) | $spl410),
  inference(cnf_transformation,[],[f1151_D])).
fof(f1151_D,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | ~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X1)) ) <=> ~$spl410),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl410])])).
fof(f1134,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | ($spl18 | $spl158)),
  inference(resolution,[],[f516,f433])).
fof(f3780,plain,(
  ( ! [X8,X9] : (s(bool,X8) != s(bool,X9) | s(bool,'T_0') = s(bool,X9) | s(bool,'F_0') = s(bool,X8)) ) | $spl1394),
  inference(cnf_transformation,[],[f3780_D])).
fof(f3780_D,plain,(
  ( ! [X8,X9] : (s(bool,X8) != s(bool,X9) | s(bool,'T_0') = s(bool,X9) | s(bool,'F_0') = s(bool,X8)) ) <=> ~$spl1394),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1394])])).
fof(f3779,plain,(
  ( ! [X8,X9] : (s(bool,X8) != s(bool,X9) | s(bool,'T_0') = s(bool,X9) | s(bool,'F_0') = s(bool,X8)) ) | ($spl6 | $spl168 | $spl410 | $spl1266)),
  inference(subsumption_resolution,[],[f3778,f421])).
fof(f3778,plain,(
  ( ! [X8,X9] : (~p(s(bool,'T_0')) | s(bool,X8) != s(bool,X9) | s(bool,'T_0') = s(bool,X9) | s(bool,'F_0') = s(bool,X8)) ) | ($spl168 | $spl410 | $spl1266)),
  inference(forward_demodulation,[],[f3754,f529])).
fof(f3754,plain,(
  ( ! [X8,X9] : (s(bool,X8) != s(bool,X9) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | s(bool,'T_0') = s(bool,X9) | s(bool,'F_0') = s(bool,X8)) ) | ($spl410 | $spl1266)),
  inference(superposition,[],[f1151,f3066])).
fof(f447,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) = s(bool,X1)) ) | $spl30),
  inference(cnf_transformation,[],[f447_D])).
fof(f447_D,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) = s(bool,X1)) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f293,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f119])).
fof(f435,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) | $spl22),
  inference(cnf_transformation,[],[f435_D])).
fof(f435_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f249,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f156])).
fof(f156,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f62])).
fof(f62,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f61])).
fof(f61,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X13] : (s(bool,'F_0') = s(bool,X13) <=> ~p(s(bool,X13)))),
  file('thm.divides.divides_|49|6_prove_ltb','thm.bool.EQ_CLAUSES|split|2')).
fof(f955,plain,(
  ( ! [X12,X11] : (~p(s(bool,X11)) | s(bool,X11) = s(bool,X12) | s(bool,'F_0') = s(bool,X12)) ) | $spl330),
  inference(cnf_transformation,[],[f955_D])).
fof(f955_D,plain,(
  ( ! [X12,X11] : (~p(s(bool,X11)) | s(bool,X11) = s(bool,X12) | s(bool,'F_0') = s(bool,X12)) ) <=> ~$spl330),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl330])])).
fof(f933,plain,(
  ( ! [X12,X11] : (~p(s(bool,X11)) | s(bool,X11) = s(bool,X12) | s(bool,'F_0') = s(bool,X12)) ) | ($spl22 | $spl30)),
  inference(resolution,[],[f447,f435])).
fof(f433,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl18),
  inference(cnf_transformation,[],[f433_D])).
fof(f433_D,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f245,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f154])).
fof(f154,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X13] : (s(bool,'T_0') = s(bool,X13) <=> p(s(bool,X13)))),
  file('thm.divides.divides_|49|6_prove_ltb','thm.bool.EQ_CLAUSES|split|1')).
fof(f419,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) | $spl4),
  inference(cnf_transformation,[],[f419_D])).
fof(f419_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f224,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))))),
  inference(cnf_transformation,[],[f150])).
fof(f529,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) | $spl168),
  inference(cnf_transformation,[],[f529_D])).
fof(f529_D,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f524,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) | ($spl4 | $spl18)),
  inference(resolution,[],[f433,f419])).
fof(f3066,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) | s(bool,'F_0') = s(bool,X4)) ) | $spl1266),
  inference(cnf_transformation,[],[f3066_D])).
fof(f3066_D,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) | s(bool,'F_0') = s(bool,X4)) ) <=> ~$spl1266),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1266])])).
fof(f3065,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) | s(bool,'F_0') = s(bool,X4)) ) | ($spl6 | $spl168 | $spl330)),
  inference(subsumption_resolution,[],[f3035,f421])).
fof(f3035,plain,(
  ( ! [X4] : (~p(s(bool,'T_0')) | s(bool,X4) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) | s(bool,'F_0') = s(bool,X4)) ) | ($spl168 | $spl330)),
  inference(superposition,[],[f955,f529])).
fof(f3814,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl1396),
  inference(cnf_transformation,[],[f3814_D])).
fof(f3814_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl1396),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1396])])).
fof(f3816,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) | s(bool,'T_0') = s(bool,X3)) ) | $spl1400),
  inference(cnf_transformation,[],[f3816_D])).
fof(f3816_D,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) | s(bool,'T_0') = s(bool,X3)) ) <=> ~$spl1400),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1400])])).
fof(f3812,plain,(
  ( ! [X2,X3] : (s(bool,X3) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) | s(bool,'T_0') = s(bool,X3) | s(bool,'F_0') = s(bool,X2)) ) | ($spl1266 | $spl1394)),
  inference(duplicate_literal_removal,[],[f3789])).
fof(f3789,plain,(
  ( ! [X2,X3] : (s(bool,X3) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12))) | s(bool,'T_0') = s(bool,X3) | s(bool,'F_0') = s(bool,X2) | s(bool,'F_0') = s(bool,X2)) ) | ($spl1266 | $spl1394)),
  inference(superposition,[],[f3780,f3066])).
fof(f421,plain,(
  p(s(bool,'T_0')) | $spl6),
  inference(cnf_transformation,[],[f421_D])).
fof(f421_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f227,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f10])).
fof(f10,axiom,(
  p(s(bool,'T_0'))),
  file('thm.divides.divides_|49|6_prove_ltb','thm.bool.TRUTH')).
fof(f2135,plain,(
  $false | ($spl6 | $spl738)),
  inference(resolution,[],[f2111,f421])).
fof(f436,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) ) | $spl24),
  inference(cnf_transformation,[],[f436_D])).
fof(f436_D,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f248,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f156])).
fof(f8168,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl24 | $spl1396)),
  inference(subsumption_resolution,[],[f436,f3814])).
% SZS output end Proof for thm.divides
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 315
% Passive clauses: 1226
% Generated clauses: 8820
% Final active clauses: 198
% Final passive clauses: 439
% Input formulas: 45
% Initial clauses: 170
% 
% Duplicate literals: 199
% Trivial inequalities: 55
% Fw subsumption resolutions: 932
% Fw demodulations: 871
% Bw demodulations: 563
% Condensations: 243
% Global subsumptions: 18
% 
% Simple tautologies: 244
% Equational tautologies: 18
% Forward subsumptions: 3146
% Bw demodulations to eq. taut.: 23
% 
% Binary resolution: 1568
% Factoring: 106
% Forward superposition: 982
% Backward superposition: 1109
% Self superposition: 32
% Equality factoring: 30
% Equality resolution: 16
% 
% Split clauses: 35
% Split components: 75
% SAT solver clauses: 1642
% SAT solver unit clauses: 187
% SAT solver binary clauses: 66
% SAT solver learnt clauses: 37
% SAT solver learnt literals: 10
% 
% Sat splits: 35
% Sat splitting refutations: 7
% 
% TWLsolver clauses: 1542
% TWLsolver calls for satisfiability: 1451
% 
% Memory used [KB]: 6524
% Time elapsed: 0.317 s
% ------------------------------
% ------------------------------
% Success in time 0.941 s
