Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.quotient
% (118172)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 815
% Passive clauses: 8663
% Generated clauses: 24759
% Final active clauses: 800
% Final passive clauses: 7462
% Input formulas: 69
% Initial clauses: 206
% 
% Duplicate literals: 383
% Trivial inequalities: 52
% Fw subsumption resolutions: 1481
% Fw demodulations: 2616
% Bw demodulations: 10
% 
% Simple tautologies: 235
% Equational tautologies: 42
% Forward subsumptions: 9750
% 
% Binary resolution: 8259
% Factoring: 999
% Forward superposition: 3794
% Backward superposition: 6753
% Self superposition: 56
% Equality factoring: 28
% Equality resolution: 11
% 
% Split clauses: 293
% Split components: 588
% SAT solver clauses: 457
% SAT solver unit clauses: 15
% SAT solver binary clauses: 83
% 
% Sat splits: 293
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 457
% TWLsolver calls for satisfiability: 309
% 
% Memory used [KB]: 15479
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 144 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.quotient
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.quotient
% SZS output start Proof for thm.quotient
fof(f5345,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f5343,f584,f586,f740,f742,f406,f609,f415,f616,f771,f776,f353,f578,f370,f591,f2895,f2915,f2489,f2490,f351,f581,f354,f583,f2474,f2475,f398,f604,f755,f759,f352,f580,f366,f589,f2890,f2892,f416,f615])).
fof(f615,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK19(X2,X1,X0)))))) ) | $spl56),
  inference(cnf_transformation,[],[f615_D])).
fof(f615_D,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK19(X2,X1,X0)))))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f416,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK19(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f279])).
fof(f279,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & ((p(s(bool,'const.bool.IN_2'(s(X0,sK19(X2,X1,X0)),s(fun(X0,bool),X1)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK19(X2,X1,X0)))))) | p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK19])],[f278])).
fof(f278,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f277])).
fof(f277,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))) & (? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(nnf_transformation,[],[f204])).
fof(f204,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(ennf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) => p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(flattening,[],[f122])).
fof(f122,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) => p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X6] : ! [X7,X17] : (p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X17)))) <=> ! [X10] : (p(s(bool,'const.bool.IN_2'(s(X6,X10),s(fun(X6,bool),X7)))) => p(s(bool,happ(s(fun(X6,bool),X17),s(X6,X10))))))),
  file('thm.quotient.RIGHT_RES_FORALL_REGULAR','thm.bool.RES_FORALL_THM')).
fof(f2892,plain,(
  ~p(s(bool,happ(s(fun(sK10,bool),sK13),s(sK10,sK19(sK13,sK12,sK10))))) | $spl1167),
  inference(cnf_transformation,[],[f2892_D])).
fof(f2892_D,plain,(
  ~p(s(bool,happ(s(fun(sK10,bool),sK13),s(sK10,sK19(sK13,sK12,sK10))))) <=> ~$spl1167),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1167])])).
fof(f2890,plain,(
  ~p(s(bool,happ(s(fun(sK10,bool),sK13),s(sK10,sK19(sK13,sK12,sK10))))) | ($spl9 | $spl56 | $spl294)),
  inference(subsumption_resolution,[],[f2867,f586])).
fof(f2867,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,happ(s(fun(sK10,bool),sK13),s(sK10,sK19(sK13,sK12,sK10))))) | ($spl56 | $spl294)),
  inference(superposition,[],[f615,f776])).
fof(f589,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl12),
  inference(cnf_transformation,[],[f589_D])).
fof(f589_D,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f366,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f248])).
fof(f248,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X12] : (s(bool,'T_0') = s(bool,X12) <=> p(s(bool,X12)))),
  file('thm.quotient.RIGHT_RES_FORALL_REGULAR','thm.bool.EQ_CLAUSES|split|1')).
fof(f580,plain,(
  p(s(bool,'|dollar|forall_1'(s(fun(sK10,bool),sK11)))) | $spl2),
  inference(cnf_transformation,[],[f580_D])).
fof(f580_D,plain,(
  p(s(bool,'|dollar|forall_1'(s(fun(sK10,bool),sK11)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f352,plain,(
  p(s(bool,'|dollar|forall_1'(s(fun(sK10,bool),sK11))))),
  inference(cnf_transformation,[],[f246])).
fof(f246,plain,(
  ! [X4] : (~p(s(bool,happ(s(fun(sK10,bool),sK12),s(sK10,X4)))) | ~p(s(bool,happ(s(fun(sK10,bool),sK11),s(sK10,X4)))) | p(s(bool,happ(s(fun(sK10,bool),sK13),s(sK10,X4))))) & p(s(bool,'|dollar|forall_1'(s(fun(sK10,bool),sK11)))) & ~p(s(bool,'const.bool.RES_FORALL_2'(s(fun(sK10,bool),sK12),s(fun(sK10,bool),sK13))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK10,sK11,sK12,sK13])],[f190])).
fof(f190,plain,(
  ? [X0,X1,X2,X3] : (! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4))))) & p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))))),
  inference(flattening,[],[f189])).
fof(f189,plain,(
  ? [X0,X1,X2,X3] : (! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))))) & (p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))))),
  inference(ennf_transformation,[],[f72])).
fof(f72,plain,(
  ~! [X0,X1,X2,X3] : (! [X4] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) => (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) => p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))))) => (p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),X1)))) => p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ~! [X0] : ! [X1,X2,X3] : (! [X4] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) => (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) => p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))))) => (p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),X1)))) => p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : (! [X10] : (p(s(bool,happ(s(fun(X6,bool),X8),s(X6,X10)))) => (p(s(bool,happ(s(fun(X6,bool),X7),s(X6,X10)))) => p(s(bool,happ(s(fun(X6,bool),X9),s(X6,X10)))))) => (p(s(bool,'|dollar|forall_1'(s(fun(X6,bool),X7)))) => p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X9))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : (! [X10] : (p(s(bool,happ(s(fun(X6,bool),X8),s(X6,X10)))) => (p(s(bool,happ(s(fun(X6,bool),X7),s(X6,X10)))) => p(s(bool,happ(s(fun(X6,bool),X9),s(X6,X10)))))) => (p(s(bool,'|dollar|forall_1'(s(fun(X6,bool),X7)))) => p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X9))))))),
  file('thm.quotient.RIGHT_RES_FORALL_REGULAR','thm.quotient.RIGHT_RES_FORALL_REGULAR')).
fof(f759,plain,(
  s(bool,'T_0') = s(bool,'|dollar|forall_1'(s(fun(sK10,bool),sK11))) | $spl292),
  inference(cnf_transformation,[],[f759_D])).
fof(f759_D,plain,(
  s(bool,'T_0') = s(bool,'|dollar|forall_1'(s(fun(sK10,bool),sK11))) <=> ~$spl292),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl292])])).
fof(f755,plain,(
  s(bool,'T_0') = s(bool,'|dollar|forall_1'(s(fun(sK10,bool),sK11))) | ($spl2 | $spl12)),
  inference(resolution,[],[f589,f580])).
fof(f604,plain,(
  ( ! [X0,X3,X1] : (~p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),X1)))) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) ) | $spl34),
  inference(cnf_transformation,[],[f604_D])).
fof(f604_D,plain,(
  ( ! [X0,X3,X1] : (~p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),X1)))) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f398,plain,(
  ( ! [X0,X3,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))) | ~p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f269])).
fof(f269,plain,(
  ! [X0,X1] : ((~p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),X1)))) | ! [X3] : s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & (s(bool,'T_0') != s(bool,happ(s(fun(X0,bool),X1),s(X0,sK15(X1,X0)))) | p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK15])],[f268])).
fof(f268,plain,(
  ! [X0,X1] : ((~p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),X1)))) | ! [X3] : s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & (? [X2] : s(bool,'T_0') != s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) | p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),X1))))))),
  inference(rectify,[],[f267])).
fof(f267,plain,(
  ! [X0,X1] : ((~p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),X1)))) | ! [X2] : s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) & (? [X2] : s(bool,'T_0') != s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) | p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),X1))))))),
  inference(nnf_transformation,[],[f106])).
fof(f106,plain,(
  ! [X0,X1] : (p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),X1)))) <=> ! [X2] : s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))),
  inference(flattening,[],[f105])).
fof(f105,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'|dollar|forall_1'(s(fun(X0,bool),X1)))) <=> ! [X2] : s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X6] : ! [X5] : (p(s(bool,'|dollar|forall_1'(s(fun(X6,bool),X5)))) <=> ! [X37] : s(bool,'T_0') = s(bool,happ(s(fun(X6,bool),X5),s(X6,X37))))),
  file('thm.quotient.RIGHT_RES_FORALL_REGULAR','thm.bool.FORALL_DEF')).
fof(f2475,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun(sK10,bool),sK11),s(sK10,X0)))) ) | $spl1010),
  inference(cnf_transformation,[],[f2475_D])).
fof(f2475_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun(sK10,bool),sK11),s(sK10,X0)))) ) <=> ~$spl1010),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1010])])).
fof(f2474,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun(sK10,bool),sK11),s(sK10,X0)))) ) | ($spl6 | $spl34 | $spl292)),
  inference(subsumption_resolution,[],[f2455,f583])).
fof(f2455,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,happ(s(fun(sK10,bool),sK11),s(sK10,X0)))) ) | ($spl34 | $spl292)),
  inference(superposition,[],[f604,f759])).
fof(f583,plain,(
  p(s(bool,'T_0')) | $spl6),
  inference(cnf_transformation,[],[f583_D])).
fof(f583_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f354,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f13])).
fof(f13,axiom,(
  p(s(bool,'T_0'))),
  file('thm.quotient.RIGHT_RES_FORALL_REGULAR','thm.bool.TRUTH')).
fof(f581,plain,(
  ( ! [X4] : (p(s(bool,happ(s(fun(sK10,bool),sK13),s(sK10,X4)))) | ~p(s(bool,happ(s(fun(sK10,bool),sK11),s(sK10,X4)))) | ~p(s(bool,happ(s(fun(sK10,bool),sK12),s(sK10,X4))))) ) | $spl4),
  inference(cnf_transformation,[],[f581_D])).
fof(f581_D,plain,(
  ( ! [X4] : (p(s(bool,happ(s(fun(sK10,bool),sK13),s(sK10,X4)))) | ~p(s(bool,happ(s(fun(sK10,bool),sK11),s(sK10,X4)))) | ~p(s(bool,happ(s(fun(sK10,bool),sK12),s(sK10,X4))))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f351,plain,(
  ( ! [X4] : (p(s(bool,happ(s(fun(sK10,bool),sK13),s(sK10,X4)))) | ~p(s(bool,happ(s(fun(sK10,bool),sK11),s(sK10,X4)))) | ~p(s(bool,happ(s(fun(sK10,bool),sK12),s(sK10,X4))))) )),
  inference(cnf_transformation,[],[f246])).
fof(f2490,plain,(
  ( ! [X4] : (p(s(bool,happ(s(fun(sK10,bool),sK13),s(sK10,X4)))) | ~p(s(bool,happ(s(fun(sK10,bool),sK12),s(sK10,X4))))) ) | $spl1012),
  inference(cnf_transformation,[],[f2490_D])).
fof(f2490_D,plain,(
  ( ! [X4] : (p(s(bool,happ(s(fun(sK10,bool),sK13),s(sK10,X4)))) | ~p(s(bool,happ(s(fun(sK10,bool),sK12),s(sK10,X4))))) ) <=> ~$spl1012),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1012])])).
fof(f2489,plain,(
  ( ! [X4] : (p(s(bool,happ(s(fun(sK10,bool),sK13),s(sK10,X4)))) | ~p(s(bool,happ(s(fun(sK10,bool),sK12),s(sK10,X4))))) ) | ($spl4 | $spl6 | $spl1010)),
  inference(subsumption_resolution,[],[f2476,f583])).
fof(f2476,plain,(
  ( ! [X4] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK10,bool),sK13),s(sK10,X4)))) | ~p(s(bool,happ(s(fun(sK10,bool),sK12),s(sK10,X4))))) ) | ($spl4 | $spl1010)),
  inference(backward_demodulation,[],[f2475,f581])).
fof(f2915,plain,(
  ~p(s(bool,happ(s(fun(sK10,bool),sK12),s(sK10,sK19(sK13,sK12,sK10))))) | $spl1171),
  inference(cnf_transformation,[],[f2915_D])).
fof(f2915_D,plain,(
  ~p(s(bool,happ(s(fun(sK10,bool),sK12),s(sK10,sK19(sK13,sK12,sK10))))) <=> ~$spl1171),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1171])])).
fof(f2895,plain,(
  ~p(s(bool,happ(s(fun(sK10,bool),sK12),s(sK10,sK19(sK13,sK12,sK10))))) | ($spl1012 | $spl1167)),
  inference(resolution,[],[f2892,f2490])).
fof(f591,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) | $spl16),
  inference(cnf_transformation,[],[f591_D])).
fof(f591_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f370,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f250])).
fof(f250,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f81])).
fof(f81,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X12] : (s(bool,'F_0') = s(bool,X12) <=> ~p(s(bool,X12)))),
  file('thm.quotient.RIGHT_RES_FORALL_REGULAR','thm.bool.EQ_CLAUSES|split|3')).
fof(f578,plain,(
  ~p(s(bool,'const.bool.RES_FORALL_2'(s(fun(sK10,bool),sK12),s(fun(sK10,bool),sK13)))) | $spl1),
  inference(cnf_transformation,[],[f578_D])).
fof(f578_D,plain,(
  ~p(s(bool,'const.bool.RES_FORALL_2'(s(fun(sK10,bool),sK12),s(fun(sK10,bool),sK13)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f353,plain,(
  ~p(s(bool,'const.bool.RES_FORALL_2'(s(fun(sK10,bool),sK12),s(fun(sK10,bool),sK13))))),
  inference(cnf_transformation,[],[f246])).
fof(f776,plain,(
  s(bool,'F_0') = s(bool,'const.bool.RES_FORALL_2'(s(fun(sK10,bool),sK12),s(fun(sK10,bool),sK13))) | $spl294),
  inference(cnf_transformation,[],[f776_D])).
fof(f776_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.RES_FORALL_2'(s(fun(sK10,bool),sK12),s(fun(sK10,bool),sK13))) <=> ~$spl294),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl294])])).
fof(f771,plain,(
  s(bool,'F_0') = s(bool,'const.bool.RES_FORALL_2'(s(fun(sK10,bool),sK12),s(fun(sK10,bool),sK13))) | ($spl1 | $spl16)),
  inference(resolution,[],[f591,f578])).
fof(f616,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK19(X2,X1,X0)),s(fun(X0,bool),X1))))) ) | $spl58),
  inference(cnf_transformation,[],[f616_D])).
fof(f616_D,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK19(X2,X1,X0)),s(fun(X0,bool),X1))))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f415,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK19(X2,X1,X0)),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f279])).
fof(f609,plain,(
  ( ! [X2,X0,X1] : (s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) = s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) ) | $spl44),
  inference(cnf_transformation,[],[f609_D])).
fof(f609_D,plain,(
  ( ! [X2,X0,X1] : (s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) = s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f406,plain,(
  ( ! [X2,X0,X1] : (s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) = s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) )),
  inference(cnf_transformation,[],[f111])).
fof(f111,plain,(
  ! [X0,X1,X2] : s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) = s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))),
  inference(flattening,[],[f110])).
fof(f110,plain,(
  ! [X0] : ! [X1,X2] : s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) = s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X6] : ! [X7,X10] : s(bool,happ(s(fun(X6,bool),X7),s(X6,X10))) = s(bool,'const.bool.IN_2'(s(X6,X10),s(fun(X6,bool),X7)))),
  file('thm.quotient.RIGHT_RES_FORALL_REGULAR','thm.pred_set.SPECIFICATION')).
fof(f742,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK19(X2,X1,X0))))) | p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) ) | $spl280),
  inference(cnf_transformation,[],[f742_D])).
fof(f742_D,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK19(X2,X1,X0))))) | p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) ) <=> ~$spl280),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl280])])).
fof(f740,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK19(X2,X1,X0))))) | p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) ) | ($spl44 | $spl58)),
  inference(backward_demodulation,[],[f609,f616])).
fof(f586,plain,(
  ~p(s(bool,'F_0')) | $spl9),
  inference(cnf_transformation,[],[f586_D])).
fof(f586_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f584,plain,(
  ~p(s(bool,'F_0'))),
  inference(global_subsumption,[],[f357,f354])).
fof(f357,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f247])).
fof(f247,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f73])).
fof(f73,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f70])).
fof(f70,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.quotient.RIGHT_RES_FORALL_REGULAR','thm.bool.NOT_CLAUSES|split|2')).
fof(f5343,plain,(
  $false | ($spl9 | $spl280 | $spl294 | $spl1171)),
  inference(subsumption_resolution,[],[f5342,f586])).
fof(f5342,plain,(
  p(s(bool,'F_0')) | ($spl280 | $spl294 | $spl1171)),
  inference(forward_demodulation,[],[f5322,f776])).
fof(f5322,plain,(
  p(s(bool,'const.bool.RES_FORALL_2'(s(fun(sK10,bool),sK12),s(fun(sK10,bool),sK13)))) | ($spl280 | $spl1171)),
  inference(resolution,[],[f2915,f742])).
% SZS output end Proof for thm.quotient
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 265
% Passive clauses: 1163
% Generated clauses: 4938
% Final active clauses: 251
% Final passive clauses: 860
% Input formulas: 69
% Initial clauses: 206
% 
% Duplicate literals: 63
% Trivial inequalities: 5
% Fw subsumption resolutions: 216
% Fw demodulations: 497
% Bw demodulations: 50
% Condensations: 36
% Global subsumptions: 11
% 
% Simple tautologies: 144
% Equational tautologies: 6
% Forward subsumptions: 1477
% 
% Binary resolution: 1219
% Factoring: 43
% Forward superposition: 698
% Backward superposition: 691
% Self superposition: 16
% Equality factoring: 2
% Equality resolution: 1
% 
% Split clauses: 4
% Split components: 8
% SAT solver clauses: 1508
% SAT solver unit clauses: 222
% SAT solver binary clauses: 79
% SAT solver learnt clauses: 23
% SAT solver learnt literals: 10
% 
% Sat splits: 4
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 1463
% TWLsolver calls for satisfiability: 1745
% 
% Memory used [KB]: 7675
% Time elapsed: 0.139 s
% ------------------------------
% ------------------------------
% Success in time 0.745 s
