Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.quotient
% (74176)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 891
% Passive clauses: 7203
% Generated clauses: 24146
% Final active clauses: 815
% Final passive clauses: 5743
% Input formulas: 54
% Initial clauses: 213
% 
% Duplicate literals: 567
% Trivial inequalities: 136
% Fw subsumption resolutions: 2452
% Fw demodulations: 1859
% Bw demodulations: 73
% 
% Simple tautologies: 298
% Equational tautologies: 61
% Forward subsumptions: 10674
% 
% Binary resolution: 8580
% Factoring: 1424
% Forward superposition: 3743
% Backward superposition: 4718
% Self superposition: 35
% Equality factoring: 30
% Equality resolution: 23
% 
% Split clauses: 794
% Split components: 1605
% SAT solver clauses: 1275
% SAT solver unit clauses: 16
% SAT solver binary clauses: 187
% 
% Sat splits: 794
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 1273
% TWLsolver calls for satisfiability: 439
% 
% Memory used [KB]: 9466
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.quotient
% (74192)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 416
% Passive clauses: 2246
% Generated clauses: 11146
% Final active clauses: 409
% Final passive clauses: 1827
% Input formulas: 54
% Initial clauses: 213
% 
% Duplicate literals: 218
% Trivial inequalities: 82
% Fw subsumption resolutions: 772
% Fw demodulations: 812
% Bw demodulations: 9
% Condensations: 297
% Global subsumptions: 22
% 
% Simple tautologies: 294
% Equational tautologies: 21
% Forward subsumptions: 3691
% 
% Binary resolution: 3033
% Factoring: 177
% Forward superposition: 1896
% Backward superposition: 1279
% Self superposition: 26
% Equality factoring: 26
% Equality resolution: 8
% 
% Split clauses: 32
% Split components: 69
% SAT solver clauses: 2911
% SAT solver unit clauses: 219
% SAT solver binary clauses: 86
% SAT solver learnt clauses: 44
% SAT solver learnt literals: 10
% 
% Sat splits: 32
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 2823
% TWLsolver calls for satisfiability: 1761
% 
% Memory used [KB]: 12665
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.quotient
% (74213)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 2353
% Passive clauses: 9138
% Generated clauses: 158211
% Final active clauses: 1549
% Final passive clauses: 3987
% Input formulas: 54
% Initial clauses: 213
% 
% Duplicate literals: 9137
% Trivial inequalities: 289
% Fw subsumption resolutions: 9026
% Fw demodulations: 20846
% Bw demodulations: 2233
% Condensations: 12759
% Global subsumptions: 433
% 
% Simple tautologies: 293
% Equational tautologies: 1749
% Forward subsumptions: 92391
% Fw demodulations to eq. taut.: 2
% Bw demodulations to eq. taut.: 124
% 
% Unit resulting resolution: 29034
% Factoring: 1
% Forward superposition: 40506
% Backward superposition: 30386
% Self superposition: 211
% Equality factoring: 173
% Equality resolution: 7
% 
% Split clauses: 630
% Split components: 1268
% SAT solver clauses: 12445
% SAT solver unit clauses: 1226
% SAT solver binary clauses: 2984
% 
% Sat splits: 630
% Sat splitting refutations: 1612
% 
% Memory used [KB]: 73303
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.quotient
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.quotient
% SZS output start Proof for thm.quotient
fof(f25520,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f615,f733,f616,f735,f731,f773,f889,f732,f734,f1018,f1972,f2554,f2563,f2577,f2578,f2579,f2664,f2667,f2566,f2544,f2529,f2404,f1974,f2784,f2770,f2839,f2840,f2841,f2842,f2960,f2959,f2958,f2946,f2945,f2944,f3029,f3030,f3031,f3032,f3033,f3034,f3590,f3591,f3821,f3830,f3831,f3832,f3939,f3940,f4135,f4161,f4171,f4162,f4163,f4252,f4253,f4254,f4409,f4438,f4439,f4440,f4441,f4449,f2565,f4483,f4486,f4507,f4508,f4509,f4513,f4514,f4515,f4591,f4594,f4595,f4596,f4592,f4597,f4599,f4598,f4600,f4603,f4604,f4605,f4601,f4658,f4671,f4672,f4388,f4387,f4674,f4675,f4676,f4677,f4678,f4679,f4146,f4126,f4068,f4067,f4682,f4683,f4684,f4685,f4673,f4688,f4690,f4689,f4693,f4694,f2665,f4739,f4740,f5046,f3836,f5047,f3837,f5048,f3838,f5088,f5089,f5090,f5094,f5095,f5096,f5241,f5243,f5244,f5245,f5242,f5246,f5248,f5247,f5249,f5251,f4602,f5252,f5253,f5250,f5238,f5237,f5256,f5259,f5254,f5325,f5324,f5328,f5329,f5330,f5331,f5444,f5445,f4593,f5446,f5447,f5448,f5443,f5449,f5442,f5257,f5468,f5469,f5470,f5608,f5099,f5609,f5100,f5610,f5101,f5614,f5052,f5615,f5054,f5616,f5056,f5633,f4527,f5634,f4528,f5635,f4529,f5639,f2585,f4523,f5640,f2587,f4524,f5641,f2589,f4525,f5833,f5835,f5834,f5836,f5839,f5840,f5841,f5837,f5844,f5845,f5826,f5825,f5847,f5848,f5846,f5922,f5961,f5966,f6398,f6399,f6509,f6513,f6535,f4411,f6539,f6613,f6615,f6614,f6616,f6618,f5838,f6619,f6620,f6617,f6623,f6624,f6625,f6626,f6718,f2558,f6769,f3824,f6806,f6520,f4492,f6810,f6518,f6933,f4172,f6934,f4173,f6947,f4141,f7425,f7426,f7427,f7428,f8091,f8092,f8093,f8094,f8662,f8663,f8664,f8665,f8853,f8856,f8854,f8857,f8860,f8861,f8864,f8865,f8868,f8859,f8809,f8810,f8811,f8869,f8812,f8863,f8870,f8813,f8814,f8873,f8872,f8874,f8877,f8858,f8817,f8818,f8880,f8879,f8881,f8885,f8882,f8946,f8949,f9374,f9376,f9757,f9758,f10003,f9982,f10009,f9978,f10295,f10300,f10304,f10305,f10309,f10310,f10311,f10363,f10364,f10390,f10391,f10708,f10709,f10814,f10815,f10816,f11149,f10334,f11154,f11158,f11159,f11163,f11164,f11165,f11242,f11243,f11244,f11246,f8884,f11295,f11296,f11322,f11323,f11582,f11583,f11720,f11721,f11722,f11752,f11753,f11754,f11755,f11983,f11986,f12320,f12321,f12623,f12626,f12627,f8867,f2666,f12628,f12629,f13080,f13085,f13088,f13089,f13090,f13094,f13095,f13096,f13140,f13141,f13163,f13164,f13403,f13404,f13497,f13498,f13499,f13828,f13833,f13836,f13837,f13838,f13842,f13843,f13844,f13913,f13914,f13915,f13916,f13917,f13958,f13959,f14011,f14012,f14287,f14288,f14408,f14409,f14410,f14436,f14437,f14438,f14439,f14728,f14729,f14730,f14731,f15112,f15114,f15468,f15473,f15476,f15477,f15478,f15482,f15483,f15484,f15553,f15554,f15555,f15556,f15557,f15582,f15583,f15584,f15585,f15586,f15587,f16261,f16263,f16317,f16318,f16378,f16379,f11985,f16692,f12343,f16693,f16820,f16821,f16822,f16848,f12344,f16849,f13521,f16850,f13520,f16851,f8875,f17541,f15113,f17542,f15115,f17610,f17611,f17612,f17613,f17608,f17606,f8947,f17912,f17918,f10336,f17919,f17999,f18000,f18093,f10335,f18094,f10337,f18097,f11269,f18692,f9375,f18693,f9377,f19927,f19929,f19928,f20284,f20289,f20292,f13115,f20293,f20294,f20298,f20299,f13117,f20300,f20444,f20445,f20446,f20447,f20448,f20475,f13116,f20476,f13118,f20477,f8866,f20478,f13519,f20479,f13936,f20480,f8876,f20938,f16262,f20939,f16264,f21826,f21827,f21930,f21931,f23676,f23679,f23678,f23805,f23808,f23807,f24466,f4444,f24470,f24520,f24475,f8855,f24477,f24521,f24552,f24553,f24559,f24684,f24560,f24685,f24686,f24736,f24737,f8948,f25141,f9784,f25142,f2556,f25275,f25276,f25277,f25309,f9785,f25310,f10842,f25311,f10841,f25312,f1973,f8883,f10840,f25515])).
fof(f25515,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f25492,f484])).
fof(f484,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.quotient.QUOTIENT_SYM',aHL_TRUTH)).
fof(f25492,plain,(
  ~p(s(bool,'T_0')) | $spl6),
  inference(superposition,[],[f25392,f581])).
fof(f581,plain,(
  s(bool,'T_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK13,fun(sK13,bool)),sK14),s(fun(sK13,sK12),sK15),s(fun(sK12,sK13),sK16)))),
  inference(resolution,[],[f276,f465])).
fof(f465,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f271])).
fof(f271,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f112])).
fof(f112,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X13] : (s(bool,'T_0') = s(bool,X13) <=> p(s(bool,X13)))),
  file('thm.quotient.QUOTIENT_SYM','thm.bool.EQ_CLAUSES|split|0')).
fof(f276,plain,(
  p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK13,fun(sK13,bool)),sK14),s(fun(sK13,sK12),sK15),s(fun(sK12,sK13),sK16))))),
  inference(cnf_transformation,[],[f175])).
fof(f175,plain,(
  p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK13,fun(sK13,bool)),sK14),s(fun(sK13,sK12),sK15),s(fun(sK12,sK13),sK16)))) & (p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK18)))) & ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK17)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13,sK14,sK15,sK16,sK17,sK18])],[f131])).
fof(f131,plain,(
  ? [X0,X1,X2,X3,X4] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,X0),X3),s(fun(X0,X1),X4)))) & ? [X5,X6] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X6)))) & ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X6))),s(X1,X5))))))),
  inference(ennf_transformation,[],[f57])).
fof(f57,plain,(
  ~! [X0,X1,X2,X3,X4] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,X0),X3),s(fun(X0,X1),X4)))) => ! [X5,X6] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X6)))) => p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X6))),s(X1,X5))))))),
  inference(flattening,[],[f56])).
fof(f56,plain,(
  ~! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,X0),X3),s(fun(X0,X1),X4)))) => ! [X5,X6] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X6)))) => p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X6))),s(X1,X5))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X7,fun(X7,bool)),X8),s(fun(X7,X6),X9),s(fun(X6,X7),X10)))) => ! [X11,X12] : (p(s(bool,happ(s(fun(X7,bool),happ(s(fun(X7,fun(X7,bool)),X8),s(X7,X11))),s(X7,X12)))) => p(s(bool,happ(s(fun(X7,bool),happ(s(fun(X7,fun(X7,bool)),X8),s(X7,X12))),s(X7,X11))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X7,fun(X7,bool)),X8),s(fun(X7,X6),X9),s(fun(X6,X7),X10)))) => ! [X11,X12] : (p(s(bool,happ(s(fun(X7,bool),happ(s(fun(X7,fun(X7,bool)),X8),s(X7,X11))),s(X7,X12)))) => p(s(bool,happ(s(fun(X7,bool),happ(s(fun(X7,fun(X7,bool)),X8),s(X7,X12))),s(X7,X11))))))),
  file('thm.quotient.QUOTIENT_SYM','thm.quotient.QUOTIENT_SYM')).
fof(f25392,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK13,fun(sK13,bool)),sK14),s(fun(sK13,X0),X1),s(fun(X0,sK13),X2))))) ) | $spl6),
  inference(subsumption_resolution,[],[f889,f617])).
fof(f617,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK13,fun(sK13,bool)),sK14),s(fun(sK13,X6),X7),s(fun(X6,sK13),X8)))) | s(X6,happ(s(fun(sK13,X6),X7),s(sK13,sK17))) = s(X6,happ(s(fun(sK13,X6),X7),s(sK13,sK18)))) )),
  inference(resolution,[],[f277,f281])).
fof(f281,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),X3),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X6))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X6)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,X0),X3),s(fun(X0,X1),X4))))) )),
  inference(cnf_transformation,[],[f177])).
fof(f177,plain,(
  ! [X0,X1,X2,X3,X4] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,X0),X3),s(fun(X0,X1),X4)))) | ! [X5,X6] : ((~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X6)))) | (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X5)))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X6))),s(X1,X6)))) & s(X0,happ(s(fun(X1,X0),X3),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X6))))) & (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X5)))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X6))),s(X1,X6)))) | s(X0,happ(s(fun(X1,X0),X3),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X6))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X6)))))))),
  inference(flattening,[],[f176])).
fof(f176,plain,(
  ! [X0,X1,X2,X3,X4] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,X0),X3),s(fun(X0,X1),X4)))) | ! [X5,X6] : ((~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X6)))) | (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X5)))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X6))),s(X1,X6)))) & s(X0,happ(s(fun(X1,X0),X3),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X6))))) & ((~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X5)))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X6))),s(X1,X6)))) | s(X0,happ(s(fun(X1,X0),X3),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X6)))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X6)))))))),
  inference(nnf_transformation,[],[f132])).
fof(f132,plain,(
  ! [X0,X1,X2,X3,X4] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,X0),X3),s(fun(X0,X1),X4)))) | ! [X5,X6] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X6)))) <=> (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X5)))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X6))),s(X1,X6)))) & s(X0,happ(s(fun(X1,X0),X3),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X6))))))),
  inference(ennf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,X0),X3),s(fun(X0,X1),X4)))) => ! [X5,X6] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X6)))) <=> (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X5)))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X6))),s(X1,X6)))) & s(X0,happ(s(fun(X1,X0),X3),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X6))))))),
  inference(flattening,[],[f58])).
fof(f58,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,X0),X3),s(fun(X0,X1),X4)))) => ! [X5,X6] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X6)))) <=> (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X5)))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X6))),s(X1,X6)))) & s(X0,happ(s(fun(X1,X0),X3),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X6))))))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X6,X7] : ! [X8,X9,X10] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X7,fun(X7,bool)),X8),s(fun(X7,X6),X9),s(fun(X6,X7),X10)))) => ! [X18,X19] : (p(s(bool,happ(s(fun(X7,bool),happ(s(fun(X7,fun(X7,bool)),X8),s(X7,X18))),s(X7,X19)))) <=> (p(s(bool,happ(s(fun(X7,bool),happ(s(fun(X7,fun(X7,bool)),X8),s(X7,X18))),s(X7,X18)))) & p(s(bool,happ(s(fun(X7,bool),happ(s(fun(X7,fun(X7,bool)),X8),s(X7,X19))),s(X7,X19)))) & s(X6,happ(s(fun(X7,X6),X9),s(X7,X18))) = s(X6,happ(s(fun(X7,X6),X9),s(X7,X19))))))),
  file('thm.quotient.QUOTIENT_SYM','thm.quotient.QUOTIENT_REL')).
fof(f277,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK18))))),
  inference(cnf_transformation,[],[f175])).
fof(f10840,plain,(
  ( ! [X11] : (s(bool,X11) != s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | $spl190),
  inference(cnf_transformation,[],[f10840_D])).
fof(f10840_D,plain,(
  ( ! [X11] : (s(bool,X11) != s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f8883,plain,(
  s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | $spl168),
  inference(cnf_transformation,[],[f8883_D])).
fof(f8883_D,plain,(
  s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'F_0'))) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f1973,plain,(
  ( ! [X3] : (p(s(bool,X3))) ) | $spl8),
  inference(cnf_transformation,[],[f1973_D])).
fof(f1973_D,plain,(
  ( ! [X3] : (p(s(bool,X3))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f25312,plain,(
  ( ! [X14,X15] : (p(s(bool,X14)) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X15) != s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | ($spl94 | $spl149)),
  inference(resolution,[],[f9460,f5575])).
fof(f5575,plain,(
  ( ! [X17,X15,X16] : (~sP6(X15,X17,X16) | p(s(bool,X16)) | s(bool,'T_0') = s(bool,X15) | s(bool,X15) != s(bool,X17)) ) | $spl94),
  inference(duplicate_literal_removal,[],[f5564])).
fof(f5564,plain,(
  ( ! [X17,X15,X16] : (s(bool,'T_0') = s(bool,X15) | p(s(bool,X16)) | ~sP6(X15,X17,X16) | p(s(bool,X16)) | s(bool,X15) != s(bool,X17)) ) | $spl94),
  inference(resolution,[],[f5528,f359])).
fof(f359,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | s(bool,X0) != s(bool,X1) | ~sP9(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f215])).
fof(f215,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP9(X2,X1,X0)) & (~sP9(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f170])).
fof(f170,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP9(X2,X1,X0))),
  inference(definition_folding,[],[f71,f169,f168,f167,f166])).
fof(f166,plain,(
  ! [X0,X1,X2] : (sP6(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f167,plain,(
  ! [X1,X0,X2] : (sP7(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f168,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f169,plain,(
  ! [X2,X1,X0] : (sP9(X2,X1,X0) <=> (sP6(X0,X1,X2) & sP7(X1,X0,X2) & sP8(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f71,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f50])).
fof(f50,axiom,(
  ! [X18,X20,X21] : ((p(s(bool,X21)) <=> s(bool,X18) = s(bool,X20)) <=> ((p(s(bool,X21)) | p(s(bool,X20)) | p(s(bool,X18))) & (p(s(bool,X21)) | ~p(s(bool,X18)) | ~p(s(bool,X20))) & (p(s(bool,X20)) | ~p(s(bool,X18)) | ~p(s(bool,X21))) & (p(s(bool,X18)) | ~p(s(bool,X20)) | ~p(s(bool,X21)))))),
  file('thm.quotient.QUOTIENT_SYM','thm.sat.sat_2|49|_prove_ltb')).
fof(f5528,plain,(
  ( ! [X17,X15,X16] : (sP9(X16,X17,X15) | s(bool,'T_0') = s(bool,X15) | p(s(bool,X16)) | ~sP6(X15,X17,X16)) ) | $spl94),
  inference(subsumption_resolution,[],[f5527,f347])).
fof(f347,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f209])).
fof(f209,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f208])).
fof(f208,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f207])).
fof(f207,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f168])).
fof(f5527,plain,(
  ( ! [X17,X15,X16] : (s(bool,'T_0') = s(bool,X15) | sP9(X16,X17,X15) | p(s(bool,X16)) | ~sP8(X16,X15,X17) | ~sP6(X15,X17,X16)) ) | $spl94),
  inference(resolution,[],[f5519,f343])).
fof(f343,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | p(s(bool,X0)) | ~sP8(X0,X2,X1) | ~sP7(X1,X2,X0) | ~sP6(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f206])).
fof(f206,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | (sP6(X2,X1,X0) & sP7(X1,X2,X0) & sP8(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP6(X2,X1,X0) | ~sP7(X1,X2,X0) | ~sP8(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f205])).
fof(f205,plain,(
  ! [X2,X1,X0] : ((~sP9(X2,X1,X0) | (sP6(X0,X1,X2) & sP7(X1,X0,X2) & sP8(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP6(X0,X1,X2) | ~sP7(X1,X0,X2) | ~sP8(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP9(X2,X1,X0)))),
  inference(flattening,[],[f204])).
fof(f204,plain,(
  ! [X2,X1,X0] : ((~sP9(X2,X1,X0) | (sP6(X0,X1,X2) & sP7(X1,X0,X2) & sP8(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP6(X0,X1,X2) | ~sP7(X1,X0,X2) | ~sP8(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP9(X2,X1,X0)))),
  inference(nnf_transformation,[],[f169])).
fof(f5519,plain,(
  ( ! [X4,X5,X3] : (sP7(X5,X3,X4) | s(bool,'T_0') = s(bool,X3)) ) | $spl94),
  inference(subsumption_resolution,[],[f5518,f349])).
fof(f349,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f212])).
fof(f212,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f211])).
fof(f211,plain,(
  ! [X1,X0,X2] : ((~sP7(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP7(X1,X0,X2)))),
  inference(flattening,[],[f210])).
fof(f210,plain,(
  ! [X1,X0,X2] : ((~sP7(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP7(X1,X0,X2)))),
  inference(nnf_transformation,[],[f167])).
fof(f5518,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') = s(bool,X3) | p(s(bool,X4)) | sP7(X5,X3,X4)) ) | $spl94),
  inference(subsumption_resolution,[],[f5501,f351])).
fof(f351,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f212])).
fof(f5501,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') = s(bool,X3) | p(s(bool,X4)) | ~p(s(bool,X5)) | sP7(X5,X3,X4)) ) | $spl94),
  inference(resolution,[],[f5247,f338])).
fof(f338,plain,(
  ( ! [X2,X0,X1] : (sP7(X1,X2,X0) | ~sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f206])).
fof(f9460,plain,(
  ( ! [X112,X113] : (sP6('const.combin.I_1'(s(bool,'F_0')),X112,X113)) ) | $spl149),
  inference(resolution,[],[f8942,f355])).
fof(f355,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f214])).
fof(f214,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(flattening,[],[f213])).
fof(f213,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(nnf_transformation,[],[f166])).
fof(f8942,plain,(
  p(s(bool,'const.combin.I_1'(s(bool,'F_0')))) | $spl149),
  inference(trivial_inequality_removal,[],[f8933])).
fof(f8933,plain,(
  s(bool,'F_0') != s(bool,'F_0') | p(s(bool,'const.combin.I_1'(s(bool,'F_0')))) | $spl149),
  inference(superposition,[],[f8863,f463])).
fof(f463,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f270])).
fof(f270,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f111])).
fof(f111,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f110])).
fof(f110,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X13] : (s(bool,'F_0') = s(bool,X13) <=> ~p(s(bool,X13)))),
  file('thm.quotient.QUOTIENT_SYM','thm.bool.EQ_CLAUSES|split|3')).
fof(f10841,plain,(
  ( ! [X15] : (s(bool,'T_0') = s(bool,X15) | s(bool,X15) != s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | $spl192),
  inference(cnf_transformation,[],[f10841_D])).
fof(f10841_D,plain,(
  ( ! [X15] : (s(bool,'T_0') = s(bool,X15) | s(bool,X15) != s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f25311,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | s(bool,'T_0') = s(bool,X13) | s(bool,X13) != s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | ($spl70 | $spl149)),
  inference(resolution,[],[f9460,f5073])).
fof(f5073,plain,(
  ( ! [X17,X15,X16] : (~sP6(X17,X15,X16) | p(s(bool,X16)) | s(bool,'T_0') = s(bool,X15) | s(bool,X15) != s(bool,X17)) ) | $spl70),
  inference(duplicate_literal_removal,[],[f5062])).
fof(f5062,plain,(
  ( ! [X17,X15,X16] : (s(bool,'T_0') = s(bool,X15) | p(s(bool,X16)) | ~sP6(X17,X15,X16) | p(s(bool,X16)) | s(bool,X15) != s(bool,X17)) ) | $spl70),
  inference(resolution,[],[f5040,f359])).
fof(f5040,plain,(
  ( ! [X17,X15,X16] : (sP9(X16,X15,X17) | s(bool,'T_0') = s(bool,X15) | p(s(bool,X16)) | ~sP6(X17,X15,X16)) ) | $spl70),
  inference(subsumption_resolution,[],[f5038,f347])).
fof(f5038,plain,(
  ( ! [X17,X15,X16] : (s(bool,'T_0') = s(bool,X15) | sP9(X16,X15,X17) | p(s(bool,X16)) | ~sP8(X16,X17,X15) | ~sP6(X17,X15,X16)) ) | $spl70),
  inference(resolution,[],[f5015,f343])).
fof(f5015,plain,(
  ( ! [X4,X5,X3] : (sP7(X3,X5,X4) | s(bool,'T_0') = s(bool,X3)) ) | $spl70),
  inference(subsumption_resolution,[],[f5014,f349])).
fof(f5014,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') = s(bool,X3) | p(s(bool,X4)) | sP7(X3,X5,X4)) ) | $spl70),
  inference(subsumption_resolution,[],[f4996,f350])).
fof(f350,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f212])).
fof(f4996,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') = s(bool,X3) | p(s(bool,X4)) | ~p(s(bool,X5)) | sP7(X3,X5,X4)) ) | $spl70),
  inference(resolution,[],[f4598,f338])).
fof(f10842,plain,(
  ( ! [X17] : (s(bool,'F_0') != s(bool,X17) | s(bool,X17) != s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | $spl194),
  inference(cnf_transformation,[],[f10842_D])).
fof(f10842_D,plain,(
  ( ! [X17] : (s(bool,'F_0') != s(bool,X17) | s(bool,X17) != s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f25310,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | s(bool,'F_0') != s(bool,X11) | s(bool,X11) != s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | ($spl10 | $spl149)),
  inference(resolution,[],[f9460,f4121])).
fof(f4121,plain,(
  ( ! [X4,X5,X3] : (~sP6(X5,X3,X4) | p(s(bool,X4)) | s(bool,'F_0') != s(bool,X3) | s(bool,X3) != s(bool,X5)) ) | $spl10),
  inference(duplicate_literal_removal,[],[f4108])).
fof(f4108,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X3) | p(s(bool,X4)) | ~sP6(X5,X3,X4) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5)) ) | $spl10),
  inference(resolution,[],[f3948,f359])).
fof(f3948,plain,(
  ( ! [X10,X11,X9] : (sP9(X10,X9,X11) | s(bool,'F_0') != s(bool,X9) | p(s(bool,X10)) | ~sP6(X11,X9,X10)) ) | $spl10),
  inference(subsumption_resolution,[],[f3946,f347])).
fof(f3946,plain,(
  ( ! [X10,X11,X9] : (s(bool,'F_0') != s(bool,X9) | sP9(X10,X9,X11) | p(s(bool,X10)) | ~sP8(X10,X11,X9) | ~sP6(X11,X9,X10)) ) | $spl10),
  inference(resolution,[],[f3942,f343])).
fof(f3942,plain,(
  ( ! [X10,X11,X9] : (sP7(X10,X9,X11) | s(bool,'F_0') != s(bool,X10)) ) | $spl10),
  inference(subsumption_resolution,[],[f3941,f350])).
fof(f3941,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | s(bool,'F_0') != s(bool,X10) | sP7(X10,X9,X11)) ) | $spl10),
  inference(subsumption_resolution,[],[f3928,f349])).
fof(f3928,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | s(bool,'F_0') != s(bool,X10) | p(s(bool,X11)) | sP7(X10,X9,X11)) ) | $spl10),
  inference(resolution,[],[f3827,f338])).
fof(f3827,plain,(
  ( ! [X6,X8,X7] : (sP9(X6,X7,X8) | ~p(s(bool,X8)) | s(bool,'F_0') != s(bool,X7) | p(s(bool,X6))) ) | $spl10),
  inference(resolution,[],[f3794,f347])).
fof(f3794,plain,(
  ( ! [X4,X5,X3] : (~sP8(X5,X4,X3) | sP9(X5,X3,X4) | ~p(s(bool,X4)) | s(bool,'F_0') != s(bool,X3)) ) | $spl10),
  inference(subsumption_resolution,[],[f3793,f465])).
fof(f3793,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X3) | s(bool,'T_0') != s(bool,X4) | sP9(X5,X3,X4) | ~p(s(bool,X4)) | ~sP8(X5,X4,X3)) ) | $spl10),
  inference(subsumption_resolution,[],[f3790,f3776])).
fof(f3776,plain,(
  ( ! [X6,X8,X7] : (sP6(X6,X7,X8) | s(bool,'F_0') != s(bool,X7) | s(bool,'T_0') != s(bool,X6)) ) | $spl10),
  inference(subsumption_resolution,[],[f3762,f353])).
fof(f353,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f214])).
fof(f3762,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X6) | s(bool,'F_0') != s(bool,X7) | p(s(bool,X8)) | sP6(X6,X7,X8)) ) | $spl10),
  inference(resolution,[],[f3059,f337])).
fof(f337,plain,(
  ( ! [X2,X0,X1] : (sP6(X2,X1,X0) | ~sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f206])).
fof(f3059,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X1,X0) | s(bool,'T_0') != s(bool,X0) | s(bool,'F_0') != s(bool,X1) | p(s(bool,X2))) ) | $spl10),
  inference(superposition,[],[f1974,f356])).
fof(f356,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f215])).
fof(f3790,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X3) | s(bool,'T_0') != s(bool,X4) | sP9(X5,X3,X4) | ~p(s(bool,X4)) | ~sP8(X5,X4,X3) | ~sP6(X4,X3,X5)) ) | $spl10),
  inference(resolution,[],[f3777,f341])).
fof(f341,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~p(s(bool,X2)) | ~sP8(X0,X2,X1) | ~sP7(X1,X2,X0) | ~sP6(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f206])).
fof(f3777,plain,(
  ( ! [X10,X11,X9] : (sP7(X10,X9,X11) | s(bool,'F_0') != s(bool,X10) | s(bool,'T_0') != s(bool,X9)) ) | $spl10),
  inference(subsumption_resolution,[],[f3763,f349])).
fof(f3763,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') != s(bool,X9) | s(bool,'F_0') != s(bool,X10) | p(s(bool,X11)) | sP7(X10,X9,X11)) ) | $spl10),
  inference(resolution,[],[f3059,f338])).
fof(f9785,plain,(
  ( ! [X22] : (p(s(bool,X22)) | s(bool,X22) != s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | $spl178),
  inference(cnf_transformation,[],[f9785_D])).
fof(f9785_D,plain,(
  ( ! [X22] : (p(s(bool,X22)) | s(bool,X22) != s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f25309,plain,(
  ( ! [X8,X9] : (p(s(bool,X8)) | p(s(bool,X9)) | s(bool,X8) != s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | ($spl10 | $spl149)),
  inference(resolution,[],[f9460,f4011])).
fof(f4011,plain,(
  ( ! [X4,X5,X3] : (~sP6(X4,X3,X5) | p(s(bool,X3)) | p(s(bool,X5)) | s(bool,X3) != s(bool,X4)) ) | $spl10),
  inference(duplicate_literal_removal,[],[f3999])).
fof(f3999,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~sP6(X4,X3,X5) | p(s(bool,X5)) | p(s(bool,X5)) | s(bool,X3) != s(bool,X4)) ) | $spl10),
  inference(resolution,[],[f3967,f359])).
fof(f3967,plain,(
  ( ! [X6,X8,X7] : (sP9(X7,X6,X8) | p(s(bool,X6)) | ~sP6(X8,X6,X7) | p(s(bool,X7))) ) | $spl10),
  inference(resolution,[],[f3947,f347])).
fof(f3947,plain,(
  ( ! [X6,X8,X7] : (~sP8(X7,X8,X6) | p(s(bool,X6)) | sP9(X7,X6,X8) | ~sP6(X8,X6,X7)) ) | $spl10),
  inference(subsumption_resolution,[],[f3945,f461])).
fof(f461,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f269])).
fof(f269,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f108])).
fof(f108,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X13] : (s(bool,'F_0') = s(bool,X13) <=> ~p(s(bool,X13)))),
  file('thm.quotient.QUOTIENT_SYM','thm.bool.EQ_CLAUSES|split|2')).
fof(f3945,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') != s(bool,X6) | sP9(X7,X6,X8) | p(s(bool,X6)) | ~sP8(X7,X8,X6) | ~sP6(X8,X6,X7)) ) | $spl10),
  inference(resolution,[],[f3942,f342])).
fof(f342,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | p(s(bool,X1)) | ~sP8(X0,X2,X1) | ~sP7(X1,X2,X0) | ~sP6(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f206])).
fof(f25277,plain,(
  ( ! [X17,X16] : (p(s(bool,X16)) | s(bool,'F_0') != s(bool,X17) | s(bool,X17) != s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | ($spl98 | $spl149)),
  inference(resolution,[],[f9459,f5593])).
fof(f5593,plain,(
  ( ! [X17,X15,X16] : (~sP6(X15,X17,X16) | p(s(bool,X16)) | s(bool,'F_0') != s(bool,X15) | s(bool,X15) != s(bool,X17)) ) | $spl98),
  inference(duplicate_literal_removal,[],[f5581])).
fof(f5581,plain,(
  ( ! [X17,X15,X16] : (s(bool,'F_0') != s(bool,X15) | p(s(bool,X16)) | ~sP6(X15,X17,X16) | p(s(bool,X16)) | s(bool,X15) != s(bool,X17)) ) | $spl98),
  inference(resolution,[],[f5558,f359])).
fof(f5558,plain,(
  ( ! [X17,X15,X16] : (sP9(X16,X17,X15) | s(bool,'F_0') != s(bool,X15) | p(s(bool,X16)) | ~sP6(X15,X17,X16)) ) | $spl98),
  inference(subsumption_resolution,[],[f5557,f347])).
fof(f5557,plain,(
  ( ! [X17,X15,X16] : (s(bool,'F_0') != s(bool,X15) | sP9(X16,X17,X15) | p(s(bool,X16)) | ~sP8(X16,X15,X17) | ~sP6(X15,X17,X16)) ) | $spl98),
  inference(resolution,[],[f5549,f343])).
fof(f5549,plain,(
  ( ! [X4,X5,X3] : (sP7(X5,X4,X3) | s(bool,'F_0') != s(bool,X4)) ) | $spl98),
  inference(subsumption_resolution,[],[f5548,f349])).
fof(f5548,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | s(bool,'F_0') != s(bool,X4) | sP7(X5,X4,X3)) ) | $spl98),
  inference(subsumption_resolution,[],[f5530,f351])).
fof(f5530,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | s(bool,'F_0') != s(bool,X4) | ~p(s(bool,X5)) | sP7(X5,X4,X3)) ) | $spl98),
  inference(resolution,[],[f5254,f338])).
fof(f9459,plain,(
  ( ! [X111,X110] : (sP6(X110,'const.combin.I_1'(s(bool,'F_0')),X111)) ) | $spl149),
  inference(resolution,[],[f8942,f354])).
fof(f354,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f214])).
fof(f25276,plain,(
  ( ! [X14,X15] : (p(s(bool,X14)) | s(bool,'T_0') = s(bool,X15) | s(bool,X15) != s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | ($spl94 | $spl149)),
  inference(resolution,[],[f9459,f5575])).
fof(f25275,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X13) != s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | ($spl70 | $spl149)),
  inference(resolution,[],[f9459,f5073])).
fof(f2556,plain,(
  ( ! [X8] : (~p(s(bool,X8))) ) | $spl12),
  inference(cnf_transformation,[],[f2556_D])).
fof(f2556_D,plain,(
  ( ! [X8] : (~p(s(bool,X8))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f25142,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | ~p(s(bool,X24))) ) | ($spl24 | $spl149)),
  inference(resolution,[],[f9454,f5693])).
fof(f5693,plain,(
  ( ! [X2,X0,X3,X1] : (~sP8(X2,X0,X3) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | ~p(s(bool,X0))) ) | $spl24),
  inference(subsumption_resolution,[],[f5692,f355])).
fof(f5692,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | ~sP8(X2,X0,X3) | ~sP6(X0,X3,X2)) ) | $spl24),
  inference(subsumption_resolution,[],[f5691,f349])).
fof(f5691,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | ~sP8(X2,X0,X3) | ~sP7(X3,X0,X2) | ~sP6(X0,X3,X2)) ) | $spl24),
  inference(duplicate_literal_removal,[],[f5650])).
fof(f5650,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | ~p(s(bool,X0)) | ~sP8(X2,X0,X3) | ~sP7(X3,X0,X2) | ~sP6(X0,X3,X2)) ) | $spl24),
  inference(resolution,[],[f4773,f341])).
fof(f4773,plain,(
  ( ! [X10,X8,X11,X9] : (~sP9(X11,X9,X8) | ~p(s(bool,X8)) | p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X9) != s(bool,X10)) ) | $spl24),
  inference(superposition,[],[f2665,f358])).
fof(f358,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X2)) | ~sP9(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f215])).
fof(f9454,plain,(
  ( ! [X101,X100] : (sP8(X100,X101,'const.combin.I_1'(s(bool,'F_0')))) ) | $spl149),
  inference(resolution,[],[f8942,f345])).
fof(f345,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f209])).
fof(f9784,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.combin.I_1'(s(bool,'F_0')),X11) | ~p(s(bool,X11))) ) | $spl176),
  inference(cnf_transformation,[],[f9784_D])).
fof(f9784_D,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.combin.I_1'(s(bool,'F_0')),X11) | ~p(s(bool,X11))) ) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f25141,plain,(
  ( ! [X21,X20] : (sP9(X20,'const.combin.I_1'(s(bool,'F_0')),X21) | ~p(s(bool,X21)) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | ($spl70 | $spl149)),
  inference(resolution,[],[f9454,f5039])).
fof(f5039,plain,(
  ( ! [X10,X11,X9] : (~sP8(X10,X11,X9) | sP9(X10,X9,X11) | ~p(s(bool,X11)) | s(bool,'T_0') = s(bool,X9)) ) | $spl70),
  inference(subsumption_resolution,[],[f5036,f355])).
fof(f5036,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') = s(bool,X9) | sP9(X10,X9,X11) | ~p(s(bool,X11)) | ~sP8(X10,X11,X9) | ~sP6(X11,X9,X10)) ) | $spl70),
  inference(resolution,[],[f5015,f341])).
fof(f8948,plain,(
  p(s(bool,'const.combin.I_1'(s(bool,'F_0')))) | $spl170),
  inference(cnf_transformation,[],[f8948_D])).
fof(f8948_D,plain,(
  p(s(bool,'const.combin.I_1'(s(bool,'F_0')))) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f24737,plain,(
  ( ! [X1] : (p(s(bool,'const.combin.I_1'(s(bool,'F_0')))) | p(s(bool,X1))) ) | $spl149),
  inference(subsumption_resolution,[],[f24713,f436])).
fof(f436,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f254])).
fof(f254,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f253])).
fof(f253,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f100])).
fof(f100,plain,(
  ! [X0] : ((p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X13] : ((p(s(bool,'F_0')) | p(s(bool,X13))) <=> p(s(bool,X13)))),
  file('thm.quotient.QUOTIENT_SYM','thm.bool.OR_CLAUSES|split|2')).
fof(f24713,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.combin.I_1'(s(bool,'F_0')))) | p(s(bool,X1))) ) | $spl149),
  inference(resolution,[],[f9472,f390])).
fof(f390,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP11(X0,X1)) )),
  inference(cnf_transformation,[],[f229])).
fof(f229,plain,(
  ! [X0,X1] : ((~sP11(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP11(X0,X1)))),
  inference(flattening,[],[f228])).
fof(f228,plain,(
  ! [X0,X1] : ((~sP11(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP11(X0,X1)))),
  inference(nnf_transformation,[],[f173])).
fof(f173,plain,(
  ! [X0,X1] : (sP11(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f9472,plain,(
  ( ! [X127] : (sP11('const.combin.I_1'(s(bool,'F_0')),X127)) ) | $spl149),
  inference(resolution,[],[f8942,f392])).
fof(f392,plain,(
  ( ! [X0,X1] : (sP11(X0,X1) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f229])).
fof(f24736,plain,(
  ( ! [X0] : (p(s(bool,'const.combin.I_1'(s(bool,'F_0')))) | p(s(bool,X0))) ) | $spl149),
  inference(subsumption_resolution,[],[f24712,f436])).
fof(f24712,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.combin.I_1'(s(bool,'F_0')))) | p(s(bool,X0))) ) | $spl149),
  inference(resolution,[],[f9472,f394])).
fof(f394,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP11(X0,X1)) )),
  inference(cnf_transformation,[],[f231])).
fof(f231,plain,(
  ! [X0,X1] : ((~sP11(X0,X1) | (p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP11(X0,X1)))),
  inference(flattening,[],[f230])).
fof(f230,plain,(
  ! [X0,X1] : ((~sP11(X0,X1) | ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0')))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP11(X0,X1)))),
  inference(nnf_transformation,[],[f174])).
fof(f174,plain,(
  ! [X0,X1] : (sP11(X0,X1) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(definition_folding,[],[f140,f173])).
fof(f140,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f139])).
fof(f139,plain,(
  ! [X0,X1] : (((p(s(bool,X1)) | p(s(bool,X0))) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | (p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(ennf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(flattening,[],[f81])).
fof(f81,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X6,X7] : ((~(p(s(bool,X7)) | p(s(bool,X6))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X7)) => p(s(bool,'F_0'))) => (~p(s(bool,X6)) => p(s(bool,'F_0')))))),
  file('thm.quotient.QUOTIENT_SYM','thm.sat.sat_9_prove_ltb')).
fof(f24686,plain,(
  ( ! [X18] : (~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,X18))),s(sK13,sK20(happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,X18)),sK13)))))) ) | $spl246),
  inference(cnf_transformation,[],[f24686_D])).
fof(f24686_D,plain,(
  ( ! [X18] : (~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,X18))),s(sK13,sK20(happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,X18)),sK13)))))) ) <=> ~$spl246),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl246])])).
fof(f24685,plain,(
  ( ! [X15] : (p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,X15))),s(sK13,X15))))) ) | $spl244),
  inference(cnf_transformation,[],[f24685_D])).
fof(f24685_D,plain,(
  ( ! [X15] : (p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,X15))),s(sK13,X15))))) ) <=> ~$spl244),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl244])])).
fof(f24560,plain,(
  ( ! [X17,X18] : (p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,X17))),s(sK13,X17)))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,X18))),s(sK13,sK20(happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,X18)),sK13)))))) )),
  inference(resolution,[],[f496,f284])).
fof(f284,plain,(
  ( ! [X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK20(X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f180])).
fof(f180,plain,(
  ! [X0,X1] : ((! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK19(X1,X0)))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK20(X1,X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK19,sK20])],[f179])).
fof(f179,plain,(
  ! [X0,X1] : ((! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(rectify,[],[f178])).
fof(f178,plain,(
  ! [X0,X1] : ((! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(nnf_transformation,[],[f133])).
fof(f133,plain,(
  ! [X0,X1] : (? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(ennf_transformation,[],[f61])).
fof(f61,plain,(
  ! [X0,X1] : (~! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(flattening,[],[f60])).
fof(f60,plain,(
  ! [X0] : ! [X1] : (~! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X7] : ! [X16] : (~! [X11] : p(s(bool,happ(s(fun(X7,bool),X16),s(X7,X11)))) <=> ? [X11] : ~p(s(bool,happ(s(fun(X7,bool),X16),s(X7,X11)))))),
  file('thm.quotient.QUOTIENT_SYM','thm.bool.NOT_FORALL_THM')).
fof(f496,plain,(
  ( ! [X2,X3] : (~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,X3))),s(sK13,X2)))) | p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,X2))),s(sK13,X2))))) )),
  inference(resolution,[],[f276,f280])).
fof(f280,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X6))),s(X1,X6)))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X6)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,X0),X3),s(fun(X0,X1),X4))))) )),
  inference(cnf_transformation,[],[f177])).
fof(f24684,plain,(
  ( ! [X16] : (~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,X16))),s(sK13,sK19(happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,X16)),sK13)))))) ) | $spl242),
  inference(cnf_transformation,[],[f24684_D])).
fof(f24684_D,plain,(
  ( ! [X16] : (~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,X16))),s(sK13,sK19(happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,X16)),sK13)))))) ) <=> ~$spl242),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl242])])).
fof(f24559,plain,(
  ( ! [X15,X16] : (p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,X15))),s(sK13,X15)))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,X16))),s(sK13,sK19(happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,X16)),sK13)))))) )),
  inference(resolution,[],[f496,f283])).
fof(f283,plain,(
  ( ! [X0,X5,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK19(X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))) )),
  inference(cnf_transformation,[],[f180])).
fof(f24553,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.combin.I_1'(s(bool,'F_0'))))) ) | $spl149),
  inference(subsumption_resolution,[],[f24529,f436])).
fof(f24529,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.combin.I_1'(s(bool,'F_0'))))) ) | $spl149),
  inference(resolution,[],[f9471,f390])).
fof(f9471,plain,(
  ( ! [X126] : (sP11(X126,'const.combin.I_1'(s(bool,'F_0')))) ) | $spl149),
  inference(resolution,[],[f8942,f391])).
fof(f391,plain,(
  ( ! [X0,X1] : (sP11(X0,X1) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f229])).
fof(f24552,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.combin.I_1'(s(bool,'F_0'))))) ) | $spl149),
  inference(subsumption_resolution,[],[f24528,f436])).
fof(f24528,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.combin.I_1'(s(bool,'F_0'))))) ) | $spl149),
  inference(resolution,[],[f9471,f394])).
fof(f24521,plain,(
  ( ! [X59,X57] : (p(s(bool,X57)) | s(bool,'T_0') != s(bool,X59) | s(bool,X59) != s(bool,'const.combin.I_1'(s(bool,X57)))) ) | $spl240),
  inference(cnf_transformation,[],[f24521_D])).
fof(f24521_D,plain,(
  ( ! [X59,X57] : (p(s(bool,X57)) | s(bool,'T_0') != s(bool,X59) | s(bool,X59) != s(bool,'const.combin.I_1'(s(bool,X57)))) ) <=> ~$spl240),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl240])])).
fof(f24477,plain,(
  ( ! [X59,X57,X58] : (p(s(bool,X57)) | ~p(s(bool,X58)) | s(bool,X59) != s(bool,'const.combin.I_1'(s(bool,X57))) | s(bool,'T_0') != s(bool,X59)) ) | ($spl10 | $spl148)),
  inference(resolution,[],[f17923,f3753])).
fof(f3753,plain,(
  ( ! [X12,X10,X11] : (~sP8(X11,X12,'F_0') | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,'T_0') != s(bool,X10)) ) | $spl10),
  inference(subsumption_resolution,[],[f3743,f3041])).
fof(f3041,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X5) | s(bool,X4) != s(bool,X5) | p(s(bool,X4))) ) | $spl10),
  inference(superposition,[],[f1974,f463])).
fof(f3743,plain,(
  ( ! [X12,X10,X11] : (s(bool,'T_0') != s(bool,X10) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | ~p(s(bool,X12)) | ~sP8(X11,X12,'F_0')) ) | $spl10),
  inference(resolution,[],[f3046,f2419])).
fof(f2419,plain,(
  ( ! [X2,X3] : (sP9(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP8(X3,X2,'F_0')) )),
  inference(subsumption_resolution,[],[f2418,f465])).
fof(f2418,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | sP9(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP8(X3,X2,'F_0')) )),
  inference(subsumption_resolution,[],[f2415,f2405])).
fof(f2405,plain,(
  ( ! [X4,X5] : (sP6(X4,'F_0',X5) | s(bool,'T_0') != s(bool,X4)) )),
  inference(subsumption_resolution,[],[f2400,f353])).
fof(f2400,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | p(s(bool,X5)) | sP6(X4,'F_0',X5)) )),
  inference(resolution,[],[f1949,f337])).
fof(f1949,plain,(
  ( ! [X6,X7] : (sP9(X7,'F_0',X6) | s(bool,'T_0') != s(bool,X6) | p(s(bool,X7))) )),
  inference(superposition,[],[f1940,f356])).
fof(f1940,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(subsumption_resolution,[],[f1928,f276])).
fof(f1928,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK13,fun(sK13,bool)),sK14),s(fun(sK13,sK12),sK15),s(fun(sK12,sK13),sK16))))),
  inference(superposition,[],[f579,f465])).
fof(f579,plain,(
  s(bool,'F_0') != s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK13,fun(sK13,bool)),sK14),s(fun(sK13,sK12),sK15),s(fun(sK12,sK13),sK16)))),
  inference(resolution,[],[f276,f460])).
fof(f460,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f269])).
fof(f2415,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | sP9(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP8(X3,X2,'F_0') | ~sP6(X2,'F_0',X3)) )),
  inference(resolution,[],[f2406,f341])).
fof(f2406,plain,(
  ( ! [X6,X7] : (sP7('F_0',X6,X7) | s(bool,'T_0') != s(bool,X6)) )),
  inference(subsumption_resolution,[],[f2401,f349])).
fof(f2401,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X6) | p(s(bool,X7)) | sP7('F_0',X6,X7)) )),
  inference(resolution,[],[f1949,f338])).
fof(f3046,plain,(
  ( ! [X19,X17,X18] : (~sP9(X19,'F_0',X17) | s(bool,'T_0') != s(bool,X18) | ~p(s(bool,X19)) | s(bool,X17) != s(bool,X18)) ) | $spl10),
  inference(superposition,[],[f1974,f358])).
fof(f17923,plain,(
  ( ! [X6,X8,X7] : (sP8(X7,'const.combin.I_1'(s(bool,X6)),X8) | p(s(bool,X6))) ) | $spl148),
  inference(superposition,[],[f9313,f463])).
fof(f9313,plain,(
  ( ! [X200,X201] : (sP8(X200,'const.combin.I_1'(s(bool,'F_0')),X201)) ) | $spl148),
  inference(subsumption_resolution,[],[f9083,f478])).
fof(f478,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f125])).
fof(f125,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.quotient.QUOTIENT_SYM',aHL_FALSITY)).
fof(f9083,plain,(
  ( ! [X200,X201] : (p(s(bool,'F_0')) | sP8(X200,'const.combin.I_1'(s(bool,'F_0')),X201)) ) | $spl148),
  inference(superposition,[],[f346,f8862])).
fof(f8862,plain,(
  s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | $spl148),
  inference(cnf_transformation,[],[f8862_D])).
fof(f8862_D,plain,(
  s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,'F_0'))) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f346,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f209])).
fof(f8855,plain,(
  ( ! [X12] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,X12))) | p(s(bool,X12))) ) | $spl142),
  inference(cnf_transformation,[],[f8855_D])).
fof(f8855_D,plain,(
  ( ! [X12] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,X12))) | p(s(bool,X12))) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f24475,plain,(
  ( ! [X52,X53] : (p(s(bool,X52)) | s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,X52))) | ~p(s(bool,X53))) ) | $spl148),
  inference(resolution,[],[f17923,f2549])).
fof(f2549,plain,(
  ( ! [X14,X15] : (~sP8(X14,X15,'F_0') | s(bool,'T_0') != s(bool,X15) | ~p(s(bool,X14))) )),
  inference(subsumption_resolution,[],[f2548,f353])).
fof(f2548,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | ~sP8(X14,X15,'F_0') | ~sP6(X15,'F_0',X14)) )),
  inference(subsumption_resolution,[],[f2547,f416])).
fof(f416,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f240])).
fof(f240,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f239])).
fof(f239,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f143])).
fof(f143,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f93])).
fof(f93,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f92])).
fof(f92,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X13] : ((p(s(bool,X13)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X13)))),
  file('thm.quotient.QUOTIENT_SYM','thm.bool.IMP_CLAUSES|split|4')).
fof(f2547,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | p(s(bool,'F_0')) | ~sP8(X14,X15,'F_0') | ~sP6(X15,'F_0',X14)) )),
  inference(subsumption_resolution,[],[f2540,f351])).
fof(f2540,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | p(s(bool,'F_0')) | ~sP8(X14,X15,'F_0') | ~sP7('F_0',X15,X14) | ~sP6(X15,'F_0',X14)) )),
  inference(resolution,[],[f1951,f342])).
fof(f1951,plain,(
  ( ! [X10,X11] : (~sP9(X11,'F_0',X10) | ~p(s(bool,X11)) | s(bool,'T_0') != s(bool,X10)) )),
  inference(superposition,[],[f1940,f358])).
fof(f24520,plain,(
  ( ! [X43,X40] : (p(s(bool,X40)) | ~p(s(bool,X43)) | s(bool,X43) != s(bool,'const.combin.I_1'(s(bool,X40)))) ) | $spl238),
  inference(cnf_transformation,[],[f24520_D])).
fof(f24520_D,plain,(
  ( ! [X43,X40] : (p(s(bool,X40)) | ~p(s(bool,X43)) | s(bool,X43) != s(bool,'const.combin.I_1'(s(bool,X40)))) ) <=> ~$spl238),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl238])])).
fof(f24470,plain,(
  ( ! [X43,X41,X42,X40] : (p(s(bool,X40)) | p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X43) != s(bool,'const.combin.I_1'(s(bool,X40))) | ~p(s(bool,X43))) ) | ($spl24 | $spl148)),
  inference(resolution,[],[f17923,f6924])).
fof(f6924,plain,(
  ( ! [X6,X4,X7,X5] : (~sP8(X6,X7,X5) | p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X4) != s(bool,X7) | ~p(s(bool,X4))) ) | $spl24),
  inference(subsumption_resolution,[],[f6923,f353])).
fof(f6923,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,X4)) | p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X4) != s(bool,X7) | ~sP8(X6,X7,X5) | ~sP6(X7,X5,X6)) ) | $spl24),
  inference(subsumption_resolution,[],[f6919,f349])).
fof(f6919,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,X4)) | p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X4) != s(bool,X7) | ~sP8(X6,X7,X5) | ~sP7(X5,X7,X6) | ~sP6(X7,X5,X6)) ) | $spl24),
  inference(duplicate_literal_removal,[],[f6870])).
fof(f6870,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,X4)) | p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X4) != s(bool,X7) | p(s(bool,X5)) | ~sP8(X6,X7,X5) | ~sP7(X5,X7,X6) | ~sP6(X7,X5,X6)) ) | $spl24),
  inference(resolution,[],[f4802,f342])).
fof(f4802,plain,(
  ( ! [X14,X12,X15,X13] : (~sP9(X15,X12,X13) | ~p(s(bool,X14)) | p(s(bool,X12)) | ~p(s(bool,X15)) | s(bool,X13) != s(bool,X14)) ) | $spl24),
  inference(superposition,[],[f2665,f358])).
fof(f4444,plain,(
  ( ! [X21] : (s(bool,'F_0') != s(bool,X21)) ) | $spl50),
  inference(cnf_transformation,[],[f4444_D])).
fof(f4444_D,plain,(
  ( ! [X21] : (s(bool,'F_0') != s(bool,X21)) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f24466,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,'F_0') != s(bool,X29) | s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,X27)))) ) | ($spl10 | $spl148)),
  inference(resolution,[],[f17923,f4242])).
fof(f4242,plain,(
  ( ! [X4,X5,X3] : (~sP8(X4,X3,X5) | ~p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5) | s(bool,'T_0') != s(bool,X3)) ) | $spl10),
  inference(subsumption_resolution,[],[f4241,f460])).
fof(f4241,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5) | p(s(bool,X5)) | ~sP8(X4,X3,X5)) ) | $spl10),
  inference(subsumption_resolution,[],[f4240,f3776])).
fof(f4240,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5) | p(s(bool,X5)) | ~sP8(X4,X3,X5) | ~sP6(X3,X5,X4)) ) | $spl10),
  inference(subsumption_resolution,[],[f4207,f3777])).
fof(f4207,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5) | p(s(bool,X5)) | ~sP8(X4,X3,X5) | ~sP7(X5,X3,X4) | ~sP6(X3,X5,X4)) ) | $spl10),
  inference(resolution,[],[f3061,f342])).
fof(f3061,plain,(
  ( ! [X6,X8,X7] : (~sP9(X8,X7,X6) | s(bool,'T_0') != s(bool,X6) | ~p(s(bool,X8)) | s(bool,'F_0') != s(bool,X7)) ) | $spl10),
  inference(superposition,[],[f1974,f358])).
fof(f23807,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK20(happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18)),sK13)))),s(sK13,sK20(happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18)),sK13))))) | $spl235),
  inference(cnf_transformation,[],[f23807_D])).
fof(f23807_D,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK20(happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18)),sK13)))),s(sK13,sK20(happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18)),sK13))))) <=> ~$spl235),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl235])])).
fof(f23808,plain,(
  ( ! [X2,X0,X1] : (s(X0,happ(s(fun(sK13,X0),X1),s(sK13,sK18))) != s(X0,happ(s(fun(sK13,X0),X1),s(sK13,sK20(happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18)),sK13)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK13,fun(sK13,bool)),sK14),s(fun(sK13,X0),X1),s(fun(X0,sK13),X2))))) ) | $spl236),
  inference(cnf_transformation,[],[f23808_D])).
fof(f23808_D,plain,(
  ( ! [X2,X0,X1] : (s(X0,happ(s(fun(sK13,X0),X1),s(sK13,sK18))) != s(X0,happ(s(fun(sK13,X0),X1),s(sK13,sK20(happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18)),sK13)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK13,fun(sK13,bool)),sK14),s(fun(sK13,X0),X1),s(fun(X0,sK13),X2))))) ) <=> ~$spl236),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl236])])).
fof(f23805,plain,(
  ( ! [X2,X0,X1] : (s(X0,happ(s(fun(sK13,X0),X1),s(sK13,sK18))) != s(X0,happ(s(fun(sK13,X0),X1),s(sK13,sK20(happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18)),sK13)))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK20(happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18)),sK13)))),s(sK13,sK20(happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18)),sK13))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK13,fun(sK13,bool)),sK14),s(fun(sK13,X0),X1),s(fun(X0,sK13),X2))))) ) | $spl4),
  inference(subsumption_resolution,[],[f23804,f491])).
fof(f491,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f409])).
fof(f409,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f236])).
fof(f236,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f87])).
fof(f87,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f86])).
fof(f86,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X7] : ! [X11] : (s(X7,X11) = s(X7,X11) <=> p(s(bool,'T_0')))),
  file('thm.quotient.QUOTIENT_SYM','thm.bool.REFL_CLAUSE')).
fof(f23804,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s(X0,happ(s(fun(sK13,X0),X1),s(sK13,sK18))) != s(X0,happ(s(fun(sK13,X0),X1),s(sK13,sK20(happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18)),sK13)))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK20(happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18)),sK13)))),s(sK13,sK20(happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18)),sK13))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK13,fun(sK13,bool)),sK14),s(fun(sK13,X0),X1),s(fun(X0,sK13),X2))))) ) | $spl4),
  inference(forward_demodulation,[],[f23688,f1460])).
fof(f1460,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK18))) | $spl4),
  inference(resolution,[],[f735,f465])).
fof(f23688,plain,(
  ( ! [X2,X0,X1] : (s(X0,happ(s(fun(sK13,X0),X1),s(sK13,sK18))) != s(X0,happ(s(fun(sK13,X0),X1),s(sK13,sK20(happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18)),sK13)))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK20(happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18)),sK13)))),s(sK13,sK20(happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18)),sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK18)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK13,fun(sK13,bool)),sK14),s(fun(sK13,X0),X1),s(fun(X0,sK13),X2))))) )),
  inference(resolution,[],[f775,f282])).
fof(f282,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X6)))) | s(X0,happ(s(fun(X1,X0),X3),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X6))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X6))),s(X1,X6)))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X5)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,X0),X3),s(fun(X0,X1),X4))))) )),
  inference(cnf_transformation,[],[f177])).
fof(f775,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK20(happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18)),sK13)))))),
  inference(resolution,[],[f278,f284])).
fof(f278,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK17))))),
  inference(cnf_transformation,[],[f175])).
fof(f23678,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK19(happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18)),sK13)))),s(sK13,sK19(happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18)),sK13))))) | $spl231),
  inference(cnf_transformation,[],[f23678_D])).
fof(f23678_D,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK19(happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18)),sK13)))),s(sK13,sK19(happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18)),sK13))))) <=> ~$spl231),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl231])])).
fof(f23679,plain,(
  ( ! [X2,X0,X1] : (s(X0,happ(s(fun(sK13,X0),X1),s(sK13,sK18))) != s(X0,happ(s(fun(sK13,X0),X1),s(sK13,sK19(happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18)),sK13)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK13,fun(sK13,bool)),sK14),s(fun(sK13,X0),X1),s(fun(X0,sK13),X2))))) ) | $spl232),
  inference(cnf_transformation,[],[f23679_D])).
fof(f23679_D,plain,(
  ( ! [X2,X0,X1] : (s(X0,happ(s(fun(sK13,X0),X1),s(sK13,sK18))) != s(X0,happ(s(fun(sK13,X0),X1),s(sK13,sK19(happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18)),sK13)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK13,fun(sK13,bool)),sK14),s(fun(sK13,X0),X1),s(fun(X0,sK13),X2))))) ) <=> ~$spl232),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl232])])).
fof(f23676,plain,(
  ( ! [X2,X0,X1] : (s(X0,happ(s(fun(sK13,X0),X1),s(sK13,sK18))) != s(X0,happ(s(fun(sK13,X0),X1),s(sK13,sK19(happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18)),sK13)))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK19(happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18)),sK13)))),s(sK13,sK19(happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18)),sK13))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK13,fun(sK13,bool)),sK14),s(fun(sK13,X0),X1),s(fun(X0,sK13),X2))))) ) | $spl4),
  inference(subsumption_resolution,[],[f23675,f491])).
fof(f23675,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s(X0,happ(s(fun(sK13,X0),X1),s(sK13,sK18))) != s(X0,happ(s(fun(sK13,X0),X1),s(sK13,sK19(happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18)),sK13)))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK19(happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18)),sK13)))),s(sK13,sK19(happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18)),sK13))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK13,fun(sK13,bool)),sK14),s(fun(sK13,X0),X1),s(fun(X0,sK13),X2))))) ) | $spl4),
  inference(forward_demodulation,[],[f23559,f1460])).
fof(f23559,plain,(
  ( ! [X2,X0,X1] : (s(X0,happ(s(fun(sK13,X0),X1),s(sK13,sK18))) != s(X0,happ(s(fun(sK13,X0),X1),s(sK13,sK19(happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18)),sK13)))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK19(happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18)),sK13)))),s(sK13,sK19(happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18)),sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK18)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK13,fun(sK13,bool)),sK14),s(fun(sK13,X0),X1),s(fun(X0,sK13),X2))))) )),
  inference(resolution,[],[f774,f282])).
fof(f774,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK19(happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18)),sK13)))))),
  inference(resolution,[],[f278,f283])).
fof(f21931,plain,(
  $false | ($spl150 | $spl212)),
  inference(trivial_inequality_removal,[],[f21896])).
fof(f21896,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl150 | $spl212)),
  inference(superposition,[],[f13519,f8866])).
fof(f21930,plain,(
  $false | $spl212),
  inference(equality_resolution,[],[f13519])).
fof(f21827,plain,(
  $false | ($spl148 | $spl190)),
  inference(trivial_inequality_removal,[],[f21788])).
fof(f21788,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl148 | $spl190)),
  inference(superposition,[],[f10840,f8862])).
fof(f21826,plain,(
  $false | $spl190),
  inference(equality_resolution,[],[f10840])).
fof(f16264,plain,(
  ( ! [X316,X317] : (~p(s(bool,X316)) | p(s(bool,X317)) | sP9(X317,X316,'const.combin.I_1'(s(bool,'T_0')))) ) | $spl226),
  inference(cnf_transformation,[],[f16264_D])).
fof(f16264_D,plain,(
  ( ! [X316,X317] : (~p(s(bool,X316)) | p(s(bool,X317)) | sP9(X317,X316,'const.combin.I_1'(s(bool,'T_0')))) ) <=> ~$spl226),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl226])])).
fof(f20939,plain,(
  ( ! [X315,X316,X317] : (s(bool,'F_0') != s(bool,X315) | ~p(s(bool,X316)) | sP9(X317,X316,'const.combin.I_1'(s(bool,'T_0'))) | p(s(bool,X317))) ) | ($spl24 | $spl150)),
  inference(subsumption_resolution,[],[f20767,f460])).
fof(f20767,plain,(
  ( ! [X315,X316,X317] : (s(bool,'F_0') != s(bool,X315) | ~p(s(bool,X316)) | p(s(bool,X315)) | sP9(X317,X316,'const.combin.I_1'(s(bool,'T_0'))) | p(s(bool,X317))) ) | ($spl24 | $spl150)),
  inference(superposition,[],[f4772,f8866])).
fof(f4772,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X5) != s(bool,X6) | ~p(s(bool,X4)) | p(s(bool,X6)) | sP9(X7,X4,X5) | p(s(bool,X7))) ) | $spl24),
  inference(superposition,[],[f2665,f356])).
fof(f16262,plain,(
  ( ! [X310,X311] : (~p(s(bool,X310)) | p(s(bool,X311)) | sP9(X311,'const.combin.I_1'(s(bool,'T_0')),X310)) ) | $spl224),
  inference(cnf_transformation,[],[f16262_D])).
fof(f16262_D,plain,(
  ( ! [X310,X311] : (~p(s(bool,X310)) | p(s(bool,X311)) | sP9(X311,'const.combin.I_1'(s(bool,'T_0')),X310)) ) <=> ~$spl224),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl224])])).
fof(f20938,plain,(
  ( ! [X310,X309,X311] : (s(bool,'F_0') != s(bool,X309) | ~p(s(bool,X310)) | sP9(X311,'const.combin.I_1'(s(bool,'T_0')),X310) | p(s(bool,X311))) ) | ($spl24 | $spl150)),
  inference(subsumption_resolution,[],[f20765,f460])).
fof(f20765,plain,(
  ( ! [X310,X309,X311] : (s(bool,'F_0') != s(bool,X309) | ~p(s(bool,X310)) | p(s(bool,X309)) | sP9(X311,'const.combin.I_1'(s(bool,'T_0')),X310) | p(s(bool,X311))) ) | ($spl24 | $spl150)),
  inference(superposition,[],[f4771,f8866])).
fof(f4771,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X1) != s(bool,X2) | ~p(s(bool,X0)) | p(s(bool,X2)) | sP9(X3,X1,X0) | p(s(bool,X3))) ) | $spl24),
  inference(superposition,[],[f2665,f356])).
fof(f8876,plain,(
  s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | $spl161),
  inference(cnf_transformation,[],[f8876_D])).
fof(f8876_D,plain,(
  s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'T_0'))) <=> ~$spl161),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl161])])).
fof(f20480,plain,(
  ( ! [X18] : (p(s(bool,X18)) | s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl122 | $spl199)),
  inference(resolution,[],[f12706,f6702])).
fof(f6702,plain,(
  ( ! [X10,X11] : (~sP7(X11,'T_0',X10) | p(s(bool,X10)) | s(bool,'T_0') != s(bool,X11)) ) | $spl122),
  inference(duplicate_literal_removal,[],[f6690])).
fof(f6690,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | ~sP7(X11,'T_0',X10) | p(s(bool,X10)) | s(bool,'T_0') != s(bool,X11)) ) | $spl122),
  inference(resolution,[],[f6673,f359])).
fof(f6673,plain,(
  ( ! [X4,X5] : (sP9(X4,X5,'T_0') | p(s(bool,X4)) | ~sP7(X5,'T_0',X4)) ) | $spl122),
  inference(subsumption_resolution,[],[f6663,f347])).
fof(f6663,plain,(
  ( ! [X4,X5] : (sP9(X4,X5,'T_0') | p(s(bool,X4)) | ~sP8(X4,'T_0',X5) | ~sP7(X5,'T_0',X4)) ) | $spl122),
  inference(resolution,[],[f6659,f343])).
fof(f6659,plain,(
  ( ! [X0,X1] : (sP6('T_0',X1,X0)) ) | $spl122),
  inference(subsumption_resolution,[],[f6658,f353])).
fof(f6658,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | sP6('T_0',X1,X0)) ) | $spl122),
  inference(subsumption_resolution,[],[f6640,f354])).
fof(f6640,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | sP6('T_0',X1,X0)) ) | $spl122),
  inference(resolution,[],[f6617,f337])).
fof(f12706,plain,(
  ( ! [X107,X106] : (sP7('const.combin.I_1'(s(bool,'T_0')),X106,X107)) ) | $spl199),
  inference(resolution,[],[f11984,f351])).
fof(f11984,plain,(
  ~p(s(bool,'const.combin.I_1'(s(bool,'T_0')))) | $spl199),
  inference(cnf_transformation,[],[f11984_D])).
fof(f11984_D,plain,(
  ~p(s(bool,'const.combin.I_1'(s(bool,'T_0')))) <=> ~$spl199),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl199])])).
fof(f13936,plain,(
  ( ! [X13] : (s(bool,'F_0') = s(bool,X13) | s(bool,X13) != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | $spl218),
  inference(cnf_transformation,[],[f13936_D])).
fof(f13936_D,plain,(
  ( ! [X13] : (s(bool,'F_0') = s(bool,X13) | s(bool,X13) != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) <=> ~$spl218),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl218])])).
fof(f20479,plain,(
  ( ! [X17,X16] : (p(s(bool,X16)) | s(bool,'F_0') = s(bool,X17) | s(bool,X17) != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl120 | $spl199)),
  inference(resolution,[],[f12706,f6791])).
fof(f6791,plain,(
  ( ! [X17,X15,X16] : (~sP7(X17,X15,X16) | p(s(bool,X16)) | s(bool,'F_0') = s(bool,X15) | s(bool,X15) != s(bool,X17)) ) | $spl120),
  inference(duplicate_literal_removal,[],[f6778])).
fof(f6778,plain,(
  ( ! [X17,X15,X16] : (s(bool,'F_0') = s(bool,X15) | p(s(bool,X16)) | ~sP7(X17,X15,X16) | p(s(bool,X16)) | s(bool,X15) != s(bool,X17)) ) | $spl120),
  inference(resolution,[],[f6760,f359])).
fof(f6760,plain,(
  ( ! [X6,X8,X7] : (sP9(X7,X8,X6) | s(bool,'F_0') = s(bool,X6) | p(s(bool,X7)) | ~sP7(X8,X6,X7)) ) | $spl120),
  inference(subsumption_resolution,[],[f6749,f347])).
fof(f6749,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') = s(bool,X6) | sP9(X7,X8,X6) | p(s(bool,X7)) | ~sP8(X7,X6,X8) | ~sP7(X8,X6,X7)) ) | $spl120),
  inference(resolution,[],[f6746,f343])).
fof(f6746,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X2,X1) | s(bool,'F_0') = s(bool,X0)) ) | $spl120),
  inference(subsumption_resolution,[],[f6745,f353])).
fof(f6745,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X1)) | sP6(X0,X2,X1)) ) | $spl120),
  inference(subsumption_resolution,[],[f6726,f354])).
fof(f6726,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | sP6(X0,X2,X1)) ) | $spl120),
  inference(resolution,[],[f6614,f337])).
fof(f13519,plain,(
  ( ! [X11] : (s(bool,X11) != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | $spl212),
  inference(cnf_transformation,[],[f13519_D])).
fof(f13519_D,plain,(
  ( ! [X11] : (s(bool,X11) != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) <=> ~$spl212),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl212])])).
fof(f20478,plain,(
  ( ! [X14,X15] : (p(s(bool,X14)) | s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X15) != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl110 | $spl199)),
  inference(resolution,[],[f12706,f6440])).
fof(f6440,plain,(
  ( ! [X17,X15,X16] : (~sP7(X15,X17,X16) | p(s(bool,X16)) | s(bool,'T_0') != s(bool,X15) | s(bool,X15) != s(bool,X17)) ) | $spl110),
  inference(duplicate_literal_removal,[],[f6425])).
fof(f6425,plain,(
  ( ! [X17,X15,X16] : (s(bool,'T_0') != s(bool,X15) | p(s(bool,X16)) | ~sP7(X15,X17,X16) | p(s(bool,X16)) | s(bool,X15) != s(bool,X17)) ) | $spl110),
  inference(resolution,[],[f6003,f359])).
fof(f6003,plain,(
  ( ! [X6,X8,X7] : (sP9(X7,X6,X8) | s(bool,'T_0') != s(bool,X6) | p(s(bool,X7)) | ~sP7(X6,X8,X7)) ) | $spl110),
  inference(subsumption_resolution,[],[f5993,f347])).
fof(f5993,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X6) | sP9(X7,X6,X8) | p(s(bool,X7)) | ~sP8(X7,X8,X6) | ~sP7(X6,X8,X7)) ) | $spl110),
  inference(resolution,[],[f5990,f343])).
fof(f5990,plain,(
  ( ! [X2,X0,X1] : (sP6(X2,X1,X0) | s(bool,'T_0') != s(bool,X1)) ) | $spl110),
  inference(subsumption_resolution,[],[f5989,f353])).
fof(f5989,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X1) | sP6(X2,X1,X0)) ) | $spl110),
  inference(subsumption_resolution,[],[f5967,f355])).
fof(f5967,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X1) | p(s(bool,X2)) | sP6(X2,X1,X0)) ) | $spl110),
  inference(resolution,[],[f5846,f337])).
fof(f8866,plain,(
  s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | $spl150),
  inference(cnf_transformation,[],[f8866_D])).
fof(f8866_D,plain,(
  s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,'T_0'))) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f20477,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X13) != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl102 | $spl199)),
  inference(resolution,[],[f12706,f6022])).
fof(f6022,plain,(
  ( ! [X17,X15,X16] : (~sP7(X15,X17,X16) | p(s(bool,X16)) | s(bool,'F_0') = s(bool,X15) | s(bool,X15) != s(bool,X17)) ) | $spl102),
  inference(duplicate_literal_removal,[],[f6009])).
fof(f6009,plain,(
  ( ! [X17,X15,X16] : (s(bool,'F_0') = s(bool,X15) | p(s(bool,X16)) | ~sP7(X15,X17,X16) | p(s(bool,X16)) | s(bool,X15) != s(bool,X17)) ) | $spl102),
  inference(resolution,[],[f5965,f359])).
fof(f5965,plain,(
  ( ! [X6,X8,X7] : (sP9(X7,X6,X8) | s(bool,'F_0') = s(bool,X6) | p(s(bool,X7)) | ~sP7(X6,X8,X7)) ) | $spl102),
  inference(subsumption_resolution,[],[f5955,f347])).
fof(f5955,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') = s(bool,X6) | sP9(X7,X6,X8) | p(s(bool,X7)) | ~sP8(X7,X8,X6) | ~sP7(X6,X8,X7)) ) | $spl102),
  inference(resolution,[],[f5952,f343])).
fof(f5952,plain,(
  ( ! [X2,X0,X1] : (sP6(X2,X0,X1) | s(bool,'F_0') = s(bool,X0)) ) | $spl102),
  inference(subsumption_resolution,[],[f5951,f353])).
fof(f5951,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X1)) | sP6(X2,X0,X1)) ) | $spl102),
  inference(subsumption_resolution,[],[f5931,f355])).
fof(f5931,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | sP6(X2,X0,X1)) ) | $spl102),
  inference(resolution,[],[f5834,f337])).
fof(f13118,plain,(
  ( ! [X35] : (s(bool,'T_0') != s(bool,X35) | s(bool,X35) != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | $spl210),
  inference(cnf_transformation,[],[f13118_D])).
fof(f13118_D,plain,(
  ( ! [X35] : (s(bool,'T_0') != s(bool,X35) | s(bool,X35) != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f20476,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | s(bool,'T_0') != s(bool,X11) | s(bool,X11) != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl10 | $spl199)),
  inference(resolution,[],[f12706,f4084])).
fof(f4084,plain,(
  ( ! [X4,X5,X3] : (~sP7(X5,X3,X4) | p(s(bool,X4)) | s(bool,'T_0') != s(bool,X3) | s(bool,X3) != s(bool,X5)) ) | $spl10),
  inference(duplicate_literal_removal,[],[f4071])).
fof(f4071,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,X3) | p(s(bool,X4)) | ~sP7(X5,X3,X4) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5)) ) | $spl10),
  inference(resolution,[],[f3863,f359])).
fof(f3863,plain,(
  ( ! [X10,X11,X9] : (sP9(X10,X11,X9) | s(bool,'T_0') != s(bool,X9) | p(s(bool,X10)) | ~sP7(X11,X9,X10)) ) | $spl10),
  inference(subsumption_resolution,[],[f3859,f347])).
fof(f3859,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') != s(bool,X9) | sP9(X10,X11,X9) | p(s(bool,X10)) | ~sP8(X10,X9,X11) | ~sP7(X11,X9,X10)) ) | $spl10),
  inference(resolution,[],[f3855,f343])).
fof(f3855,plain,(
  ( ! [X6,X8,X7] : (sP6(X7,X6,X8) | s(bool,'T_0') != s(bool,X7)) ) | $spl10),
  inference(subsumption_resolution,[],[f3854,f353])).
fof(f3854,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X7) | p(s(bool,X8)) | sP6(X7,X6,X8)) ) | $spl10),
  inference(subsumption_resolution,[],[f3841,f354])).
fof(f3841,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | s(bool,'T_0') != s(bool,X7) | p(s(bool,X8)) | sP6(X7,X6,X8)) ) | $spl10),
  inference(resolution,[],[f3815,f337])).
fof(f3815,plain,(
  ( ! [X6,X8,X7] : (sP9(X6,X7,X8) | p(s(bool,X7)) | s(bool,'T_0') != s(bool,X8) | p(s(bool,X6))) ) | $spl10),
  inference(resolution,[],[f3788,f347])).
fof(f3788,plain,(
  ( ! [X6,X8,X7] : (~sP8(X8,X7,X6) | sP9(X8,X6,X7) | p(s(bool,X6)) | s(bool,'T_0') != s(bool,X7)) ) | $spl10),
  inference(subsumption_resolution,[],[f3787,f351])).
fof(f3787,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X7) | sP9(X8,X6,X7) | p(s(bool,X6)) | ~sP8(X8,X7,X6) | ~sP7(X6,X7,X8)) ) | $spl10),
  inference(subsumption_resolution,[],[f3780,f461])).
fof(f3780,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,'T_0') != s(bool,X7) | sP9(X8,X6,X7) | p(s(bool,X6)) | ~sP8(X8,X7,X6) | ~sP7(X6,X7,X8)) ) | $spl10),
  inference(resolution,[],[f3776,f342])).
fof(f13116,plain,(
  ( ! [X29] : (s(bool,X29) != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | ~p(s(bool,X29))) ) | $spl206),
  inference(cnf_transformation,[],[f13116_D])).
fof(f13116_D,plain,(
  ( ! [X29] : (s(bool,X29) != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | ~p(s(bool,X29))) ) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f20475,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | p(s(bool,X9)) | s(bool,X8) != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl10 | $spl199)),
  inference(resolution,[],[f12706,f3924])).
fof(f3924,plain,(
  ( ! [X4,X5,X3] : (~sP7(X4,X3,X5) | ~p(s(bool,X3)) | p(s(bool,X5)) | s(bool,X3) != s(bool,X4)) ) | $spl10),
  inference(duplicate_literal_removal,[],[f3911])).
fof(f3911,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | ~sP7(X4,X3,X5) | p(s(bool,X5)) | p(s(bool,X5)) | s(bool,X3) != s(bool,X4)) ) | $spl10),
  inference(resolution,[],[f3887,f359])).
fof(f3887,plain,(
  ( ! [X6,X8,X7] : (sP9(X7,X8,X6) | ~p(s(bool,X6)) | ~sP7(X8,X6,X7) | p(s(bool,X7))) ) | $spl10),
  inference(resolution,[],[f3862,f347])).
fof(f3862,plain,(
  ( ! [X4,X5,X3] : (~sP8(X4,X3,X5) | ~p(s(bool,X3)) | sP9(X4,X5,X3) | ~sP7(X5,X3,X4)) ) | $spl10),
  inference(subsumption_resolution,[],[f3857,f465])).
fof(f3857,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,X3) | sP9(X4,X5,X3) | ~p(s(bool,X3)) | ~sP8(X4,X3,X5) | ~sP7(X5,X3,X4)) ) | $spl10),
  inference(resolution,[],[f3855,f341])).
fof(f20448,plain,(
  ( ! [X18] : (p(s(bool,X18)) | s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl104 | $spl199)),
  inference(resolution,[],[f12705,f5907])).
fof(f5907,plain,(
  ( ! [X10,X11] : (~sP7('T_0',X11,X10) | p(s(bool,X10)) | s(bool,'T_0') != s(bool,X11)) ) | $spl104),
  inference(duplicate_literal_removal,[],[f5898])).
fof(f5898,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | ~sP7('T_0',X11,X10) | p(s(bool,X10)) | s(bool,'T_0') != s(bool,X11)) ) | $spl104),
  inference(resolution,[],[f5892,f359])).
fof(f5892,plain,(
  ( ! [X4,X5] : (sP9(X4,'T_0',X5) | p(s(bool,X4)) | ~sP7('T_0',X5,X4)) ) | $spl104),
  inference(subsumption_resolution,[],[f5883,f347])).
fof(f5883,plain,(
  ( ! [X4,X5] : (sP9(X4,'T_0',X5) | p(s(bool,X4)) | ~sP8(X4,X5,'T_0') | ~sP7('T_0',X5,X4)) ) | $spl104),
  inference(resolution,[],[f5880,f343])).
fof(f5880,plain,(
  ( ! [X0,X1] : (sP6(X1,'T_0',X0)) ) | $spl104),
  inference(subsumption_resolution,[],[f5879,f353])).
fof(f5879,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | sP6(X1,'T_0',X0)) ) | $spl104),
  inference(subsumption_resolution,[],[f5863,f355])).
fof(f5863,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | sP6(X1,'T_0',X0)) ) | $spl104),
  inference(resolution,[],[f5837,f337])).
fof(f12705,plain,(
  ( ! [X105,X104] : (sP7(X104,'const.combin.I_1'(s(bool,'T_0')),X105)) ) | $spl199),
  inference(resolution,[],[f11984,f350])).
fof(f20447,plain,(
  ( ! [X17,X16] : (p(s(bool,X16)) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X17) != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl120 | $spl199)),
  inference(resolution,[],[f12705,f6791])).
fof(f20446,plain,(
  ( ! [X14,X15] : (p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | s(bool,X15) != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl110 | $spl199)),
  inference(resolution,[],[f12705,f6440])).
fof(f20445,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | s(bool,'F_0') = s(bool,X13) | s(bool,X13) != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl102 | $spl199)),
  inference(resolution,[],[f12705,f6022])).
fof(f20444,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X11) != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl10 | $spl199)),
  inference(resolution,[],[f12705,f4084])).
fof(f20300,plain,(
  ( ! [X37,X38] : (~p(s(bool,X37)) | s(bool,X38) != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') != s(bool,X38)) ) | ($spl10 | $spl199)),
  inference(resolution,[],[f12702,f3753])).
fof(f12702,plain,(
  ( ! [X99,X98] : (sP8(X98,'const.combin.I_1'(s(bool,'T_0')),X99)) ) | $spl199),
  inference(resolution,[],[f11984,f346])).
fof(f13117,plain,(
  ( ! [X30] : (sP9(X30,'F_0','const.combin.I_1'(s(bool,'T_0')))) ) | $spl208),
  inference(cnf_transformation,[],[f13117_D])).
fof(f13117_D,plain,(
  ( ! [X30] : (sP9(X30,'F_0','const.combin.I_1'(s(bool,'T_0')))) ) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f20299,plain,(
  ( ! [X35,X36] : (s(bool,'T_0') != s(bool,X35) | sP9(X36,'F_0','const.combin.I_1'(s(bool,'T_0'))) | s(bool,X35) != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl10 | $spl199)),
  inference(resolution,[],[f12702,f3705])).
fof(f3705,plain,(
  ( ! [X10,X11,X9] : (~sP8(X11,X10,'F_0') | s(bool,'T_0') != s(bool,X9) | sP9(X11,'F_0',X10) | s(bool,X9) != s(bool,X10)) ) | $spl10),
  inference(subsumption_resolution,[],[f3704,f2436])).
fof(f2436,plain,(
  ( ! [X4,X5] : (sP7('F_0',X4,X5)) )),
  inference(subsumption_resolution,[],[f2432,f349])).
fof(f2432,plain,(
  ( ! [X4,X5] : (sP7('F_0',X4,X5) | p(s(bool,X5))) )),
  inference(resolution,[],[f2429,f347])).
fof(f2429,plain,(
  ( ! [X6,X7] : (~sP8(X7,X6,'F_0') | sP7('F_0',X6,X7)) )),
  inference(subsumption_resolution,[],[f2425,f350])).
fof(f2425,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | ~sP8(X7,X6,'F_0') | sP7('F_0',X6,X7)) )),
  inference(resolution,[],[f2419,f338])).
fof(f3704,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,X10) | s(bool,'T_0') != s(bool,X9) | sP9(X11,'F_0',X10) | ~sP8(X11,X10,'F_0') | ~sP7('F_0',X10,X11)) ) | $spl10),
  inference(subsumption_resolution,[],[f3700,f3041])).
fof(f3700,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,X10) | s(bool,'T_0') != s(bool,X9) | sP9(X11,'F_0',X10) | ~p(s(bool,X10)) | ~sP8(X11,X10,'F_0') | ~sP7('F_0',X10,X11)) ) | $spl10),
  inference(resolution,[],[f3696,f341])).
fof(f3696,plain,(
  ( ! [X17,X15,X16] : (sP6(X16,'F_0',X17) | s(bool,X15) != s(bool,X16) | s(bool,'T_0') != s(bool,X15)) ) | $spl10),
  inference(subsumption_resolution,[],[f3689,f353])).
fof(f3689,plain,(
  ( ! [X17,X15,X16] : (s(bool,'T_0') != s(bool,X15) | s(bool,X15) != s(bool,X16) | p(s(bool,X17)) | sP6(X16,'F_0',X17)) ) | $spl10),
  inference(resolution,[],[f3044,f337])).
fof(f3044,plain,(
  ( ! [X12,X13,X11] : (sP9(X13,'F_0',X11) | s(bool,'T_0') != s(bool,X12) | s(bool,X11) != s(bool,X12) | p(s(bool,X13))) ) | $spl10),
  inference(superposition,[],[f1974,f356])).
fof(f20298,plain,(
  ( ! [X34] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | ~p(s(bool,X34))) ) | $spl199),
  inference(resolution,[],[f12702,f2549])).
fof(f20294,plain,(
  ( ! [X30] : (sP9(X30,'F_0','const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | $spl199),
  inference(resolution,[],[f12702,f2413])).
fof(f2413,plain,(
  ( ! [X4,X5] : (~sP8(X5,X4,'F_0') | sP9(X5,'F_0',X4) | s(bool,'T_0') != s(bool,X4)) )),
  inference(subsumption_resolution,[],[f2412,f478])).
fof(f2412,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | sP9(X5,'F_0',X4) | p(s(bool,'F_0')) | ~sP8(X5,X4,'F_0')) )),
  inference(subsumption_resolution,[],[f2409,f351])).
fof(f2409,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | sP9(X5,'F_0',X4) | p(s(bool,'F_0')) | ~sP8(X5,X4,'F_0') | ~sP7('F_0',X4,X5)) )),
  inference(resolution,[],[f2405,f342])).
fof(f20293,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X29) != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | ~p(s(bool,X29))) ) | ($spl24 | $spl199)),
  inference(resolution,[],[f12702,f6924])).
fof(f13115,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,'const.combin.I_1'(s(bool,'T_0'))) | p(s(bool,X9))) ) | $spl204),
  inference(cnf_transformation,[],[f13115_D])).
fof(f13115_D,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,'const.combin.I_1'(s(bool,'T_0'))) | p(s(bool,X9))) ) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f20292,plain,(
  ( ! [X26,X25] : (sP9(X25,X26,'const.combin.I_1'(s(bool,'T_0'))) | p(s(bool,X26)) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl120 | $spl199)),
  inference(resolution,[],[f12702,f6759])).
fof(f6759,plain,(
  ( ! [X4,X5,X3] : (~sP8(X4,X3,X5) | sP9(X4,X5,X3) | p(s(bool,X5)) | s(bool,'F_0') = s(bool,X3)) ) | $spl120),
  inference(subsumption_resolution,[],[f6748,f351])).
fof(f6748,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') = s(bool,X3) | sP9(X4,X5,X3) | p(s(bool,X5)) | ~sP8(X4,X3,X5) | ~sP7(X5,X3,X4)) ) | $spl120),
  inference(resolution,[],[f6746,f342])).
fof(f20289,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'F_0') != s(bool,X19) | s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl10 | $spl199)),
  inference(resolution,[],[f12702,f4242])).
fof(f20284,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,'const.combin.I_1'(s(bool,'T_0'))) | p(s(bool,X9)) | s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl10 | $spl199)),
  inference(resolution,[],[f12702,f3788])).
fof(f19928,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | p(s(bool,'const.combin.I_1'(s(bool,X1))))) ) | $spl228),
  inference(cnf_transformation,[],[f19928_D])).
fof(f19928_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | p(s(bool,'const.combin.I_1'(s(bool,X1))))) ) <=> ~$spl228),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl228])])).
fof(f19929,plain,(
  ( ! [X4,X3] : (~p(s(bool,X3)) | p(s(bool,'const.combin.I_1'(s(bool,X3)))) | p(s(bool,X4))) ) | $spl160),
  inference(subsumption_resolution,[],[f19885,f436])).
fof(f19885,plain,(
  ( ! [X4,X3] : (~p(s(bool,X3)) | p(s(bool,'F_0')) | p(s(bool,'const.combin.I_1'(s(bool,X3)))) | p(s(bool,X4))) ) | $spl160),
  inference(resolution,[],[f16303,f390])).
fof(f16303,plain,(
  ( ! [X4,X5] : (sP11('const.combin.I_1'(s(bool,X4)),X5) | ~p(s(bool,X4))) ) | $spl160),
  inference(superposition,[],[f16294,f467])).
fof(f467,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f272])).
fof(f272,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f113])).
fof(f113,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X13] : (s(bool,'T_0') = s(bool,X13) <=> p(s(bool,X13)))),
  file('thm.quotient.QUOTIENT_SYM','thm.bool.EQ_CLAUSES|split|1')).
fof(f16294,plain,(
  ( ! [X275] : (sP11('const.combin.I_1'(s(bool,'T_0')),X275)) ) | $spl160),
  inference(subsumption_resolution,[],[f14912,f491])).
fof(f14912,plain,(
  ( ! [X275] : (~p(s(bool,'T_0')) | sP11('const.combin.I_1'(s(bool,'T_0')),X275)) ) | $spl160),
  inference(superposition,[],[f397,f8875])).
fof(f397,plain,(
  ( ! [X0,X1] : (sP11(X0,X1) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f231])).
fof(f19927,plain,(
  ( ! [X2,X1] : (~p(s(bool,X1)) | p(s(bool,'const.combin.I_1'(s(bool,X1)))) | p(s(bool,X2))) ) | $spl160),
  inference(subsumption_resolution,[],[f19884,f436])).
fof(f19884,plain,(
  ( ! [X2,X1] : (~p(s(bool,X1)) | p(s(bool,'F_0')) | p(s(bool,'const.combin.I_1'(s(bool,X1)))) | p(s(bool,X2))) ) | $spl160),
  inference(resolution,[],[f16303,f394])).
fof(f9377,plain,(
  ( ! [X316,X317] : (~p(s(bool,X316)) | p(s(bool,X317)) | sP9(X317,X316,'const.combin.I_1'(s(bool,'F_0')))) ) | $spl174),
  inference(cnf_transformation,[],[f9377_D])).
fof(f9377_D,plain,(
  ( ! [X316,X317] : (~p(s(bool,X316)) | p(s(bool,X317)) | sP9(X317,X316,'const.combin.I_1'(s(bool,'F_0')))) ) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f18693,plain,(
  ( ! [X315,X316,X317] : (s(bool,'F_0') != s(bool,X315) | ~p(s(bool,X316)) | sP9(X317,X316,'const.combin.I_1'(s(bool,'F_0'))) | p(s(bool,X317))) ) | ($spl24 | $spl148)),
  inference(subsumption_resolution,[],[f18523,f460])).
fof(f18523,plain,(
  ( ! [X315,X316,X317] : (s(bool,'F_0') != s(bool,X315) | ~p(s(bool,X316)) | p(s(bool,X315)) | sP9(X317,X316,'const.combin.I_1'(s(bool,'F_0'))) | p(s(bool,X317))) ) | ($spl24 | $spl148)),
  inference(superposition,[],[f4772,f8862])).
fof(f9375,plain,(
  ( ! [X310,X311] : (~p(s(bool,X310)) | p(s(bool,X311)) | sP9(X311,'const.combin.I_1'(s(bool,'F_0')),X310)) ) | $spl172),
  inference(cnf_transformation,[],[f9375_D])).
fof(f9375_D,plain,(
  ( ! [X310,X311] : (~p(s(bool,X310)) | p(s(bool,X311)) | sP9(X311,'const.combin.I_1'(s(bool,'F_0')),X310)) ) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f18692,plain,(
  ( ! [X310,X309,X311] : (s(bool,'F_0') != s(bool,X309) | ~p(s(bool,X310)) | sP9(X311,'const.combin.I_1'(s(bool,'F_0')),X310) | p(s(bool,X311))) ) | ($spl24 | $spl148)),
  inference(subsumption_resolution,[],[f18521,f460])).
fof(f18521,plain,(
  ( ! [X310,X309,X311] : (s(bool,'F_0') != s(bool,X309) | ~p(s(bool,X310)) | p(s(bool,X309)) | sP9(X311,'const.combin.I_1'(s(bool,'F_0')),X310) | p(s(bool,X311))) ) | ($spl24 | $spl148)),
  inference(superposition,[],[f4771,f8862])).
fof(f11269,plain,(
  ( ! [X13] : (s(bool,'F_0') = s(bool,X13) | s(bool,X13) != s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | $spl196),
  inference(cnf_transformation,[],[f11269_D])).
fof(f11269_D,plain,(
  ( ! [X13] : (s(bool,'F_0') = s(bool,X13) | s(bool,X13) != s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f18097,plain,(
  ( ! [X17,X16] : (p(s(bool,X16)) | s(bool,'F_0') = s(bool,X17) | s(bool,X17) != s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | ($spl120 | $spl148)),
  inference(resolution,[],[f9317,f6791])).
fof(f9317,plain,(
  ( ! [X215,X214] : (sP7('const.combin.I_1'(s(bool,'F_0')),X214,X215)) ) | $spl148),
  inference(subsumption_resolution,[],[f9090,f478])).
fof(f9090,plain,(
  ( ! [X215,X214] : (p(s(bool,'F_0')) | sP7('const.combin.I_1'(s(bool,'F_0')),X214,X215)) ) | $spl148),
  inference(superposition,[],[f351,f8862])).
fof(f10337,plain,(
  ( ! [X35] : (s(bool,'T_0') != s(bool,X35) | s(bool,X35) != s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | $spl188),
  inference(cnf_transformation,[],[f10337_D])).
fof(f10337_D,plain,(
  ( ! [X35] : (s(bool,'T_0') != s(bool,X35) | s(bool,X35) != s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f18094,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | s(bool,'T_0') != s(bool,X11) | s(bool,X11) != s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | ($spl10 | $spl148)),
  inference(resolution,[],[f9317,f4084])).
fof(f10335,plain,(
  ( ! [X29] : (s(bool,X29) != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | ~p(s(bool,X29))) ) | $spl184),
  inference(cnf_transformation,[],[f10335_D])).
fof(f10335_D,plain,(
  ( ! [X29] : (s(bool,X29) != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | ~p(s(bool,X29))) ) <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f18093,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | p(s(bool,X9)) | s(bool,X8) != s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | ($spl10 | $spl148)),
  inference(resolution,[],[f9317,f3924])).
fof(f18000,plain,(
  ( ! [X14,X15] : (p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | s(bool,X15) != s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | ($spl110 | $spl148)),
  inference(resolution,[],[f9316,f6440])).
fof(f9316,plain,(
  ( ! [X212,X213] : (sP7(X212,'const.combin.I_1'(s(bool,'F_0')),X213)) ) | $spl148),
  inference(subsumption_resolution,[],[f9089,f478])).
fof(f9089,plain,(
  ( ! [X212,X213] : (p(s(bool,'F_0')) | sP7(X212,'const.combin.I_1'(s(bool,'F_0')),X213)) ) | $spl148),
  inference(superposition,[],[f350,f8862])).
fof(f17999,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | s(bool,'F_0') = s(bool,X13) | s(bool,X13) != s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | ($spl102 | $spl148)),
  inference(resolution,[],[f9316,f6022])).
fof(f17919,plain,(
  ( ! [X37,X38] : (~p(s(bool,X37)) | s(bool,X38) != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') != s(bool,X38)) ) | ($spl10 | $spl148)),
  inference(resolution,[],[f9313,f3753])).
fof(f10336,plain,(
  ( ! [X30] : (sP9(X30,'F_0','const.combin.I_1'(s(bool,'F_0')))) ) | $spl186),
  inference(cnf_transformation,[],[f10336_D])).
fof(f10336_D,plain,(
  ( ! [X30] : (sP9(X30,'F_0','const.combin.I_1'(s(bool,'F_0')))) ) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f17918,plain,(
  ( ! [X35,X36] : (s(bool,'T_0') != s(bool,X35) | sP9(X36,'F_0','const.combin.I_1'(s(bool,'F_0'))) | s(bool,X35) != s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | ($spl10 | $spl148)),
  inference(resolution,[],[f9313,f3705])).
fof(f17912,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X29) != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | ~p(s(bool,X29))) ) | ($spl24 | $spl148)),
  inference(resolution,[],[f9313,f6924])).
fof(f8947,plain,(
  ~p(s(bool,'const.combin.I_1'(s(bool,'F_0')))) | $spl171),
  inference(cnf_transformation,[],[f8947_D])).
fof(f8947_D,plain,(
  ~p(s(bool,'const.combin.I_1'(s(bool,'F_0')))) <=> ~$spl171),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl171])])).
fof(f17606,plain,(
  ( ! [X10] : (~p(s(bool,'const.combin.I_1'(s(bool,'F_0')))) | ~p(s(bool,X10))) ) | $spl169),
  inference(subsumption_resolution,[],[f17590,f465])).
fof(f17590,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,X10) | ~p(s(bool,'const.combin.I_1'(s(bool,'F_0')))) | ~p(s(bool,X10))) ) | $spl169),
  inference(superposition,[],[f8884,f406])).
fof(f406,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f142])).
fof(f142,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f141])).
fof(f141,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X14,X15] : ((p(s(bool,X14)) => p(s(bool,X15))) => ((p(s(bool,X15)) => p(s(bool,X14))) => s(bool,X14) = s(bool,X15)))),
  file('thm.quotient.QUOTIENT_SYM','thm.bool.IMP_ANTISYM_AX')).
fof(f17608,plain,(
  ( ! [X11] : (~p(s(bool,X11)) | ~p(s(bool,'const.combin.I_1'(s(bool,'F_0'))))) ) | $spl169),
  inference(subsumption_resolution,[],[f17591,f465])).
fof(f17591,plain,(
  ( ! [X11] : (s(bool,'T_0') != s(bool,X11) | ~p(s(bool,X11)) | ~p(s(bool,'const.combin.I_1'(s(bool,'F_0'))))) ) | $spl169),
  inference(superposition,[],[f8884,f406])).
fof(f17613,plain,(
  $false | ($spl149 | $spl169)),
  inference(subsumption_resolution,[],[f17600,f8863])).
fof(f17600,plain,(
  s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | $spl169),
  inference(trivial_inequality_removal,[],[f17598])).
fof(f17598,plain,(
  s(bool,'T_0') != s(bool,'T_0') | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | $spl169),
  inference(superposition,[],[f8884,f486])).
fof(f486,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X13] : (s(bool,'T_0') = s(bool,X13) | s(bool,'F_0') = s(bool,X13))),
  file('thm.quotient.QUOTIENT_SYM','thm.bool.BOOL_CASES_AX')).
fof(f17612,plain,(
  $false | ($spl149 | $spl169)),
  inference(subsumption_resolution,[],[f17601,f8863])).
fof(f17601,plain,(
  s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | $spl169),
  inference(trivial_inequality_removal,[],[f17596])).
fof(f17596,plain,(
  s(bool,'T_0') != s(bool,'T_0') | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | $spl169),
  inference(superposition,[],[f8884,f485])).
fof(f485,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.quotient.QUOTIENT_SYM',aHL_BOOL_CASES)).
fof(f17611,plain,(
  $false | ($spl149 | $spl169)),
  inference(subsumption_resolution,[],[f17602,f8942])).
fof(f17602,plain,(
  ~p(s(bool,'const.combin.I_1'(s(bool,'F_0')))) | $spl169),
  inference(trivial_inequality_removal,[],[f17595])).
fof(f17595,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ~p(s(bool,'const.combin.I_1'(s(bool,'F_0')))) | $spl169),
  inference(superposition,[],[f8884,f467])).
fof(f17610,plain,(
  $false | ($spl149 | $spl169)),
  inference(subsumption_resolution,[],[f17603,f8942])).
fof(f17603,plain,(
  ~p(s(bool,'const.combin.I_1'(s(bool,'F_0')))) | $spl169),
  inference(trivial_inequality_removal,[],[f17594])).
fof(f17594,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ~p(s(bool,'const.combin.I_1'(s(bool,'F_0')))) | $spl169),
  inference(superposition,[],[f8884,f465])).
fof(f15115,plain,(
  ( ! [X329,X328] : (p(s(bool,X328)) | p(s(bool,X329)) | sP9(X329,X328,'const.combin.I_1'(s(bool,'T_0')))) ) | $spl222),
  inference(cnf_transformation,[],[f15115_D])).
fof(f15115_D,plain,(
  ( ! [X329,X328] : (p(s(bool,X328)) | p(s(bool,X329)) | sP9(X329,X328,'const.combin.I_1'(s(bool,'T_0')))) ) <=> ~$spl222),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl222])])).
fof(f17542,plain,(
  ( ! [X327,X329,X328] : (~p(s(bool,X327)) | p(s(bool,X328)) | sP9(X329,X328,'const.combin.I_1'(s(bool,'T_0'))) | p(s(bool,X329))) ) | ($spl24 | $spl160)),
  inference(subsumption_resolution,[],[f17409,f465])).
fof(f17409,plain,(
  ( ! [X327,X329,X328] : (s(bool,'T_0') != s(bool,X327) | ~p(s(bool,X327)) | p(s(bool,X328)) | sP9(X329,X328,'const.combin.I_1'(s(bool,'T_0'))) | p(s(bool,X329))) ) | ($spl24 | $spl160)),
  inference(superposition,[],[f4800,f8875])).
fof(f4800,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X5) != s(bool,X6) | ~p(s(bool,X6)) | p(s(bool,X4)) | sP9(X7,X4,X5) | p(s(bool,X7))) ) | $spl24),
  inference(superposition,[],[f2665,f356])).
fof(f15113,plain,(
  ( ! [X323,X322] : (p(s(bool,X322)) | p(s(bool,X323)) | sP9(X323,'const.combin.I_1'(s(bool,'T_0')),X322)) ) | $spl220),
  inference(cnf_transformation,[],[f15113_D])).
fof(f15113_D,plain,(
  ( ! [X323,X322] : (p(s(bool,X322)) | p(s(bool,X323)) | sP9(X323,'const.combin.I_1'(s(bool,'T_0')),X322)) ) <=> ~$spl220),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl220])])).
fof(f17541,plain,(
  ( ! [X321,X323,X322] : (~p(s(bool,X321)) | p(s(bool,X322)) | sP9(X323,'const.combin.I_1'(s(bool,'T_0')),X322) | p(s(bool,X323))) ) | ($spl24 | $spl160)),
  inference(subsumption_resolution,[],[f17407,f465])).
fof(f17407,plain,(
  ( ! [X321,X323,X322] : (s(bool,'T_0') != s(bool,X321) | ~p(s(bool,X321)) | p(s(bool,X322)) | sP9(X323,'const.combin.I_1'(s(bool,'T_0')),X322) | p(s(bool,X323))) ) | ($spl24 | $spl160)),
  inference(superposition,[],[f4799,f8875])).
fof(f4799,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X1) != s(bool,X2) | ~p(s(bool,X2)) | p(s(bool,X0)) | sP9(X3,X1,X0) | p(s(bool,X3))) ) | $spl24),
  inference(superposition,[],[f2665,f356])).
fof(f8875,plain,(
  s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | $spl160),
  inference(cnf_transformation,[],[f8875_D])).
fof(f8875_D,plain,(
  s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'T_0'))) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f16851,plain,(
  ( ! [X14,X15] : (p(s(bool,X14)) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X15) != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl94 | $spl151)),
  inference(resolution,[],[f12043,f5575])).
fof(f12043,plain,(
  ( ! [X112,X113] : (sP6('const.combin.I_1'(s(bool,'T_0')),X112,X113)) ) | $spl151),
  inference(resolution,[],[f11979,f355])).
fof(f11979,plain,(
  p(s(bool,'const.combin.I_1'(s(bool,'T_0')))) | $spl151),
  inference(trivial_inequality_removal,[],[f11970])).
fof(f11970,plain,(
  s(bool,'F_0') != s(bool,'F_0') | p(s(bool,'const.combin.I_1'(s(bool,'T_0')))) | $spl151),
  inference(superposition,[],[f8867,f463])).
fof(f13520,plain,(
  ( ! [X15] : (s(bool,'T_0') = s(bool,X15) | s(bool,X15) != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | $spl214),
  inference(cnf_transformation,[],[f13520_D])).
fof(f13520_D,plain,(
  ( ! [X15] : (s(bool,'T_0') = s(bool,X15) | s(bool,X15) != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) <=> ~$spl214),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl214])])).
fof(f16850,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | s(bool,'T_0') = s(bool,X13) | s(bool,X13) != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl70 | $spl151)),
  inference(resolution,[],[f12043,f5073])).
fof(f13521,plain,(
  ( ! [X17] : (s(bool,'F_0') != s(bool,X17) | s(bool,X17) != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | $spl216),
  inference(cnf_transformation,[],[f13521_D])).
fof(f13521_D,plain,(
  ( ! [X17] : (s(bool,'F_0') != s(bool,X17) | s(bool,X17) != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) <=> ~$spl216),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl216])])).
fof(f16849,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | s(bool,'F_0') != s(bool,X11) | s(bool,X11) != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl10 | $spl151)),
  inference(resolution,[],[f12043,f4121])).
fof(f12344,plain,(
  ( ! [X22] : (p(s(bool,X22)) | s(bool,X22) != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | $spl202),
  inference(cnf_transformation,[],[f12344_D])).
fof(f12344_D,plain,(
  ( ! [X22] : (p(s(bool,X22)) | s(bool,X22) != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f16848,plain,(
  ( ! [X8,X9] : (p(s(bool,X8)) | p(s(bool,X9)) | s(bool,X8) != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl10 | $spl151)),
  inference(resolution,[],[f12043,f4011])).
fof(f16822,plain,(
  ( ! [X17,X16] : (p(s(bool,X16)) | s(bool,'F_0') != s(bool,X17) | s(bool,X17) != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl98 | $spl151)),
  inference(resolution,[],[f12042,f5593])).
fof(f12042,plain,(
  ( ! [X111,X110] : (sP6(X110,'const.combin.I_1'(s(bool,'T_0')),X111)) ) | $spl151),
  inference(resolution,[],[f11979,f354])).
fof(f16821,plain,(
  ( ! [X14,X15] : (p(s(bool,X14)) | s(bool,'T_0') = s(bool,X15) | s(bool,X15) != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl94 | $spl151)),
  inference(resolution,[],[f12042,f5575])).
fof(f16820,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X13) != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl70 | $spl151)),
  inference(resolution,[],[f12042,f5073])).
fof(f16693,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | ~p(s(bool,X24))) ) | ($spl24 | $spl151)),
  inference(resolution,[],[f12037,f5693])).
fof(f12037,plain,(
  ( ! [X101,X100] : (sP8(X100,X101,'const.combin.I_1'(s(bool,'T_0')))) ) | $spl151),
  inference(resolution,[],[f11979,f345])).
fof(f12343,plain,(
  ( ! [X21,X20] : (sP9(X20,'const.combin.I_1'(s(bool,'T_0')),X21) | ~p(s(bool,X21))) ) | $spl200),
  inference(cnf_transformation,[],[f12343_D])).
fof(f12343_D,plain,(
  ( ! [X21,X20] : (sP9(X20,'const.combin.I_1'(s(bool,'T_0')),X21) | ~p(s(bool,X21))) ) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f16692,plain,(
  ( ! [X21,X20] : (sP9(X20,'const.combin.I_1'(s(bool,'T_0')),X21) | ~p(s(bool,X21)) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl70 | $spl151)),
  inference(resolution,[],[f12037,f5039])).
fof(f11985,plain,(
  p(s(bool,'const.combin.I_1'(s(bool,'T_0')))) | $spl198),
  inference(cnf_transformation,[],[f11985_D])).
fof(f11985_D,plain,(
  p(s(bool,'const.combin.I_1'(s(bool,'T_0')))) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f16379,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.combin.I_1'(s(bool,'T_0'))))) ) | $spl160),
  inference(subsumption_resolution,[],[f16360,f436])).
fof(f16360,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.combin.I_1'(s(bool,'T_0'))))) ) | $spl160),
  inference(resolution,[],[f16296,f390])).
fof(f16296,plain,(
  ( ! [X270] : (sP11(X270,'const.combin.I_1'(s(bool,'T_0')))) ) | $spl160),
  inference(subsumption_resolution,[],[f14907,f491])).
fof(f14907,plain,(
  ( ! [X270] : (~p(s(bool,'T_0')) | sP11(X270,'const.combin.I_1'(s(bool,'T_0')))) ) | $spl160),
  inference(superposition,[],[f391,f8875])).
fof(f16378,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.combin.I_1'(s(bool,'T_0'))))) ) | $spl160),
  inference(subsumption_resolution,[],[f16359,f436])).
fof(f16359,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.combin.I_1'(s(bool,'T_0'))))) ) | $spl160),
  inference(resolution,[],[f16296,f394])).
fof(f16318,plain,(
  ( ! [X1] : (p(s(bool,'const.combin.I_1'(s(bool,'T_0')))) | p(s(bool,X1))) ) | $spl160),
  inference(subsumption_resolution,[],[f16299,f436])).
fof(f16299,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.combin.I_1'(s(bool,'T_0')))) | p(s(bool,X1))) ) | $spl160),
  inference(resolution,[],[f16294,f390])).
fof(f16317,plain,(
  ( ! [X0] : (p(s(bool,'const.combin.I_1'(s(bool,'T_0')))) | p(s(bool,X0))) ) | $spl160),
  inference(subsumption_resolution,[],[f16298,f436])).
fof(f16298,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.combin.I_1'(s(bool,'T_0')))) | p(s(bool,X0))) ) | $spl160),
  inference(resolution,[],[f16294,f394])).
fof(f16263,plain,(
  ( ! [X315,X316,X317] : (s(bool,'F_0') != s(bool,X315) | ~p(s(bool,X316)) | sP9(X317,X316,'const.combin.I_1'(s(bool,'T_0'))) | p(s(bool,X317))) ) | ($spl24 | $spl150)),
  inference(subsumption_resolution,[],[f16092,f460])).
fof(f16092,plain,(
  ( ! [X315,X316,X317] : (s(bool,'F_0') != s(bool,X315) | ~p(s(bool,X316)) | p(s(bool,X315)) | sP9(X317,X316,'const.combin.I_1'(s(bool,'T_0'))) | p(s(bool,X317))) ) | ($spl24 | $spl150)),
  inference(superposition,[],[f4772,f8866])).
fof(f16261,plain,(
  ( ! [X310,X309,X311] : (s(bool,'F_0') != s(bool,X309) | ~p(s(bool,X310)) | sP9(X311,'const.combin.I_1'(s(bool,'T_0')),X310) | p(s(bool,X311))) ) | ($spl24 | $spl150)),
  inference(subsumption_resolution,[],[f16090,f460])).
fof(f16090,plain,(
  ( ! [X310,X309,X311] : (s(bool,'F_0') != s(bool,X309) | ~p(s(bool,X310)) | p(s(bool,X309)) | sP9(X311,'const.combin.I_1'(s(bool,'T_0')),X310) | p(s(bool,X311))) ) | ($spl24 | $spl150)),
  inference(superposition,[],[f4771,f8866])).
fof(f15587,plain,(
  ( ! [X18] : (p(s(bool,X18)) | s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl122 | $spl199)),
  inference(resolution,[],[f12706,f6702])).
fof(f15586,plain,(
  ( ! [X17,X16] : (p(s(bool,X16)) | s(bool,'F_0') = s(bool,X17) | s(bool,X17) != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl120 | $spl199)),
  inference(resolution,[],[f12706,f6791])).
fof(f15585,plain,(
  ( ! [X14,X15] : (p(s(bool,X14)) | s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X15) != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl110 | $spl199)),
  inference(resolution,[],[f12706,f6440])).
fof(f15584,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X13) != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl102 | $spl199)),
  inference(resolution,[],[f12706,f6022])).
fof(f15583,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | s(bool,'T_0') != s(bool,X11) | s(bool,X11) != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl10 | $spl199)),
  inference(resolution,[],[f12706,f4084])).
fof(f15582,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | p(s(bool,X9)) | s(bool,X8) != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl10 | $spl199)),
  inference(resolution,[],[f12706,f3924])).
fof(f15557,plain,(
  ( ! [X18] : (p(s(bool,X18)) | s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl104 | $spl199)),
  inference(resolution,[],[f12705,f5907])).
fof(f15556,plain,(
  ( ! [X17,X16] : (p(s(bool,X16)) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X17) != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl120 | $spl199)),
  inference(resolution,[],[f12705,f6791])).
fof(f15555,plain,(
  ( ! [X14,X15] : (p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | s(bool,X15) != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl110 | $spl199)),
  inference(resolution,[],[f12705,f6440])).
fof(f15554,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | s(bool,'F_0') = s(bool,X13) | s(bool,X13) != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl102 | $spl199)),
  inference(resolution,[],[f12705,f6022])).
fof(f15553,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X11) != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl10 | $spl199)),
  inference(resolution,[],[f12705,f4084])).
fof(f15484,plain,(
  ( ! [X37,X38] : (~p(s(bool,X37)) | s(bool,X38) != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') != s(bool,X38)) ) | ($spl10 | $spl199)),
  inference(resolution,[],[f12702,f3753])).
fof(f15483,plain,(
  ( ! [X35,X36] : (s(bool,'T_0') != s(bool,X35) | sP9(X36,'F_0','const.combin.I_1'(s(bool,'T_0'))) | s(bool,X35) != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl10 | $spl199)),
  inference(resolution,[],[f12702,f3705])).
fof(f15482,plain,(
  ( ! [X34] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | ~p(s(bool,X34))) ) | $spl199),
  inference(resolution,[],[f12702,f2549])).
fof(f15478,plain,(
  ( ! [X30] : (sP9(X30,'F_0','const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | $spl199),
  inference(resolution,[],[f12702,f2413])).
fof(f15477,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X29) != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | ~p(s(bool,X29))) ) | ($spl24 | $spl199)),
  inference(resolution,[],[f12702,f6924])).
fof(f15476,plain,(
  ( ! [X26,X25] : (sP9(X25,X26,'const.combin.I_1'(s(bool,'T_0'))) | p(s(bool,X26)) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl120 | $spl199)),
  inference(resolution,[],[f12702,f6759])).
fof(f15473,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'F_0') != s(bool,X19) | s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl10 | $spl199)),
  inference(resolution,[],[f12702,f4242])).
fof(f15468,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,'const.combin.I_1'(s(bool,'T_0'))) | p(s(bool,X9)) | s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl10 | $spl199)),
  inference(resolution,[],[f12702,f3788])).
fof(f15114,plain,(
  ( ! [X327,X329,X328] : (~p(s(bool,X327)) | p(s(bool,X328)) | sP9(X329,X328,'const.combin.I_1'(s(bool,'T_0'))) | p(s(bool,X329))) ) | ($spl24 | $spl160)),
  inference(subsumption_resolution,[],[f14980,f465])).
fof(f14980,plain,(
  ( ! [X327,X329,X328] : (s(bool,'T_0') != s(bool,X327) | ~p(s(bool,X327)) | p(s(bool,X328)) | sP9(X329,X328,'const.combin.I_1'(s(bool,'T_0'))) | p(s(bool,X329))) ) | ($spl24 | $spl160)),
  inference(superposition,[],[f4800,f8875])).
fof(f15112,plain,(
  ( ! [X321,X323,X322] : (~p(s(bool,X321)) | p(s(bool,X322)) | sP9(X323,'const.combin.I_1'(s(bool,'T_0')),X322) | p(s(bool,X323))) ) | ($spl24 | $spl160)),
  inference(subsumption_resolution,[],[f14978,f465])).
fof(f14978,plain,(
  ( ! [X321,X323,X322] : (s(bool,'T_0') != s(bool,X321) | ~p(s(bool,X321)) | p(s(bool,X322)) | sP9(X323,'const.combin.I_1'(s(bool,'T_0')),X322) | p(s(bool,X323))) ) | ($spl24 | $spl160)),
  inference(superposition,[],[f4799,f8875])).
fof(f14731,plain,(
  $false | ($spl151 | $spl161)),
  inference(subsumption_resolution,[],[f14718,f8867])).
fof(f14718,plain,(
  s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | $spl161),
  inference(trivial_inequality_removal,[],[f14716])).
fof(f14716,plain,(
  s(bool,'T_0') != s(bool,'T_0') | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | $spl161),
  inference(superposition,[],[f8876,f486])).
fof(f14730,plain,(
  $false | ($spl151 | $spl161)),
  inference(subsumption_resolution,[],[f14719,f8867])).
fof(f14719,plain,(
  s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | $spl161),
  inference(trivial_inequality_removal,[],[f14714])).
fof(f14714,plain,(
  s(bool,'T_0') != s(bool,'T_0') | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | $spl161),
  inference(superposition,[],[f8876,f485])).
fof(f14729,plain,(
  $false | ($spl151 | $spl161)),
  inference(subsumption_resolution,[],[f14720,f11979])).
fof(f14720,plain,(
  ~p(s(bool,'const.combin.I_1'(s(bool,'T_0')))) | $spl161),
  inference(trivial_inequality_removal,[],[f14713])).
fof(f14713,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ~p(s(bool,'const.combin.I_1'(s(bool,'T_0')))) | $spl161),
  inference(superposition,[],[f8876,f467])).
fof(f14728,plain,(
  $false | ($spl151 | $spl161)),
  inference(subsumption_resolution,[],[f14721,f11979])).
fof(f14721,plain,(
  ~p(s(bool,'const.combin.I_1'(s(bool,'T_0')))) | $spl161),
  inference(trivial_inequality_removal,[],[f14712])).
fof(f14712,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ~p(s(bool,'const.combin.I_1'(s(bool,'T_0')))) | $spl161),
  inference(superposition,[],[f8876,f465])).
fof(f14439,plain,(
  ( ! [X14,X15] : (p(s(bool,X14)) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X15) != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl94 | $spl151)),
  inference(resolution,[],[f12043,f5575])).
fof(f14438,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | s(bool,'T_0') = s(bool,X13) | s(bool,X13) != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl70 | $spl151)),
  inference(resolution,[],[f12043,f5073])).
fof(f14437,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | s(bool,'F_0') != s(bool,X11) | s(bool,X11) != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl10 | $spl151)),
  inference(resolution,[],[f12043,f4121])).
fof(f14436,plain,(
  ( ! [X8,X9] : (p(s(bool,X8)) | p(s(bool,X9)) | s(bool,X8) != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl10 | $spl151)),
  inference(resolution,[],[f12043,f4011])).
fof(f14410,plain,(
  ( ! [X17,X16] : (p(s(bool,X16)) | s(bool,'F_0') != s(bool,X17) | s(bool,X17) != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl98 | $spl151)),
  inference(resolution,[],[f12042,f5593])).
fof(f14409,plain,(
  ( ! [X14,X15] : (p(s(bool,X14)) | s(bool,'T_0') = s(bool,X15) | s(bool,X15) != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl94 | $spl151)),
  inference(resolution,[],[f12042,f5575])).
fof(f14408,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X13) != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl70 | $spl151)),
  inference(resolution,[],[f12042,f5073])).
fof(f14288,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | ~p(s(bool,X24))) ) | ($spl24 | $spl151)),
  inference(resolution,[],[f12037,f5693])).
fof(f14287,plain,(
  ( ! [X21,X20] : (sP9(X20,'const.combin.I_1'(s(bool,'T_0')),X21) | ~p(s(bool,X21)) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl70 | $spl151)),
  inference(resolution,[],[f12037,f5039])).
fof(f14012,plain,(
  ( ! [X1] : (p(s(bool,'const.combin.I_1'(s(bool,'T_0')))) | p(s(bool,X1))) ) | $spl151),
  inference(subsumption_resolution,[],[f13993,f436])).
fof(f13993,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.combin.I_1'(s(bool,'T_0')))) | p(s(bool,X1))) ) | $spl151),
  inference(resolution,[],[f12055,f390])).
fof(f12055,plain,(
  ( ! [X127] : (sP11('const.combin.I_1'(s(bool,'T_0')),X127)) ) | $spl151),
  inference(resolution,[],[f11979,f392])).
fof(f14011,plain,(
  ( ! [X0] : (p(s(bool,'const.combin.I_1'(s(bool,'T_0')))) | p(s(bool,X0))) ) | $spl151),
  inference(subsumption_resolution,[],[f13992,f436])).
fof(f13992,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.combin.I_1'(s(bool,'T_0')))) | p(s(bool,X0))) ) | $spl151),
  inference(resolution,[],[f12055,f394])).
fof(f13959,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.combin.I_1'(s(bool,'T_0'))))) ) | $spl151),
  inference(subsumption_resolution,[],[f13940,f436])).
fof(f13940,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.combin.I_1'(s(bool,'T_0'))))) ) | $spl151),
  inference(resolution,[],[f12054,f390])).
fof(f12054,plain,(
  ( ! [X126] : (sP11(X126,'const.combin.I_1'(s(bool,'T_0')))) ) | $spl151),
  inference(resolution,[],[f11979,f391])).
fof(f13958,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.combin.I_1'(s(bool,'T_0'))))) ) | $spl151),
  inference(subsumption_resolution,[],[f13939,f436])).
fof(f13939,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.combin.I_1'(s(bool,'T_0'))))) ) | $spl151),
  inference(resolution,[],[f12054,f394])).
fof(f13917,plain,(
  ( ! [X18] : (p(s(bool,X18)) | s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl104 | $spl199)),
  inference(resolution,[],[f12705,f5907])).
fof(f13916,plain,(
  ( ! [X17,X16] : (p(s(bool,X16)) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X17) != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl120 | $spl199)),
  inference(resolution,[],[f12705,f6791])).
fof(f13915,plain,(
  ( ! [X14,X15] : (p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | s(bool,X15) != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl110 | $spl199)),
  inference(resolution,[],[f12705,f6440])).
fof(f13914,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | s(bool,'F_0') = s(bool,X13) | s(bool,X13) != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl102 | $spl199)),
  inference(resolution,[],[f12705,f6022])).
fof(f13913,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X11) != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl10 | $spl199)),
  inference(resolution,[],[f12705,f4084])).
fof(f13844,plain,(
  ( ! [X37,X38] : (~p(s(bool,X37)) | s(bool,X38) != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') != s(bool,X38)) ) | ($spl10 | $spl199)),
  inference(resolution,[],[f12702,f3753])).
fof(f13843,plain,(
  ( ! [X35,X36] : (s(bool,'T_0') != s(bool,X35) | sP9(X36,'F_0','const.combin.I_1'(s(bool,'T_0'))) | s(bool,X35) != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl10 | $spl199)),
  inference(resolution,[],[f12702,f3705])).
fof(f13842,plain,(
  ( ! [X34] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | ~p(s(bool,X34))) ) | $spl199),
  inference(resolution,[],[f12702,f2549])).
fof(f13838,plain,(
  ( ! [X30] : (sP9(X30,'F_0','const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | $spl199),
  inference(resolution,[],[f12702,f2413])).
fof(f13837,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X29) != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | ~p(s(bool,X29))) ) | ($spl24 | $spl199)),
  inference(resolution,[],[f12702,f6924])).
fof(f13836,plain,(
  ( ! [X26,X25] : (sP9(X25,X26,'const.combin.I_1'(s(bool,'T_0'))) | p(s(bool,X26)) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl120 | $spl199)),
  inference(resolution,[],[f12702,f6759])).
fof(f13833,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'F_0') != s(bool,X19) | s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl10 | $spl199)),
  inference(resolution,[],[f12702,f4242])).
fof(f13828,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,'const.combin.I_1'(s(bool,'T_0'))) | p(s(bool,X9)) | s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl10 | $spl199)),
  inference(resolution,[],[f12702,f3788])).
fof(f13499,plain,(
  ( ! [X17,X16] : (p(s(bool,X16)) | s(bool,'F_0') != s(bool,X17) | s(bool,X17) != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl98 | $spl151)),
  inference(resolution,[],[f12042,f5593])).
fof(f13498,plain,(
  ( ! [X14,X15] : (p(s(bool,X14)) | s(bool,'T_0') = s(bool,X15) | s(bool,X15) != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl94 | $spl151)),
  inference(resolution,[],[f12042,f5575])).
fof(f13497,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,X13) != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl70 | $spl151)),
  inference(resolution,[],[f12042,f5073])).
fof(f13404,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | ~p(s(bool,X24))) ) | ($spl24 | $spl151)),
  inference(resolution,[],[f12037,f5693])).
fof(f13403,plain,(
  ( ! [X21,X20] : (sP9(X20,'const.combin.I_1'(s(bool,'T_0')),X21) | ~p(s(bool,X21)) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl70 | $spl151)),
  inference(resolution,[],[f12037,f5039])).
fof(f13164,plain,(
  ( ! [X1] : (p(s(bool,'const.combin.I_1'(s(bool,'T_0')))) | p(s(bool,X1))) ) | $spl151),
  inference(subsumption_resolution,[],[f13145,f436])).
fof(f13145,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.combin.I_1'(s(bool,'T_0')))) | p(s(bool,X1))) ) | $spl151),
  inference(resolution,[],[f12055,f390])).
fof(f13163,plain,(
  ( ! [X0] : (p(s(bool,'const.combin.I_1'(s(bool,'T_0')))) | p(s(bool,X0))) ) | $spl151),
  inference(subsumption_resolution,[],[f13144,f436])).
fof(f13144,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.combin.I_1'(s(bool,'T_0')))) | p(s(bool,X0))) ) | $spl151),
  inference(resolution,[],[f12055,f394])).
fof(f13141,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.combin.I_1'(s(bool,'T_0'))))) ) | $spl151),
  inference(subsumption_resolution,[],[f13122,f436])).
fof(f13122,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.combin.I_1'(s(bool,'T_0'))))) ) | $spl151),
  inference(resolution,[],[f12054,f390])).
fof(f13140,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.combin.I_1'(s(bool,'T_0'))))) ) | $spl151),
  inference(subsumption_resolution,[],[f13121,f436])).
fof(f13121,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.combin.I_1'(s(bool,'T_0'))))) ) | $spl151),
  inference(resolution,[],[f12054,f394])).
fof(f13096,plain,(
  ( ! [X37,X38] : (~p(s(bool,X37)) | s(bool,X38) != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') != s(bool,X38)) ) | ($spl10 | $spl199)),
  inference(resolution,[],[f12702,f3753])).
fof(f13095,plain,(
  ( ! [X35,X36] : (s(bool,'T_0') != s(bool,X35) | sP9(X36,'F_0','const.combin.I_1'(s(bool,'T_0'))) | s(bool,X35) != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl10 | $spl199)),
  inference(resolution,[],[f12702,f3705])).
fof(f13094,plain,(
  ( ! [X34] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | ~p(s(bool,X34))) ) | $spl199),
  inference(resolution,[],[f12702,f2549])).
fof(f13090,plain,(
  ( ! [X30] : (sP9(X30,'F_0','const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | $spl199),
  inference(resolution,[],[f12702,f2413])).
fof(f13089,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X29) != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | ~p(s(bool,X29))) ) | ($spl24 | $spl199)),
  inference(resolution,[],[f12702,f6924])).
fof(f13088,plain,(
  ( ! [X26,X25] : (sP9(X25,X26,'const.combin.I_1'(s(bool,'T_0'))) | p(s(bool,X26)) | s(bool,'F_0') = s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl120 | $spl199)),
  inference(resolution,[],[f12702,f6759])).
fof(f13085,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'F_0') != s(bool,X19) | s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl10 | $spl199)),
  inference(resolution,[],[f12702,f4242])).
fof(f13080,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,'const.combin.I_1'(s(bool,'T_0'))) | p(s(bool,X9)) | s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl10 | $spl199)),
  inference(resolution,[],[f12702,f3788])).
fof(f12629,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | ~p(s(bool,X24))) ) | ($spl24 | $spl198)),
  inference(resolution,[],[f12397,f5693])).
fof(f12397,plain,(
  ( ! [X101,X100] : (sP8(X100,X101,'const.combin.I_1'(s(bool,'T_0')))) ) | $spl198),
  inference(resolution,[],[f11985,f345])).
fof(f12628,plain,(
  ( ! [X21,X20] : (sP9(X20,'const.combin.I_1'(s(bool,'T_0')),X21) | ~p(s(bool,X21)) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl70 | $spl198)),
  inference(resolution,[],[f12397,f5039])).
fof(f2666,plain,(
  ( ! [X23] : (s(bool,'T_0') != s(bool,X23)) ) | $spl26),
  inference(cnf_transformation,[],[f2666_D])).
fof(f2666_D,plain,(
  ( ! [X23] : (s(bool,'T_0') != s(bool,X23)) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f8867,plain,(
  s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | $spl151),
  inference(cnf_transformation,[],[f8867_D])).
fof(f8867_D,plain,(
  s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,'T_0'))) <=> ~$spl151),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl151])])).
fof(f12627,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') != s(bool,X19)) ) | ($spl10 | $spl198)),
  inference(resolution,[],[f12397,f4242])).
fof(f12626,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | ~p(s(bool,X16)) | ~p(s(bool,X17))) ) | ($spl10 | $spl198)),
  inference(resolution,[],[f12397,f4239])).
fof(f4239,plain,(
  ( ! [X2,X0,X1] : (~sP8(X1,X0,X2) | s(bool,'F_0') != s(bool,X2) | ~p(s(bool,X0)) | ~p(s(bool,X1))) ) | $spl10),
  inference(subsumption_resolution,[],[f4238,f465])).
fof(f4238,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X1)) | s(bool,'F_0') != s(bool,X2) | ~p(s(bool,X0)) | ~sP8(X1,X0,X2)) ) | $spl10),
  inference(subsumption_resolution,[],[f4237,f3776])).
fof(f4237,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X1)) | s(bool,'F_0') != s(bool,X2) | ~p(s(bool,X0)) | ~sP8(X1,X0,X2) | ~sP6(X0,X2,X1)) ) | $spl10),
  inference(subsumption_resolution,[],[f4206,f3777])).
fof(f4206,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X1)) | s(bool,'F_0') != s(bool,X2) | ~p(s(bool,X0)) | ~sP8(X1,X0,X2) | ~sP7(X2,X0,X1) | ~sP6(X0,X2,X1)) ) | $spl10),
  inference(resolution,[],[f3061,f341])).
fof(f12623,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.combin.I_1'(s(bool,'T_0')),X11) | ~p(s(bool,X11)) | s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl10 | $spl198)),
  inference(resolution,[],[f12397,f3794])).
fof(f12321,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | ~p(s(bool,X24))) ) | ($spl24 | $spl151)),
  inference(resolution,[],[f12037,f5693])).
fof(f12320,plain,(
  ( ! [X21,X20] : (sP9(X20,'const.combin.I_1'(s(bool,'T_0')),X21) | ~p(s(bool,X21)) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'T_0')))) ) | ($spl70 | $spl151)),
  inference(resolution,[],[f12037,f5039])).
fof(f11986,plain,(
  ( ! [X9] : (p(s(bool,'const.combin.I_1'(s(bool,'T_0')))) | p(s(bool,X9))) ) | $spl151),
  inference(subsumption_resolution,[],[f11966,f461])).
fof(f11966,plain,(
  ( ! [X9] : (s(bool,'F_0') != s(bool,X9) | p(s(bool,'const.combin.I_1'(s(bool,'T_0')))) | p(s(bool,X9))) ) | $spl151),
  inference(superposition,[],[f8867,f403])).
fof(f403,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f142])).
fof(f11983,plain,(
  ( ! [X8] : (p(s(bool,X8)) | p(s(bool,'const.combin.I_1'(s(bool,'T_0'))))) ) | $spl151),
  inference(subsumption_resolution,[],[f11965,f461])).
fof(f11965,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,X8) | p(s(bool,X8)) | p(s(bool,'const.combin.I_1'(s(bool,'T_0'))))) ) | $spl151),
  inference(superposition,[],[f8867,f403])).
fof(f11755,plain,(
  ( ! [X14,X15] : (p(s(bool,X14)) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X15) != s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | ($spl94 | $spl149)),
  inference(resolution,[],[f9460,f5575])).
fof(f11754,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | s(bool,'T_0') = s(bool,X13) | s(bool,X13) != s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | ($spl70 | $spl149)),
  inference(resolution,[],[f9460,f5073])).
fof(f11753,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | s(bool,'F_0') != s(bool,X11) | s(bool,X11) != s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | ($spl10 | $spl149)),
  inference(resolution,[],[f9460,f4121])).
fof(f11752,plain,(
  ( ! [X8,X9] : (p(s(bool,X8)) | p(s(bool,X9)) | s(bool,X8) != s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | ($spl10 | $spl149)),
  inference(resolution,[],[f9460,f4011])).
fof(f11722,plain,(
  ( ! [X17,X16] : (p(s(bool,X16)) | s(bool,'F_0') != s(bool,X17) | s(bool,X17) != s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | ($spl98 | $spl149)),
  inference(resolution,[],[f9459,f5593])).
fof(f11721,plain,(
  ( ! [X14,X15] : (p(s(bool,X14)) | s(bool,'T_0') = s(bool,X15) | s(bool,X15) != s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | ($spl94 | $spl149)),
  inference(resolution,[],[f9459,f5575])).
fof(f11720,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X13) != s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | ($spl70 | $spl149)),
  inference(resolution,[],[f9459,f5073])).
fof(f11583,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | ~p(s(bool,X24))) ) | ($spl24 | $spl149)),
  inference(resolution,[],[f9454,f5693])).
fof(f11582,plain,(
  ( ! [X21,X20] : (sP9(X20,'const.combin.I_1'(s(bool,'F_0')),X21) | ~p(s(bool,X21)) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | ($spl70 | $spl149)),
  inference(resolution,[],[f9454,f5039])).
fof(f11323,plain,(
  ( ! [X1] : (p(s(bool,'const.combin.I_1'(s(bool,'F_0')))) | p(s(bool,X1))) ) | $spl149),
  inference(subsumption_resolution,[],[f11300,f436])).
fof(f11300,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.combin.I_1'(s(bool,'F_0')))) | p(s(bool,X1))) ) | $spl149),
  inference(resolution,[],[f9472,f390])).
fof(f11322,plain,(
  ( ! [X0] : (p(s(bool,'const.combin.I_1'(s(bool,'F_0')))) | p(s(bool,X0))) ) | $spl149),
  inference(subsumption_resolution,[],[f11299,f436])).
fof(f11299,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.combin.I_1'(s(bool,'F_0')))) | p(s(bool,X0))) ) | $spl149),
  inference(resolution,[],[f9472,f394])).
fof(f11296,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.combin.I_1'(s(bool,'F_0'))))) ) | $spl149),
  inference(subsumption_resolution,[],[f11273,f436])).
fof(f11273,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.combin.I_1'(s(bool,'F_0'))))) ) | $spl149),
  inference(resolution,[],[f9471,f390])).
fof(f11295,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.combin.I_1'(s(bool,'F_0'))))) ) | $spl149),
  inference(subsumption_resolution,[],[f11272,f436])).
fof(f11272,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.combin.I_1'(s(bool,'F_0'))))) ) | $spl149),
  inference(resolution,[],[f9471,f394])).
fof(f8884,plain,(
  s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | $spl169),
  inference(cnf_transformation,[],[f8884_D])).
fof(f8884_D,plain,(
  s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'F_0'))) <=> ~$spl169),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl169])])).
fof(f11246,plain,(
  ( ! [X18] : (p(s(bool,X18)) | s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | ($spl104 | $spl148)),
  inference(resolution,[],[f9316,f5907])).
fof(f11244,plain,(
  ( ! [X14,X15] : (p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | s(bool,X15) != s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | ($spl110 | $spl148)),
  inference(resolution,[],[f9316,f6440])).
fof(f11243,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | s(bool,'F_0') = s(bool,X13) | s(bool,X13) != s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | ($spl102 | $spl148)),
  inference(resolution,[],[f9316,f6022])).
fof(f11242,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X11) != s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | ($spl10 | $spl148)),
  inference(resolution,[],[f9316,f4084])).
fof(f11165,plain,(
  ( ! [X37,X38] : (~p(s(bool,X37)) | s(bool,X38) != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') != s(bool,X38)) ) | ($spl10 | $spl148)),
  inference(resolution,[],[f9313,f3753])).
fof(f11164,plain,(
  ( ! [X35,X36] : (s(bool,'T_0') != s(bool,X35) | sP9(X36,'F_0','const.combin.I_1'(s(bool,'F_0'))) | s(bool,X35) != s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | ($spl10 | $spl148)),
  inference(resolution,[],[f9313,f3705])).
fof(f11163,plain,(
  ( ! [X34] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | ~p(s(bool,X34))) ) | $spl148),
  inference(resolution,[],[f9313,f2549])).
fof(f11159,plain,(
  ( ! [X30] : (sP9(X30,'F_0','const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | $spl148),
  inference(resolution,[],[f9313,f2413])).
fof(f11158,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X29) != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | ~p(s(bool,X29))) ) | ($spl24 | $spl148)),
  inference(resolution,[],[f9313,f6924])).
fof(f11154,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'F_0') != s(bool,X19) | s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | ($spl10 | $spl148)),
  inference(resolution,[],[f9313,f4242])).
fof(f10334,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,'const.combin.I_1'(s(bool,'F_0'))) | p(s(bool,X9))) ) | $spl182),
  inference(cnf_transformation,[],[f10334_D])).
fof(f10334_D,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,'const.combin.I_1'(s(bool,'F_0'))) | p(s(bool,X9))) ) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f11149,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,'const.combin.I_1'(s(bool,'F_0'))) | p(s(bool,X9)) | s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | ($spl10 | $spl148)),
  inference(resolution,[],[f9313,f3788])).
fof(f10816,plain,(
  ( ! [X17,X16] : (p(s(bool,X16)) | s(bool,'F_0') != s(bool,X17) | s(bool,X17) != s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | ($spl98 | $spl149)),
  inference(resolution,[],[f9459,f5593])).
fof(f10815,plain,(
  ( ! [X14,X15] : (p(s(bool,X14)) | s(bool,'T_0') = s(bool,X15) | s(bool,X15) != s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | ($spl94 | $spl149)),
  inference(resolution,[],[f9459,f5575])).
fof(f10814,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X13) != s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | ($spl70 | $spl149)),
  inference(resolution,[],[f9459,f5073])).
fof(f10709,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | ~p(s(bool,X24))) ) | ($spl24 | $spl149)),
  inference(resolution,[],[f9454,f5693])).
fof(f10708,plain,(
  ( ! [X21,X20] : (sP9(X20,'const.combin.I_1'(s(bool,'F_0')),X21) | ~p(s(bool,X21)) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | ($spl70 | $spl149)),
  inference(resolution,[],[f9454,f5039])).
fof(f10391,plain,(
  ( ! [X1] : (p(s(bool,'const.combin.I_1'(s(bool,'F_0')))) | p(s(bool,X1))) ) | $spl149),
  inference(subsumption_resolution,[],[f10368,f436])).
fof(f10368,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.combin.I_1'(s(bool,'F_0')))) | p(s(bool,X1))) ) | $spl149),
  inference(resolution,[],[f9472,f390])).
fof(f10390,plain,(
  ( ! [X0] : (p(s(bool,'const.combin.I_1'(s(bool,'F_0')))) | p(s(bool,X0))) ) | $spl149),
  inference(subsumption_resolution,[],[f10367,f436])).
fof(f10367,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.combin.I_1'(s(bool,'F_0')))) | p(s(bool,X0))) ) | $spl149),
  inference(resolution,[],[f9472,f394])).
fof(f10364,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.combin.I_1'(s(bool,'F_0'))))) ) | $spl149),
  inference(subsumption_resolution,[],[f10341,f436])).
fof(f10341,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.combin.I_1'(s(bool,'F_0'))))) ) | $spl149),
  inference(resolution,[],[f9471,f390])).
fof(f10363,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.combin.I_1'(s(bool,'F_0'))))) ) | $spl149),
  inference(subsumption_resolution,[],[f10340,f436])).
fof(f10340,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.combin.I_1'(s(bool,'F_0'))))) ) | $spl149),
  inference(resolution,[],[f9471,f394])).
fof(f10311,plain,(
  ( ! [X37,X38] : (~p(s(bool,X37)) | s(bool,X38) != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') != s(bool,X38)) ) | ($spl10 | $spl148)),
  inference(resolution,[],[f9313,f3753])).
fof(f10310,plain,(
  ( ! [X35,X36] : (s(bool,'T_0') != s(bool,X35) | sP9(X36,'F_0','const.combin.I_1'(s(bool,'F_0'))) | s(bool,X35) != s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | ($spl10 | $spl148)),
  inference(resolution,[],[f9313,f3705])).
fof(f10309,plain,(
  ( ! [X34] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | ~p(s(bool,X34))) ) | $spl148),
  inference(resolution,[],[f9313,f2549])).
fof(f10305,plain,(
  ( ! [X30] : (sP9(X30,'F_0','const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | $spl148),
  inference(resolution,[],[f9313,f2413])).
fof(f10304,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X29) != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | ~p(s(bool,X29))) ) | ($spl24 | $spl148)),
  inference(resolution,[],[f9313,f6924])).
fof(f10300,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'F_0') != s(bool,X19) | s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | ($spl10 | $spl148)),
  inference(resolution,[],[f9313,f4242])).
fof(f10295,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,'const.combin.I_1'(s(bool,'F_0'))) | p(s(bool,X9)) | s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | ($spl10 | $spl148)),
  inference(resolution,[],[f9313,f3788])).
fof(f9978,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,'const.combin.I_1'(s(bool,'F_0'))))) ) | $spl148),
  inference(resolution,[],[f9252,f293])).
fof(f293,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP0(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f186])).
fof(f186,plain,(
  ! [X0,X1,X2] : ((~sP0(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP0(X0,X1,X2)))),
  inference(rectify,[],[f185])).
fof(f185,plain,(
  ! [X2,X0,X1] : ((~sP0(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP0(X2,X0,X1)))),
  inference(flattening,[],[f184])).
fof(f184,plain,(
  ! [X2,X0,X1] : ((~sP0(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP0(X2,X0,X1)))),
  inference(nnf_transformation,[],[f157])).
fof(f157,plain,(
  ! [X2,X0,X1] : (sP0(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f9252,plain,(
  ( ! [X26,X27] : (sP0(X26,X27,'const.combin.I_1'(s(bool,'F_0')))) ) | $spl148),
  inference(subsumption_resolution,[],[f8996,f478])).
fof(f8996,plain,(
  ( ! [X26,X27] : (p(s(bool,'F_0')) | sP0(X26,X27,'const.combin.I_1'(s(bool,'F_0')))) ) | $spl148),
  inference(superposition,[],[f294,f8862])).
fof(f294,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f186])).
fof(f10009,plain,(
  ( ! [X8,X9] : (sP1(X8,X9,'const.combin.I_1'(s(bool,'F_0'))) | p(s(bool,X9))) ) | $spl180),
  inference(cnf_transformation,[],[f10009_D])).
fof(f10009_D,plain,(
  ( ! [X8,X9] : (sP1(X8,X9,'const.combin.I_1'(s(bool,'F_0'))) | p(s(bool,X9))) ) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f9982,plain,(
  ( ! [X8,X9] : (sP1(X8,X9,'const.combin.I_1'(s(bool,'F_0'))) | p(s(bool,X9)) | ~p(s(bool,'const.combin.I_1'(s(bool,'F_0'))))) ) | $spl148),
  inference(resolution,[],[f9252,f292])).
fof(f292,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | ~sP0(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f183])).
fof(f183,plain,(
  ! [X0,X1,X2] : ((~sP1(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP0(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP0(X0,X1,X2) | sP1(X0,X1,X2)))),
  inference(rectify,[],[f182])).
fof(f182,plain,(
  ! [X2,X0,X1] : ((~sP1(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP0(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP0(X2,X0,X1) | sP1(X2,X0,X1)))),
  inference(flattening,[],[f181])).
fof(f181,plain,(
  ! [X2,X0,X1] : ((~sP1(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP0(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP0(X2,X0,X1)) | sP1(X2,X0,X1)))),
  inference(nnf_transformation,[],[f158])).
fof(f158,plain,(
  ! [X2,X0,X1] : (sP1(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP0(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f10003,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1))) ) | ($spl148 | $spl170)),
  inference(subsumption_resolution,[],[f9978,f8948])).
fof(f9758,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | ~p(s(bool,X24))) ) | ($spl24 | $spl149)),
  inference(resolution,[],[f9454,f5693])).
fof(f9757,plain,(
  ( ! [X21,X20] : (sP9(X20,'const.combin.I_1'(s(bool,'F_0')),X21) | ~p(s(bool,X21)) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | ($spl70 | $spl149)),
  inference(resolution,[],[f9454,f5039])).
fof(f9376,plain,(
  ( ! [X315,X316,X317] : (s(bool,'F_0') != s(bool,X315) | ~p(s(bool,X316)) | sP9(X317,X316,'const.combin.I_1'(s(bool,'F_0'))) | p(s(bool,X317))) ) | ($spl24 | $spl148)),
  inference(subsumption_resolution,[],[f9196,f460])).
fof(f9196,plain,(
  ( ! [X315,X316,X317] : (s(bool,'F_0') != s(bool,X315) | ~p(s(bool,X316)) | p(s(bool,X315)) | sP9(X317,X316,'const.combin.I_1'(s(bool,'F_0'))) | p(s(bool,X317))) ) | ($spl24 | $spl148)),
  inference(superposition,[],[f4772,f8862])).
fof(f9374,plain,(
  ( ! [X310,X309,X311] : (s(bool,'F_0') != s(bool,X309) | ~p(s(bool,X310)) | sP9(X311,'const.combin.I_1'(s(bool,'F_0')),X310) | p(s(bool,X311))) ) | ($spl24 | $spl148)),
  inference(subsumption_resolution,[],[f9194,f460])).
fof(f9194,plain,(
  ( ! [X310,X309,X311] : (s(bool,'F_0') != s(bool,X309) | ~p(s(bool,X310)) | p(s(bool,X309)) | sP9(X311,'const.combin.I_1'(s(bool,'F_0')),X310) | p(s(bool,X311))) ) | ($spl24 | $spl148)),
  inference(superposition,[],[f4771,f8862])).
fof(f8949,plain,(
  ( ! [X9] : (p(s(bool,'const.combin.I_1'(s(bool,'F_0')))) | p(s(bool,X9))) ) | $spl149),
  inference(subsumption_resolution,[],[f8929,f461])).
fof(f8929,plain,(
  ( ! [X9] : (s(bool,'F_0') != s(bool,X9) | p(s(bool,'const.combin.I_1'(s(bool,'F_0')))) | p(s(bool,X9))) ) | $spl149),
  inference(superposition,[],[f8863,f403])).
fof(f8946,plain,(
  ( ! [X8] : (p(s(bool,X8)) | p(s(bool,'const.combin.I_1'(s(bool,'F_0'))))) ) | $spl149),
  inference(subsumption_resolution,[],[f8928,f461])).
fof(f8928,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,X8) | p(s(bool,X8)) | p(s(bool,'const.combin.I_1'(s(bool,'F_0'))))) ) | $spl149),
  inference(superposition,[],[f8863,f403])).
fof(f8882,plain,(
  ( ! [X34] : (s(bool,X34) != s(bool,'const.combin.I_1'(s(bool,X34))) | p(s(bool,X34))) ) | $spl166),
  inference(cnf_transformation,[],[f8882_D])).
fof(f8882_D,plain,(
  ( ! [X34] : (s(bool,X34) != s(bool,'const.combin.I_1'(s(bool,X34))) | p(s(bool,X34))) ) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f8885,plain,(
  ( ! [X35] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X35) != s(bool,'const.combin.I_1'(s(bool,X35))) | p(s(bool,X35))) ) | $spl10),
  inference(forward_demodulation,[],[f8820,f285])).
fof(f285,plain,(
  ( ! [X0,X1] : (s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f63])).
fof(f63,plain,(
  ! [X0,X1] : s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))),
  inference(flattening,[],[f62])).
fof(f62,plain,(
  ! [X0] : ! [X1] : s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))),
  inference(rectify,[],[f55])).
fof(f55,axiom,(
  ! [X7] : ! [X22] : s(X7,'const.combin.I_1'(s(X7,X22))) = s(X7,happ(s(fun(X7,X7),'const.combin.I_0'),s(X7,X22)))),
  file('thm.quotient.QUOTIENT_SYM','ARITY_0_1_const.combin.I')).
fof(f8820,plain,(
  ( ! [X35] : (s(bool,X35) != s(bool,'const.combin.I_1'(s(bool,X35))) | s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | p(s(bool,X35))) ) | $spl10),
  inference(superposition,[],[f3082,f461])).
fof(f3082,plain,(
  ( ! [X48] : (s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,X48))) | s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X48)))) ) | $spl10),
  inference(superposition,[],[f1974,f285])).
fof(f8881,plain,(
  ( ! [X34] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X34) != s(bool,'const.combin.I_1'(s(bool,X34))) | p(s(bool,X34))) ) | $spl10),
  inference(forward_demodulation,[],[f8819,f285])).
fof(f8819,plain,(
  ( ! [X34] : (s(bool,X34) != s(bool,'const.combin.I_1'(s(bool,X34))) | s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | p(s(bool,X34))) ) | $spl10),
  inference(superposition,[],[f3082,f463])).
fof(f8879,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | $spl163),
  inference(cnf_transformation,[],[f8879_D])).
fof(f8879_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) <=> ~$spl163),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl163])])).
fof(f8880,plain,(
  ( ! [X32] : (s(bool,X32) != s(bool,'const.combin.I_1'(s(bool,X32))) | s(bool,'T_0') = s(bool,X32)) ) | $spl164),
  inference(cnf_transformation,[],[f8880_D])).
fof(f8880_D,plain,(
  ( ! [X32] : (s(bool,X32) != s(bool,'const.combin.I_1'(s(bool,X32))) | s(bool,'T_0') = s(bool,X32)) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f8818,plain,(
  ( ! [X33] : (s(bool,X33) != s(bool,'const.combin.I_1'(s(bool,X33))) | s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X33)) ) | $spl10),
  inference(superposition,[],[f3082,f485])).
fof(f8817,plain,(
  ( ! [X32] : (s(bool,X32) != s(bool,'const.combin.I_1'(s(bool,X32))) | s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X32)) ) | $spl10),
  inference(superposition,[],[f3082,f486])).
fof(f8858,plain,(
  ( ! [X17] : (s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,X17))) | ~p(s(bool,X17))) ) | $spl144),
  inference(cnf_transformation,[],[f8858_D])).
fof(f8858_D,plain,(
  ( ! [X17] : (s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,X17))) | ~p(s(bool,X17))) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f8877,plain,(
  ( ! [X31] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,X31))) | ~p(s(bool,X31))) ) | $spl10),
  inference(forward_demodulation,[],[f8816,f285])).
fof(f8816,plain,(
  ( ! [X31] : (s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,X31))) | s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'T_0'))) | ~p(s(bool,X31))) ) | $spl10),
  inference(superposition,[],[f3082,f465])).
fof(f8874,plain,(
  ( ! [X30] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,X30))) | ~p(s(bool,X30))) ) | $spl10),
  inference(forward_demodulation,[],[f8815,f285])).
fof(f8815,plain,(
  ( ! [X30] : (s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,X30))) | s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'T_0'))) | ~p(s(bool,X30))) ) | $spl10),
  inference(superposition,[],[f3082,f467])).
fof(f8872,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'T_0'))) | $spl157),
  inference(cnf_transformation,[],[f8872_D])).
fof(f8872_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'T_0'))) <=> ~$spl157),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl157])])).
fof(f8873,plain,(
  ( ! [X28] : (s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,X28))) | s(bool,'F_0') = s(bool,X28)) ) | $spl158),
  inference(cnf_transformation,[],[f8873_D])).
fof(f8873_D,plain,(
  ( ! [X28] : (s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,X28))) | s(bool,'F_0') = s(bool,X28)) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f8814,plain,(
  ( ! [X29] : (s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,X29))) | s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X29)) ) | $spl10),
  inference(superposition,[],[f3082,f485])).
fof(f8813,plain,(
  ( ! [X28] : (s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,X28))) | s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X28)) ) | $spl10),
  inference(superposition,[],[f3082,f486])).
fof(f8870,plain,(
  ( ! [X25] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X25))) | s(bool,'T_0') = s(bool,X25)) ) | $spl154),
  inference(cnf_transformation,[],[f8870_D])).
fof(f8870_D,plain,(
  ( ! [X25] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X25))) | s(bool,'T_0') = s(bool,X25)) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f8863,plain,(
  s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | $spl149),
  inference(cnf_transformation,[],[f8863_D])).
fof(f8863_D,plain,(
  s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,'F_0'))) <=> ~$spl149),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl149])])).
fof(f8812,plain,(
  ( ! [X27] : (s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X27))) | s(bool,'T_0') = s(bool,X27)) ) | $spl10),
  inference(superposition,[],[f3082,f486])).
fof(f8869,plain,(
  ( ! [X24] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X24))) | s(bool,'F_0') = s(bool,X24)) ) | $spl152),
  inference(cnf_transformation,[],[f8869_D])).
fof(f8869_D,plain,(
  ( ! [X24] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X24))) | s(bool,'F_0') = s(bool,X24)) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f8811,plain,(
  ( ! [X26] : (s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X26))) | s(bool,'F_0') = s(bool,X26)) ) | $spl10),
  inference(superposition,[],[f3082,f486])).
fof(f8810,plain,(
  ( ! [X25] : (s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X25))) | s(bool,'T_0') = s(bool,X25)) ) | $spl10),
  inference(superposition,[],[f3082,f485])).
fof(f8809,plain,(
  ( ! [X24] : (s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X24))) | s(bool,'F_0') = s(bool,X24)) ) | $spl10),
  inference(superposition,[],[f3082,f485])).
fof(f8859,plain,(
  ( ! [X16] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,X16))) | ~p(s(bool,X16))) ) | $spl146),
  inference(cnf_transformation,[],[f8859_D])).
fof(f8859_D,plain,(
  ( ! [X16] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,X16))) | ~p(s(bool,X16))) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f8868,plain,(
  ( ! [X23] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,X23))) | s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | ~p(s(bool,X23))) ) | $spl10),
  inference(forward_demodulation,[],[f8808,f285])).
fof(f8808,plain,(
  ( ! [X23] : (s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X23))) | ~p(s(bool,X23))) ) | $spl10),
  inference(superposition,[],[f3082,f467])).
fof(f8865,plain,(
  ( ! [X22] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,X22))) | s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | ~p(s(bool,X22))) ) | $spl10),
  inference(forward_demodulation,[],[f8807,f285])).
fof(f8807,plain,(
  ( ! [X22] : (s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X22))) | ~p(s(bool,X22))) ) | $spl10),
  inference(superposition,[],[f3082,f465])).
fof(f8864,plain,(
  ( ! [X21] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,X21))) | s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | p(s(bool,X21))) ) | $spl10),
  inference(forward_demodulation,[],[f8806,f285])).
fof(f8806,plain,(
  ( ! [X21] : (s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X21))) | p(s(bool,X21))) ) | $spl10),
  inference(superposition,[],[f3082,f463])).
fof(f8861,plain,(
  ( ! [X20] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,X20))) | s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | p(s(bool,X20))) ) | $spl10),
  inference(forward_demodulation,[],[f8805,f285])).
fof(f8805,plain,(
  ( ! [X20] : (s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X20))) | p(s(bool,X20))) ) | $spl10),
  inference(superposition,[],[f3082,f461])).
fof(f8860,plain,(
  ( ! [X19,X18] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,X18))) | s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,X19))) | ~p(s(bool,X19)) | ~p(s(bool,X18))) ) | $spl10),
  inference(forward_demodulation,[],[f8804,f285])).
fof(f8804,plain,(
  ( ! [X19,X18] : (s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,X19))) | s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X18))) | ~p(s(bool,X19)) | ~p(s(bool,X18))) ) | $spl10),
  inference(superposition,[],[f3082,f406])).
fof(f8857,plain,(
  ( ! [X17,X16] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,X16))) | s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,X17))) | ~p(s(bool,X16)) | ~p(s(bool,X17))) ) | $spl10),
  inference(forward_demodulation,[],[f8803,f285])).
fof(f8803,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,X17))) | s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X16))) | ~p(s(bool,X16)) | ~p(s(bool,X17))) ) | $spl10),
  inference(superposition,[],[f3082,f406])).
fof(f8854,plain,(
  ( ! [X13] : (s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,X13))) | p(s(bool,X13))) ) | $spl140),
  inference(cnf_transformation,[],[f8854_D])).
fof(f8854_D,plain,(
  ( ! [X13] : (s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,X13))) | p(s(bool,X13))) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f8856,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,X14))) | s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,X15))) | p(s(bool,X14)) | p(s(bool,X15))) ) | $spl10),
  inference(forward_demodulation,[],[f8802,f285])).
fof(f8802,plain,(
  ( ! [X14,X15] : (s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,X15))) | s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X14))) | p(s(bool,X14)) | p(s(bool,X15))) ) | $spl10),
  inference(superposition,[],[f3082,f403])).
fof(f8853,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,X12))) | s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,X13))) | p(s(bool,X13)) | p(s(bool,X12))) ) | $spl10),
  inference(forward_demodulation,[],[f8801,f285])).
fof(f8801,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,X13))) | s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X12))) | p(s(bool,X13)) | p(s(bool,X12))) ) | $spl10),
  inference(superposition,[],[f3082,f403])).
fof(f8665,plain,(
  ( ! [X368,X367] : (p(s(bool,X367)) | p(s(bool,X368)) | sP9(X368,X367,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK18)))) ) | $spl138),
  inference(cnf_transformation,[],[f8665_D])).
fof(f8665_D,plain,(
  ( ! [X368,X367] : (p(s(bool,X367)) | p(s(bool,X368)) | sP9(X368,X367,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK18)))) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f8664,plain,(
  ( ! [X366,X368,X367] : (~p(s(bool,X366)) | p(s(bool,X367)) | sP9(X368,X367,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK18))) | p(s(bool,X368))) ) | ($spl4 | $spl24)),
  inference(subsumption_resolution,[],[f8525,f465])).
fof(f8525,plain,(
  ( ! [X366,X368,X367] : (s(bool,'T_0') != s(bool,X366) | ~p(s(bool,X366)) | p(s(bool,X367)) | sP9(X368,X367,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK18))) | p(s(bool,X368))) ) | ($spl4 | $spl24)),
  inference(superposition,[],[f4800,f1460])).
fof(f8663,plain,(
  ( ! [X362,X361] : (p(s(bool,X361)) | p(s(bool,X362)) | sP9(X362,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK18)),X361)) ) | $spl136),
  inference(cnf_transformation,[],[f8663_D])).
fof(f8663_D,plain,(
  ( ! [X362,X361] : (p(s(bool,X361)) | p(s(bool,X362)) | sP9(X362,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK18)),X361)) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f8662,plain,(
  ( ! [X360,X362,X361] : (~p(s(bool,X360)) | p(s(bool,X361)) | sP9(X362,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK18)),X361) | p(s(bool,X362))) ) | ($spl4 | $spl24)),
  inference(subsumption_resolution,[],[f8523,f465])).
fof(f8523,plain,(
  ( ! [X360,X362,X361] : (s(bool,'T_0') != s(bool,X360) | ~p(s(bool,X360)) | p(s(bool,X361)) | sP9(X362,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK18)),X361) | p(s(bool,X362))) ) | ($spl4 | $spl24)),
  inference(superposition,[],[f4799,f1460])).
fof(f8094,plain,(
  ( ! [X368,X367] : (p(s(bool,X367)) | p(s(bool,X368)) | sP9(X368,X367,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK17)))) ) | $spl134),
  inference(cnf_transformation,[],[f8094_D])).
fof(f8094_D,plain,(
  ( ! [X368,X367] : (p(s(bool,X367)) | p(s(bool,X368)) | sP9(X368,X367,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK17)))) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f8093,plain,(
  ( ! [X366,X368,X367] : (~p(s(bool,X366)) | p(s(bool,X367)) | sP9(X368,X367,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK17))) | p(s(bool,X368))) ) | ($spl2 | $spl24)),
  inference(subsumption_resolution,[],[f7954,f465])).
fof(f7954,plain,(
  ( ! [X366,X368,X367] : (s(bool,'T_0') != s(bool,X366) | ~p(s(bool,X366)) | p(s(bool,X367)) | sP9(X368,X367,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK17))) | p(s(bool,X368))) ) | ($spl2 | $spl24)),
  inference(superposition,[],[f4800,f1337])).
fof(f1337,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK17))) | $spl2),
  inference(resolution,[],[f733,f465])).
fof(f8092,plain,(
  ( ! [X362,X361] : (p(s(bool,X361)) | p(s(bool,X362)) | sP9(X362,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK17)),X361)) ) | $spl132),
  inference(cnf_transformation,[],[f8092_D])).
fof(f8092_D,plain,(
  ( ! [X362,X361] : (p(s(bool,X361)) | p(s(bool,X362)) | sP9(X362,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK17)),X361)) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f8091,plain,(
  ( ! [X360,X362,X361] : (~p(s(bool,X360)) | p(s(bool,X361)) | sP9(X362,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK17)),X361) | p(s(bool,X362))) ) | ($spl2 | $spl24)),
  inference(subsumption_resolution,[],[f7952,f465])).
fof(f7952,plain,(
  ( ! [X360,X362,X361] : (s(bool,'T_0') != s(bool,X360) | ~p(s(bool,X360)) | p(s(bool,X361)) | sP9(X362,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK17)),X361) | p(s(bool,X362))) ) | ($spl2 | $spl24)),
  inference(superposition,[],[f4799,f1337])).
fof(f7428,plain,(
  ( ! [X341,X340] : (~p(s(bool,X340)) | p(s(bool,X341)) | sP9(X341,X340,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK17)))) ) | $spl130),
  inference(cnf_transformation,[],[f7428_D])).
fof(f7428_D,plain,(
  ( ! [X341,X340] : (~p(s(bool,X340)) | p(s(bool,X341)) | sP9(X341,X340,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK17)))) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f7427,plain,(
  ( ! [X341,X339,X340] : (s(bool,'F_0') != s(bool,X339) | ~p(s(bool,X340)) | sP9(X341,X340,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK17))) | p(s(bool,X341))) ) | $spl24),
  inference(subsumption_resolution,[],[f7262,f460])).
fof(f7262,plain,(
  ( ! [X341,X339,X340] : (s(bool,'F_0') != s(bool,X339) | ~p(s(bool,X340)) | p(s(bool,X339)) | sP9(X341,X340,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK17))) | p(s(bool,X341))) ) | $spl24),
  inference(superposition,[],[f4772,f857])).
fof(f857,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK17)))),
  inference(resolution,[],[f278,f461])).
fof(f7426,plain,(
  ( ! [X335,X334] : (~p(s(bool,X334)) | p(s(bool,X335)) | sP9(X335,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK17)),X334)) ) | $spl128),
  inference(cnf_transformation,[],[f7426_D])).
fof(f7426_D,plain,(
  ( ! [X335,X334] : (~p(s(bool,X334)) | p(s(bool,X335)) | sP9(X335,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK17)),X334)) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f7425,plain,(
  ( ! [X333,X335,X334] : (s(bool,'F_0') != s(bool,X333) | ~p(s(bool,X334)) | sP9(X335,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK17)),X334) | p(s(bool,X335))) ) | $spl24),
  inference(subsumption_resolution,[],[f7260,f460])).
fof(f7260,plain,(
  ( ! [X333,X335,X334] : (s(bool,'F_0') != s(bool,X333) | ~p(s(bool,X334)) | p(s(bool,X333)) | sP9(X335,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK17)),X334) | p(s(bool,X335))) ) | $spl24),
  inference(superposition,[],[f4771,f857])).
fof(f4141,plain,(
  ( ! [X33] : (~p(s(bool,X33)) | s(bool,X33) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK17)))) ) | $spl40),
  inference(cnf_transformation,[],[f4141_D])).
fof(f4141_D,plain,(
  ( ! [X33] : (~p(s(bool,X33)) | s(bool,X33) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK17)))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f6947,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X30) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK17))) | ~p(s(bool,X30))) ) | $spl24),
  inference(resolution,[],[f6924,f825])).
fof(f825,plain,(
  ( ! [X101,X102] : (sP8(X101,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK17)),X102)) )),
  inference(resolution,[],[f278,f346])).
fof(f4173,plain,(
  ( ! [X50] : (p(s(bool,X50)) | s(bool,X50) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK18)))) ) | $spl46),
  inference(cnf_transformation,[],[f4173_D])).
fof(f4173_D,plain,(
  ( ! [X50] : (p(s(bool,X50)) | s(bool,X50) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK18)))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f6934,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK18))) | ~p(s(bool,X27))) ) | ($spl4 | $spl24)),
  inference(resolution,[],[f5693,f1428])).
fof(f1428,plain,(
  ( ! [X118,X117] : (sP8(X117,X118,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK18)))) ) | $spl4),
  inference(resolution,[],[f735,f345])).
fof(f4172,plain,(
  ( ! [X48] : (p(s(bool,X48)) | s(bool,X48) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK17)))) ) | $spl44),
  inference(cnf_transformation,[],[f4172_D])).
fof(f4172_D,plain,(
  ( ! [X48] : (p(s(bool,X48)) | s(bool,X48) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK17)))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f6933,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK17))) | ~p(s(bool,X24))) ) | ($spl2 | $spl24)),
  inference(resolution,[],[f5693,f1305])).
fof(f1305,plain,(
  ( ! [X118,X117] : (sP8(X117,X118,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK17)))) ) | $spl2),
  inference(resolution,[],[f733,f345])).
fof(f6518,plain,(
  ( ! [X39] : (s(bool,'F_0') = s(bool,X39) | s(bool,X39) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK17)))) ) | $spl116),
  inference(cnf_transformation,[],[f6518_D])).
fof(f6518_D,plain,(
  ( ! [X39] : (s(bool,'F_0') = s(bool,X39) | s(bool,X39) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK17)))) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f6810,plain,(
  ( ! [X47,X48] : (p(s(bool,X47)) | s(bool,'F_0') = s(bool,X48) | s(bool,X48) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK17)))) ) | $spl120),
  inference(resolution,[],[f6791,f829])).
fof(f829,plain,(
  ( ! [X109,X110] : (sP7(happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK17)),X109,X110)) )),
  inference(resolution,[],[f278,f351])).
fof(f4492,plain,(
  ( ! [X30] : (s(bool,X30) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK17)))) ) | $spl52),
  inference(cnf_transformation,[],[f4492_D])).
fof(f4492_D,plain,(
  ( ! [X30] : (s(bool,X30) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK17)))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f6520,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK17))) | $spl118),
  inference(cnf_transformation,[],[f6520_D])).
fof(f6520_D,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK17))) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f6806,plain,(
  ( ! [X39,X38] : (p(s(bool,X38)) | s(bool,'F_0') = s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK17))) | s(bool,X39) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK17)))) ) | $spl120),
  inference(resolution,[],[f6791,f828])).
fof(f828,plain,(
  ( ! [X107,X108] : (sP7(X107,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK17)),X108)) )),
  inference(resolution,[],[f278,f350])).
fof(f3824,plain,(
  ( ! [X21,X20] : (sP9(X20,X21,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK17))) | p(s(bool,X21))) ) | $spl30),
  inference(cnf_transformation,[],[f3824_D])).
fof(f3824_D,plain,(
  ( ! [X21,X20] : (sP9(X20,X21,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK17))) | p(s(bool,X21))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f6769,plain,(
  ( ! [X21,X20] : (sP9(X20,X21,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK17))) | p(s(bool,X21)) | s(bool,'F_0') = s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK17)))) ) | $spl120),
  inference(resolution,[],[f6759,f825])).
fof(f2558,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK17))) | $spl15),
  inference(cnf_transformation,[],[f2558_D])).
fof(f2558_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK17))) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f6718,plain,(
  ( ! [X24] : (p(s(bool,X24)) | s(bool,'T_0') != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK17)))) ) | $spl122),
  inference(resolution,[],[f6702,f829])).
fof(f6626,plain,(
  ( ! [X112,X113] : (p(s(bool,X112)) | p(s(bool,X113)) | sP9(X113,X112,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK18)))) ) | $spl126),
  inference(cnf_transformation,[],[f6626_D])).
fof(f6626_D,plain,(
  ( ! [X112,X113] : (p(s(bool,X112)) | p(s(bool,X113)) | sP9(X113,X112,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK18)))) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f6625,plain,(
  ( ! [X111,X112,X113] : (~p(s(bool,X111)) | p(s(bool,X112)) | sP9(X113,X112,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK18))) | p(s(bool,X113))) ) | $spl24),
  inference(subsumption_resolution,[],[f6571,f465])).
fof(f6571,plain,(
  ( ! [X111,X112,X113] : (s(bool,'T_0') != s(bool,X111) | ~p(s(bool,X111)) | p(s(bool,X112)) | sP9(X113,X112,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK18))) | p(s(bool,X113))) ) | $spl24),
  inference(superposition,[],[f4800,f700])).
fof(f700,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK18)))),
  inference(resolution,[],[f277,f465])).
fof(f6624,plain,(
  ( ! [X109,X110] : (p(s(bool,X109)) | p(s(bool,X110)) | sP9(X110,X109,'const.quotient.QUOTIENT_3'(s(fun(sK13,fun(sK13,bool)),sK14),s(fun(sK13,sK12),sK15),s(fun(sK12,sK13),sK16)))) ) | $spl124),
  inference(cnf_transformation,[],[f6624_D])).
fof(f6624_D,plain,(
  ( ! [X109,X110] : (p(s(bool,X109)) | p(s(bool,X110)) | sP9(X110,X109,'const.quotient.QUOTIENT_3'(s(fun(sK13,fun(sK13,bool)),sK14),s(fun(sK13,sK12),sK15),s(fun(sK12,sK13),sK16)))) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f6623,plain,(
  ( ! [X109,X110,X108] : (~p(s(bool,X108)) | p(s(bool,X109)) | sP9(X110,X109,'const.quotient.QUOTIENT_3'(s(fun(sK13,fun(sK13,bool)),sK14),s(fun(sK13,sK12),sK15),s(fun(sK12,sK13),sK16))) | p(s(bool,X110))) ) | $spl24),
  inference(subsumption_resolution,[],[f6570,f465])).
fof(f6570,plain,(
  ( ! [X109,X110,X108] : (s(bool,'T_0') != s(bool,X108) | ~p(s(bool,X108)) | p(s(bool,X109)) | sP9(X110,X109,'const.quotient.QUOTIENT_3'(s(fun(sK13,fun(sK13,bool)),sK14),s(fun(sK13,sK12),sK15),s(fun(sK12,sK13),sK16))) | p(s(bool,X110))) ) | $spl24),
  inference(superposition,[],[f4800,f581])).
fof(f6617,plain,(
  ( ! [X78,X79] : (sP9(X79,X78,'T_0') | p(s(bool,X79)) | p(s(bool,X78))) ) | $spl122),
  inference(cnf_transformation,[],[f6617_D])).
fof(f6617_D,plain,(
  ( ! [X78,X79] : (sP9(X79,X78,'T_0') | p(s(bool,X79)) | p(s(bool,X78))) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f6620,plain,(
  ( ! [X90,X88,X91,X89] : (~p(s(bool,X89)) | p(s(bool,X90)) | sP9(X91,X90,'T_0') | p(s(bool,X91)) | ~p(s(bool,X88))) ) | $spl24),
  inference(subsumption_resolution,[],[f6565,f406])).
fof(f6565,plain,(
  ( ! [X90,X88,X91,X89] : (s(bool,X88) != s(bool,X89) | ~p(s(bool,X89)) | p(s(bool,X90)) | sP9(X91,X90,'T_0') | p(s(bool,X91)) | ~p(s(bool,X88))) ) | $spl24),
  inference(superposition,[],[f4800,f465])).
fof(f6619,plain,(
  ( ! [X87,X85,X86,X84] : (~p(s(bool,X85)) | p(s(bool,X86)) | sP9(X87,X86,'T_0') | p(s(bool,X87)) | ~p(s(bool,X84))) ) | $spl24),
  inference(subsumption_resolution,[],[f6564,f406])).
fof(f6564,plain,(
  ( ! [X87,X85,X86,X84] : (s(bool,X84) != s(bool,X85) | ~p(s(bool,X85)) | p(s(bool,X86)) | sP9(X87,X86,'T_0') | p(s(bool,X87)) | ~p(s(bool,X84))) ) | $spl24),
  inference(superposition,[],[f4800,f467])).
fof(f5838,plain,(
  ( ! [X76,X77] : (s(bool,X76) != s(bool,X77) | ~p(s(bool,X77))) ) | $spl106),
  inference(cnf_transformation,[],[f5838_D])).
fof(f5838_D,plain,(
  ( ! [X76,X77] : (s(bool,X76) != s(bool,X77) | ~p(s(bool,X77))) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f6618,plain,(
  ( ! [X80,X83,X81,X82] : (s(bool,X80) != s(bool,X81) | ~p(s(bool,X81)) | p(s(bool,X82)) | sP9(X83,X82,'T_0') | p(s(bool,X83))) ) | ($spl24 | $spl38)),
  inference(subsumption_resolution,[],[f6563,f3940])).
fof(f6563,plain,(
  ( ! [X80,X83,X81,X82] : (s(bool,X80) != s(bool,X81) | ~p(s(bool,X81)) | p(s(bool,X82)) | sP9(X83,X82,'T_0') | p(s(bool,X83)) | s(bool,'F_0') = s(bool,X80)) ) | $spl24),
  inference(superposition,[],[f4800,f485])).
fof(f6616,plain,(
  ( ! [X78,X76,X79,X77] : (s(bool,X76) != s(bool,X77) | ~p(s(bool,X77)) | p(s(bool,X78)) | sP9(X79,X78,'T_0') | p(s(bool,X79))) ) | ($spl24 | $spl38)),
  inference(subsumption_resolution,[],[f6562,f3940])).
fof(f6562,plain,(
  ( ! [X78,X76,X79,X77] : (s(bool,X76) != s(bool,X77) | ~p(s(bool,X77)) | p(s(bool,X78)) | sP9(X79,X78,'T_0') | p(s(bool,X79)) | s(bool,'F_0') = s(bool,X76)) ) | $spl24),
  inference(superposition,[],[f4800,f486])).
fof(f6614,plain,(
  ( ! [X62,X60,X63] : (sP9(X63,X62,X60) | s(bool,'F_0') = s(bool,X60) | p(s(bool,X63)) | p(s(bool,X62))) ) | $spl120),
  inference(cnf_transformation,[],[f6614_D])).
fof(f6614_D,plain,(
  ( ! [X62,X60,X63] : (sP9(X63,X62,X60) | s(bool,'F_0') = s(bool,X60) | p(s(bool,X63)) | p(s(bool,X62))) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f6615,plain,(
  ( ! [X70,X68,X71,X69] : (~p(s(bool,X69)) | p(s(bool,X70)) | sP9(X71,X70,X68) | p(s(bool,X71)) | s(bool,'F_0') = s(bool,X68)) ) | $spl24),
  inference(subsumption_resolution,[],[f6560,f465])).
fof(f6560,plain,(
  ( ! [X70,X68,X71,X69] : (s(bool,'T_0') != s(bool,X69) | ~p(s(bool,X69)) | p(s(bool,X70)) | sP9(X71,X70,X68) | p(s(bool,X71)) | s(bool,'F_0') = s(bool,X68)) ) | $spl24),
  inference(superposition,[],[f4800,f486])).
fof(f6613,plain,(
  ( ! [X61,X62,X60,X63] : (~p(s(bool,X61)) | p(s(bool,X62)) | sP9(X63,X62,X60) | p(s(bool,X63)) | s(bool,'F_0') = s(bool,X60)) ) | $spl24),
  inference(subsumption_resolution,[],[f6558,f465])).
fof(f6558,plain,(
  ( ! [X61,X62,X60,X63] : (s(bool,'T_0') != s(bool,X61) | ~p(s(bool,X61)) | p(s(bool,X62)) | sP9(X63,X62,X60) | p(s(bool,X63)) | s(bool,'F_0') = s(bool,X60)) ) | $spl24),
  inference(superposition,[],[f4800,f485])).
fof(f6539,plain,(
  ( ! [X47,X48] : (p(s(bool,X47)) | s(bool,'T_0') != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK17))) | s(bool,X48) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK17)))) ) | $spl110),
  inference(resolution,[],[f6440,f829])).
fof(f4411,plain,(
  ( ! [X13] : (s(bool,X13) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK17))) | s(bool,'T_0') != s(bool,X13)) ) | $spl48),
  inference(cnf_transformation,[],[f4411_D])).
fof(f4411_D,plain,(
  ( ! [X13] : (s(bool,X13) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK17))) | s(bool,'T_0') != s(bool,X13)) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f6535,plain,(
  ( ! [X39,X38] : (p(s(bool,X38)) | s(bool,'T_0') != s(bool,X39) | s(bool,X39) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK17)))) ) | $spl110),
  inference(resolution,[],[f6440,f828])).
fof(f6513,plain,(
  ( ! [X47,X48] : (p(s(bool,X47)) | s(bool,'F_0') = s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK17))) | s(bool,X48) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK17)))) ) | $spl102),
  inference(resolution,[],[f6022,f829])).
fof(f6509,plain,(
  ( ! [X39,X38] : (p(s(bool,X38)) | s(bool,'F_0') = s(bool,X39) | s(bool,X39) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK17)))) ) | $spl102),
  inference(resolution,[],[f6022,f828])).
fof(f6399,plain,(
  ( ! [X346,X347] : (p(s(bool,X346)) | p(s(bool,X347)) | sP9(X347,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK18)),X346)) ) | $spl114),
  inference(cnf_transformation,[],[f6399_D])).
fof(f6399_D,plain,(
  ( ! [X346,X347] : (p(s(bool,X346)) | p(s(bool,X347)) | sP9(X347,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK18)),X346)) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f6398,plain,(
  ( ! [X346,X345,X347] : (~p(s(bool,X345)) | p(s(bool,X346)) | sP9(X347,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK18)),X346) | p(s(bool,X347))) ) | $spl24),
  inference(subsumption_resolution,[],[f6269,f465])).
fof(f6269,plain,(
  ( ! [X346,X345,X347] : (s(bool,'T_0') != s(bool,X345) | ~p(s(bool,X345)) | p(s(bool,X346)) | sP9(X347,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK18)),X346) | p(s(bool,X347))) ) | $spl24),
  inference(superposition,[],[f4799,f700])).
fof(f5966,plain,(
  ( ! [X26,X24] : (s(bool,'F_0') != s(bool,X26) | s(bool,X24) != s(bool,X26) | s(bool,'F_0') = s(bool,X24)) ) | $spl112),
  inference(cnf_transformation,[],[f5966_D])).
fof(f5966_D,plain,(
  ( ! [X26,X24] : (s(bool,'F_0') != s(bool,X26) | s(bool,X24) != s(bool,X26) | s(bool,'F_0') = s(bool,X24)) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f5961,plain,(
  ( ! [X26,X24,X25] : (s(bool,'F_0') = s(bool,X24) | p(s(bool,X25)) | s(bool,'F_0') != s(bool,X26) | s(bool,X24) != s(bool,X26)) ) | ($spl98 | $spl102)),
  inference(resolution,[],[f5952,f5593])).
fof(f5922,plain,(
  ( ! [X24] : (p(s(bool,X24)) | s(bool,'T_0') != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK17)))) ) | $spl104),
  inference(resolution,[],[f5907,f828])).
fof(f5846,plain,(
  ( ! [X54,X55,X53] : (sP9(X55,X53,X54) | p(s(bool,X55)) | s(bool,'T_0') != s(bool,X53) | p(s(bool,X54))) ) | $spl110),
  inference(cnf_transformation,[],[f5846_D])).
fof(f5846_D,plain,(
  ( ! [X54,X55,X53] : (sP9(X55,X53,X54) | p(s(bool,X55)) | s(bool,'T_0') != s(bool,X53) | p(s(bool,X54))) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f5848,plain,(
  ( ! [X70,X68,X71,X69] : (s(bool,'T_0') != s(bool,X69) | ~p(s(bool,X68)) | p(s(bool,X70)) | sP9(X71,X69,X70) | p(s(bool,X71))) ) | $spl24),
  inference(subsumption_resolution,[],[f5810,f460])).
fof(f5810,plain,(
  ( ! [X70,X68,X71,X69] : (s(bool,'T_0') != s(bool,X69) | ~p(s(bool,X68)) | p(s(bool,X70)) | sP9(X71,X69,X70) | p(s(bool,X71)) | s(bool,'F_0') = s(bool,X68)) ) | $spl24),
  inference(superposition,[],[f4799,f486])).
fof(f5847,plain,(
  ( ! [X61,X62,X60,X63] : (s(bool,'T_0') != s(bool,X61) | ~p(s(bool,X60)) | p(s(bool,X62)) | sP9(X63,X61,X62) | p(s(bool,X63))) ) | $spl24),
  inference(subsumption_resolution,[],[f5808,f460])).
fof(f5808,plain,(
  ( ! [X61,X62,X60,X63] : (s(bool,'T_0') != s(bool,X61) | ~p(s(bool,X60)) | p(s(bool,X62)) | sP9(X63,X61,X62) | p(s(bool,X63)) | s(bool,'F_0') = s(bool,X60)) ) | $spl24),
  inference(superposition,[],[f4799,f485])).
fof(f5825,plain,(
  ( ! [X59,X57,X58,X56] : (s(bool,'T_0') != s(bool,X57) | ~p(s(bool,X56)) | p(s(bool,X58)) | sP9(X59,X57,X58) | p(s(bool,X59))) ) | $spl24),
  inference(duplicate_literal_removal,[],[f5807])).
fof(f5807,plain,(
  ( ! [X59,X57,X58,X56] : (s(bool,'T_0') != s(bool,X57) | ~p(s(bool,X56)) | p(s(bool,X58)) | sP9(X59,X57,X58) | p(s(bool,X59)) | ~p(s(bool,X56))) ) | $spl24),
  inference(superposition,[],[f4799,f467])).
fof(f5826,plain,(
  ( ! [X54,X52,X55,X53] : (s(bool,'T_0') != s(bool,X53) | ~p(s(bool,X52)) | p(s(bool,X54)) | sP9(X55,X53,X54) | p(s(bool,X55))) ) | $spl24),
  inference(duplicate_literal_removal,[],[f5806])).
fof(f5806,plain,(
  ( ! [X54,X52,X55,X53] : (s(bool,'T_0') != s(bool,X53) | ~p(s(bool,X52)) | p(s(bool,X54)) | sP9(X55,X53,X54) | p(s(bool,X55)) | ~p(s(bool,X52))) ) | $spl24),
  inference(superposition,[],[f4799,f465])).
fof(f5845,plain,(
  ( ! [X109,X110] : (p(s(bool,X109)) | p(s(bool,X110)) | sP9(X110,'const.quotient.QUOTIENT_3'(s(fun(sK13,fun(sK13,bool)),sK14),s(fun(sK13,sK12),sK15),s(fun(sK12,sK13),sK16)),X109)) ) | $spl108),
  inference(cnf_transformation,[],[f5845_D])).
fof(f5845_D,plain,(
  ( ! [X109,X110] : (p(s(bool,X109)) | p(s(bool,X110)) | sP9(X110,'const.quotient.QUOTIENT_3'(s(fun(sK13,fun(sK13,bool)),sK14),s(fun(sK13,sK12),sK15),s(fun(sK12,sK13),sK16)),X109)) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f5844,plain,(
  ( ! [X109,X110,X108] : (~p(s(bool,X108)) | p(s(bool,X109)) | sP9(X110,'const.quotient.QUOTIENT_3'(s(fun(sK13,fun(sK13,bool)),sK14),s(fun(sK13,sK12),sK15),s(fun(sK12,sK13),sK16)),X109) | p(s(bool,X110))) ) | $spl24),
  inference(subsumption_resolution,[],[f5792,f465])).
fof(f5792,plain,(
  ( ! [X109,X110,X108] : (s(bool,'T_0') != s(bool,X108) | ~p(s(bool,X108)) | p(s(bool,X109)) | sP9(X110,'const.quotient.QUOTIENT_3'(s(fun(sK13,fun(sK13,bool)),sK14),s(fun(sK13,sK12),sK15),s(fun(sK12,sK13),sK16)),X109) | p(s(bool,X110))) ) | $spl24),
  inference(superposition,[],[f4799,f581])).
fof(f5837,plain,(
  ( ! [X78,X79] : (sP9(X79,'T_0',X78) | p(s(bool,X79)) | p(s(bool,X78))) ) | $spl104),
  inference(cnf_transformation,[],[f5837_D])).
fof(f5837_D,plain,(
  ( ! [X78,X79] : (sP9(X79,'T_0',X78) | p(s(bool,X79)) | p(s(bool,X78))) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f5841,plain,(
  ( ! [X90,X88,X91,X89] : (~p(s(bool,X89)) | p(s(bool,X90)) | sP9(X91,'T_0',X90) | p(s(bool,X91)) | ~p(s(bool,X88))) ) | $spl24),
  inference(subsumption_resolution,[],[f5787,f406])).
fof(f5787,plain,(
  ( ! [X90,X88,X91,X89] : (s(bool,X88) != s(bool,X89) | ~p(s(bool,X89)) | p(s(bool,X90)) | sP9(X91,'T_0',X90) | p(s(bool,X91)) | ~p(s(bool,X88))) ) | $spl24),
  inference(superposition,[],[f4799,f465])).
fof(f5840,plain,(
  ( ! [X87,X85,X86,X84] : (~p(s(bool,X85)) | p(s(bool,X86)) | sP9(X87,'T_0',X86) | p(s(bool,X87)) | ~p(s(bool,X84))) ) | $spl24),
  inference(subsumption_resolution,[],[f5786,f406])).
fof(f5786,plain,(
  ( ! [X87,X85,X86,X84] : (s(bool,X84) != s(bool,X85) | ~p(s(bool,X85)) | p(s(bool,X86)) | sP9(X87,'T_0',X86) | p(s(bool,X87)) | ~p(s(bool,X84))) ) | $spl24),
  inference(superposition,[],[f4799,f467])).
fof(f5839,plain,(
  ( ! [X80,X83,X81,X82] : (s(bool,X80) != s(bool,X81) | ~p(s(bool,X81)) | p(s(bool,X82)) | sP9(X83,'T_0',X82) | p(s(bool,X83))) ) | ($spl24 | $spl38)),
  inference(subsumption_resolution,[],[f5785,f3940])).
fof(f5785,plain,(
  ( ! [X80,X83,X81,X82] : (s(bool,X80) != s(bool,X81) | ~p(s(bool,X81)) | p(s(bool,X82)) | sP9(X83,'T_0',X82) | p(s(bool,X83)) | s(bool,'F_0') = s(bool,X80)) ) | $spl24),
  inference(superposition,[],[f4799,f485])).
fof(f5836,plain,(
  ( ! [X78,X76,X79,X77] : (s(bool,X76) != s(bool,X77) | ~p(s(bool,X77)) | p(s(bool,X78)) | sP9(X79,'T_0',X78) | p(s(bool,X79))) ) | ($spl24 | $spl38)),
  inference(subsumption_resolution,[],[f5784,f3940])).
fof(f5784,plain,(
  ( ! [X78,X76,X79,X77] : (s(bool,X76) != s(bool,X77) | ~p(s(bool,X77)) | p(s(bool,X78)) | sP9(X79,'T_0',X78) | p(s(bool,X79)) | s(bool,'F_0') = s(bool,X76)) ) | $spl24),
  inference(superposition,[],[f4799,f486])).
fof(f5834,plain,(
  ( ! [X62,X60,X63] : (sP9(X63,X60,X62) | s(bool,'F_0') = s(bool,X60) | p(s(bool,X63)) | p(s(bool,X62))) ) | $spl102),
  inference(cnf_transformation,[],[f5834_D])).
fof(f5834_D,plain,(
  ( ! [X62,X60,X63] : (sP9(X63,X60,X62) | s(bool,'F_0') = s(bool,X60) | p(s(bool,X63)) | p(s(bool,X62))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f5835,plain,(
  ( ! [X70,X68,X71,X69] : (~p(s(bool,X69)) | p(s(bool,X70)) | sP9(X71,X68,X70) | p(s(bool,X71)) | s(bool,'F_0') = s(bool,X68)) ) | $spl24),
  inference(subsumption_resolution,[],[f5782,f465])).
fof(f5782,plain,(
  ( ! [X70,X68,X71,X69] : (s(bool,'T_0') != s(bool,X69) | ~p(s(bool,X69)) | p(s(bool,X70)) | sP9(X71,X68,X70) | p(s(bool,X71)) | s(bool,'F_0') = s(bool,X68)) ) | $spl24),
  inference(superposition,[],[f4799,f486])).
fof(f5833,plain,(
  ( ! [X61,X62,X60,X63] : (~p(s(bool,X61)) | p(s(bool,X62)) | sP9(X63,X60,X62) | p(s(bool,X63)) | s(bool,'F_0') = s(bool,X60)) ) | $spl24),
  inference(subsumption_resolution,[],[f5780,f465])).
fof(f5780,plain,(
  ( ! [X61,X62,X60,X63] : (s(bool,'T_0') != s(bool,X61) | ~p(s(bool,X61)) | p(s(bool,X62)) | sP9(X63,X60,X62) | p(s(bool,X63)) | s(bool,'F_0') = s(bool,X60)) ) | $spl24),
  inference(superposition,[],[f4799,f485])).
fof(f4525,plain,(
  ( ! [X41] : (s(bool,X41) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK18)))) ) | $spl58),
  inference(cnf_transformation,[],[f4525_D])).
fof(f4525_D,plain,(
  ( ! [X41] : (s(bool,X41) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK18)))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f2589,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK18))) | $spl23),
  inference(cnf_transformation,[],[f2589_D])).
fof(f2589_D,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK18))) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f5641,plain,(
  ( ! [X54,X53] : (p(s(bool,X53)) | s(bool,'F_0') != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK18))) | s(bool,X54) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK18)))) ) | ($spl4 | $spl98)),
  inference(resolution,[],[f5593,f1434])).
fof(f1434,plain,(
  ( ! [X130,X129] : (sP6(happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK18)),X129,X130)) ) | $spl4),
  inference(resolution,[],[f735,f355])).
fof(f4524,plain,(
  ( ! [X39] : (s(bool,X39) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK17)))) ) | $spl56),
  inference(cnf_transformation,[],[f4524_D])).
fof(f4524_D,plain,(
  ( ! [X39] : (s(bool,X39) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK17)))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f2587,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK17))) | $spl21),
  inference(cnf_transformation,[],[f2587_D])).
fof(f2587_D,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK17))) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f5640,plain,(
  ( ! [X52,X51] : (p(s(bool,X51)) | s(bool,'F_0') != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK17))) | s(bool,X52) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK17)))) ) | ($spl2 | $spl98)),
  inference(resolution,[],[f5593,f1311])).
fof(f1311,plain,(
  ( ! [X130,X129] : (sP6(happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK17)),X129,X130)) ) | $spl2),
  inference(resolution,[],[f733,f355])).
fof(f4523,plain,(
  ( ! [X37] : (s(bool,X37) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK18)))) ) | $spl54),
  inference(cnf_transformation,[],[f4523_D])).
fof(f4523_D,plain,(
  ( ! [X37] : (s(bool,X37) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK18)))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f2585,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK18))) | $spl19),
  inference(cnf_transformation,[],[f2585_D])).
fof(f2585_D,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK18))) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f5639,plain,(
  ( ! [X50,X49] : (p(s(bool,X49)) | s(bool,'F_0') != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK18))) | s(bool,X50) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK18)))) ) | $spl98),
  inference(resolution,[],[f5593,f674])).
fof(f674,plain,(
  ( ! [X121,X122] : (sP6(happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK18)),X121,X122)) )),
  inference(resolution,[],[f277,f355])).
fof(f4529,plain,(
  ( ! [X54] : (s(bool,'F_0') != s(bool,X54) | s(bool,X54) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK18)))) ) | $spl64),
  inference(cnf_transformation,[],[f4529_D])).
fof(f4529_D,plain,(
  ( ! [X54] : (s(bool,'F_0') != s(bool,X54) | s(bool,X54) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK18)))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f5635,plain,(
  ( ! [X41,X40] : (p(s(bool,X40)) | s(bool,'F_0') != s(bool,X41) | s(bool,X41) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK18)))) ) | ($spl4 | $spl98)),
  inference(resolution,[],[f5593,f1433])).
fof(f1433,plain,(
  ( ! [X127,X128] : (sP6(X127,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK18)),X128)) ) | $spl4),
  inference(resolution,[],[f735,f354])).
fof(f4528,plain,(
  ( ! [X52] : (s(bool,'F_0') != s(bool,X52) | s(bool,X52) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK17)))) ) | $spl62),
  inference(cnf_transformation,[],[f4528_D])).
fof(f4528_D,plain,(
  ( ! [X52] : (s(bool,'F_0') != s(bool,X52) | s(bool,X52) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK17)))) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f5634,plain,(
  ( ! [X39,X38] : (p(s(bool,X38)) | s(bool,'F_0') != s(bool,X39) | s(bool,X39) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK17)))) ) | ($spl2 | $spl98)),
  inference(resolution,[],[f5593,f1310])).
fof(f1310,plain,(
  ( ! [X127,X128] : (sP6(X127,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK17)),X128)) ) | $spl2),
  inference(resolution,[],[f733,f354])).
fof(f4527,plain,(
  ( ! [X50] : (s(bool,'F_0') != s(bool,X50) | s(bool,X50) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK18)))) ) | $spl60),
  inference(cnf_transformation,[],[f4527_D])).
fof(f4527_D,plain,(
  ( ! [X50] : (s(bool,'F_0') != s(bool,X50) | s(bool,X50) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK18)))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f5633,plain,(
  ( ! [X37,X36] : (p(s(bool,X36)) | s(bool,'F_0') != s(bool,X37) | s(bool,X37) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK18)))) ) | $spl98),
  inference(resolution,[],[f5593,f673])).
fof(f673,plain,(
  ( ! [X120,X119] : (sP6(X119,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK18)),X120)) )),
  inference(resolution,[],[f277,f354])).
fof(f5056,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK18))) | $spl84),
  inference(cnf_transformation,[],[f5056_D])).
fof(f5056_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK18))) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f5616,plain,(
  ( ! [X54,X53] : (p(s(bool,X53)) | s(bool,'T_0') = s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK18))) | s(bool,X54) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK18)))) ) | ($spl4 | $spl94)),
  inference(resolution,[],[f5575,f1434])).
fof(f5054,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK17))) | $spl82),
  inference(cnf_transformation,[],[f5054_D])).
fof(f5054_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK17))) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f5615,plain,(
  ( ! [X52,X51] : (p(s(bool,X51)) | s(bool,'T_0') = s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK17))) | s(bool,X52) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK17)))) ) | ($spl2 | $spl94)),
  inference(resolution,[],[f5575,f1311])).
fof(f5052,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK18))) | $spl80),
  inference(cnf_transformation,[],[f5052_D])).
fof(f5052_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK18))) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f5614,plain,(
  ( ! [X50,X49] : (p(s(bool,X49)) | s(bool,'T_0') = s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK18))) | s(bool,X50) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK18)))) ) | $spl94),
  inference(resolution,[],[f5575,f674])).
fof(f5101,plain,(
  ( ! [X54] : (s(bool,'T_0') = s(bool,X54) | s(bool,X54) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK18)))) ) | $spl90),
  inference(cnf_transformation,[],[f5101_D])).
fof(f5101_D,plain,(
  ( ! [X54] : (s(bool,'T_0') = s(bool,X54) | s(bool,X54) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK18)))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f5610,plain,(
  ( ! [X41,X40] : (p(s(bool,X40)) | s(bool,'T_0') = s(bool,X41) | s(bool,X41) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK18)))) ) | ($spl4 | $spl94)),
  inference(resolution,[],[f5575,f1433])).
fof(f5100,plain,(
  ( ! [X52] : (s(bool,'T_0') = s(bool,X52) | s(bool,X52) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK17)))) ) | $spl88),
  inference(cnf_transformation,[],[f5100_D])).
fof(f5100_D,plain,(
  ( ! [X52] : (s(bool,'T_0') = s(bool,X52) | s(bool,X52) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK17)))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f5609,plain,(
  ( ! [X39,X38] : (p(s(bool,X38)) | s(bool,'T_0') = s(bool,X39) | s(bool,X39) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK17)))) ) | ($spl2 | $spl94)),
  inference(resolution,[],[f5575,f1310])).
fof(f5099,plain,(
  ( ! [X50] : (s(bool,'T_0') = s(bool,X50) | s(bool,X50) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK18)))) ) | $spl86),
  inference(cnf_transformation,[],[f5099_D])).
fof(f5099_D,plain,(
  ( ! [X50] : (s(bool,'T_0') = s(bool,X50) | s(bool,X50) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK18)))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f5608,plain,(
  ( ! [X37,X36] : (p(s(bool,X36)) | s(bool,'T_0') = s(bool,X37) | s(bool,X37) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK18)))) ) | $spl94),
  inference(resolution,[],[f5575,f673])).
fof(f5470,plain,(
  ( ! [X29] : (p(s(bool,X29)) | s(bool,'F_0') != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK18)))) ) | ($spl4 | $spl96)),
  inference(resolution,[],[f5384,f1433])).
fof(f5384,plain,(
  ( ! [X15,X16] : (~sP6('F_0',X16,X15) | p(s(bool,X15)) | s(bool,'F_0') != s(bool,X16)) ) | $spl96),
  inference(duplicate_literal_removal,[],[f5375])).
fof(f5375,plain,(
  ( ! [X15,X16] : (p(s(bool,X15)) | ~sP6('F_0',X16,X15) | p(s(bool,X15)) | s(bool,'F_0') != s(bool,X16)) ) | $spl96),
  inference(resolution,[],[f5367,f359])).
fof(f5367,plain,(
  ( ! [X10,X11] : (sP9(X10,X11,'F_0') | p(s(bool,X10)) | ~sP6('F_0',X11,X10)) ) | $spl96),
  inference(subsumption_resolution,[],[f5364,f347])).
fof(f5364,plain,(
  ( ! [X10,X11] : (sP9(X10,X11,'F_0') | p(s(bool,X10)) | ~sP8(X10,'F_0',X11) | ~sP6('F_0',X11,X10)) ) | $spl96),
  inference(resolution,[],[f5358,f343])).
fof(f5358,plain,(
  ( ! [X8,X7] : (sP7(X8,'F_0',X7)) ) | $spl96),
  inference(subsumption_resolution,[],[f5357,f349])).
fof(f5357,plain,(
  ( ! [X8,X7] : (p(s(bool,X7)) | sP7(X8,'F_0',X7)) ) | $spl96),
  inference(subsumption_resolution,[],[f5341,f351])).
fof(f5341,plain,(
  ( ! [X8,X7] : (p(s(bool,X7)) | ~p(s(bool,X8)) | sP7(X8,'F_0',X7)) ) | $spl96),
  inference(resolution,[],[f5250,f338])).
fof(f5469,plain,(
  ( ! [X28] : (p(s(bool,X28)) | s(bool,'F_0') != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK17)))) ) | ($spl2 | $spl96)),
  inference(resolution,[],[f5384,f1310])).
fof(f5468,plain,(
  ( ! [X27] : (p(s(bool,X27)) | s(bool,'F_0') != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK18)))) ) | $spl96),
  inference(resolution,[],[f5384,f673])).
fof(f5257,plain,(
  ( ! [X64] : (s(bool,'T_0') = s(bool,X64)) ) | $spl100),
  inference(cnf_transformation,[],[f5257_D])).
fof(f5257_D,plain,(
  ( ! [X64] : (s(bool,'T_0') = s(bool,X64)) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f5442,plain,(
  ( ! [X43,X42] : (s(bool,'F_0') != s(bool,X43) | s(bool,'T_0') = s(bool,X42)) ) | $spl74),
  inference(duplicate_literal_removal,[],[f5400])).
fof(f5400,plain,(
  ( ! [X43,X42] : (s(bool,'F_0') != s(bool,X43) | s(bool,'T_0') = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) ) | $spl74),
  inference(superposition,[],[f4602,f486])).
fof(f5449,plain,(
  ( ! [X41,X40] : (s(bool,'T_0') != s(bool,X41) | s(bool,'T_0') = s(bool,X40)) ) | $spl74),
  inference(subsumption_resolution,[],[f5399,f4602])).
fof(f5399,plain,(
  ( ! [X41,X40] : (s(bool,'T_0') != s(bool,X41) | s(bool,'T_0') = s(bool,X40) | s(bool,'F_0') = s(bool,X40)) ) | $spl74),
  inference(superposition,[],[f4602,f486])).
fof(f5443,plain,(
  ( ! [X39,X38] : (s(bool,'F_0') != s(bool,X39) | s(bool,'T_0') = s(bool,X38)) ) | $spl74),
  inference(duplicate_literal_removal,[],[f5398])).
fof(f5398,plain,(
  ( ! [X39,X38] : (s(bool,'F_0') != s(bool,X39) | s(bool,'T_0') = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) ) | $spl74),
  inference(superposition,[],[f4602,f485])).
fof(f5448,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,X37) | s(bool,'T_0') = s(bool,X36)) ) | $spl74),
  inference(subsumption_resolution,[],[f5397,f4602])).
fof(f5397,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,X37) | s(bool,'T_0') = s(bool,X36) | s(bool,'F_0') = s(bool,X36)) ) | $spl74),
  inference(superposition,[],[f4602,f485])).
fof(f5447,plain,(
  ( ! [X30,X31] : (s(bool,'F_0') != s(bool,X31) | p(s(bool,X30))) ) | $spl74),
  inference(subsumption_resolution,[],[f5394,f464])).
fof(f464,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f271])).
fof(f5394,plain,(
  ( ! [X30,X31] : (s(bool,'F_0') != s(bool,X31) | s(bool,'T_0') = s(bool,X30) | p(s(bool,X30))) ) | $spl74),
  inference(superposition,[],[f4602,f463])).
fof(f5446,plain,(
  ( ! [X28,X29] : (s(bool,'F_0') != s(bool,X29) | p(s(bool,X28))) ) | $spl74),
  inference(subsumption_resolution,[],[f5393,f464])).
fof(f5393,plain,(
  ( ! [X28,X29] : (s(bool,'F_0') != s(bool,X29) | s(bool,'T_0') = s(bool,X28) | p(s(bool,X28))) ) | $spl74),
  inference(superposition,[],[f4602,f461])).
fof(f4593,plain,(
  ( ! [X26,X25] : (s(bool,X25) != s(bool,X26) | p(s(bool,X26))) ) | $spl68),
  inference(cnf_transformation,[],[f4593_D])).
fof(f4593_D,plain,(
  ( ! [X26,X25] : (s(bool,X25) != s(bool,X26) | p(s(bool,X26))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f5445,plain,(
  ( ! [X21,X19,X20] : (s(bool,X20) != s(bool,X21) | p(s(bool,X19)) | p(s(bool,X20))) ) | $spl74),
  inference(subsumption_resolution,[],[f5390,f464])).
fof(f5390,plain,(
  ( ! [X21,X19,X20] : (s(bool,X20) != s(bool,X21) | s(bool,'T_0') = s(bool,X19) | p(s(bool,X19)) | p(s(bool,X20))) ) | $spl74),
  inference(superposition,[],[f4602,f403])).
fof(f5444,plain,(
  ( ! [X17,X18,X16] : (s(bool,X17) != s(bool,X18) | p(s(bool,X17)) | p(s(bool,X16))) ) | $spl74),
  inference(subsumption_resolution,[],[f5389,f464])).
fof(f5389,plain,(
  ( ! [X17,X18,X16] : (s(bool,X17) != s(bool,X18) | s(bool,'T_0') = s(bool,X16) | p(s(bool,X17)) | p(s(bool,X16))) ) | $spl74),
  inference(superposition,[],[f4602,f403])).
fof(f5331,plain,(
  ( ! [X43,X42] : (s(bool,'F_0') != s(bool,X43) | p(s(bool,X42))) ) | $spl68),
  inference(subsumption_resolution,[],[f5310,f4593])).
fof(f5310,plain,(
  ( ! [X43,X42] : (s(bool,'F_0') != s(bool,X43) | p(s(bool,X42)) | s(bool,'T_0') = s(bool,X42)) ) | $spl68),
  inference(superposition,[],[f4593,f486])).
fof(f5330,plain,(
  ( ! [X41,X40] : (s(bool,'T_0') != s(bool,X41) | p(s(bool,X40))) ) | $spl68),
  inference(subsumption_resolution,[],[f5309,f4593])).
fof(f5309,plain,(
  ( ! [X41,X40] : (s(bool,'T_0') != s(bool,X41) | p(s(bool,X40)) | s(bool,'F_0') = s(bool,X40)) ) | $spl68),
  inference(superposition,[],[f4593,f486])).
fof(f5329,plain,(
  ( ! [X39,X38] : (s(bool,'F_0') != s(bool,X39) | p(s(bool,X38))) ) | $spl68),
  inference(subsumption_resolution,[],[f5308,f4593])).
fof(f5308,plain,(
  ( ! [X39,X38] : (s(bool,'F_0') != s(bool,X39) | p(s(bool,X38)) | s(bool,'T_0') = s(bool,X38)) ) | $spl68),
  inference(superposition,[],[f4593,f485])).
fof(f5328,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,X37) | p(s(bool,X36))) ) | $spl68),
  inference(subsumption_resolution,[],[f5307,f4593])).
fof(f5307,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,X37) | p(s(bool,X36)) | s(bool,'F_0') = s(bool,X36)) ) | $spl68),
  inference(superposition,[],[f4593,f485])).
fof(f5324,plain,(
  ( ! [X30,X31] : (s(bool,'F_0') != s(bool,X31) | p(s(bool,X30))) ) | $spl68),
  inference(duplicate_literal_removal,[],[f5304])).
fof(f5304,plain,(
  ( ! [X30,X31] : (s(bool,'F_0') != s(bool,X31) | p(s(bool,X30)) | p(s(bool,X30))) ) | $spl68),
  inference(superposition,[],[f4593,f463])).
fof(f5325,plain,(
  ( ! [X28,X29] : (s(bool,'F_0') != s(bool,X29) | p(s(bool,X28))) ) | $spl68),
  inference(duplicate_literal_removal,[],[f5303])).
fof(f5303,plain,(
  ( ! [X28,X29] : (s(bool,'F_0') != s(bool,X29) | p(s(bool,X28)) | p(s(bool,X28))) ) | $spl68),
  inference(superposition,[],[f4593,f461])).
fof(f5254,plain,(
  ( ! [X47,X45,X46] : (sP9(X47,X46,X45) | p(s(bool,X47)) | s(bool,'F_0') != s(bool,X45) | ~p(s(bool,X46))) ) | $spl98),
  inference(cnf_transformation,[],[f5254_D])).
fof(f5254_D,plain,(
  ( ! [X47,X45,X46] : (sP9(X47,X46,X45) | p(s(bool,X47)) | s(bool,'F_0') != s(bool,X45) | ~p(s(bool,X46))) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f5259,plain,(
  ( ! [X74,X72,X75,X73] : (s(bool,'F_0') != s(bool,X73) | ~p(s(bool,X74)) | sP9(X75,X74,X73) | p(s(bool,X75)) | s(bool,'T_0') = s(bool,X72)) ) | $spl24),
  inference(subsumption_resolution,[],[f5223,f465])).
fof(f5223,plain,(
  ( ! [X74,X72,X75,X73] : (s(bool,'F_0') != s(bool,X73) | ~p(s(bool,X74)) | p(s(bool,X72)) | sP9(X75,X74,X73) | p(s(bool,X75)) | s(bool,'T_0') = s(bool,X72)) ) | $spl24),
  inference(superposition,[],[f4772,f486])).
fof(f5256,plain,(
  ( ! [X66,X64,X67,X65] : (s(bool,'F_0') != s(bool,X65) | ~p(s(bool,X66)) | sP9(X67,X66,X65) | p(s(bool,X67)) | s(bool,'T_0') = s(bool,X64)) ) | $spl24),
  inference(subsumption_resolution,[],[f5221,f465])).
fof(f5221,plain,(
  ( ! [X66,X64,X67,X65] : (s(bool,'F_0') != s(bool,X65) | ~p(s(bool,X66)) | p(s(bool,X64)) | sP9(X67,X66,X65) | p(s(bool,X67)) | s(bool,'T_0') = s(bool,X64)) ) | $spl24),
  inference(superposition,[],[f4772,f485])).
fof(f5237,plain,(
  ( ! [X50,X48,X51,X49] : (s(bool,'F_0') != s(bool,X49) | ~p(s(bool,X50)) | p(s(bool,X48)) | sP9(X51,X50,X49) | p(s(bool,X51))) ) | $spl24),
  inference(duplicate_literal_removal,[],[f5217])).
fof(f5217,plain,(
  ( ! [X50,X48,X51,X49] : (s(bool,'F_0') != s(bool,X49) | ~p(s(bool,X50)) | p(s(bool,X48)) | sP9(X51,X50,X49) | p(s(bool,X51)) | p(s(bool,X48))) ) | $spl24),
  inference(superposition,[],[f4772,f463])).
fof(f5238,plain,(
  ( ! [X47,X45,X46,X44] : (s(bool,'F_0') != s(bool,X45) | ~p(s(bool,X46)) | p(s(bool,X44)) | sP9(X47,X46,X45) | p(s(bool,X47))) ) | $spl24),
  inference(duplicate_literal_removal,[],[f5216])).
fof(f5216,plain,(
  ( ! [X47,X45,X46,X44] : (s(bool,'F_0') != s(bool,X45) | ~p(s(bool,X46)) | p(s(bool,X44)) | sP9(X47,X46,X45) | p(s(bool,X47)) | p(s(bool,X44))) ) | $spl24),
  inference(superposition,[],[f4772,f461])).
fof(f5250,plain,(
  ( ! [X94,X95] : (sP9(X95,X94,'F_0') | p(s(bool,X95)) | ~p(s(bool,X94))) ) | $spl96),
  inference(cnf_transformation,[],[f5250_D])).
fof(f5250_D,plain,(
  ( ! [X94,X95] : (sP9(X95,X94,'F_0') | p(s(bool,X95)) | ~p(s(bool,X94))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f5253,plain,(
  ( ! [X107,X105,X106,X104] : (s(bool,X104) != s(bool,X105) | ~p(s(bool,X106)) | p(s(bool,X105)) | sP9(X107,X106,'F_0') | p(s(bool,X107))) ) | $spl24),
  inference(subsumption_resolution,[],[f5203,f2665])).
fof(f5203,plain,(
  ( ! [X107,X105,X106,X104] : (s(bool,X104) != s(bool,X105) | ~p(s(bool,X106)) | p(s(bool,X105)) | sP9(X107,X106,'F_0') | p(s(bool,X107)) | p(s(bool,X104))) ) | $spl24),
  inference(superposition,[],[f4772,f461])).
fof(f5252,plain,(
  ( ! [X103,X101,X102,X100] : (s(bool,X100) != s(bool,X101) | ~p(s(bool,X102)) | p(s(bool,X101)) | sP9(X103,X102,'F_0') | p(s(bool,X103))) ) | $spl24),
  inference(subsumption_resolution,[],[f5202,f2665])).
fof(f5202,plain,(
  ( ! [X103,X101,X102,X100] : (s(bool,X100) != s(bool,X101) | ~p(s(bool,X102)) | p(s(bool,X101)) | sP9(X103,X102,'F_0') | p(s(bool,X103)) | p(s(bool,X100))) ) | $spl24),
  inference(superposition,[],[f4772,f463])).
fof(f4602,plain,(
  ( ! [X92,X93] : (s(bool,X92) != s(bool,X93) | s(bool,'T_0') = s(bool,X92)) ) | $spl74),
  inference(cnf_transformation,[],[f4602_D])).
fof(f4602_D,plain,(
  ( ! [X92,X93] : (s(bool,X92) != s(bool,X93) | s(bool,'T_0') = s(bool,X92)) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f5251,plain,(
  ( ! [X99,X97,X98,X96] : (s(bool,X96) != s(bool,X97) | ~p(s(bool,X98)) | sP9(X99,X98,'F_0') | p(s(bool,X99)) | s(bool,'T_0') = s(bool,X96)) ) | ($spl10 | $spl24)),
  inference(subsumption_resolution,[],[f5201,f3499])).
fof(f3499,plain,(
  ( ! [X35,X34] : (s(bool,X34) != s(bool,X35) | s(bool,'T_0') = s(bool,X35) | ~p(s(bool,X34))) ) | $spl10),
  inference(trivial_inequality_removal,[],[f3478])).
fof(f3478,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X34) != s(bool,X35) | s(bool,'T_0') = s(bool,X35) | ~p(s(bool,X34))) ) | $spl10),
  inference(superposition,[],[f3039,f467])).
fof(f3039,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | $spl10),
  inference(superposition,[],[f1974,f486])).
fof(f5201,plain,(
  ( ! [X99,X97,X98,X96] : (s(bool,X96) != s(bool,X97) | ~p(s(bool,X98)) | p(s(bool,X97)) | sP9(X99,X98,'F_0') | p(s(bool,X99)) | s(bool,'T_0') = s(bool,X96)) ) | $spl24),
  inference(superposition,[],[f4772,f485])).
fof(f5249,plain,(
  ( ! [X94,X92,X95,X93] : (s(bool,X92) != s(bool,X93) | ~p(s(bool,X94)) | sP9(X95,X94,'F_0') | p(s(bool,X95)) | s(bool,'T_0') = s(bool,X92)) ) | ($spl10 | $spl24)),
  inference(subsumption_resolution,[],[f5200,f3499])).
fof(f5200,plain,(
  ( ! [X94,X92,X95,X93] : (s(bool,X92) != s(bool,X93) | ~p(s(bool,X94)) | p(s(bool,X93)) | sP9(X95,X94,'F_0') | p(s(bool,X95)) | s(bool,'T_0') = s(bool,X92)) ) | $spl24),
  inference(superposition,[],[f4772,f486])).
fof(f5247,plain,(
  ( ! [X66,X64,X67] : (sP9(X67,X66,X64) | s(bool,'T_0') = s(bool,X64) | p(s(bool,X67)) | ~p(s(bool,X66))) ) | $spl94),
  inference(cnf_transformation,[],[f5247_D])).
fof(f5247_D,plain,(
  ( ! [X66,X64,X67] : (sP9(X67,X66,X64) | s(bool,'T_0') = s(bool,X64) | p(s(bool,X67)) | ~p(s(bool,X66))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f5248,plain,(
  ( ! [X74,X72,X75,X73] : (s(bool,'F_0') != s(bool,X73) | ~p(s(bool,X74)) | sP9(X75,X74,X72) | p(s(bool,X75)) | s(bool,'T_0') = s(bool,X72)) ) | $spl24),
  inference(subsumption_resolution,[],[f5195,f460])).
fof(f5195,plain,(
  ( ! [X74,X72,X75,X73] : (s(bool,'F_0') != s(bool,X73) | ~p(s(bool,X74)) | p(s(bool,X73)) | sP9(X75,X74,X72) | p(s(bool,X75)) | s(bool,'T_0') = s(bool,X72)) ) | $spl24),
  inference(superposition,[],[f4772,f486])).
fof(f5246,plain,(
  ( ! [X66,X64,X67,X65] : (s(bool,'F_0') != s(bool,X65) | ~p(s(bool,X66)) | sP9(X67,X66,X64) | p(s(bool,X67)) | s(bool,'T_0') = s(bool,X64)) ) | $spl24),
  inference(subsumption_resolution,[],[f5193,f460])).
fof(f5193,plain,(
  ( ! [X66,X64,X67,X65] : (s(bool,'F_0') != s(bool,X65) | ~p(s(bool,X66)) | p(s(bool,X65)) | sP9(X67,X66,X64) | p(s(bool,X67)) | s(bool,'T_0') = s(bool,X64)) ) | $spl24),
  inference(superposition,[],[f4772,f485])).
fof(f5242,plain,(
  ( ! [X28,X24,X27] : (~p(s(bool,X27)) | p(s(bool,X24)) | p(s(bool,X28)) | sP9(X28,X27,X24)) ) | $spl92),
  inference(cnf_transformation,[],[f5242_D])).
fof(f5242_D,plain,(
  ( ! [X28,X24,X27] : (~p(s(bool,X27)) | p(s(bool,X24)) | p(s(bool,X28)) | sP9(X28,X27,X24)) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f5245,plain,(
  ( ! [X50,X48,X51,X49] : (s(bool,'F_0') != s(bool,X49) | ~p(s(bool,X50)) | sP9(X51,X50,X48) | p(s(bool,X51)) | p(s(bool,X48))) ) | $spl24),
  inference(subsumption_resolution,[],[f5189,f460])).
fof(f5189,plain,(
  ( ! [X50,X48,X51,X49] : (s(bool,'F_0') != s(bool,X49) | ~p(s(bool,X50)) | p(s(bool,X49)) | sP9(X51,X50,X48) | p(s(bool,X51)) | p(s(bool,X48))) ) | $spl24),
  inference(superposition,[],[f4772,f463])).
fof(f5244,plain,(
  ( ! [X47,X45,X46,X44] : (s(bool,'F_0') != s(bool,X45) | ~p(s(bool,X46)) | sP9(X47,X46,X44) | p(s(bool,X47)) | p(s(bool,X44))) ) | $spl24),
  inference(subsumption_resolution,[],[f5188,f460])).
fof(f5188,plain,(
  ( ! [X47,X45,X46,X44] : (s(bool,'F_0') != s(bool,X45) | ~p(s(bool,X46)) | p(s(bool,X45)) | sP9(X47,X46,X44) | p(s(bool,X47)) | p(s(bool,X44))) ) | $spl24),
  inference(superposition,[],[f4772,f461])).
fof(f5243,plain,(
  ( ! [X30,X33,X31,X29,X32] : (s(bool,X30) != s(bool,X31) | ~p(s(bool,X32)) | p(s(bool,X31)) | sP9(X33,X32,X29) | p(s(bool,X33)) | p(s(bool,X29))) ) | $spl24),
  inference(subsumption_resolution,[],[f5185,f2665])).
fof(f5185,plain,(
  ( ! [X30,X33,X31,X29,X32] : (s(bool,X30) != s(bool,X31) | ~p(s(bool,X32)) | p(s(bool,X31)) | sP9(X33,X32,X29) | p(s(bool,X33)) | p(s(bool,X29)) | p(s(bool,X30))) ) | $spl24),
  inference(superposition,[],[f4772,f403])).
fof(f5241,plain,(
  ( ! [X28,X26,X24,X27,X25] : (s(bool,X25) != s(bool,X26) | ~p(s(bool,X27)) | p(s(bool,X26)) | sP9(X28,X27,X24) | p(s(bool,X28)) | p(s(bool,X24))) ) | $spl24),
  inference(subsumption_resolution,[],[f5184,f2665])).
fof(f5184,plain,(
  ( ! [X28,X26,X24,X27,X25] : (s(bool,X25) != s(bool,X26) | ~p(s(bool,X27)) | p(s(bool,X26)) | sP9(X28,X27,X24) | p(s(bool,X28)) | p(s(bool,X25)) | p(s(bool,X24))) ) | $spl24),
  inference(superposition,[],[f4772,f403])).
fof(f5096,plain,(
  ( ! [X54,X53] : (p(s(bool,X53)) | s(bool,'T_0') = s(bool,X54) | s(bool,X54) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK18)))) ) | ($spl4 | $spl70)),
  inference(resolution,[],[f5073,f1434])).
fof(f5095,plain,(
  ( ! [X52,X51] : (p(s(bool,X51)) | s(bool,'T_0') = s(bool,X52) | s(bool,X52) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK17)))) ) | ($spl2 | $spl70)),
  inference(resolution,[],[f5073,f1311])).
fof(f5094,plain,(
  ( ! [X50,X49] : (p(s(bool,X49)) | s(bool,'T_0') = s(bool,X50) | s(bool,X50) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK18)))) ) | $spl70),
  inference(resolution,[],[f5073,f674])).
fof(f5090,plain,(
  ( ! [X41,X40] : (p(s(bool,X40)) | s(bool,'T_0') = s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK18))) | s(bool,X41) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK18)))) ) | ($spl4 | $spl70)),
  inference(resolution,[],[f5073,f1433])).
fof(f5089,plain,(
  ( ! [X39,X38] : (p(s(bool,X38)) | s(bool,'T_0') = s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK17))) | s(bool,X39) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK17)))) ) | ($spl2 | $spl70)),
  inference(resolution,[],[f5073,f1310])).
fof(f5088,plain,(
  ( ! [X37,X36] : (p(s(bool,X36)) | s(bool,'T_0') = s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK18))) | s(bool,X37) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK18)))) ) | $spl70),
  inference(resolution,[],[f5073,f673])).
fof(f3838,plain,(
  ( ! [X19,X18] : (sP9(X18,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK18)),X19) | ~p(s(bool,X19))) ) | $spl36),
  inference(cnf_transformation,[],[f3838_D])).
fof(f3838_D,plain,(
  ( ! [X19,X18] : (sP9(X18,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK18)),X19) | ~p(s(bool,X19))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f5048,plain,(
  ( ! [X19,X18] : (sP9(X18,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK18)),X19) | ~p(s(bool,X19)) | s(bool,'T_0') = s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK18)))) ) | ($spl4 | $spl70)),
  inference(resolution,[],[f5039,f1428])).
fof(f3837,plain,(
  ( ! [X17,X16] : (sP9(X16,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK17)),X17) | ~p(s(bool,X17))) ) | $spl34),
  inference(cnf_transformation,[],[f3837_D])).
fof(f3837_D,plain,(
  ( ! [X17,X16] : (sP9(X16,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK17)),X17) | ~p(s(bool,X17))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f5047,plain,(
  ( ! [X17,X16] : (sP9(X16,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK17)),X17) | ~p(s(bool,X17)) | s(bool,'T_0') = s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK17)))) ) | ($spl2 | $spl70)),
  inference(resolution,[],[f5039,f1305])).
fof(f3836,plain,(
  ( ! [X14,X15] : (sP9(X14,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK18)),X15) | ~p(s(bool,X15))) ) | $spl32),
  inference(cnf_transformation,[],[f3836_D])).
fof(f3836_D,plain,(
  ( ! [X14,X15] : (sP9(X14,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK18)),X15) | ~p(s(bool,X15))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f5046,plain,(
  ( ! [X14,X15] : (sP9(X14,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK18)),X15) | ~p(s(bool,X15)) | s(bool,'T_0') = s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK18)))) ) | $spl70),
  inference(resolution,[],[f5039,f668])).
fof(f668,plain,(
  ( ! [X109,X110] : (sP8(X109,X110,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK18)))) )),
  inference(resolution,[],[f277,f345])).
fof(f4740,plain,(
  $false | $spl26),
  inference(trivial_inequality_removal,[],[f4735])).
fof(f4735,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl26),
  inference(superposition,[],[f2666,f581])).
fof(f4739,plain,(
  $false | $spl26),
  inference(equality_resolution,[],[f2666])).
fof(f2665,plain,(
  ( ! [X24,X22] : (s(bool,X22) != s(bool,X24) | ~p(s(bool,X22)) | p(s(bool,X24))) ) | $spl24),
  inference(cnf_transformation,[],[f2665_D])).
fof(f2665_D,plain,(
  ( ! [X24,X22] : (s(bool,X22) != s(bool,X24) | ~p(s(bool,X22)) | p(s(bool,X24))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f4694,plain,(
  ( ! [X24,X23,X22] : (s(bool,'T_0') != s(bool,X23) | s(bool,X22) != s(bool,X24) | p(s(bool,X24)) | ~p(s(bool,X22))) ) | $spl10),
  inference(subsumption_resolution,[],[f3223,f464])).
fof(f3223,plain,(
  ( ! [X24,X23,X22] : (s(bool,'T_0') != s(bool,X23) | s(bool,X22) != s(bool,X24) | p(s(bool,X24)) | ~p(s(bool,X22)) | ~p(s(bool,X23))) ) | $spl10),
  inference(superposition,[],[f3041,f406])).
fof(f4693,plain,(
  ( ! [X26,X27,X25] : (s(bool,'T_0') != s(bool,X26) | s(bool,X25) != s(bool,X27) | p(s(bool,X27)) | ~p(s(bool,X25))) ) | $spl10),
  inference(subsumption_resolution,[],[f3224,f464])).
fof(f3224,plain,(
  ( ! [X26,X27,X25] : (s(bool,'T_0') != s(bool,X26) | s(bool,X25) != s(bool,X27) | p(s(bool,X27)) | ~p(s(bool,X26)) | ~p(s(bool,X25))) ) | $spl10),
  inference(superposition,[],[f3041,f406])).
fof(f4689,plain,(
  ( ! [X25] : (s(bool,'F_0') = s(bool,X25)) ) | $spl78),
  inference(cnf_transformation,[],[f4689_D])).
fof(f4689_D,plain,(
  ( ! [X25] : (s(bool,'F_0') = s(bool,X25)) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f4690,plain,(
  ( ! [X24,X23,X22] : (s(bool,X23) != s(bool,X24) | p(s(bool,X24)) | s(bool,'F_0') = s(bool,X22) | ~p(s(bool,X23))) ) | $spl10),
  inference(subsumption_resolution,[],[f3282,f461])).
fof(f3282,plain,(
  ( ! [X24,X23,X22] : (s(bool,X23) != s(bool,X24) | p(s(bool,X24)) | s(bool,'F_0') = s(bool,X22) | ~p(s(bool,X22)) | ~p(s(bool,X23))) ) | $spl10),
  inference(superposition,[],[f3251,f406])).
fof(f3251,plain,(
  ( ! [X41,X40] : (s(bool,X40) != s(bool,X41) | p(s(bool,X41)) | s(bool,'F_0') = s(bool,X40)) ) | $spl10),
  inference(trivial_inequality_removal,[],[f3231])).
fof(f3231,plain,(
  ( ! [X41,X40] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X40) != s(bool,X41) | p(s(bool,X41)) | s(bool,'F_0') = s(bool,X40)) ) | $spl10),
  inference(superposition,[],[f3041,f486])).
fof(f4688,plain,(
  ( ! [X26,X27,X25] : (s(bool,X26) != s(bool,X27) | p(s(bool,X27)) | s(bool,'F_0') = s(bool,X25) | ~p(s(bool,X26))) ) | $spl10),
  inference(subsumption_resolution,[],[f3283,f461])).
fof(f3283,plain,(
  ( ! [X26,X27,X25] : (s(bool,X26) != s(bool,X27) | p(s(bool,X27)) | s(bool,'F_0') = s(bool,X25) | ~p(s(bool,X26)) | ~p(s(bool,X25))) ) | $spl10),
  inference(superposition,[],[f3251,f406])).
fof(f4673,plain,(
  ( ! [X59,X57] : (s(bool,X57) != s(bool,X59) | s(bool,'T_0') = s(bool,X57) | ~p(s(bool,X59))) ) | $spl76),
  inference(cnf_transformation,[],[f4673_D])).
fof(f4673_D,plain,(
  ( ! [X59,X57] : (s(bool,X57) != s(bool,X59) | s(bool,'T_0') = s(bool,X57) | ~p(s(bool,X59))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f4685,plain,(
  ( ! [X24,X23,X22] : (s(bool,X22) != s(bool,X24) | s(bool,'T_0') = s(bool,X24) | ~p(s(bool,X22)) | ~p(s(bool,X23))) ) | $spl10),
  inference(subsumption_resolution,[],[f3473,f465])).
fof(f3473,plain,(
  ( ! [X24,X23,X22] : (s(bool,'T_0') != s(bool,X23) | s(bool,X22) != s(bool,X24) | s(bool,'T_0') = s(bool,X24) | ~p(s(bool,X22)) | ~p(s(bool,X23))) ) | $spl10),
  inference(superposition,[],[f3039,f406])).
fof(f4684,plain,(
  ( ! [X26,X27,X25] : (s(bool,X25) != s(bool,X27) | s(bool,'T_0') = s(bool,X27) | ~p(s(bool,X26)) | ~p(s(bool,X25))) ) | $spl10),
  inference(subsumption_resolution,[],[f3474,f465])).
fof(f3474,plain,(
  ( ! [X26,X27,X25] : (s(bool,'T_0') != s(bool,X26) | s(bool,X25) != s(bool,X27) | s(bool,'T_0') = s(bool,X27) | ~p(s(bool,X26)) | ~p(s(bool,X25))) ) | $spl10),
  inference(superposition,[],[f3039,f406])).
fof(f4683,plain,(
  ( ! [X24,X23,X22] : (s(bool,X23) != s(bool,X24) | s(bool,'T_0') = s(bool,X24) | ~p(s(bool,X22)) | ~p(s(bool,X23))) ) | $spl10),
  inference(subsumption_resolution,[],[f3513,f460])).
fof(f3513,plain,(
  ( ! [X24,X23,X22] : (s(bool,X23) != s(bool,X24) | s(bool,'T_0') = s(bool,X24) | s(bool,'F_0') = s(bool,X22) | ~p(s(bool,X22)) | ~p(s(bool,X23))) ) | $spl10),
  inference(superposition,[],[f3497,f406])).
fof(f3497,plain,(
  ( ! [X41,X40] : (s(bool,X40) != s(bool,X41) | s(bool,'T_0') = s(bool,X41) | s(bool,'F_0') = s(bool,X40)) ) | $spl10),
  inference(trivial_inequality_removal,[],[f3481])).
fof(f3481,plain,(
  ( ! [X41,X40] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X40) != s(bool,X41) | s(bool,'T_0') = s(bool,X41) | s(bool,'F_0') = s(bool,X40)) ) | $spl10),
  inference(superposition,[],[f3039,f486])).
fof(f4682,plain,(
  ( ! [X26,X27,X25] : (s(bool,X26) != s(bool,X27) | s(bool,'T_0') = s(bool,X27) | ~p(s(bool,X26)) | ~p(s(bool,X25))) ) | $spl10),
  inference(subsumption_resolution,[],[f3514,f460])).
fof(f3514,plain,(
  ( ! [X26,X27,X25] : (s(bool,X26) != s(bool,X27) | s(bool,'T_0') = s(bool,X27) | s(bool,'F_0') = s(bool,X25) | ~p(s(bool,X26)) | ~p(s(bool,X25))) ) | $spl10),
  inference(superposition,[],[f3497,f406])).
fof(f4067,plain,(
  ( ! [X17,X18,X16] : (s(bool,'F_0') != s(bool,X17) | s(bool,X16) != s(bool,X18) | ~p(s(bool,X18)) | p(s(bool,X16))) ) | $spl38),
  inference(subsumption_resolution,[],[f4036,f460])).
fof(f4036,plain,(
  ( ! [X17,X18,X16] : (s(bool,'F_0') != s(bool,X17) | s(bool,X16) != s(bool,X18) | ~p(s(bool,X18)) | p(s(bool,X17)) | p(s(bool,X16))) ) | $spl38),
  inference(superposition,[],[f3940,f403])).
fof(f4068,plain,(
  ( ! [X21,X19,X20] : (s(bool,'F_0') != s(bool,X20) | s(bool,X19) != s(bool,X21) | ~p(s(bool,X21)) | p(s(bool,X19))) ) | $spl38),
  inference(subsumption_resolution,[],[f4037,f460])).
fof(f4037,plain,(
  ( ! [X21,X19,X20] : (s(bool,'F_0') != s(bool,X20) | s(bool,X19) != s(bool,X21) | ~p(s(bool,X21)) | p(s(bool,X19)) | p(s(bool,X20))) ) | $spl38),
  inference(superposition,[],[f3940,f403])).
fof(f4126,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | p(s(bool,X14))) ) | $spl10),
  inference(resolution,[],[f3924,f351])).
fof(f4146,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | ~p(s(bool,X14))) ) | $spl10),
  inference(resolution,[],[f4011,f355])).
fof(f4679,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,X20) != s(bool,X23) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X21)) | p(s(bool,X20))) ) | $spl10),
  inference(subsumption_resolution,[],[f4326,f403])).
fof(f4326,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,X21) != s(bool,X22) | s(bool,X20) != s(bool,X23) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X21)) | p(s(bool,X20))) ) | $spl10),
  inference(superposition,[],[f3203,f403])).
fof(f3203,plain,(
  ( ! [X6,X8,X7] : (s(bool,X7) != s(bool,X8) | s(bool,X6) != s(bool,X7) | p(s(bool,X8)) | ~p(s(bool,X6))) ) | $spl10),
  inference(superposition,[],[f3041,f467])).
fof(f4678,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,X24) != s(bool,X27) | p(s(bool,X26)) | ~p(s(bool,X27)) | p(s(bool,X24)) | p(s(bool,X25))) ) | $spl10),
  inference(subsumption_resolution,[],[f4327,f403])).
fof(f4327,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,X25) != s(bool,X26) | s(bool,X24) != s(bool,X27) | p(s(bool,X26)) | ~p(s(bool,X27)) | p(s(bool,X24)) | p(s(bool,X25))) ) | $spl10),
  inference(superposition,[],[f3203,f403])).
fof(f4677,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X29) != s(bool,X30) | p(s(bool,X30)) | ~p(s(bool,X31)) | ~p(s(bool,X28)) | ~p(s(bool,X29))) ) | $spl10),
  inference(subsumption_resolution,[],[f4328,f406])).
fof(f4328,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X29) != s(bool,X30) | s(bool,X28) != s(bool,X31) | p(s(bool,X30)) | ~p(s(bool,X31)) | ~p(s(bool,X28)) | ~p(s(bool,X29))) ) | $spl10),
  inference(superposition,[],[f3203,f406])).
fof(f4676,plain,(
  ( ! [X35,X33,X34,X32] : (s(bool,X33) != s(bool,X34) | p(s(bool,X34)) | ~p(s(bool,X35)) | ~p(s(bool,X33)) | ~p(s(bool,X32))) ) | $spl10),
  inference(subsumption_resolution,[],[f4329,f406])).
fof(f4329,plain,(
  ( ! [X35,X33,X34,X32] : (s(bool,X33) != s(bool,X34) | s(bool,X32) != s(bool,X35) | p(s(bool,X34)) | ~p(s(bool,X35)) | ~p(s(bool,X33)) | ~p(s(bool,X32))) ) | $spl10),
  inference(superposition,[],[f3203,f406])).
fof(f4675,plain,(
  ( ! [X37,X38,X36] : (s(bool,X36) != s(bool,X38) | p(s(bool,X37)) | ~p(s(bool,X38)) | p(s(bool,X36))) ) | $spl10),
  inference(subsumption_resolution,[],[f4330,f461])).
fof(f4330,plain,(
  ( ! [X37,X38,X36] : (s(bool,'F_0') != s(bool,X37) | s(bool,X36) != s(bool,X38) | p(s(bool,X37)) | ~p(s(bool,X38)) | p(s(bool,X36))) ) | $spl10),
  inference(superposition,[],[f3203,f461])).
fof(f4674,plain,(
  ( ! [X39,X41,X40] : (s(bool,X39) != s(bool,X41) | p(s(bool,X40)) | ~p(s(bool,X41)) | p(s(bool,X39))) ) | $spl10),
  inference(subsumption_resolution,[],[f4331,f461])).
fof(f4331,plain,(
  ( ! [X39,X41,X40] : (s(bool,'F_0') != s(bool,X40) | s(bool,X39) != s(bool,X41) | p(s(bool,X40)) | ~p(s(bool,X41)) | p(s(bool,X39))) ) | $spl10),
  inference(superposition,[],[f3203,f463])).
fof(f4387,plain,(
  ( ! [X52,X53,X51] : (s(bool,X51) != s(bool,X53) | p(s(bool,X52)) | ~p(s(bool,X53)) | s(bool,'T_0') = s(bool,X51)) ) | $spl10),
  inference(subsumption_resolution,[],[f4335,f461])).
fof(f4335,plain,(
  ( ! [X52,X53,X51] : (s(bool,'F_0') != s(bool,X52) | s(bool,X51) != s(bool,X53) | p(s(bool,X52)) | ~p(s(bool,X53)) | s(bool,'T_0') = s(bool,X51)) ) | $spl10),
  inference(superposition,[],[f3203,f485])).
fof(f4388,plain,(
  ( ! [X59,X57,X58] : (s(bool,X57) != s(bool,X59) | p(s(bool,X58)) | ~p(s(bool,X59)) | s(bool,'T_0') = s(bool,X57)) ) | $spl10),
  inference(subsumption_resolution,[],[f4337,f461])).
fof(f4337,plain,(
  ( ! [X59,X57,X58] : (s(bool,'F_0') != s(bool,X58) | s(bool,X57) != s(bool,X59) | p(s(bool,X58)) | ~p(s(bool,X59)) | s(bool,'T_0') = s(bool,X57)) ) | $spl10),
  inference(superposition,[],[f3203,f486])).
fof(f4672,plain,(
  ( ! [X68,X66,X67] : (s(bool,X66) != s(bool,X67) | s(bool,'T_0') != s(bool,X68) | p(s(bool,X67)) | ~p(s(bool,X66))) ) | $spl10),
  inference(subsumption_resolution,[],[f4340,f464])).
fof(f4340,plain,(
  ( ! [X68,X66,X67] : (s(bool,X66) != s(bool,X67) | s(bool,'T_0') != s(bool,X68) | p(s(bool,X67)) | ~p(s(bool,X68)) | ~p(s(bool,X66))) ) | $spl10),
  inference(superposition,[],[f3203,f467])).
fof(f4671,plain,(
  ( ! [X70,X71,X69] : (s(bool,X69) != s(bool,X70) | s(bool,'T_0') != s(bool,X71) | p(s(bool,X70)) | ~p(s(bool,X69))) ) | $spl10),
  inference(subsumption_resolution,[],[f4341,f464])).
fof(f4341,plain,(
  ( ! [X70,X71,X69] : (s(bool,X69) != s(bool,X70) | s(bool,'T_0') != s(bool,X71) | p(s(bool,X70)) | ~p(s(bool,X71)) | ~p(s(bool,X69))) ) | $spl10),
  inference(superposition,[],[f3203,f465])).
fof(f4658,plain,(
  $false | $spl50),
  inference(equality_resolution,[],[f4444])).
fof(f4601,plain,(
  ( ! [X94,X95] : (sP9(X95,'F_0',X94) | p(s(bool,X95)) | ~p(s(bool,X94))) ) | $spl72),
  inference(cnf_transformation,[],[f4601_D])).
fof(f4601_D,plain,(
  ( ! [X94,X95] : (sP9(X95,'F_0',X94) | p(s(bool,X95)) | ~p(s(bool,X94))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f4605,plain,(
  ( ! [X107,X105,X106,X104] : (s(bool,X104) != s(bool,X105) | ~p(s(bool,X106)) | p(s(bool,X105)) | sP9(X107,'F_0',X106) | p(s(bool,X107))) ) | $spl24),
  inference(subsumption_resolution,[],[f4553,f2665])).
fof(f4553,plain,(
  ( ! [X107,X105,X106,X104] : (s(bool,X104) != s(bool,X105) | ~p(s(bool,X106)) | p(s(bool,X105)) | sP9(X107,'F_0',X106) | p(s(bool,X107)) | p(s(bool,X104))) ) | $spl24),
  inference(superposition,[],[f3114,f461])).
fof(f3114,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X1) != s(bool,X2) | ~p(s(bool,X0)) | p(s(bool,X2)) | sP9(X3,X1,X0) | p(s(bool,X3))) ) | $spl24),
  inference(superposition,[],[f2665,f356])).
fof(f4604,plain,(
  ( ! [X103,X101,X102,X100] : (s(bool,X100) != s(bool,X101) | ~p(s(bool,X102)) | p(s(bool,X101)) | sP9(X103,'F_0',X102) | p(s(bool,X103))) ) | $spl24),
  inference(subsumption_resolution,[],[f4552,f2665])).
fof(f4552,plain,(
  ( ! [X103,X101,X102,X100] : (s(bool,X100) != s(bool,X101) | ~p(s(bool,X102)) | p(s(bool,X101)) | sP9(X103,'F_0',X102) | p(s(bool,X103)) | p(s(bool,X100))) ) | $spl24),
  inference(superposition,[],[f3114,f463])).
fof(f4603,plain,(
  ( ! [X99,X97,X98,X96] : (s(bool,X96) != s(bool,X97) | ~p(s(bool,X98)) | sP9(X99,'F_0',X98) | p(s(bool,X99)) | s(bool,'T_0') = s(bool,X96)) ) | $spl24),
  inference(subsumption_resolution,[],[f4551,f3196])).
fof(f3196,plain,(
  ( ! [X52,X53] : (s(bool,X52) != s(bool,X53) | ~p(s(bool,X53)) | s(bool,'T_0') = s(bool,X52)) ) | $spl24),
  inference(subsumption_resolution,[],[f3162,f416])).
fof(f3162,plain,(
  ( ! [X52,X53] : (s(bool,X52) != s(bool,X53) | ~p(s(bool,X53)) | p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,X52)) ) | $spl24),
  inference(superposition,[],[f2665,f486])).
fof(f4551,plain,(
  ( ! [X99,X97,X98,X96] : (s(bool,X96) != s(bool,X97) | ~p(s(bool,X98)) | p(s(bool,X97)) | sP9(X99,'F_0',X98) | p(s(bool,X99)) | s(bool,'T_0') = s(bool,X96)) ) | $spl24),
  inference(superposition,[],[f3114,f485])).
fof(f4600,plain,(
  ( ! [X94,X92,X95,X93] : (s(bool,X92) != s(bool,X93) | ~p(s(bool,X94)) | sP9(X95,'F_0',X94) | p(s(bool,X95)) | s(bool,'T_0') = s(bool,X92)) ) | $spl24),
  inference(subsumption_resolution,[],[f4550,f3196])).
fof(f4550,plain,(
  ( ! [X94,X92,X95,X93] : (s(bool,X92) != s(bool,X93) | ~p(s(bool,X94)) | p(s(bool,X93)) | sP9(X95,'F_0',X94) | p(s(bool,X95)) | s(bool,'T_0') = s(bool,X92)) ) | $spl24),
  inference(superposition,[],[f3114,f486])).
fof(f4598,plain,(
  ( ! [X66,X64,X67] : (sP9(X67,X64,X66) | s(bool,'T_0') = s(bool,X64) | p(s(bool,X67)) | ~p(s(bool,X66))) ) | $spl70),
  inference(cnf_transformation,[],[f4598_D])).
fof(f4598_D,plain,(
  ( ! [X66,X64,X67] : (sP9(X67,X64,X66) | s(bool,'T_0') = s(bool,X64) | p(s(bool,X67)) | ~p(s(bool,X66))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f4599,plain,(
  ( ! [X74,X72,X75,X73] : (s(bool,'F_0') != s(bool,X73) | ~p(s(bool,X74)) | sP9(X75,X72,X74) | p(s(bool,X75)) | s(bool,'T_0') = s(bool,X72)) ) | $spl24),
  inference(subsumption_resolution,[],[f4545,f460])).
fof(f4545,plain,(
  ( ! [X74,X72,X75,X73] : (s(bool,'F_0') != s(bool,X73) | ~p(s(bool,X74)) | p(s(bool,X73)) | sP9(X75,X72,X74) | p(s(bool,X75)) | s(bool,'T_0') = s(bool,X72)) ) | $spl24),
  inference(superposition,[],[f3114,f486])).
fof(f4597,plain,(
  ( ! [X66,X64,X67,X65] : (s(bool,'F_0') != s(bool,X65) | ~p(s(bool,X66)) | sP9(X67,X64,X66) | p(s(bool,X67)) | s(bool,'T_0') = s(bool,X64)) ) | $spl24),
  inference(subsumption_resolution,[],[f4543,f460])).
fof(f4543,plain,(
  ( ! [X66,X64,X67,X65] : (s(bool,'F_0') != s(bool,X65) | ~p(s(bool,X66)) | p(s(bool,X65)) | sP9(X67,X64,X66) | p(s(bool,X67)) | s(bool,'T_0') = s(bool,X64)) ) | $spl24),
  inference(superposition,[],[f3114,f485])).
fof(f4592,plain,(
  ( ! [X28,X24,X27] : (sP9(X28,X24,X27) | p(s(bool,X24)) | p(s(bool,X28)) | ~p(s(bool,X27))) ) | $spl66),
  inference(cnf_transformation,[],[f4592_D])).
fof(f4592_D,plain,(
  ( ! [X28,X24,X27] : (sP9(X28,X24,X27) | p(s(bool,X24)) | p(s(bool,X28)) | ~p(s(bool,X27))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f4596,plain,(
  ( ! [X50,X48,X51,X49] : (s(bool,'F_0') != s(bool,X49) | ~p(s(bool,X50)) | sP9(X51,X48,X50) | p(s(bool,X51)) | p(s(bool,X48))) ) | $spl24),
  inference(subsumption_resolution,[],[f4539,f460])).
fof(f4539,plain,(
  ( ! [X50,X48,X51,X49] : (s(bool,'F_0') != s(bool,X49) | ~p(s(bool,X50)) | p(s(bool,X49)) | sP9(X51,X48,X50) | p(s(bool,X51)) | p(s(bool,X48))) ) | $spl24),
  inference(superposition,[],[f3114,f463])).
fof(f4595,plain,(
  ( ! [X47,X45,X46,X44] : (s(bool,'F_0') != s(bool,X45) | ~p(s(bool,X46)) | sP9(X47,X44,X46) | p(s(bool,X47)) | p(s(bool,X44))) ) | $spl24),
  inference(subsumption_resolution,[],[f4538,f460])).
fof(f4538,plain,(
  ( ! [X47,X45,X46,X44] : (s(bool,'F_0') != s(bool,X45) | ~p(s(bool,X46)) | p(s(bool,X45)) | sP9(X47,X44,X46) | p(s(bool,X47)) | p(s(bool,X44))) ) | $spl24),
  inference(superposition,[],[f3114,f461])).
fof(f4594,plain,(
  ( ! [X30,X33,X31,X29,X32] : (s(bool,X30) != s(bool,X31) | ~p(s(bool,X32)) | p(s(bool,X31)) | sP9(X33,X29,X32) | p(s(bool,X33)) | p(s(bool,X29))) ) | $spl24),
  inference(subsumption_resolution,[],[f4535,f2665])).
fof(f4535,plain,(
  ( ! [X30,X33,X31,X29,X32] : (s(bool,X30) != s(bool,X31) | ~p(s(bool,X32)) | p(s(bool,X31)) | sP9(X33,X29,X32) | p(s(bool,X33)) | p(s(bool,X29)) | p(s(bool,X30))) ) | $spl24),
  inference(superposition,[],[f3114,f403])).
fof(f4591,plain,(
  ( ! [X28,X26,X24,X27,X25] : (s(bool,X25) != s(bool,X26) | ~p(s(bool,X27)) | p(s(bool,X26)) | sP9(X28,X24,X27) | p(s(bool,X28)) | p(s(bool,X24))) ) | $spl24),
  inference(subsumption_resolution,[],[f4534,f2665])).
fof(f4534,plain,(
  ( ! [X28,X26,X24,X27,X25] : (s(bool,X25) != s(bool,X26) | ~p(s(bool,X27)) | p(s(bool,X26)) | sP9(X28,X24,X27) | p(s(bool,X28)) | p(s(bool,X25)) | p(s(bool,X24))) ) | $spl24),
  inference(superposition,[],[f3114,f403])).
fof(f4515,plain,(
  ( ! [X54,X53] : (p(s(bool,X53)) | s(bool,'F_0') != s(bool,X54) | s(bool,X54) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK18)))) ) | ($spl4 | $spl10)),
  inference(resolution,[],[f4121,f1434])).
fof(f4514,plain,(
  ( ! [X52,X51] : (p(s(bool,X51)) | s(bool,'F_0') != s(bool,X52) | s(bool,X52) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK17)))) ) | ($spl2 | $spl10)),
  inference(resolution,[],[f4121,f1311])).
fof(f4513,plain,(
  ( ! [X50,X49] : (p(s(bool,X49)) | s(bool,'F_0') != s(bool,X50) | s(bool,X50) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK18)))) ) | $spl10),
  inference(resolution,[],[f4121,f674])).
fof(f4509,plain,(
  ( ! [X41,X40] : (p(s(bool,X40)) | s(bool,'F_0') != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK18))) | s(bool,X41) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK18)))) ) | ($spl4 | $spl10)),
  inference(resolution,[],[f4121,f1433])).
fof(f4508,plain,(
  ( ! [X39,X38] : (p(s(bool,X38)) | s(bool,'F_0') != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK17))) | s(bool,X39) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK17)))) ) | ($spl2 | $spl10)),
  inference(resolution,[],[f4121,f1310])).
fof(f4507,plain,(
  ( ! [X37,X36] : (p(s(bool,X36)) | s(bool,'F_0') != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK18))) | s(bool,X37) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK18)))) ) | $spl10),
  inference(resolution,[],[f4121,f673])).
fof(f4486,plain,(
  ( ! [X37,X36] : (p(s(bool,X36)) | s(bool,'T_0') != s(bool,X37) | s(bool,X37) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK17)))) ) | $spl10),
  inference(resolution,[],[f4084,f829])).
fof(f4483,plain,(
  ( ! [X30,X29] : (p(s(bool,X29)) | s(bool,'T_0') != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK17))) | s(bool,X30) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK17)))) ) | $spl10),
  inference(resolution,[],[f4084,f828])).
fof(f2565,plain,(
  ( ! [X8] : (sP9(X8,'F_0',happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK17)))) ) | $spl16),
  inference(cnf_transformation,[],[f2565_D])).
fof(f2565_D,plain,(
  ( ! [X8] : (sP9(X8,'F_0',happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK17)))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f4449,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') != s(bool,X12) | sP9(X13,'F_0',happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK17))) | s(bool,X12) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK17)))) ) | $spl10),
  inference(resolution,[],[f3705,f825])).
fof(f4441,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | s(bool,'F_0') != s(bool,X21) | s(bool,'T_0') != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK17)))) ) | $spl10),
  inference(resolution,[],[f4242,f825])).
fof(f4440,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'F_0') != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK18))) | s(bool,'T_0') != s(bool,X19)) ) | ($spl4 | $spl10)),
  inference(resolution,[],[f4242,f1428])).
fof(f4439,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | s(bool,'F_0') != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK17))) | s(bool,'T_0') != s(bool,X17)) ) | ($spl2 | $spl10)),
  inference(resolution,[],[f4242,f1305])).
fof(f4438,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'F_0') != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK18))) | s(bool,'T_0') != s(bool,X15)) ) | $spl10),
  inference(resolution,[],[f4242,f668])).
fof(f4409,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | s(bool,X13) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK17))) | s(bool,'T_0') != s(bool,X13)) ) | $spl10),
  inference(resolution,[],[f3753,f825])).
fof(f4254,plain,(
  ( ! [X19,X18] : (s(bool,'F_0') != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK18))) | ~p(s(bool,X18)) | ~p(s(bool,X19))) ) | ($spl4 | $spl10)),
  inference(resolution,[],[f4239,f1428])).
fof(f4253,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK17))) | ~p(s(bool,X16)) | ~p(s(bool,X17))) ) | ($spl2 | $spl10)),
  inference(resolution,[],[f4239,f1305])).
fof(f4252,plain,(
  ( ! [X14,X15] : (s(bool,'F_0') != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK18))) | ~p(s(bool,X14)) | ~p(s(bool,X15))) ) | $spl10),
  inference(resolution,[],[f4239,f668])).
fof(f4163,plain,(
  ( ! [X50,X51] : (p(s(bool,X50)) | p(s(bool,X51)) | s(bool,X50) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK18)))) ) | ($spl4 | $spl10)),
  inference(resolution,[],[f4011,f1434])).
fof(f4162,plain,(
  ( ! [X48,X49] : (p(s(bool,X48)) | p(s(bool,X49)) | s(bool,X48) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK17)))) ) | ($spl2 | $spl10)),
  inference(resolution,[],[f4011,f1311])).
fof(f4171,plain,(
  ( ! [X46] : (p(s(bool,X46)) | s(bool,X46) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK18)))) ) | $spl42),
  inference(cnf_transformation,[],[f4171_D])).
fof(f4171_D,plain,(
  ( ! [X46] : (p(s(bool,X46)) | s(bool,X46) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK18)))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f4161,plain,(
  ( ! [X47,X46] : (p(s(bool,X46)) | p(s(bool,X47)) | s(bool,X46) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK18)))) ) | $spl10),
  inference(resolution,[],[f4011,f674])).
fof(f4135,plain,(
  ( ! [X33,X34] : (~p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X33) != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK17)))) ) | $spl10),
  inference(resolution,[],[f3924,f829])).
fof(f3940,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') != s(bool,X4) | s(bool,X3) != s(bool,X4) | ~p(s(bool,X3))) ) | $spl38),
  inference(cnf_transformation,[],[f3940_D])).
fof(f3940_D,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') != s(bool,X4) | s(bool,X3) != s(bool,X4) | ~p(s(bool,X3))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f3939,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | s(bool,'F_0') != s(bool,X4) | p(s(bool,X5)) | s(bool,X3) != s(bool,X4)) ) | $spl10),
  inference(duplicate_literal_removal,[],[f3926])).
fof(f3926,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | s(bool,'F_0') != s(bool,X4) | p(s(bool,X5)) | p(s(bool,X5)) | s(bool,X3) != s(bool,X4)) ) | $spl10),
  inference(resolution,[],[f3827,f359])).
fof(f3832,plain,(
  ( ! [X19,X18] : (sP9(X18,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK18)),X19) | ~p(s(bool,X19)) | s(bool,'F_0') != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK18)))) ) | ($spl4 | $spl10)),
  inference(resolution,[],[f3794,f1428])).
fof(f3831,plain,(
  ( ! [X17,X16] : (sP9(X16,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK17)),X17) | ~p(s(bool,X17)) | s(bool,'F_0') != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK17)))) ) | ($spl2 | $spl10)),
  inference(resolution,[],[f3794,f1305])).
fof(f3830,plain,(
  ( ! [X14,X15] : (sP9(X14,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK18)),X15) | ~p(s(bool,X15)) | s(bool,'F_0') != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK18)))) ) | $spl10),
  inference(resolution,[],[f3794,f668])).
fof(f3821,plain,(
  ( ! [X21,X20] : (sP9(X20,X21,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK17))) | p(s(bool,X21)) | s(bool,'T_0') != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK17)))) ) | $spl10),
  inference(resolution,[],[f3788,f825])).
fof(f3591,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X7) | s(bool,'T_0') != s(bool,X6) | s(bool,X6) != s(bool,X7)) ) | $spl28),
  inference(cnf_transformation,[],[f3591_D])).
fof(f3591_D,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X7) | s(bool,'T_0') != s(bool,X6) | s(bool,X6) != s(bool,X7)) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f3590,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X6) | s(bool,X6) != s(bool,X7) | p(s(bool,X8)) | s(bool,'F_0') != s(bool,X7)) ) | $spl10),
  inference(duplicate_literal_removal,[],[f3581])).
fof(f3581,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X6) | s(bool,X6) != s(bool,X7) | p(s(bool,X8)) | p(s(bool,X8)) | s(bool,'F_0') != s(bool,X7)) ) | $spl10),
  inference(resolution,[],[f3043,f359])).
fof(f3043,plain,(
  ( ! [X10,X8,X9] : (sP9(X10,X8,'F_0') | s(bool,'T_0') != s(bool,X9) | s(bool,X8) != s(bool,X9) | p(s(bool,X10))) ) | $spl10),
  inference(superposition,[],[f1974,f356])).
fof(f3034,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f2943,f2556])).
fof(f2943,plain,(
  p(s(bool,'F_0')) | $spl12),
  inference(resolution,[],[f2556,f481])).
fof(f481,plain,(
  p(s(bool,'F_0')) | p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f274])).
fof(f274,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f127])).
fof(f127,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f31])).
fof(f31,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.quotient.QUOTIENT_SYM','thm.bool.NOT_CLAUSES|split|1')).
fof(f3033,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f2942,f2556])).
fof(f2942,plain,(
  p(s(bool,'F_0')) | $spl12),
  inference(resolution,[],[f2556,f479])).
fof(f479,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f273])).
fof(f273,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f126])).
fof(f126,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f32])).
fof(f32,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.quotient.QUOTIENT_SYM','thm.bool.NOT_CLAUSES|split|2')).
fof(f3032,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f2937,f2556])).
fof(f2937,plain,(
  ( ! [X216] : (p(s(bool,X216))) ) | $spl12),
  inference(resolution,[],[f2556,f427])).
fof(f427,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f248])).
fof(f248,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,'T_0'))))),
  inference(flattening,[],[f247])).
fof(f247,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X13] : ((p(s(bool,X13)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  file('thm.quotient.QUOTIENT_SYM','thm.bool.IMP_CLAUSES|split|1')).
fof(f3031,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f2936,f2556])).
fof(f2936,plain,(
  ( ! [X215] : (p(s(bool,X215))) ) | $spl12),
  inference(resolution,[],[f2556,f424])).
fof(f424,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f246])).
fof(f246,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,'T_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f245])).
fof(f245,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (~p(s(bool,'T_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0] : ((~p(s(bool,'T_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(ennf_transformation,[],[f96])).
fof(f96,plain,(
  ! [X0] : ((p(s(bool,'T_0')) => p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X13] : ((p(s(bool,'T_0')) => p(s(bool,X13))) <=> p(s(bool,X13)))),
  file('thm.quotient.QUOTIENT_SYM','thm.bool.IMP_CLAUSES|split|0')).
fof(f3030,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f2934,f2556])).
fof(f2934,plain,(
  p(s(bool,'F_0')) | $spl12),
  inference(resolution,[],[f2556,f421])).
fof(f421,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f244])).
fof(f244,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f243])).
fof(f243,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f95])).
fof(f95,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X13] : ((p(s(bool,'F_0')) => p(s(bool,X13))) <=> p(s(bool,'T_0')))),
  file('thm.quotient.QUOTIENT_SYM','thm.bool.IMP_CLAUSES|split|2')).
fof(f3029,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f2932,f2556])).
fof(f2932,plain,(
  ( ! [X212] : (p(s(bool,X212))) ) | $spl12),
  inference(resolution,[],[f2556,f418])).
fof(f418,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f242])).
fof(f242,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))))),
  inference(flattening,[],[f241])).
fof(f241,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f144])).
fof(f144,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f94])).
fof(f94,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X13] : ((p(s(bool,X13)) => p(s(bool,X13))) <=> p(s(bool,'T_0')))),
  file('thm.quotient.QUOTIENT_SYM','thm.bool.IMP_CLAUSES|split|3')).
fof(f2944,plain,(
  $false | $spl12),
  inference(resolution,[],[f2556,f483])).
fof(f483,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f8])).
fof(f8,axiom,(
  p(s(bool,'T_0'))),
  file('thm.quotient.QUOTIENT_SYM','thm.bool.TRUTH')).
fof(f2945,plain,(
  $false | $spl12),
  inference(resolution,[],[f2556,f484])).
fof(f2946,plain,(
  $false | $spl12),
  inference(resolution,[],[f2556,f491])).
fof(f2958,plain,(
  $false | $spl12),
  inference(resolution,[],[f2556,f277])).
fof(f2959,plain,(
  $false | ($spl2 | $spl12)),
  inference(resolution,[],[f2556,f733])).
fof(f2960,plain,(
  $false | ($spl4 | $spl12)),
  inference(resolution,[],[f2556,f735])).
fof(f2842,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f2772,f1973])).
fof(f2772,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(resolution,[],[f1973,f482])).
fof(f482,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f274])).
fof(f2841,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f2771,f1973])).
fof(f2771,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(resolution,[],[f1973,f480])).
fof(f480,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f273])).
fof(f2840,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f2768,f1973])).
fof(f2768,plain,(
  ( ! [X226] : (~p(s(bool,X226))) ) | $spl8),
  inference(resolution,[],[f1973,f474])).
fof(f474,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f152])).
fof(f152,plain,(
  ! [X0] : ((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) => ~p(s(bool,X0)))),
  inference(flattening,[],[f118])).
fof(f118,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) => ~p(s(bool,X0)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X13] : ((p(s(bool,X13)) => p(s(bool,'F_0'))) => ~p(s(bool,X13)))),
  file('thm.quotient.QUOTIENT_SYM','thm.bool.IMP_F')).
fof(f2839,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f2762,f1973])).
fof(f2762,plain,(
  ( ! [X220] : (~p(s(bool,X220))) ) | $spl8),
  inference(resolution,[],[f1973,f416])).
fof(f2770,plain,(
  $false | $spl8),
  inference(resolution,[],[f1973,f478])).
fof(f2784,plain,(
  $false | $spl8),
  inference(resolution,[],[f1973,f278])).
fof(f1974,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | s(bool,'T_0') != s(bool,X2)) ) | $spl10),
  inference(cnf_transformation,[],[f1974_D])).
fof(f1974_D,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | s(bool,'T_0') != s(bool,X2)) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f2404,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) )),
  inference(duplicate_literal_removal,[],[f2399])).
fof(f2399,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) )),
  inference(resolution,[],[f1949,f359])).
fof(f2529,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | s(bool,'T_0') != s(bool,X5) | s(bool,'F_0') != s(bool,X5)) )),
  inference(duplicate_literal_removal,[],[f2520])).
fof(f2520,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | s(bool,'T_0') != s(bool,X5) | s(bool,'F_0') != s(bool,X5) | ~p(s(bool,X4))) )),
  inference(resolution,[],[f1950,f357])).
fof(f357,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f215])).
fof(f1950,plain,(
  ( ! [X8,X9] : (~sP9(X9,X8,'F_0') | ~p(s(bool,X9)) | s(bool,'T_0') != s(bool,X8)) )),
  inference(superposition,[],[f1940,f358])).
fof(f2544,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | s(bool,'T_0') != s(bool,X11) | s(bool,'F_0') != s(bool,X11)) )),
  inference(duplicate_literal_removal,[],[f2538])).
fof(f2538,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | s(bool,'T_0') != s(bool,X11) | s(bool,'F_0') != s(bool,X11) | ~p(s(bool,X10))) )),
  inference(resolution,[],[f1951,f357])).
fof(f2566,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'F_0') != s(bool,X1) | s(bool,'T_0') != s(bool,X1)) )),
  inference(resolution,[],[f2474,f2405])).
fof(f2474,plain,(
  ( ! [X4,X5] : (~sP6(X5,'F_0',X4) | p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5)) )),
  inference(duplicate_literal_removal,[],[f2469])).
fof(f2469,plain,(
  ( ! [X4,X5] : (p(s(bool,X4)) | ~sP6(X5,'F_0',X4) | p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5)) )),
  inference(resolution,[],[f2442,f359])).
fof(f2442,plain,(
  ( ! [X6,X7] : (sP9(X6,'F_0',X7) | p(s(bool,X6)) | ~sP6(X7,'F_0',X6)) )),
  inference(subsumption_resolution,[],[f2440,f347])).
fof(f2440,plain,(
  ( ! [X6,X7] : (sP9(X6,'F_0',X7) | p(s(bool,X6)) | ~sP8(X6,X7,'F_0') | ~sP6(X7,'F_0',X6)) )),
  inference(resolution,[],[f2436,f343])).
fof(f2667,plain,(
  ( ! [X26,X27,X25] : (s(bool,'T_0') != s(bool,X26) | s(bool,X25) != s(bool,X27) | p(s(bool,X27)) | ~p(s(bool,X25))) ) | $spl10),
  inference(subsumption_resolution,[],[f2631,f464])).
fof(f2631,plain,(
  ( ! [X26,X27,X25] : (s(bool,'T_0') != s(bool,X26) | s(bool,X25) != s(bool,X27) | p(s(bool,X27)) | ~p(s(bool,X26)) | ~p(s(bool,X25))) ) | $spl10),
  inference(superposition,[],[f1990,f406])).
fof(f1990,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X5) | s(bool,X4) != s(bool,X5) | p(s(bool,X4))) ) | $spl10),
  inference(superposition,[],[f1974,f463])).
fof(f2664,plain,(
  ( ! [X24,X23,X22] : (s(bool,'T_0') != s(bool,X23) | s(bool,X22) != s(bool,X24) | p(s(bool,X24)) | ~p(s(bool,X22))) ) | $spl10),
  inference(subsumption_resolution,[],[f2630,f464])).
fof(f2630,plain,(
  ( ! [X24,X23,X22] : (s(bool,'T_0') != s(bool,X23) | s(bool,X22) != s(bool,X24) | p(s(bool,X24)) | ~p(s(bool,X22)) | ~p(s(bool,X23))) ) | $spl10),
  inference(superposition,[],[f1990,f406])).
fof(f2579,plain,(
  ( ! [X18] : (p(s(bool,X18)) | s(bool,'F_0') != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK18)))) ) | $spl4),
  inference(resolution,[],[f2474,f1434])).
fof(f2578,plain,(
  ( ! [X17] : (p(s(bool,X17)) | s(bool,'F_0') != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK17)))) ) | $spl2),
  inference(resolution,[],[f2474,f1311])).
fof(f2577,plain,(
  ( ! [X16] : (p(s(bool,X16)) | s(bool,'F_0') != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK18)))) )),
  inference(resolution,[],[f2474,f674])).
fof(f2563,plain,(
  ( ! [X8] : (sP9(X8,'F_0',happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK17))) | s(bool,'T_0') != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK17)))) )),
  inference(resolution,[],[f2413,f825])).
fof(f2554,plain,(
  ( ! [X8] : (s(bool,'T_0') != s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK17))) | ~p(s(bool,X8))) )),
  inference(resolution,[],[f2549,f825])).
fof(f1972,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) )),
  inference(duplicate_literal_removal,[],[f1967])).
fof(f1967,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) )),
  inference(resolution,[],[f1948,f359])).
fof(f1948,plain,(
  ( ! [X4,X5] : (sP9(X5,X4,'F_0') | s(bool,'T_0') != s(bool,X4) | p(s(bool,X5))) )),
  inference(superposition,[],[f1940,f356])).
fof(f1018,plain,(
  $false | $spl0),
  inference(resolution,[],[f731,f276])).
fof(f734,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK18)))) | $spl5),
  inference(cnf_transformation,[],[f734_D])).
fof(f734_D,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK18)))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f732,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK17)))) | $spl3),
  inference(cnf_transformation,[],[f732_D])).
fof(f732_D,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK17)))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f889,plain,(
  ( ! [X2,X0,X1] : (s(X0,happ(s(fun(sK13,X0),X1),s(sK13,sK17))) != s(X0,happ(s(fun(sK13,X0),X1),s(sK13,sK18))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK13,fun(sK13,bool)),sK14),s(fun(sK13,X0),X1),s(fun(X0,sK13),X2))))) ) | $spl6),
  inference(cnf_transformation,[],[f889_D])).
fof(f889_D,plain,(
  ( ! [X2,X0,X1] : (s(X0,happ(s(fun(sK13,X0),X1),s(sK13,sK17))) != s(X0,happ(s(fun(sK13,X0),X1),s(sK13,sK18))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK13,fun(sK13,bool)),sK14),s(fun(sK13,X0),X1),s(fun(X0,sK13),X2))))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f773,plain,(
  ( ! [X2,X0,X1] : (s(X0,happ(s(fun(sK13,X0),X1),s(sK13,sK17))) != s(X0,happ(s(fun(sK13,X0),X1),s(sK13,sK18))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK17)))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK18)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK13,fun(sK13,bool)),sK14),s(fun(sK13,X0),X1),s(fun(X0,sK13),X2))))) )),
  inference(resolution,[],[f278,f282])).
fof(f731,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK13,fun(sK13,bool)),sK14),s(fun(sK13,X0),X1),s(fun(X0,sK13),X2))))) ) | $spl0),
  inference(cnf_transformation,[],[f731_D])).
fof(f731_D,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK13,fun(sK13,bool)),sK14),s(fun(sK13,X0),X1),s(fun(X0,sK13),X2))))) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f735,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK18)))) | $spl4),
  inference(cnf_transformation,[],[f735_D])).
fof(f735_D,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK18)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f616,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK18))),s(sK13,sK18)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK13,fun(sK13,bool)),sK14),s(fun(sK13,X3),X4),s(fun(X3,sK13),X5))))) )),
  inference(resolution,[],[f277,f280])).
fof(f733,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK17)))) | $spl2),
  inference(cnf_transformation,[],[f733_D])).
fof(f733_D,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK17)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f615,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17))),s(sK13,sK17)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK13,fun(sK13,bool)),sK14),s(fun(sK13,X0),X1),s(fun(X0,sK13),X2))))) )),
  inference(resolution,[],[f277,f279])).
fof(f279,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X5)))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X6)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,X0),X3),s(fun(X0,X1),X4))))) )),
  inference(cnf_transformation,[],[f177])).
% SZS output end Proof for thm.quotient
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 710
% Passive clauses: 8433
% Generated clauses: 25521
% Final active clauses: 507
% Final passive clauses: 3410
% Input formulas: 54
% Initial clauses: 212
% 
% Selected by SInE selection: 53
% SInE iterations: 1
% 
% Duplicate literals: 908
% Trivial inequalities: 211
% Fw subsumption resolutions: 3474
% Fw demodulations: 1757
% Bw demodulations: 16
% 
% Simple tautologies: 1951
% Equational tautologies: 27
% Forward subsumptions: 8696
% 
% Binary resolution: 8175
% Forward superposition: 7477
% Backward superposition: 2644
% Equality resolution: 26
% 
% Split clauses: 427
% Split components: 964
% SAT solver clauses: 486
% SAT solver unit clauses: 23
% SAT solver binary clauses: 24
% 
% Sat splits: 427
% Sat splitting refutations: 35
% 
% Memory used [KB]: 6652
% Time elapsed: 0.443 s
% ------------------------------
% ------------------------------
% Success in time 12.196 s
