Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.option
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.option
% SZS output start Proof for thm.option
fof(f297,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f296])).
fof(f296,plain,(
  s(sK1,sK3) != s(sK1,sK3)),
  inference(superposition,[],[f269,f108])).
fof(f108,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.option.option_CASE_3'(s('type.option.option'(X0),'const.option.NONE_0'),s(X1,X2),s(fun(X0,X1),X3)))) )),
  inference(cnf_transformation,[],[f62])).
fof(f62,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,'const.option.option_CASE_3'(s('type.option.option'(X0),'const.option.NONE_0'),s(X1,X2),s(fun(X0,X1),X3)))),
  inference(flattening,[],[f61])).
fof(f61,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,'const.option.option_CASE_3'(s('type.option.option'(X0),'const.option.NONE_0'),s(X1,X2),s(fun(X0,X1),X3)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6,X7] : ! [X11,X8] : s(X7,X11) = s(X7,'const.option.option_CASE_3'(s('type.option.option'(X6),'const.option.NONE_0'),s(X7,X11),s(fun(X6,X7),X8)))),
  file('thm.option.option_CLAUSES|split|11','thm.option.option_case_def|split|0')).
fof(f269,plain,(
  s(sK1,sK3) != s(sK1,'const.option.option_CASE_3'(s('type.option.option'(sK0),'const.option.NONE_0'),s(sK1,sK3),s(fun(sK0,sK1),sK2)))),
  inference(forward_demodulation,[],[f85,f183])).
fof(f183,plain,(
  s('type.option.option'(sK0),'const.option.NONE_0') = s('type.option.option'(sK0),sK4)),
  inference(subsumption_resolution,[],[f180,f86])).
fof(f86,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.option.option_CLAUSES|split|11',aHL_TRUTH)).
fof(f180,plain,(
  ~p(s(bool,'T_0')) | s('type.option.option'(sK0),'const.option.NONE_0') = s('type.option.option'(sK0),sK4)),
  inference(superposition,[],[f98,f125])).
fof(f125,plain,(
  s(bool,'T_0') = s(bool,'const.option.IS_NONE_1'(s('type.option.option'(sK0),sK4)))),
  inference(subsumption_resolution,[],[f121,f87])).
fof(f87,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f30])).
fof(f30,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.option.option_CLAUSES|split|11',aHL_FALSITY)).
fof(f121,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.option.IS_NONE_1'(s('type.option.option'(sK0),sK4)))),
  inference(superposition,[],[f84,f90])).
fof(f90,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.option.option_CLAUSES|split|11',aHL_BOOL_CASES)).
fof(f84,plain,(
  p(s(bool,'const.option.IS_NONE_1'(s('type.option.option'(sK0),sK4))))),
  inference(cnf_transformation,[],[f79])).
fof(f79,plain,(
  p(s(bool,'const.option.IS_NONE_1'(s('type.option.option'(sK0),sK4)))) & s(sK1,sK3) != s(sK1,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK4),s(sK1,sK3),s(fun(sK0,sK1),sK2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4])],[f73])).
fof(f73,plain,(
  ? [X0,X1,X2,X3,X4] : (p(s(bool,'const.option.IS_NONE_1'(s('type.option.option'(X0),X4)))) & s(X1,X3) != s(X1,'const.option.option_CASE_3'(s('type.option.option'(X0),X4),s(X1,X3),s(fun(X0,X1),X2))))),
  inference(ennf_transformation,[],[f29])).
fof(f29,plain,(
  ~! [X0,X1,X2,X3,X4] : (p(s(bool,'const.option.IS_NONE_1'(s('type.option.option'(X0),X4)))) => s(X1,X3) = s(X1,'const.option.option_CASE_3'(s('type.option.option'(X0),X4),s(X1,X3),s(fun(X0,X1),X2))))),
  inference(flattening,[],[f28])).
fof(f28,plain,(
  ~! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.option.IS_NONE_1'(s('type.option.option'(X0),X4)))) => s(X1,X3) = s(X1,'const.option.option_CASE_3'(s('type.option.option'(X0),X4),s(X1,X3),s(fun(X0,X1),X2))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10] : (p(s(bool,'const.option.IS_NONE_1'(s('type.option.option'(X6),X10)))) => s(X7,X9) = s(X7,'const.option.option_CASE_3'(s('type.option.option'(X6),X10),s(X7,X9),s(fun(X6,X7),X8))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10] : (p(s(bool,'const.option.IS_NONE_1'(s('type.option.option'(X6),X10)))) => s(X7,X9) = s(X7,'const.option.option_CASE_3'(s('type.option.option'(X6),X10),s(X7,X9),s(fun(X6,X7),X8))))),
  file('thm.option.option_CLAUSES|split|11','thm.option.option_CLAUSES|split|11')).
fof(f98,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.option.IS_NONE_1'(s('type.option.option'(X0),X1)))) | s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X1)) )),
  inference(cnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.option.IS_NONE_1'(s('type.option.option'(X0),X1)))) | s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X1)) & (s('type.option.option'(X0),'const.option.NONE_0') != s('type.option.option'(X0),X1) | p(s(bool,'const.option.IS_NONE_1'(s('type.option.option'(X0),X1))))))),
  inference(nnf_transformation,[],[f48])).
fof(f48,plain,(
  ! [X0,X1] : (p(s(bool,'const.option.IS_NONE_1'(s('type.option.option'(X0),X1)))) <=> s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X1))),
  inference(flattening,[],[f47])).
fof(f47,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.option.IS_NONE_1'(s('type.option.option'(X0),X1)))) <=> s('type.option.option'(X0),'const.option.NONE_0') = s('type.option.option'(X0),X1))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X6] : ! [X10] : (p(s(bool,'const.option.IS_NONE_1'(s('type.option.option'(X6),X10)))) <=> s('type.option.option'(X6),'const.option.NONE_0') = s('type.option.option'(X6),X10))),
  file('thm.option.option_CLAUSES|split|11','thm.option.IS_NONE_EQ_NONE')).
fof(f85,plain,(
  s(sK1,sK3) != s(sK1,'const.option.option_CASE_3'(s('type.option.option'(sK0),sK4),s(sK1,sK3),s(fun(sK0,sK1),sK2)))),
  inference(cnf_transformation,[],[f79])).
% SZS output end Proof for thm.option
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 26
% Passive clauses: 52
% Generated clauses: 185
% Final active clauses: 21
% Final passive clauses: 17
% Input formulas: 26
% Initial clauses: 30
% 
% Duplicate literals: 1
% Trivial inequalities: 5
% Fw subsumption resolutions: 9
% Fw demodulations: 17
% Bw demodulations: 4
% 
% Simple tautologies: 3
% Equational tautologies: 3
% Forward subsumptions: 97
% 
% Binary resolution: 5
% Forward superposition: 61
% Backward superposition: 64
% Self superposition: 9
% Equality factoring: 2
% Equality resolution: 1
% 
% Split clauses: 1
% Split components: 2
% SAT solver clauses: 2
% SAT solver unit clauses: 1
% SAT solver binary clauses: 1
% 
% Sat splits: 1
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 2
% TWLsolver calls for satisfiability: 2
% 
% Memory used [KB]: 767
% Time elapsed: 0.017 s
% ------------------------------
% ------------------------------
% Success in time 0.019 s
