Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.pred_set
% (111834)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 150
% Passive clauses: 3390
% Generated clauses: 20424
% Final active clauses: 125
% Final passive clauses: 2939
% Input formulas: 19
% Initial clauses: 42
% 
% Duplicate literals: 757
% Trivial inequalities: 314
% Fw subsumption resolutions: 3080
% Fw demodulations: 3295
% Bw demodulations: 22
% 
% Simple tautologies: 382
% Equational tautologies: 44
% Forward subsumptions: 7997
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 1071
% Factoring: 9
% Forward superposition: 6402
% Backward superposition: 5017
% Self superposition: 153
% Equality factoring: 99
% Equality resolution: 30
% 
% Split clauses: 480
% Split components: 977
% SAT solver clauses: 1179
% SAT solver unit clauses: 24
% SAT solver binary clauses: 225
% 
% Sat splits: 480
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1174
% TWLsolver calls for satisfiability: 224
% 
% Memory used [KB]: 9338
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.pred_set
% (111851)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 190
% Passive clauses: 1377
% Generated clauses: 19958
% Final active clauses: 113
% Final passive clauses: 601
% Input formulas: 19
% Initial clauses: 42
% 
% Duplicate literals: 576
% Trivial inequalities: 82
% Fw subsumption resolutions: 2441
% Fw demodulations: 3144
% Bw demodulations: 578
% Condensations: 1039
% Global subsumptions: 4
% 
% Simple tautologies: 174
% Equational tautologies: 38
% Forward subsumptions: 8784
% Bw demodulations to eq. taut.: 54
% 
% Binary resolution: 798
% Factoring: 7
% Forward superposition: 4800
% Backward superposition: 4202
% Self superposition: 148
% Equality factoring: 70
% Equality resolution: 18
% 
% Split clauses: 462
% Split components: 946
% SAT solver clauses: 2232
% SAT solver unit clauses: 44
% SAT solver binary clauses: 60
% SAT solver learnt clauses: 12
% SAT solver learnt literals: 8
% 
% Sat splits: 462
% Sat splitting refutations: 24
% 
% TWLsolver clauses: 2212
% TWLsolver calls for satisfiability: 325
% 
% Memory used [KB]: 13304
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f22943,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f75,f120,f118,f74,f119,f117,f362,f663,f735,f739,f737,f773,f777,f775,f1624,f1625,f1626,f1627,f1634,f1628,f1636,f1633,f1958,f1959,f1965,f1963,f1960,f1961,f1968,f1962,f1970,f1967,f2096,f2099,f2097,f2219,f2221,f2218,f367,f2217,f2216,f2223,f364,f2215,f2214,f2225,f365,f2220,f2445,f2447,f2446,f2449,f2453,f2455,f2457,f2459,f2461,f2463,f2465,f2467,f2469,f2471,f2475,f2477,f2479,f2481,f2483,f2485,f2487,f2489,f2491,f2493,f2495,f2497,f2503,f2532,f2533,f2615,f2617,f2618,f2620,f2621,f2622,f2623,f2624,f2625,f2626,f2627,f2628,f2629,f2630,f2631,f2632,f2633,f2634,f2635,f2636,f2637,f2638,f2639,f2640,f2641,f2893,f2895,f2894,f2900,f2901,f3093,f3096,f3099,f3102,f3104,f3107,f3190,f3192,f3198,f3203,f3204,f3207,f3219,f3060,f363,f3220,f3208,f2098,f3037,f3223,f3229,f3231,f3077,f3225,f3078,f3086,f3222,f3230,f3307,f3309,f3315,f3320,f3321,f3322,f6376,f6379,f6400,f6398,f12758,f13632,f13637,f13640,f13649,f13651,f13654,f13657,f13660,f13663,f13665,f13667,f13675,f13677,f13679,f13680,f13681,f13683,f13685,f13687,f13689,f13690,f13691,f13693,f13695,f13697,f13699,f13700,f13701,f14131,f14133,f14136,f14138,f14166,f14171,f14172,f14173,f13011,f13008,f1631,f1629,f22726,f22728,f22732,f22735,f22750,f22767,f22770,f22773,f22776,f22778,f22780,f22787,f22791,f22801,f22807,f22817,f22820,f22829,f22835,f22843,f22847,f22852,f22856,f22859,f22862,f22866,f22869,f22872,f22875,f22877,f22882,f22886,f22890,f22894,f22897,f22899,f22903,f22906,f22909,f22911,f22913,f22915,f22917,f22919,f22921,f22923,f22925,f22927,f22929,f22931,f22933,f22935,f22937,f22939,f22941])).
fof(f22941,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f22940,f76])).
fof(f76,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.SUBSET_UNION_ABSORPTION',aHL_TRUTH)).
fof(f22940,plain,(
  ~p(s(bool,'T_0')) | ($spl20 | $spl23)),
  inference(forward_demodulation,[],[f22624,f1630])).
fof(f1630,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | $spl20),
  inference(cnf_transformation,[],[f1630_D])).
fof(f1630_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f22624,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2)))) | $spl23),
  inference(unit_resulting_resolution,[],[f1632,f673])).
fof(f673,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,'const.bool.IN_2'(s(X8,X9),s(fun(X8,bool),X11)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X8,X9),s(fun(X8,bool),'const.pred_set.UNION_2'(s(fun(X8,bool),X10),s(fun(X8,bool),X11)))))) )),
  inference(subsumption_resolution,[],[f513,f78])).
fof(f78,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f23])).
fof(f23,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.pred_set.SUBSET_UNION_ABSORPTION',aHL_FALSITY)).
fof(f513,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(X8,X9),s(fun(X8,bool),X11)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X8,X9),s(fun(X8,bool),'const.pred_set.UNION_2'(s(fun(X8,bool),X10),s(fun(X8,bool),X11)))))) )),
  inference(superposition,[],[f110,f79])).
fof(f79,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.pred_set.SUBSET_UNION_ABSORPTION',aHL_BOOL_CASES)).
fof(f110,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) & ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))) & ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))))),
  inference(nnf_transformation,[],[f40])).
fof(f40,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(flattening,[],[f39])).
fof(f39,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X6] : ! [X7,X8,X11] : (p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X7)))) | p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X8))))))),
  file('thm.pred_set.SUBSET_UNION_ABSORPTION','thm.pred_set.IN_UNION')).
fof(f1632,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl23),
  inference(cnf_transformation,[],[f1632_D])).
fof(f1632_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f22939,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f22938,f121])).
fof(f121,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f78,f80])).
fof(f80,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f24])).
fof(f24,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X8] : (s(bool,'T_0') = s(bool,X8) <=> p(s(bool,X8)))),
  file('thm.pred_set.SUBSET_UNION_ABSORPTION','thm.bool.EQ_CLAUSES|split|0')).
fof(f22938,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl20 | $spl23)),
  inference(forward_demodulation,[],[f22625,f20789])).
fof(f20789,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),X1),s(fun(sK0,bool),sK2)))))) ) | $spl20),
  inference(subsumption_resolution,[],[f20697,f76])).
fof(f20697,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),X1),s(fun(sK0,bool),sK2)))))) ) | $spl20),
  inference(superposition,[],[f673,f1630])).
fof(f22625,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl23),
  inference(unit_resulting_resolution,[],[f1632,f79])).
fof(f22937,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f22936,f76])).
fof(f22936,plain,(
  ~p(s(bool,'T_0')) | ($spl20 | $spl23)),
  inference(forward_demodulation,[],[f22626,f20789])).
fof(f22626,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) | $spl23),
  inference(unit_resulting_resolution,[],[f1632,f81])).
fof(f81,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f51])).
fof(f22935,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f22934,f121])).
fof(f22934,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl20 | $spl23)),
  inference(forward_demodulation,[],[f22627,f20789])).
fof(f22627,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl23),
  inference(unit_resulting_resolution,[],[f121,f1632,f135])).
fof(f135,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f79,f79])).
fof(f22933,plain,(
  $false | ($spl20 | $spl23 | $spl47)),
  inference(subsumption_resolution,[],[f22932,f2225])).
fof(f22932,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | ($spl20 | $spl23 | $spl47)),
  inference(forward_demodulation,[],[f22628,f20789])).
fof(f22628,plain,(
  s(bool,f(s(bool,'F_0'))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl23 | $spl47)),
  inference(unit_resulting_resolution,[],[f2225,f1632,f135])).
fof(f22931,plain,(
  $false | ($spl3 | $spl13 | $spl20 | $spl23)),
  inference(subsumption_resolution,[],[f22930,f750])).
fof(f750,plain,(
  s(bool,'T_0') != s(bool,g(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) | ($spl3 | $spl13)),
  inference(backward_demodulation,[],[f744,f718])).
fof(f718,plain,(
  s(bool,f(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) != s(bool,g(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) | $spl3),
  inference(unit_resulting_resolution,[],[f120,f115])).
fof(f115,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK5(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK5(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK5(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK5(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5])],[f48])).
fof(f48,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f43])).
fof(f43,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f42])).
fof(f42,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.pred_set.SUBSET_UNION_ABSORPTION',aHL_EXT)).
fof(f744,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) | $spl13),
  inference(unit_resulting_resolution,[],[f78,f739,f171])).
fof(f171,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f166,f78])).
fof(f166,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f111,f79])).
fof(f111,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f46])).
fof(f46,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f41])).
fof(f41,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X9,X10] : ((p(s(bool,X9)) => p(s(bool,X10))) => ((p(s(bool,X10)) => p(s(bool,X9))) => s(bool,X9) = s(bool,X10)))),
  file('thm.pred_set.SUBSET_UNION_ABSORPTION','thm.bool.IMP_ANTISYM_AX')).
fof(f22930,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) | ($spl3 | $spl13 | $spl20 | $spl23)),
  inference(forward_demodulation,[],[f22629,f20789])).
fof(f22629,plain,(
  s(bool,g(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl3 | $spl13 | $spl23)),
  inference(unit_resulting_resolution,[],[f750,f1632,f135])).
fof(f22929,plain,(
  $false | ($spl3 | $spl17 | $spl20 | $spl23)),
  inference(subsumption_resolution,[],[f22928,f825])).
fof(f825,plain,(
  s(bool,'T_0') != s(bool,g(s(sK0,sK5(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | ($spl3 | $spl17)),
  inference(backward_demodulation,[],[f819,f720])).
fof(f720,plain,(
  s(bool,f(s(sK0,sK5(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) != s(bool,g(s(sK0,sK5(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl3),
  inference(unit_resulting_resolution,[],[f120,f115])).
fof(f819,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK5(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl17),
  inference(unit_resulting_resolution,[],[f78,f777,f171])).
fof(f22928,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK5(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | ($spl3 | $spl17 | $spl20 | $spl23)),
  inference(forward_demodulation,[],[f22630,f20789])).
fof(f22630,plain,(
  s(bool,g(s(sK0,sK5(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl3 | $spl17 | $spl23)),
  inference(unit_resulting_resolution,[],[f825,f1632,f135])).
fof(f22927,plain,(
  $false | ($spl3 | $spl20 | $spl23 | $spl32)),
  inference(subsumption_resolution,[],[f22926,f1971])).
fof(f1971,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl32)),
  inference(backward_demodulation,[],[f1967,f721])).
fof(f721,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl3),
  inference(unit_resulting_resolution,[],[f120,f104])).
fof(f104,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.bool.IN_2'(s(X0,sK3(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK3(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) )),
  inference(cnf_transformation,[],[f67])).
fof(f67,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (s(bool,'const.bool.IN_2'(s(X0,sK3(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK3(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3])],[f66])).
fof(f66,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(rectify,[],[f65])).
fof(f65,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(nnf_transformation,[],[f36])).
fof(f36,plain,(
  ! [X0,X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(flattening,[],[f35])).
fof(f35,plain,(
  ! [X0] : ! [X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X6] : ! [X7,X8] : (s(fun(X6,bool),X7) = s(fun(X6,bool),X8) <=> ! [X11] : s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X7))) = s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X8))))),
  file('thm.pred_set.SUBSET_UNION_ABSORPTION','thm.pred_set.EXTENSION')).
fof(f22926,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl20 | $spl23 | $spl32)),
  inference(forward_demodulation,[],[f22631,f20789])).
fof(f22631,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl3 | $spl23 | $spl32)),
  inference(unit_resulting_resolution,[],[f1971,f1632,f135])).
fof(f22925,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f22924,f121])).
fof(f22924,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl20 | $spl23)),
  inference(forward_demodulation,[],[f22634,f20789])).
fof(f22634,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl23),
  inference(unit_resulting_resolution,[],[f121,f1632,f135])).
fof(f22923,plain,(
  $false | ($spl20 | $spl23 | $spl47)),
  inference(subsumption_resolution,[],[f22922,f2225])).
fof(f22922,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | ($spl20 | $spl23 | $spl47)),
  inference(forward_demodulation,[],[f22635,f20789])).
fof(f22635,plain,(
  s(bool,f(s(bool,'F_0'))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl23 | $spl47)),
  inference(unit_resulting_resolution,[],[f2225,f1632,f135])).
fof(f22921,plain,(
  $false | ($spl3 | $spl13 | $spl20 | $spl23)),
  inference(subsumption_resolution,[],[f22920,f750])).
fof(f22920,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) | ($spl3 | $spl13 | $spl20 | $spl23)),
  inference(forward_demodulation,[],[f22636,f20789])).
fof(f22636,plain,(
  s(bool,g(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl3 | $spl13 | $spl23)),
  inference(unit_resulting_resolution,[],[f750,f1632,f135])).
fof(f22919,plain,(
  $false | ($spl3 | $spl17 | $spl20 | $spl23)),
  inference(subsumption_resolution,[],[f22918,f825])).
fof(f22918,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK5(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | ($spl3 | $spl17 | $spl20 | $spl23)),
  inference(forward_demodulation,[],[f22637,f20789])).
fof(f22637,plain,(
  s(bool,g(s(sK0,sK5(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl3 | $spl17 | $spl23)),
  inference(unit_resulting_resolution,[],[f825,f1632,f135])).
fof(f22917,plain,(
  $false | ($spl3 | $spl20 | $spl23 | $spl32)),
  inference(subsumption_resolution,[],[f22916,f1971])).
fof(f22916,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl20 | $spl23 | $spl32)),
  inference(forward_demodulation,[],[f22638,f20789])).
fof(f22638,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl3 | $spl23 | $spl32)),
  inference(unit_resulting_resolution,[],[f1971,f1632,f135])).
fof(f22915,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f22914,f121])).
fof(f22914,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl20 | $spl23)),
  inference(forward_demodulation,[],[f22641,f20789])).
fof(f22641,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl23),
  inference(unit_resulting_resolution,[],[f78,f1632,f171])).
fof(f22913,plain,(
  $false | ($spl20 | $spl23 | $spl24)),
  inference(subsumption_resolution,[],[f22912,f1630])).
fof(f22912,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl20 | $spl23 | $spl24)),
  inference(forward_demodulation,[],[f22643,f20789])).
fof(f22643,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl23 | $spl24)),
  inference(unit_resulting_resolution,[],[f1632,f1634])).
fof(f22911,plain,(
  $false | ($spl3 | $spl20 | $spl23 | $spl32)),
  inference(subsumption_resolution,[],[f22910,f1971])).
fof(f22910,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl20 | $spl23 | $spl32)),
  inference(forward_demodulation,[],[f22644,f20789])).
fof(f22644,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl3 | $spl23 | $spl32)),
  inference(unit_resulting_resolution,[],[f1632,f2026])).
fof(f2026,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl3 | $spl32)),
  inference(trivial_inequality_removal,[],[f2020])).
fof(f2020,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl3 | $spl32)),
  inference(superposition,[],[f1971,f135])).
fof(f22909,plain,(
  $false | ($spl20 | $spl23 | $spl50 | $spl65)),
  inference(subsumption_resolution,[],[f22908,f121])).
fof(f22908,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl20 | $spl23 | $spl50 | $spl65)),
  inference(forward_demodulation,[],[f22907,f6699])).
fof(f6699,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | $spl65),
  inference(unit_resulting_resolution,[],[f78,f6400,f171])).
fof(f22907,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | ($spl20 | $spl23 | $spl50)),
  inference(forward_demodulation,[],[f22645,f20789])).
fof(f22645,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))) | ($spl23 | $spl50)),
  inference(unit_resulting_resolution,[],[f1632,f2894])).
fof(f22906,plain,(
  $false | ($spl4 | $spl8 | $spl20 | $spl23 | $spl50 | $spl65)),
  inference(subsumption_resolution,[],[f22905,f22737])).
fof(f22737,plain,(
  s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | ($spl4 | $spl8 | $spl20 | $spl23 | $spl50 | $spl65)),
  inference(forward_demodulation,[],[f22736,f6699])).
fof(f22736,plain,(
  s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | ($spl4 | $spl8 | $spl20 | $spl23 | $spl50)),
  inference(forward_demodulation,[],[f22707,f20789])).
fof(f22707,plain,(
  s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))) != s(bool,g(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))) | ($spl4 | $spl8 | $spl23 | $spl50)),
  inference(unit_resulting_resolution,[],[f1632,f12522])).
fof(f12522,plain,(
  ( ! [X1] : (s(bool,f(s(bool,X1))) != s(bool,g(s(bool,X1))) | s(bool,'T_0') = s(bool,X1)) ) | ($spl4 | $spl8 | $spl50)),
  inference(duplicate_literal_removal,[],[f12490])).
fof(f12490,plain,(
  ( ! [X1] : (s(bool,f(s(bool,X1))) != s(bool,g(s(bool,X1))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | ($spl4 | $spl8 | $spl50)),
  inference(superposition,[],[f7422,f4310])).
fof(f4310,plain,(
  ( ! [X12] : (s(bool,X12) = s(bool,f(s(bool,f(s(bool,X12))))) | s(bool,'T_0') = s(bool,X12)) ) | ($spl8 | $spl50)),
  inference(subsumption_resolution,[],[f4233,f121])).
fof(f4233,plain,(
  ( ! [X12] : (s(bool,X12) = s(bool,f(s(bool,f(s(bool,X12))))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X12)) ) | ($spl8 | $spl50)),
  inference(superposition,[],[f3900,f79])).
fof(f3900,plain,(
  ( ! [X28] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,X28))))) | s(bool,'T_0') = s(bool,X28)) ) | ($spl8 | $spl50)),
  inference(subsumption_resolution,[],[f3883,f78])).
fof(f3883,plain,(
  ( ! [X28] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,X28))))) | s(bool,'T_0') = s(bool,X28)) ) | ($spl8 | $spl50)),
  inference(superposition,[],[f3594,f2894])).
fof(f3594,plain,(
  ( ! [X4] : (p(s(bool,X4)) | s(bool,'F_0') = s(bool,f(s(bool,X4)))) ) | $spl8),
  inference(subsumption_resolution,[],[f3570,f80])).
fof(f3570,plain,(
  ( ! [X4] : (p(s(bool,X4)) | s(bool,'F_0') = s(bool,f(s(bool,X4))) | s(bool,'T_0') = s(bool,X4)) ) | $spl8),
  inference(superposition,[],[f85,f3373])).
fof(f3373,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl8),
  inference(superposition,[],[f366,f79])).
fof(f366,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | $spl8),
  inference(cnf_transformation,[],[f366_D])).
fof(f366_D,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f85,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f27])).
fof(f27,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f26])).
fof(f26,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X8] : (s(bool,'F_0') = s(bool,X8) <=> ~p(s(bool,X8)))),
  file('thm.pred_set.SUBSET_UNION_ABSORPTION','thm.bool.EQ_CLAUSES|split|3')).
fof(f7422,plain,(
  ( ! [X0] : (s(bool,f(s(bool,X0))) != s(bool,g(s(bool,f(s(bool,f(s(bool,X0))))))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl4 | $spl8)),
  inference(condensation,[],[f7421])).
fof(f7421,plain,(
  ( ! [X10,X11] : (s(bool,f(s(bool,X11))) != s(bool,g(s(bool,f(s(bool,f(s(bool,X11))))))) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X11)) ) | ($spl4 | $spl8)),
  inference(duplicate_literal_removal,[],[f7367])).
fof(f7367,plain,(
  ( ! [X10,X11] : (s(bool,f(s(bool,X11))) != s(bool,g(s(bool,f(s(bool,f(s(bool,X11))))))) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X10)) ) | ($spl4 | $spl8)),
  inference(superposition,[],[f6948,f3584])).
fof(f3584,plain,(
  ( ! [X2,X1] : (s(bool,X2) = s(bool,f(s(bool,X1))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X2)) ) | $spl8),
  inference(duplicate_literal_removal,[],[f3561])).
fof(f3561,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,X2) = s(bool,f(s(bool,X1))) | s(bool,'T_0') = s(bool,X2)) ) | $spl8),
  inference(superposition,[],[f3373,f135])).
fof(f6948,plain,(
  ( ! [X14] : (s(bool,X14) != s(bool,g(s(bool,f(s(bool,X14))))) | s(bool,'T_0') = s(bool,X14)) ) | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f6868,f121])).
fof(f6868,plain,(
  ( ! [X14] : (s(bool,X14) != s(bool,g(s(bool,f(s(bool,X14))))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X14)) ) | ($spl4 | $spl8)),
  inference(superposition,[],[f6217,f79])).
fof(f6217,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,g(s(bool,f(s(bool,X0))))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl4 | $spl8)),
  inference(condensation,[],[f6216])).
fof(f6216,plain,(
  ( ! [X45,X44] : (s(bool,'F_0') != s(bool,g(s(bool,f(s(bool,X45))))) | s(bool,'T_0') = s(bool,X44) | s(bool,'T_0') = s(bool,X45)) ) | ($spl4 | $spl8)),
  inference(duplicate_literal_removal,[],[f6035])).
fof(f6035,plain,(
  ( ! [X45,X44] : (s(bool,'F_0') != s(bool,g(s(bool,f(s(bool,X45))))) | s(bool,'T_0') = s(bool,X44) | s(bool,'T_0') = s(bool,X45) | s(bool,'T_0') = s(bool,X44)) ) | ($spl4 | $spl8)),
  inference(superposition,[],[f3475,f3584])).
fof(f3475,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,g(s(bool,X3))) | s(bool,'T_0') = s(bool,X3)) ) | $spl4),
  inference(subsumption_resolution,[],[f3448,f76])).
fof(f3448,plain,(
  ( ! [X3] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,g(s(bool,X3))) | s(bool,'T_0') = s(bool,X3)) ) | $spl4),
  inference(superposition,[],[f84,f3342])).
fof(f3342,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,g(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl4),
  inference(superposition,[],[f364,f79])).
fof(f84,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f53])).
fof(f22905,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | ($spl4 | $spl20 | $spl23)),
  inference(forward_demodulation,[],[f22646,f20789])).
fof(f22646,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))) | ($spl4 | $spl23)),
  inference(unit_resulting_resolution,[],[f1632,f3342])).
fof(f22903,plain,(
  $false | ($spl8 | $spl20 | $spl23 | $spl47)),
  inference(subsumption_resolution,[],[f22902,f2225])).
fof(f22902,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | ($spl8 | $spl20 | $spl23)),
  inference(forward_demodulation,[],[f22648,f20789])).
fof(f22648,plain,(
  s(bool,f(s(bool,'F_0'))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl8 | $spl23)),
  inference(unit_resulting_resolution,[],[f1632,f3393])).
fof(f3393,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl8),
  inference(subsumption_resolution,[],[f3379,f121])).
fof(f3379,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X0) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl8),
  inference(superposition,[],[f366,f135])).
fof(f22899,plain,(
  $false | ($spl4 | $spl8 | $spl20 | $spl23 | $spl50 | $spl65)),
  inference(subsumption_resolution,[],[f22898,f22737])).
fof(f22898,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | ($spl4 | $spl20 | $spl23)),
  inference(forward_demodulation,[],[f22651,f20789])).
fof(f22651,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))) | ($spl4 | $spl23)),
  inference(unit_resulting_resolution,[],[f76,f1632,f3476])).
fof(f3476,plain,(
  ( ! [X8,X9] : (~p(s(bool,X9)) | s(bool,X9) = s(bool,g(s(bool,X8))) | s(bool,'T_0') = s(bool,X8)) ) | $spl4),
  inference(subsumption_resolution,[],[f3451,f76])).
fof(f3451,plain,(
  ( ! [X8,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,X9)) | s(bool,X9) = s(bool,g(s(bool,X8))) | s(bool,'T_0') = s(bool,X8)) ) | $spl4),
  inference(superposition,[],[f114,f3342])).
fof(f114,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f47])).
fof(f22897,plain,(
  $false | ($spl8 | $spl20 | $spl23 | $spl65)),
  inference(subsumption_resolution,[],[f22896,f121])).
fof(f22896,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl8 | $spl20 | $spl23 | $spl65)),
  inference(forward_demodulation,[],[f22895,f6699])).
fof(f22895,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | ($spl8 | $spl20 | $spl23)),
  inference(forward_demodulation,[],[f22652,f20789])).
fof(f22652,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))) | ($spl8 | $spl23)),
  inference(unit_resulting_resolution,[],[f121,f1632,f3584])).
fof(f22894,plain,(
  $false | ($spl8 | $spl20 | $spl23 | $spl47 | $spl65)),
  inference(subsumption_resolution,[],[f22893,f121])).
fof(f22893,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl8 | $spl20 | $spl23 | $spl47 | $spl65)),
  inference(forward_demodulation,[],[f22892,f6699])).
fof(f22892,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | ($spl8 | $spl20 | $spl23 | $spl47)),
  inference(forward_demodulation,[],[f22891,f366])).
fof(f22891,plain,(
  s(bool,f(s(bool,'T_0'))) = s(bool,f(s(bool,'F_0'))) | ($spl8 | $spl20 | $spl23 | $spl47)),
  inference(forward_demodulation,[],[f22653,f20789])).
fof(f22653,plain,(
  s(bool,f(s(bool,'F_0'))) = s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))) | ($spl8 | $spl23 | $spl47)),
  inference(unit_resulting_resolution,[],[f2225,f1632,f3584])).
fof(f22890,plain,(
  $false | ($spl3 | $spl8 | $spl13 | $spl20 | $spl23 | $spl65)),
  inference(subsumption_resolution,[],[f22889,f121])).
fof(f22889,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl3 | $spl8 | $spl13 | $spl20 | $spl23 | $spl65)),
  inference(forward_demodulation,[],[f22888,f6699])).
fof(f22888,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | ($spl3 | $spl8 | $spl13 | $spl20 | $spl23)),
  inference(forward_demodulation,[],[f22887,f756])).
fof(f756,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) | ($spl3 | $spl13)),
  inference(unit_resulting_resolution,[],[f78,f750,f171])).
fof(f22887,plain,(
  s(bool,f(s(bool,'T_0'))) = s(bool,g(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) | ($spl3 | $spl8 | $spl13 | $spl20 | $spl23)),
  inference(forward_demodulation,[],[f22654,f20789])).
fof(f22654,plain,(
  s(bool,g(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) = s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))) | ($spl3 | $spl8 | $spl13 | $spl23)),
  inference(unit_resulting_resolution,[],[f750,f1632,f3584])).
fof(f22886,plain,(
  $false | ($spl3 | $spl8 | $spl17 | $spl20 | $spl23 | $spl65)),
  inference(subsumption_resolution,[],[f22885,f121])).
fof(f22885,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl3 | $spl8 | $spl17 | $spl20 | $spl23 | $spl65)),
  inference(forward_demodulation,[],[f22884,f6699])).
fof(f22884,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | ($spl3 | $spl8 | $spl17 | $spl20 | $spl23)),
  inference(forward_demodulation,[],[f22883,f831])).
fof(f831,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK5(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | ($spl3 | $spl17)),
  inference(unit_resulting_resolution,[],[f78,f825,f171])).
fof(f22883,plain,(
  s(bool,f(s(bool,'T_0'))) = s(bool,g(s(sK0,sK5(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | ($spl3 | $spl8 | $spl17 | $spl20 | $spl23)),
  inference(forward_demodulation,[],[f22655,f20789])).
fof(f22655,plain,(
  s(bool,g(s(sK0,sK5(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) = s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))) | ($spl3 | $spl8 | $spl17 | $spl23)),
  inference(unit_resulting_resolution,[],[f825,f1632,f3584])).
fof(f22882,plain,(
  $false | ($spl3 | $spl8 | $spl20 | $spl23 | $spl32 | $spl65)),
  inference(subsumption_resolution,[],[f22881,f121])).
fof(f22881,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl3 | $spl8 | $spl20 | $spl23 | $spl32 | $spl65)),
  inference(forward_demodulation,[],[f22880,f6699])).
fof(f22880,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | ($spl3 | $spl8 | $spl20 | $spl23 | $spl32)),
  inference(forward_demodulation,[],[f22879,f1987])).
fof(f1987,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl32)),
  inference(unit_resulting_resolution,[],[f78,f1971,f171])).
fof(f22879,plain,(
  s(bool,f(s(bool,'T_0'))) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl8 | $spl20 | $spl23 | $spl32)),
  inference(forward_demodulation,[],[f22656,f20789])).
fof(f22656,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) = s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))) | ($spl3 | $spl8 | $spl23 | $spl32)),
  inference(unit_resulting_resolution,[],[f1971,f1632,f3584])).
fof(f22877,plain,(
  $false | ($spl8 | $spl20 | $spl23 | $spl47)),
  inference(subsumption_resolution,[],[f22876,f2225])).
fof(f22876,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | ($spl8 | $spl20 | $spl23)),
  inference(forward_demodulation,[],[f22658,f20789])).
fof(f22658,plain,(
  s(bool,f(s(bool,'F_0'))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl8 | $spl23)),
  inference(unit_resulting_resolution,[],[f121,f1632,f3584])).
fof(f22875,plain,(
  $false | ($spl8 | $spl20 | $spl23 | $spl47)),
  inference(subsumption_resolution,[],[f22874,f2225])).
fof(f22874,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | ($spl8 | $spl20 | $spl23 | $spl47)),
  inference(forward_demodulation,[],[f22873,f366])).
fof(f22873,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,f(s(bool,'F_0'))))) | ($spl8 | $spl20 | $spl23 | $spl47)),
  inference(forward_demodulation,[],[f22659,f20789])).
fof(f22659,plain,(
  s(bool,f(s(bool,f(s(bool,'F_0'))))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl8 | $spl23 | $spl47)),
  inference(unit_resulting_resolution,[],[f2225,f1632,f3584])).
fof(f22872,plain,(
  $false | ($spl3 | $spl8 | $spl13 | $spl20 | $spl23 | $spl47)),
  inference(subsumption_resolution,[],[f22871,f2225])).
fof(f22871,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | ($spl3 | $spl8 | $spl13 | $spl20 | $spl23)),
  inference(forward_demodulation,[],[f22870,f756])).
fof(f22870,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,g(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))))) | ($spl3 | $spl8 | $spl13 | $spl20 | $spl23)),
  inference(forward_demodulation,[],[f22660,f20789])).
fof(f22660,plain,(
  s(bool,f(s(bool,g(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl3 | $spl8 | $spl13 | $spl23)),
  inference(unit_resulting_resolution,[],[f750,f1632,f3584])).
fof(f22869,plain,(
  $false | ($spl3 | $spl8 | $spl17 | $spl20 | $spl23 | $spl47)),
  inference(subsumption_resolution,[],[f22868,f2225])).
fof(f22868,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | ($spl3 | $spl8 | $spl17 | $spl20 | $spl23)),
  inference(forward_demodulation,[],[f22867,f831])).
fof(f22867,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,g(s(sK0,sK5(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))))) | ($spl3 | $spl8 | $spl17 | $spl20 | $spl23)),
  inference(forward_demodulation,[],[f22661,f20789])).
fof(f22661,plain,(
  s(bool,f(s(bool,g(s(sK0,sK5(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl3 | $spl8 | $spl17 | $spl23)),
  inference(unit_resulting_resolution,[],[f825,f1632,f3584])).
fof(f22866,plain,(
  $false | ($spl3 | $spl8 | $spl20 | $spl23 | $spl32 | $spl47)),
  inference(subsumption_resolution,[],[f22865,f2225])).
fof(f22865,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | ($spl3 | $spl8 | $spl20 | $spl23 | $spl32)),
  inference(forward_demodulation,[],[f22864,f1987])).
fof(f22864,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))))) | ($spl3 | $spl8 | $spl20 | $spl23 | $spl32)),
  inference(forward_demodulation,[],[f22662,f20789])).
fof(f22662,plain,(
  s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl3 | $spl8 | $spl23 | $spl32)),
  inference(unit_resulting_resolution,[],[f1971,f1632,f3584])).
fof(f22862,plain,(
  $false | ($spl8 | $spl20 | $spl23 | $spl47)),
  inference(subsumption_resolution,[],[f22861,f2225])).
fof(f22861,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | ($spl8 | $spl20 | $spl23)),
  inference(forward_demodulation,[],[f22664,f20789])).
fof(f22664,plain,(
  s(bool,f(s(bool,'F_0'))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl8 | $spl23)),
  inference(unit_resulting_resolution,[],[f78,f1632,f3596])).
fof(f3596,plain,(
  ( ! [X19,X20] : (p(s(bool,X19)) | s(bool,X20) = s(bool,f(s(bool,X19))) | s(bool,'T_0') = s(bool,X20)) ) | $spl8),
  inference(subsumption_resolution,[],[f3578,f80])).
fof(f3578,plain,(
  ( ! [X19,X20] : (p(s(bool,X19)) | s(bool,X20) = s(bool,f(s(bool,X19))) | s(bool,'T_0') = s(bool,X20) | s(bool,'T_0') = s(bool,X19)) ) | $spl8),
  inference(superposition,[],[f171,f3373])).
fof(f22859,plain,(
  $false | ($spl20 | $spl23 | $spl50 | $spl65)),
  inference(subsumption_resolution,[],[f22858,f121])).
fof(f22858,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl20 | $spl23 | $spl50 | $spl65)),
  inference(forward_demodulation,[],[f22857,f6699])).
fof(f22857,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | ($spl20 | $spl23 | $spl50)),
  inference(forward_demodulation,[],[f22666,f20789])).
fof(f22666,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))) | ($spl23 | $spl50)),
  inference(unit_resulting_resolution,[],[f78,f1632,f3897])).
fof(f3897,plain,(
  ( ! [X8,X7] : (p(s(bool,X8)) | s(bool,X8) = s(bool,f(s(bool,X7))) | s(bool,'T_0') = s(bool,X7)) ) | $spl50),
  inference(subsumption_resolution,[],[f3870,f78])).
fof(f3870,plain,(
  ( ! [X8,X7] : (p(s(bool,'F_0')) | p(s(bool,X8)) | s(bool,X8) = s(bool,f(s(bool,X7))) | s(bool,'T_0') = s(bool,X7)) ) | $spl50),
  inference(superposition,[],[f111,f2894])).
fof(f22856,plain,(
  $false | ($spl8 | $spl20 | $spl23 | $spl50 | $spl65)),
  inference(subsumption_resolution,[],[f22855,f121])).
fof(f22855,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl8 | $spl20 | $spl23 | $spl50 | $spl65)),
  inference(forward_demodulation,[],[f22854,f6699])).
fof(f22854,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | ($spl8 | $spl20 | $spl23 | $spl50 | $spl65)),
  inference(forward_demodulation,[],[f22853,f6699])).
fof(f22853,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,'T_0'))))) | ($spl8 | $spl20 | $spl23 | $spl50)),
  inference(forward_demodulation,[],[f22667,f20789])).
fof(f22667,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))) | ($spl8 | $spl23 | $spl50)),
  inference(unit_resulting_resolution,[],[f1632,f3900])).
fof(f22852,plain,(
  $false | ($spl8 | $spl20 | $spl23 | $spl50 | $spl65)),
  inference(subsumption_resolution,[],[f22851,f121])).
fof(f22851,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl8 | $spl20 | $spl23 | $spl50 | $spl65)),
  inference(forward_demodulation,[],[f22850,f6699])).
fof(f22850,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | ($spl8 | $spl20 | $spl23 | $spl50 | $spl65)),
  inference(forward_demodulation,[],[f22849,f6699])).
fof(f22849,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,'T_0'))))) | ($spl8 | $spl20 | $spl23 | $spl50 | $spl65)),
  inference(forward_demodulation,[],[f22848,f6699])).
fof(f22848,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,'T_0'))))))) | ($spl8 | $spl20 | $spl23 | $spl50)),
  inference(forward_demodulation,[],[f22668,f20789])).
fof(f22668,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))))) | ($spl8 | $spl23 | $spl50)),
  inference(unit_resulting_resolution,[],[f1632,f3931])).
fof(f3931,plain,(
  ( ! [X16] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,X16))))))) | s(bool,'T_0') = s(bool,X16)) ) | ($spl8 | $spl50)),
  inference(subsumption_resolution,[],[f3915,f78])).
fof(f3915,plain,(
  ( ! [X16] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,X16))))))) | s(bool,'T_0') = s(bool,X16)) ) | ($spl8 | $spl50)),
  inference(superposition,[],[f3617,f2894])).
fof(f3617,plain,(
  ( ! [X14] : (p(s(bool,X14)) | s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,X14)))))) ) | $spl8),
  inference(subsumption_resolution,[],[f3609,f80])).
fof(f3609,plain,(
  ( ! [X14] : (p(s(bool,X14)) | s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,X14))))) | s(bool,'T_0') = s(bool,X14)) ) | $spl8),
  inference(superposition,[],[f3594,f3373])).
fof(f22847,plain,(
  $false | ($spl8 | $spl20 | $spl23 | $spl50 | $spl65)),
  inference(subsumption_resolution,[],[f22846,f121])).
fof(f22846,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl8 | $spl20 | $spl23 | $spl50 | $spl65)),
  inference(forward_demodulation,[],[f22845,f6699])).
fof(f22845,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | ($spl8 | $spl20 | $spl23 | $spl50 | $spl65)),
  inference(forward_demodulation,[],[f22844,f6699])).
fof(f22844,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,'T_0'))))) | ($spl8 | $spl20 | $spl23 | $spl50)),
  inference(forward_demodulation,[],[f22669,f20789])).
fof(f22669,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))) | ($spl8 | $spl23 | $spl50)),
  inference(unit_resulting_resolution,[],[f78,f1632,f4043])).
fof(f4043,plain,(
  ( ! [X30,X29] : (p(s(bool,X30)) | s(bool,X30) = s(bool,f(s(bool,f(s(bool,X29))))) | s(bool,'T_0') = s(bool,X29)) ) | ($spl8 | $spl50)),
  inference(subsumption_resolution,[],[f4006,f78])).
fof(f4006,plain,(
  ( ! [X30,X29] : (p(s(bool,'F_0')) | p(s(bool,X30)) | s(bool,X30) = s(bool,f(s(bool,f(s(bool,X29))))) | s(bool,'T_0') = s(bool,X29)) ) | ($spl8 | $spl50)),
  inference(superposition,[],[f3595,f2894])).
fof(f3595,plain,(
  ( ! [X6,X5] : (p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X6) = s(bool,f(s(bool,X5)))) ) | $spl8),
  inference(subsumption_resolution,[],[f3571,f80])).
fof(f3571,plain,(
  ( ! [X6,X5] : (p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X6) = s(bool,f(s(bool,X5))) | s(bool,'T_0') = s(bool,X5)) ) | $spl8),
  inference(superposition,[],[f111,f3373])).
fof(f22843,plain,(
  $false | ($spl8 | $spl20 | $spl23 | $spl50 | $spl65)),
  inference(subsumption_resolution,[],[f22842,f121])).
fof(f22842,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl8 | $spl20 | $spl23 | $spl50 | $spl65)),
  inference(forward_demodulation,[],[f22841,f6699])).
fof(f22841,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | ($spl8 | $spl20 | $spl23 | $spl50)),
  inference(forward_demodulation,[],[f22840,f366])).
fof(f22840,plain,(
  s(bool,f(s(bool,'T_0'))) = s(bool,f(s(bool,'F_0'))) | ($spl8 | $spl20 | $spl23 | $spl50)),
  inference(forward_demodulation,[],[f22670,f20789])).
fof(f22670,plain,(
  s(bool,f(s(bool,'F_0'))) = s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))) | ($spl8 | $spl23 | $spl50)),
  inference(unit_resulting_resolution,[],[f78,f1632,f4054])).
fof(f4054,plain,(
  ( ! [X30,X29] : (p(s(bool,X30)) | s(bool,f(s(bool,X29))) = s(bool,f(s(bool,X30))) | s(bool,'T_0') = s(bool,X29)) ) | ($spl8 | $spl50)),
  inference(subsumption_resolution,[],[f4024,f78])).
fof(f4024,plain,(
  ( ! [X30,X29] : (p(s(bool,'F_0')) | p(s(bool,X30)) | s(bool,f(s(bool,X29))) = s(bool,f(s(bool,X30))) | s(bool,'T_0') = s(bool,X29)) ) | ($spl8 | $spl50)),
  inference(superposition,[],[f3595,f2894])).
fof(f22835,plain,(
  $false | ($spl8 | $spl20 | $spl23 | $spl50 | $spl65)),
  inference(subsumption_resolution,[],[f22834,f121])).
fof(f22834,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl8 | $spl20 | $spl23 | $spl50 | $spl65)),
  inference(forward_demodulation,[],[f22833,f6699])).
fof(f22833,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | ($spl8 | $spl20 | $spl23 | $spl50 | $spl65)),
  inference(forward_demodulation,[],[f22832,f6699])).
fof(f22832,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,'T_0'))))) | ($spl8 | $spl20 | $spl23 | $spl50 | $spl65)),
  inference(forward_demodulation,[],[f22831,f6699])).
fof(f22831,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,'T_0'))))))) | ($spl8 | $spl20 | $spl23 | $spl50 | $spl65)),
  inference(forward_demodulation,[],[f22830,f6699])).
fof(f22830,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,'T_0'))))))))) | ($spl8 | $spl20 | $spl23 | $spl50)),
  inference(forward_demodulation,[],[f22673,f20789])).
fof(f22673,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))))))) | ($spl8 | $spl23 | $spl50)),
  inference(unit_resulting_resolution,[],[f1632,f4321])).
fof(f4321,plain,(
  ( ! [X34] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X34))))))))) | s(bool,'T_0') = s(bool,X34)) ) | ($spl8 | $spl50)),
  inference(subsumption_resolution,[],[f4283,f78])).
fof(f4283,plain,(
  ( ! [X34] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X34))))))))) | s(bool,'T_0') = s(bool,X34)) ) | ($spl8 | $spl50)),
  inference(superposition,[],[f3617,f3900])).
fof(f22829,plain,(
  $false | ($spl8 | $spl20 | $spl23 | $spl47 | $spl50)),
  inference(subsumption_resolution,[],[f22828,f2225])).
fof(f22828,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | ($spl8 | $spl20 | $spl23 | $spl50)),
  inference(forward_demodulation,[],[f22827,f366])).
fof(f22827,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,f(s(bool,'F_0'))))) | ($spl8 | $spl20 | $spl23 | $spl50)),
  inference(forward_demodulation,[],[f22674,f20789])).
fof(f22674,plain,(
  s(bool,f(s(bool,f(s(bool,'F_0'))))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl8 | $spl23 | $spl50)),
  inference(unit_resulting_resolution,[],[f78,f1632,f4508])).
fof(f4508,plain,(
  ( ! [X24,X23] : (p(s(bool,X23)) | s(bool,X24) = s(bool,f(s(bool,f(s(bool,X23))))) | s(bool,'T_0') = s(bool,X24)) ) | ($spl8 | $spl50)),
  inference(subsumption_resolution,[],[f4457,f80])).
fof(f4457,plain,(
  ( ! [X24,X23] : (p(s(bool,X23)) | s(bool,X24) = s(bool,f(s(bool,f(s(bool,X23))))) | s(bool,'T_0') = s(bool,X24) | s(bool,'T_0') = s(bool,X23)) ) | ($spl8 | $spl50)),
  inference(superposition,[],[f171,f4310])).
fof(f22820,plain,(
  $false | ($spl8 | $spl20 | $spl23 | $spl50 | $spl65)),
  inference(subsumption_resolution,[],[f22819,f76])).
fof(f22819,plain,(
  ~p(s(bool,'T_0')) | ($spl8 | $spl20 | $spl23 | $spl50 | $spl65)),
  inference(forward_demodulation,[],[f22818,f6699])).
fof(f22818,plain,(
  ~p(s(bool,f(s(bool,'T_0')))) | ($spl8 | $spl20 | $spl23 | $spl50)),
  inference(forward_demodulation,[],[f22677,f20789])).
fof(f22677,plain,(
  ~p(s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | ($spl8 | $spl23 | $spl50)),
  inference(unit_resulting_resolution,[],[f1632,f5203])).
fof(f5203,plain,(
  ( ! [X12] : (~p(s(bool,f(s(bool,X12)))) | s(bool,'T_0') = s(bool,X12)) ) | ($spl8 | $spl50)),
  inference(subsumption_resolution,[],[f5131,f3900])).
fof(f5131,plain,(
  ( ! [X12] : (~p(s(bool,f(s(bool,X12)))) | s(bool,'F_0') != s(bool,f(s(bool,f(s(bool,X12))))) | s(bool,'T_0') = s(bool,X12)) ) | ($spl8 | $spl50)),
  inference(superposition,[],[f84,f4114])).
fof(f4114,plain,(
  ( ! [X16] : (s(bool,f(s(bool,X16))) = s(bool,f(s(bool,f(s(bool,X16))))) | s(bool,'T_0') = s(bool,X16)) ) | ($spl8 | $spl50)),
  inference(subsumption_resolution,[],[f4103,f78])).
fof(f4103,plain,(
  ( ! [X16] : (p(s(bool,'F_0')) | s(bool,f(s(bool,X16))) = s(bool,f(s(bool,f(s(bool,X16))))) | s(bool,'T_0') = s(bool,X16)) ) | ($spl8 | $spl50)),
  inference(superposition,[],[f3992,f2894])).
fof(f3992,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,X0) = s(bool,f(s(bool,X0)))) ) | $spl8),
  inference(factoring,[],[f3595])).
fof(f22817,plain,(
  $false | ($spl8 | $spl20 | $spl23 | $spl50 | $spl65)),
  inference(subsumption_resolution,[],[f22816,f76])).
fof(f22816,plain,(
  ~p(s(bool,'T_0')) | ($spl8 | $spl20 | $spl23 | $spl50 | $spl65)),
  inference(forward_demodulation,[],[f22815,f6699])).
fof(f22815,plain,(
  ~p(s(bool,f(s(bool,'T_0')))) | ($spl8 | $spl20 | $spl23 | $spl50 | $spl65)),
  inference(forward_demodulation,[],[f22814,f6699])).
fof(f22814,plain,(
  ~p(s(bool,f(s(bool,f(s(bool,'T_0')))))) | ($spl8 | $spl20 | $spl23 | $spl50)),
  inference(forward_demodulation,[],[f22678,f20789])).
fof(f22678,plain,(
  ~p(s(bool,f(s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))) | ($spl8 | $spl23 | $spl50)),
  inference(unit_resulting_resolution,[],[f1632,f6203])).
fof(f6203,plain,(
  ( ! [X0] : (~p(s(bool,f(s(bool,f(s(bool,X0)))))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl8 | $spl50)),
  inference(condensation,[],[f6202])).
fof(f6202,plain,(
  ( ! [X101,X102] : (~p(s(bool,f(s(bool,f(s(bool,X102)))))) | s(bool,'T_0') = s(bool,X101) | s(bool,'T_0') = s(bool,X102)) ) | ($spl8 | $spl50)),
  inference(duplicate_literal_removal,[],[f6059])).
fof(f6059,plain,(
  ( ! [X101,X102] : (~p(s(bool,f(s(bool,f(s(bool,X102)))))) | s(bool,'T_0') = s(bool,X101) | s(bool,'T_0') = s(bool,X102) | s(bool,'T_0') = s(bool,X101)) ) | ($spl8 | $spl50)),
  inference(superposition,[],[f5203,f3584])).
fof(f22807,plain,(
  $false | ($spl4 | $spl8 | $spl20 | $spl23 | $spl50 | $spl65)),
  inference(subsumption_resolution,[],[f22806,f22737])).
fof(f22806,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | ($spl4 | $spl8 | $spl20 | $spl23 | $spl65)),
  inference(forward_demodulation,[],[f22805,f6699])).
fof(f22805,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,f(s(bool,'T_0'))))) | ($spl4 | $spl8 | $spl20 | $spl23)),
  inference(forward_demodulation,[],[f22682,f20789])).
fof(f22682,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))) | ($spl4 | $spl8 | $spl23)),
  inference(unit_resulting_resolution,[],[f1632,f6225])).
fof(f6225,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,g(s(bool,f(s(bool,X0))))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl4 | $spl8)),
  inference(condensation,[],[f6224])).
fof(f6224,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') = s(bool,g(s(bool,f(s(bool,X37))))) | s(bool,'T_0') = s(bool,X36) | s(bool,'T_0') = s(bool,X37)) ) | ($spl4 | $spl8)),
  inference(duplicate_literal_removal,[],[f6031])).
fof(f6031,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') = s(bool,g(s(bool,f(s(bool,X37))))) | s(bool,'T_0') = s(bool,X36) | s(bool,'T_0') = s(bool,X37) | s(bool,'T_0') = s(bool,X36)) ) | ($spl4 | $spl8)),
  inference(superposition,[],[f3342,f3584])).
fof(f22801,plain,(
  $false | ($spl8 | $spl20 | $spl23 | $spl50 | $spl65)),
  inference(subsumption_resolution,[],[f22800,f76])).
fof(f22800,plain,(
  ~p(s(bool,'T_0')) | ($spl8 | $spl20 | $spl23 | $spl50 | $spl65)),
  inference(forward_demodulation,[],[f22799,f6699])).
fof(f22799,plain,(
  ~p(s(bool,f(s(bool,'T_0')))) | ($spl8 | $spl20 | $spl23 | $spl50 | $spl65)),
  inference(forward_demodulation,[],[f22798,f6699])).
fof(f22798,plain,(
  ~p(s(bool,f(s(bool,f(s(bool,'T_0')))))) | ($spl8 | $spl20 | $spl23 | $spl50 | $spl65)),
  inference(forward_demodulation,[],[f22797,f6699])).
fof(f22797,plain,(
  ~p(s(bool,f(s(bool,f(s(bool,f(s(bool,'T_0')))))))) | ($spl8 | $spl20 | $spl23 | $spl50)),
  inference(forward_demodulation,[],[f22684,f20789])).
fof(f22684,plain,(
  ~p(s(bool,f(s(bool,f(s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))))) | ($spl8 | $spl23 | $spl50)),
  inference(unit_resulting_resolution,[],[f1632,f6590])).
fof(f6590,plain,(
  ( ! [X0] : (~p(s(bool,f(s(bool,f(s(bool,f(s(bool,X0)))))))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl8 | $spl50)),
  inference(condensation,[],[f6589])).
fof(f6589,plain,(
  ( ! [X10,X11] : (~p(s(bool,f(s(bool,f(s(bool,f(s(bool,X11)))))))) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X11)) ) | ($spl8 | $spl50)),
  inference(duplicate_literal_removal,[],[f6529])).
fof(f6529,plain,(
  ( ! [X10,X11] : (~p(s(bool,f(s(bool,f(s(bool,f(s(bool,X11)))))))) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X10)) ) | ($spl8 | $spl50)),
  inference(superposition,[],[f6203,f3584])).
fof(f22791,plain,(
  $false | ($spl4 | $spl8 | $spl20 | $spl23 | $spl50 | $spl65)),
  inference(subsumption_resolution,[],[f22790,f22737])).
fof(f22790,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | ($spl4 | $spl8 | $spl20 | $spl23 | $spl65)),
  inference(forward_demodulation,[],[f22789,f6699])).
fof(f22789,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,f(s(bool,'T_0'))))) | ($spl4 | $spl8 | $spl20 | $spl23 | $spl65)),
  inference(forward_demodulation,[],[f22788,f6699])).
fof(f22788,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,f(s(bool,f(s(bool,'T_0'))))))) | ($spl4 | $spl8 | $spl20 | $spl23)),
  inference(forward_demodulation,[],[f22687,f20789])).
fof(f22687,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,f(s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))))) | ($spl4 | $spl8 | $spl23)),
  inference(unit_resulting_resolution,[],[f1632,f7103])).
fof(f7103,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,g(s(bool,f(s(bool,f(s(bool,X0))))))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl4 | $spl8)),
  inference(condensation,[],[f7102])).
fof(f7102,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') = s(bool,g(s(bool,f(s(bool,f(s(bool,X11))))))) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X11)) ) | ($spl4 | $spl8)),
  inference(duplicate_literal_removal,[],[f6983])).
fof(f6983,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') = s(bool,g(s(bool,f(s(bool,f(s(bool,X11))))))) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X10)) ) | ($spl4 | $spl8)),
  inference(superposition,[],[f6225,f3584])).
fof(f22787,plain,(
  $false | ($spl4 | $spl8 | $spl20 | $spl23 | $spl50 | $spl65)),
  inference(subsumption_resolution,[],[f22786,f22737])).
fof(f22786,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | ($spl4 | $spl8 | $spl20 | $spl23 | $spl65)),
  inference(forward_demodulation,[],[f22785,f6699])).
fof(f22785,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,f(s(bool,'T_0'))))) | ($spl4 | $spl8 | $spl20 | $spl23)),
  inference(forward_demodulation,[],[f22688,f20789])).
fof(f22688,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))) | ($spl4 | $spl8 | $spl23)),
  inference(unit_resulting_resolution,[],[f76,f1632,f7142])).
fof(f7142,plain,(
  ( ! [X15,X16] : (~p(s(bool,X16)) | s(bool,X16) = s(bool,g(s(bool,f(s(bool,X15))))) | s(bool,'T_0') = s(bool,X15)) ) | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f7041,f76])).
fof(f7041,plain,(
  ( ! [X15,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X16)) | s(bool,X16) = s(bool,g(s(bool,f(s(bool,X15))))) | s(bool,'T_0') = s(bool,X15)) ) | ($spl4 | $spl8)),
  inference(superposition,[],[f114,f6225])).
fof(f22780,plain,(
  $false | ($spl4 | $spl8 | $spl20 | $spl23 | $spl50)),
  inference(subsumption_resolution,[],[f22779,f364])).
fof(f22779,plain,(
  s(bool,'T_0') != s(bool,g(s(bool,'F_0'))) | ($spl4 | $spl8 | $spl20 | $spl23 | $spl50)),
  inference(forward_demodulation,[],[f22690,f20789])).
fof(f22690,plain,(
  s(bool,g(s(bool,'F_0'))) != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl4 | $spl8 | $spl23 | $spl50)),
  inference(unit_resulting_resolution,[],[f1632,f7419])).
fof(f7419,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) | ($spl4 | $spl8 | $spl50)),
  inference(duplicate_literal_removal,[],[f7391])).
fof(f7391,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | ($spl4 | $spl8 | $spl50)),
  inference(superposition,[],[f6948,f2894])).
fof(f22778,plain,(
  $false | ($spl4 | $spl8 | $spl20 | $spl23)),
  inference(subsumption_resolution,[],[f22777,f364])).
fof(f22777,plain,(
  s(bool,'T_0') != s(bool,g(s(bool,'F_0'))) | ($spl4 | $spl8 | $spl20 | $spl23)),
  inference(forward_demodulation,[],[f22691,f20789])).
fof(f22691,plain,(
  s(bool,g(s(bool,'F_0'))) != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl4 | $spl8 | $spl23)),
  inference(unit_resulting_resolution,[],[f121,f1632,f7420])).
fof(f7420,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,g(s(bool,X1))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1)) ) | ($spl4 | $spl8)),
  inference(duplicate_literal_removal,[],[f7390])).
fof(f7390,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,g(s(bool,X1))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1)) ) | ($spl4 | $spl8)),
  inference(superposition,[],[f6948,f3584])).
fof(f22776,plain,(
  $false | ($spl4 | $spl8 | $spl20 | $spl23 | $spl47)),
  inference(subsumption_resolution,[],[f22775,f364])).
fof(f22775,plain,(
  s(bool,'T_0') != s(bool,g(s(bool,'F_0'))) | ($spl4 | $spl8 | $spl20 | $spl23 | $spl47)),
  inference(forward_demodulation,[],[f22774,f366])).
fof(f22774,plain,(
  s(bool,'T_0') != s(bool,g(s(bool,f(s(bool,'F_0'))))) | ($spl4 | $spl8 | $spl20 | $spl23 | $spl47)),
  inference(forward_demodulation,[],[f22692,f20789])).
fof(f22692,plain,(
  s(bool,g(s(bool,f(s(bool,'F_0'))))) != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl4 | $spl8 | $spl23 | $spl47)),
  inference(unit_resulting_resolution,[],[f2225,f1632,f7420])).
fof(f22773,plain,(
  $false | ($spl3 | $spl4 | $spl8 | $spl13 | $spl20 | $spl23)),
  inference(subsumption_resolution,[],[f22772,f364])).
fof(f22772,plain,(
  s(bool,'T_0') != s(bool,g(s(bool,'F_0'))) | ($spl3 | $spl4 | $spl8 | $spl13 | $spl20 | $spl23)),
  inference(forward_demodulation,[],[f22771,f756])).
fof(f22771,plain,(
  s(bool,'T_0') != s(bool,g(s(bool,g(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))))) | ($spl3 | $spl4 | $spl8 | $spl13 | $spl20 | $spl23)),
  inference(forward_demodulation,[],[f22693,f20789])).
fof(f22693,plain,(
  s(bool,g(s(bool,g(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))))) != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl3 | $spl4 | $spl8 | $spl13 | $spl23)),
  inference(unit_resulting_resolution,[],[f750,f1632,f7420])).
fof(f22770,plain,(
  $false | ($spl3 | $spl4 | $spl8 | $spl17 | $spl20 | $spl23)),
  inference(subsumption_resolution,[],[f22769,f364])).
fof(f22769,plain,(
  s(bool,'T_0') != s(bool,g(s(bool,'F_0'))) | ($spl3 | $spl4 | $spl8 | $spl17 | $spl20 | $spl23)),
  inference(forward_demodulation,[],[f22768,f831])).
fof(f22768,plain,(
  s(bool,'T_0') != s(bool,g(s(bool,g(s(sK0,sK5(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))))) | ($spl3 | $spl4 | $spl8 | $spl17 | $spl20 | $spl23)),
  inference(forward_demodulation,[],[f22694,f20789])).
fof(f22694,plain,(
  s(bool,g(s(bool,g(s(sK0,sK5(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))))) != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl3 | $spl4 | $spl8 | $spl17 | $spl23)),
  inference(unit_resulting_resolution,[],[f825,f1632,f7420])).
fof(f22767,plain,(
  $false | ($spl3 | $spl4 | $spl8 | $spl20 | $spl23 | $spl32)),
  inference(subsumption_resolution,[],[f22766,f364])).
fof(f22766,plain,(
  s(bool,'T_0') != s(bool,g(s(bool,'F_0'))) | ($spl3 | $spl4 | $spl8 | $spl20 | $spl23 | $spl32)),
  inference(forward_demodulation,[],[f22765,f1987])).
fof(f22765,plain,(
  s(bool,'T_0') != s(bool,g(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))))) | ($spl3 | $spl4 | $spl8 | $spl20 | $spl23 | $spl32)),
  inference(forward_demodulation,[],[f22695,f20789])).
fof(f22695,plain,(
  s(bool,g(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))))) != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl3 | $spl4 | $spl8 | $spl23 | $spl32)),
  inference(unit_resulting_resolution,[],[f1971,f1632,f7420])).
fof(f22750,plain,(
  $false | ($spl8 | $spl20 | $spl23 | $spl50 | $spl65)),
  inference(subsumption_resolution,[],[f22749,f76])).
fof(f22749,plain,(
  ~p(s(bool,'T_0')) | ($spl8 | $spl20 | $spl23 | $spl50 | $spl65)),
  inference(forward_demodulation,[],[f22748,f6699])).
fof(f22748,plain,(
  ~p(s(bool,f(s(bool,'T_0')))) | ($spl8 | $spl20 | $spl23 | $spl50 | $spl65)),
  inference(forward_demodulation,[],[f22747,f6699])).
fof(f22747,plain,(
  ~p(s(bool,f(s(bool,f(s(bool,'T_0')))))) | ($spl8 | $spl20 | $spl23 | $spl50 | $spl65)),
  inference(forward_demodulation,[],[f22746,f6699])).
fof(f22746,plain,(
  ~p(s(bool,f(s(bool,f(s(bool,f(s(bool,'T_0')))))))) | ($spl8 | $spl20 | $spl23 | $spl50 | $spl65)),
  inference(forward_demodulation,[],[f22745,f6699])).
fof(f22745,plain,(
  ~p(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,'T_0')))))))))) | ($spl8 | $spl20 | $spl23 | $spl50)),
  inference(forward_demodulation,[],[f22704,f20789])).
fof(f22704,plain,(
  ~p(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))))))))) | ($spl8 | $spl23 | $spl50)),
  inference(unit_resulting_resolution,[],[f1632,f7671])).
fof(f7671,plain,(
  ( ! [X0] : (~p(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X0)))))))))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl8 | $spl50)),
  inference(condensation,[],[f7670])).
fof(f7670,plain,(
  ( ! [X10,X11] : (~p(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X11)))))))))) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X11)) ) | ($spl8 | $spl50)),
  inference(duplicate_literal_removal,[],[f7590])).
fof(f7590,plain,(
  ( ! [X10,X11] : (~p(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X11)))))))))) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X10)) ) | ($spl8 | $spl50)),
  inference(superposition,[],[f6590,f3584])).
fof(f22735,plain,(
  $false | ($spl4 | $spl8 | $spl20 | $spl23 | $spl47 | $spl50 | $spl65)),
  inference(subsumption_resolution,[],[f22734,f6699])).
fof(f22734,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl4 | $spl8 | $spl20 | $spl23 | $spl47 | $spl50 | $spl65)),
  inference(forward_demodulation,[],[f22733,f6699])).
fof(f22733,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,'T_0'))))) | ($spl4 | $spl8 | $spl20 | $spl23 | $spl47 | $spl50)),
  inference(forward_demodulation,[],[f22708,f20789])).
fof(f22708,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))) | ($spl4 | $spl8 | $spl23 | $spl47 | $spl50)),
  inference(unit_resulting_resolution,[],[f1632,f12631])).
fof(f12631,plain,(
  ( ! [X5] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,X5))))) | s(bool,'T_0') = s(bool,X5)) ) | ($spl4 | $spl8 | $spl47 | $spl50)),
  inference(forward_demodulation,[],[f12630,f364])).
fof(f12630,plain,(
  ( ! [X5] : (s(bool,g(s(bool,'F_0'))) != s(bool,f(s(bool,f(s(bool,X5))))) | s(bool,'T_0') = s(bool,X5)) ) | ($spl4 | $spl8 | $spl47 | $spl50)),
  inference(subsumption_resolution,[],[f12494,f4148])).
fof(f4148,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl47),
  inference(superposition,[],[f2225,f79])).
fof(f12494,plain,(
  ( ! [X5] : (s(bool,g(s(bool,'F_0'))) != s(bool,f(s(bool,f(s(bool,X5))))) | s(bool,'T_0') = s(bool,f(s(bool,X5))) | s(bool,'T_0') = s(bool,X5)) ) | ($spl4 | $spl8 | $spl50)),
  inference(superposition,[],[f7422,f3931])).
fof(f22732,plain,(
  $false | ($spl4 | $spl8 | $spl20 | $spl23 | $spl47 | $spl50 | $spl65)),
  inference(subsumption_resolution,[],[f22731,f6699])).
fof(f22731,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl4 | $spl8 | $spl20 | $spl23 | $spl47 | $spl50 | $spl65)),
  inference(forward_demodulation,[],[f22730,f6699])).
fof(f22730,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,'T_0'))))) | ($spl4 | $spl8 | $spl20 | $spl23 | $spl47 | $spl50 | $spl65)),
  inference(forward_demodulation,[],[f22729,f6699])).
fof(f22729,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,f(s(bool,'T_0'))))))) | ($spl4 | $spl8 | $spl20 | $spl23 | $spl47 | $spl50)),
  inference(forward_demodulation,[],[f22709,f20789])).
fof(f22709,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))))) | ($spl4 | $spl8 | $spl23 | $spl47 | $spl50)),
  inference(unit_resulting_resolution,[],[f1632,f12633])).
fof(f12633,plain,(
  ( ! [X6] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,f(s(bool,X6))))))) | s(bool,'T_0') = s(bool,X6)) ) | ($spl4 | $spl8 | $spl47 | $spl50)),
  inference(forward_demodulation,[],[f12632,f364])).
fof(f12632,plain,(
  ( ! [X6] : (s(bool,g(s(bool,'F_0'))) != s(bool,f(s(bool,f(s(bool,f(s(bool,X6))))))) | s(bool,'T_0') = s(bool,X6)) ) | ($spl4 | $spl8 | $spl47 | $spl50)),
  inference(subsumption_resolution,[],[f12495,f12631])).
fof(f12495,plain,(
  ( ! [X6] : (s(bool,g(s(bool,'F_0'))) != s(bool,f(s(bool,f(s(bool,f(s(bool,X6))))))) | s(bool,'T_0') = s(bool,f(s(bool,f(s(bool,X6))))) | s(bool,'T_0') = s(bool,X6)) ) | ($spl4 | $spl8 | $spl50)),
  inference(superposition,[],[f7422,f4321])).
fof(f22728,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f22727,f76])).
fof(f22727,plain,(
  ~p(s(bool,'T_0')) | ($spl20 | $spl23)),
  inference(forward_demodulation,[],[f22711,f20789])).
fof(f22711,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) | $spl23),
  inference(unit_resulting_resolution,[],[f76,f1632,f114])).
fof(f22726,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f22725,f76])).
fof(f22725,plain,(
  ~p(s(bool,'T_0')) | ($spl20 | $spl23)),
  inference(forward_demodulation,[],[f22714,f20789])).
fof(f22714,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) | $spl23),
  inference(unit_resulting_resolution,[],[f76,f1632,f114])).
fof(f1629,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl18),
  inference(cnf_transformation,[],[f1629_D])).
fof(f1629_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f1631,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | $spl21),
  inference(cnf_transformation,[],[f1631_D])).
fof(f1631_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f13008,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(superposition,[],[f4513,f135])).
fof(f4513,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl3),
  inference(unit_resulting_resolution,[],[f120,f104])).
fof(f13011,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | s(bool,X3) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,X3)) ) | $spl3),
  inference(superposition,[],[f4513,f135])).
fof(f14173,plain,(
  $false | ($spl4 | $spl8 | $spl21 | $spl24)),
  inference(subsumption_resolution,[],[f14018,f13702])).
fof(f13702,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl21 | $spl24)),
  inference(subsumption_resolution,[],[f13586,f1634])).
fof(f13586,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X0)) ) | $spl21),
  inference(trivial_inequality_removal,[],[f13578])).
fof(f13578,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X0)) ) | $spl21),
  inference(superposition,[],[f1631,f135])).
fof(f14018,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl4 | $spl8 | $spl21 | $spl24)),
  inference(backward_demodulation,[],[f13702,f13601])).
fof(f13601,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))))) | ($spl4 | $spl8 | $spl21)),
  inference(backward_demodulation,[],[f13549,f13565])).
fof(f13565,plain,(
  s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))))) != s(bool,g(s(bool,f(s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))))))))) | ($spl4 | $spl8 | $spl21)),
  inference(unit_resulting_resolution,[],[f1631,f7422])).
fof(f13549,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,f(s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))))))))) | ($spl4 | $spl8 | $spl21)),
  inference(unit_resulting_resolution,[],[f1631,f7103])).
fof(f14172,plain,(
  $false | ($spl21 | $spl24)),
  inference(subsumption_resolution,[],[f14017,f76])).
fof(f14017,plain,(
  ~p(s(bool,'T_0')) | ($spl21 | $spl24)),
  inference(backward_demodulation,[],[f13702,f13576])).
fof(f13576,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2)))) | $spl21),
  inference(unit_resulting_resolution,[],[f76,f1631,f114])).
fof(f14171,plain,(
  $false | ($spl0 | $spl3 | $spl21 | $spl24 | $spl32)),
  inference(subsumption_resolution,[],[f14014,f76])).
fof(f14014,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl3 | $spl21 | $spl24 | $spl32)),
  inference(backward_demodulation,[],[f13702,f2067])).
fof(f2067,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))) | ($spl0 | $spl3 | $spl32)),
  inference(unit_resulting_resolution,[],[f1972,f1972,f108])).
fof(f108,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f72])).
fof(f1972,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1)))) | ($spl0 | $spl3 | $spl32)),
  inference(unit_resulting_resolution,[],[f1971,f733])).
fof(f733,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK2)))) ) | $spl0),
  inference(subsumption_resolution,[],[f732,f78])).
fof(f732,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK2)))) ) | $spl0),
  inference(superposition,[],[f708,f79])).
fof(f708,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK2)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK1))))) ) | $spl0),
  inference(subsumption_resolution,[],[f689,f76])).
fof(f689,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK1)))) | p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK2))))) ) | $spl0),
  inference(superposition,[],[f105,f682])).
fof(f682,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | $spl0),
  inference(unit_resulting_resolution,[],[f117,f81])).
fof(f105,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f70])).
fof(f70,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))))) & ((p(s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X2))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f69])).
fof(f69,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))))) & (? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f68])).
fof(f68,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))) & (? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(nnf_transformation,[],[f45])).
fof(f45,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(ennf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) => p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(flattening,[],[f37])).
fof(f37,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) => p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X6] : ! [X7,X8] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))) <=> ! [X11] : (p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X7)))) => p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X8))))))),
  file('thm.pred_set.SUBSET_UNION_ABSORPTION','thm.pred_set.SUBSET_DEF')).
fof(f14166,plain,(
  $false | ($spl21 | $spl24)),
  inference(global_subsumption,[],[f75,f74,f76,f79,f81,f80,f114,f111,f115,f121,f135,f171,f960,f1021,f1022,f1032,f1083,f1401,f1402,f1133,f1134,f1144,f2096,f2099,f2218,f2217,f2216,f2215,f2214,f2445,f2447,f1029,f1139,f1138,f1137,f1027,f1026,f1025,f1141,f1099,f1098,f1097,f1406,f674,f673,f110,f109,f108,f103,f104,f2100,f1407,f1404,f675,f107,f106,f105,f1143,f1142,f1031,f1030,f1028,f1140,f2448,f2312,f723,f722,f623,f620,f618,f500,f499,f413,f412,f345,f343,f288,f247,f85,f84,f78,f5559,f5561,f5564,f5566,f5595,f5738,f5740,f5743,f5745,f5776,f6376,f6379,f6394,f6685,f6851,f7189,f7485,f7749,f7750,f7751,f7762,f7763,f7764,f8056,f8057,f8058,f8070,f8515,f8516,f8517,f8921,f9089,f1631,f1634,f13576,f13702,f14132,f14165])).
fof(f14165,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | ($spl21 | $spl24)),
  inference(subsumption_resolution,[],[f14004,f13702])).
fof(f14004,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X0,sK3(X2,X1,X0)),s(fun(X0,bool),X1))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | ($spl21 | $spl24)),
  inference(backward_demodulation,[],[f13702,f104])).
fof(f14132,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl21 | $spl24)),
  inference(subsumption_resolution,[],[f13704,f13702])).
fof(f13704,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | ($spl21 | $spl24)),
  inference(backward_demodulation,[],[f13702,f84])).
fof(f9089,plain,(
  ( ! [X6,X4,X5,X3] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X4)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X5)))) | s(bool,'T_0') = s(bool,X3) | s(bool,X6) = s(bool,'const.bool.IN_2'(s(bool,X3),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X4),s(fun(bool,bool),X5))))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f9074,f76])).
fof(f9074,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X4)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X5)))) | s(bool,'T_0') = s(bool,X3) | s(bool,X6) = s(bool,'const.bool.IN_2'(s(bool,X3),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X4),s(fun(bool,bool),X5))))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f723,f135])).
fof(f8921,plain,(
  ( ! [X6,X4,X5,X3] : (p(s(bool,'const.bool.IN_2'(s(bool,X5),s(fun(bool,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(bool,X5),s(fun(bool,bool),X4)))) | s(bool,'T_0') = s(bool,X5) | s(bool,X6) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X3),s(fun(bool,bool),X4))))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f8915,f76])).
fof(f8915,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(bool,X5),s(fun(bool,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(bool,X5),s(fun(bool,bool),X4)))) | s(bool,'T_0') = s(bool,X5) | s(bool,X6) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X3),s(fun(bool,bool),X4))))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f722,f135])).
fof(f8517,plain,(
  ( ! [X6,X7,X5] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(bool,X7),s(fun(bool,bool),X6))) | s(fun(bool,bool),X5) = s(fun(bool,bool),X6) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,sK3(X5,X6,bool)),s(fun(bool,bool),X5))) | s(bool,'T_0') = s(bool,X7) | s(bool,'T_0') = s(bool,sK3(X5,X6,bool))) )),
  inference(superposition,[],[f623,f135])).
fof(f8516,plain,(
  ( ! [X4,X2,X3] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X3))) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,sK3(X2,X3,bool)),s(fun(bool,bool),X2))) | s(bool,X4) = s(bool,sK3(X2,X3,bool)) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f623,f135])).
fof(f8515,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X1))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,sK3(X0,X1,bool)),s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,sK3(X0,X1,bool))) )),
  inference(superposition,[],[f623,f79])).
fof(f8070,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,sK3(X1,X2,X0)),s(fun(X0,bool),X2))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,sK3(X1,X2,X0)),s(fun(X0,bool),X1)))) )),
  inference(trivial_inequality_removal,[],[f8063])).
fof(f8063,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,'F_0') | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,sK3(X1,X2,X0)),s(fun(X0,bool),X2))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,sK3(X1,X2,X0)),s(fun(X0,bool),X1)))) )),
  inference(superposition,[],[f618,f79])).
fof(f8058,plain,(
  ( ! [X6,X7,X5] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(bool,X7),s(fun(bool,bool),X5))) | s(fun(bool,bool),X5) = s(fun(bool,bool),X6) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,sK3(X5,X6,bool)),s(fun(bool,bool),X6))) | s(bool,'T_0') = s(bool,X7) | s(bool,'T_0') = s(bool,sK3(X5,X6,bool))) )),
  inference(superposition,[],[f618,f135])).
fof(f8057,plain,(
  ( ! [X4,X2,X3] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X2))) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,sK3(X2,X3,bool)),s(fun(bool,bool),X3))) | s(bool,X4) = s(bool,sK3(X2,X3,bool)) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f618,f135])).
fof(f8056,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X0))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,sK3(X0,X1,bool)),s(fun(bool,bool),X1))) | s(bool,'T_0') = s(bool,sK3(X0,X1,bool))) )),
  inference(superposition,[],[f618,f79])).
fof(f7764,plain,(
  ( ! [X6,X7,X5] : (s(bool,X6) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X7))) | s(fun(bool,bool),X5) = s(fun(bool,bool),X7) | s(bool,'T_0') = s(bool,sK3(X7,X5,bool)) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X5)))) )),
  inference(superposition,[],[f620,f135])).
fof(f7763,plain,(
  ( ! [X4,X2,X3] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X3))) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,sK3(X3,X2,bool)) | s(bool,X4) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X2))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f620,f135])).
fof(f7762,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X1))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK3(X1,X0,bool)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))) )),
  inference(superposition,[],[f620,f79])).
fof(f7751,plain,(
  ( ! [X6,X7,X5] : (s(bool,X6) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X7))) | s(fun(bool,bool),X5) = s(fun(bool,bool),X7) | s(bool,'T_0') = s(bool,sK3(X5,X7,bool)) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X5)))) )),
  inference(superposition,[],[f620,f135])).
fof(f7750,plain,(
  ( ! [X4,X2,X3] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X3))) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,sK3(X2,X3,bool)) | s(bool,X4) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X2))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f620,f135])).
fof(f7749,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X1))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK3(X0,X1,bool)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))) )),
  inference(superposition,[],[f620,f79])).
fof(f7485,plain,(
  ( ! [X6,X4,X5,X3] : (p(s(bool,'const.bool.IN_2'(s(bool,X4),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X5),s(fun(bool,bool),X3)))))) | s(bool,'T_0') = s(bool,X4) | s(bool,X6) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X3))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f7479,f76])).
fof(f7479,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(bool,X4),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X5),s(fun(bool,bool),X3)))))) | s(bool,'T_0') = s(bool,X4) | s(bool,X6) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X3))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f500,f135])).
fof(f7189,plain,(
  ( ! [X6,X4,X5,X3] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X5),s(fun(bool,bool),X4)))))) | s(bool,'T_0') = s(bool,X3) | s(bool,X6) = s(bool,'const.bool.IN_2'(s(bool,X3),s(fun(bool,bool),X4))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f7174,f76])).
fof(f7174,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X5),s(fun(bool,bool),X4)))))) | s(bool,'T_0') = s(bool,X3) | s(bool,X6) = s(bool,'const.bool.IN_2'(s(bool,X3),s(fun(bool,bool),X4))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f499,f135])).
fof(f6851,plain,(
  ( ! [X6,X4,X5,X3] : (p(s(bool,'const.bool.IN_2'(s(bool,X4),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X3),s(fun(bool,bool),X5)))))) | s(bool,'T_0') = s(bool,X4) | s(bool,X6) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X3))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f6845,f76])).
fof(f6845,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(bool,X4),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X3),s(fun(bool,bool),X5)))))) | s(bool,'T_0') = s(bool,X4) | s(bool,X6) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X3))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f413,f135])).
fof(f6685,plain,(
  ( ! [X6,X4,X5,X3] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X4),s(fun(bool,bool),X5)))))) | s(bool,'T_0') = s(bool,X3) | s(bool,X6) = s(bool,'const.bool.IN_2'(s(bool,X3),s(fun(bool,bool),X4))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f6670,f76])).
fof(f6670,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X4),s(fun(bool,bool),X5)))))) | s(bool,'T_0') = s(bool,X3) | s(bool,X6) = s(bool,'const.bool.IN_2'(s(bool,X3),s(fun(bool,bool),X4))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f412,f135])).
fof(f6394,plain,(
  ( ! [X8,X7,X9] : (s(fun(X7,bool),X8) = s(fun(X7,bool),X9) | s(bool,'T_0') = s(bool,g(s(X7,sK5(X8,X9,bool,X7)))) | s(bool,'T_0') = s(bool,f(s(X7,sK5(X8,X9,bool,X7))))) )),
  inference(trivial_inequality_removal,[],[f6387])).
fof(f6387,plain,(
  ( ! [X8,X7,X9] : (s(bool,'F_0') != s(bool,'F_0') | s(fun(X7,bool),X8) = s(fun(X7,bool),X9) | s(bool,'T_0') = s(bool,g(s(X7,sK5(X8,X9,bool,X7)))) | s(bool,'T_0') = s(bool,f(s(X7,sK5(X8,X9,bool,X7))))) )),
  inference(superposition,[],[f345,f79])).
fof(f5776,plain,(
  ( ! [X6,X4,X8,X7,X5] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X4,X5),s(fun(X4,bool),'const.pred_set.UNION_2'(s(fun(X4,bool),X6),s(fun(X4,bool),X7))))) | s(bool,X8) = s(bool,'const.bool.IN_2'(s(X4,X5),s(fun(X4,bool),X6))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(subsumption_resolution,[],[f5762,f76])).
fof(f5762,plain,(
  ( ! [X6,X4,X8,X7,X5] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X4,X5),s(fun(X4,bool),'const.pred_set.UNION_2'(s(fun(X4,bool),X6),s(fun(X4,bool),X7))))) | s(bool,X8) = s(bool,'const.bool.IN_2'(s(X4,X5),s(fun(X4,bool),X6))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f674,f135])).
fof(f5745,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,'const.bool.IN_2'(s(bool,X30),s(fun(bool,bool),X31)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X31),s(fun(bool,bool),X32))))) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f674,f79])).
fof(f5743,plain,(
  ( ! [X24,X23,X25,X22] : (~p(s(bool,'const.bool.IN_2'(s(bool,X22),s(fun(bool,bool),X23)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X23),s(fun(bool,bool),X24))))) | s(bool,X22) = s(bool,X25) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f674,f135])).
fof(f5740,plain,(
  ( ! [X14,X12,X13,X11] : (~p(s(bool,'const.bool.IN_2'(s(bool,X12),s(fun(bool,bool),X13)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,X11),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X13),s(fun(bool,bool),X14))))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f674,f135])).
fof(f5738,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X4)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,X3),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X4),s(fun(bool,bool),X5))))) | s(bool,X3) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f674,f135])).
fof(f5595,plain,(
  ( ! [X6,X4,X8,X7,X5] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X4,X5),s(fun(X4,bool),'const.pred_set.UNION_2'(s(fun(X4,bool),X7),s(fun(X4,bool),X6))))) | s(bool,X8) = s(bool,'const.bool.IN_2'(s(X4,X5),s(fun(X4,bool),X6))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(subsumption_resolution,[],[f5582,f76])).
fof(f5582,plain,(
  ( ! [X6,X4,X8,X7,X5] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X4,X5),s(fun(X4,bool),'const.pred_set.UNION_2'(s(fun(X4,bool),X7),s(fun(X4,bool),X6))))) | s(bool,X8) = s(bool,'const.bool.IN_2'(s(X4,X5),s(fun(X4,bool),X6))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f673,f135])).
fof(f5566,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,'const.bool.IN_2'(s(bool,X30),s(fun(bool,bool),X31)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X32),s(fun(bool,bool),X31))))) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f673,f79])).
fof(f5564,plain,(
  ( ! [X24,X23,X25,X22] : (~p(s(bool,'const.bool.IN_2'(s(bool,X22),s(fun(bool,bool),X23)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X24),s(fun(bool,bool),X23))))) | s(bool,X22) = s(bool,X25) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f673,f135])).
fof(f5561,plain,(
  ( ! [X14,X12,X13,X11] : (~p(s(bool,'const.bool.IN_2'(s(bool,X12),s(fun(bool,bool),X13)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,X11),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X14),s(fun(bool,bool),X13))))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f673,f135])).
fof(f5559,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X4)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,X3),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X5),s(fun(bool,bool),X4))))) | s(bool,X3) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f673,f135])).
fof(f247,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(bool,bool),X3),s(fun(bool,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X3)))) | s(bool,'T_0') = s(bool,sK4(X2,X3,bool))) )),
  inference(superposition,[],[f106,f79])).
fof(f288,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(bool,bool),X3),s(fun(bool,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X2)))) | s(bool,'T_0') = s(bool,sK4(X2,X3,bool))) )),
  inference(superposition,[],[f107,f79])).
fof(f343,plain,(
  ( ! [X4,X5,X3] : (s(X5,f(s(bool,'F_0'))) != s(X5,g(s(bool,'F_0'))) | s(fun(bool,X5),X3) = s(fun(bool,X5),X4) | s(bool,'T_0') = s(bool,sK5(X3,X4,X5,bool))) )),
  inference(superposition,[],[f115,f79])).
fof(f345,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,f(s(X3,sK5(X4,X5,bool,X3)))) | s(fun(X3,bool),X4) = s(fun(X3,bool),X5) | s(bool,'T_0') = s(bool,g(s(X3,sK5(X4,X5,bool,X3))))) )),
  inference(superposition,[],[f115,f79])).
fof(f412,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,'const.bool.IN_2'(s(bool,X5),s(fun(bool,bool),X6)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X6),s(fun(bool,bool),X7)))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f109,f79])).
fof(f413,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X9)))) | p(s(bool,'const.bool.IN_2'(s(bool,X8),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X9),s(fun(bool,bool),X10)))))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f109,f79])).
fof(f499,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,'const.bool.IN_2'(s(bool,X5),s(fun(bool,bool),X7)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X6),s(fun(bool,bool),X7)))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f110,f79])).
fof(f500,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X10)))) | p(s(bool,'const.bool.IN_2'(s(bool,X8),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X9),s(fun(bool,bool),X10)))))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f110,f79])).
fof(f618,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(X3,sK3(X4,X5,X3)),s(fun(X3,bool),X4))) | s(fun(X3,bool),X4) = s(fun(X3,bool),X5) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X3,sK3(X4,X5,X3)),s(fun(X3,bool),X5)))) )),
  inference(superposition,[],[f104,f79])).
fof(f620,plain,(
  ( ! [X2,X3] : (s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X2))) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X3))) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,sK3(X2,X3,bool))) )),
  inference(superposition,[],[f104,f79])).
fof(f623,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(X3,sK3(X4,X5,X3)),s(fun(X3,bool),X5))) | s(fun(X3,bool),X4) = s(fun(X3,bool),X5) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X3,sK3(X4,X5,X3)),s(fun(X3,bool),X4)))) )),
  inference(superposition,[],[f104,f79])).
fof(f722,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X1),s(fun(bool,bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),X2)))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f108,f79])).
fof(f723,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'const.bool.IN_2'(s(bool,X3),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X4),s(fun(bool,bool),X5)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X4)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X5)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f108,f79])).
fof(f2312,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))) | s(bool,'T_0') = s(bool,sK4(X1,X0,bool)) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(bool,bool),X0),s(fun(bool,bool),X1)))) )),
  inference(condensation,[],[f2311])).
fof(f2311,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X8)) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X6)))) | s(bool,'T_0') = s(bool,sK4(X7,X6,bool)) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(bool,bool),X6),s(fun(bool,bool),X7)))) )),
  inference(subsumption_resolution,[],[f2308,f80])).
fof(f2308,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X8)) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X6)))) | s(bool,'T_0') = s(bool,sK4(X7,X6,bool)) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(bool,bool),X6),s(fun(bool,bool),X7))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f247,f135])).
fof(f2448,plain,(
  ( ! [X12,X13] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X13)))) | s(bool,'T_0') = s(bool,sK4(X13,X12,bool)) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(bool,bool),X12),s(fun(bool,bool),X13)))) )),
  inference(subsumption_resolution,[],[f2444,f78])).
fof(f2444,plain,(
  ( ! [X12,X13] : (p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X13)))) | s(bool,'T_0') = s(bool,sK4(X13,X12,bool)) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(bool,bool),X12),s(fun(bool,bool),X13)))) )),
  inference(superposition,[],[f288,f79])).
fof(f1140,plain,(
  ( ! [X111,X109,X110,X108] : (s(bool,'T_0') != s(bool,f(s(X108,sK5(X109,X110,bool,X108)))) | s(fun(X108,bool),X109) = s(fun(X108,bool),X110) | s(bool,X111) = s(bool,g(s(X108,sK5(X109,X110,bool,X108)))) | s(bool,'T_0') = s(bool,X111)) )),
  inference(superposition,[],[f115,f135])).
fof(f1028,plain,(
  ( ! [X118,X120,X119,X117] : (s(bool,X120) != s(bool,f(s(X117,sK5(X118,X119,bool,X117)))) | s(fun(X117,bool),X118) = s(fun(X117,bool),X119) | s(bool,'T_0') = s(bool,X120) | s(bool,'T_0') = s(bool,g(s(X117,sK5(X118,X119,bool,X117))))) )),
  inference(superposition,[],[f115,f135])).
fof(f1030,plain,(
  ( ! [X125,X126,X124] : (~p(s(bool,'const.bool.IN_2'(s(bool,X126),s(fun(bool,bool),X124)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(bool,bool),X125),s(fun(bool,bool),X124)))) | s(bool,'T_0') = s(bool,X126) | s(bool,'T_0') = s(bool,sK4(X124,X125,bool))) )),
  inference(superposition,[],[f107,f135])).
fof(f1031,plain,(
  ( ! [X127,X128,X129] : (p(s(bool,'const.bool.IN_2'(s(bool,X129),s(fun(bool,bool),X128)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(bool,bool),X128),s(fun(bool,bool),X127)))) | s(bool,'T_0') = s(bool,X129) | s(bool,'T_0') = s(bool,sK4(X127,X128,bool))) )),
  inference(superposition,[],[f106,f135])).
fof(f1142,plain,(
  ( ! [X116,X117,X115] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X115)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(bool,bool),X116),s(fun(bool,bool),X115)))) | s(bool,X117) = s(bool,sK4(X115,X116,bool)) | s(bool,'T_0') = s(bool,X117)) )),
  inference(superposition,[],[f107,f135])).
fof(f1143,plain,(
  ( ! [X118,X120,X119] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X119)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(bool,bool),X119),s(fun(bool,bool),X118)))) | s(bool,X120) = s(bool,sK4(X118,X119,bool)) | s(bool,'T_0') = s(bool,X120)) )),
  inference(superposition,[],[f106,f135])).
fof(f106,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X1)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f70])).
fof(f107,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X2)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f70])).
fof(f675,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X3,bool),X5),s(fun(X3,bool),X4)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X3,sK4(X4,X5,X3)),s(fun(X3,bool),X5)))) )),
  inference(subsumption_resolution,[],[f250,f78])).
fof(f250,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,'F_0')) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X3,bool),X5),s(fun(X3,bool),X4)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X3,sK4(X4,X5,X3)),s(fun(X3,bool),X5)))) )),
  inference(superposition,[],[f106,f79])).
fof(f1404,plain,(
  ( ! [X57,X54,X56,X55,X53] : (~p(s(bool,'const.bool.IN_2'(s(X53,X56),s(fun(X53,bool),X54)))) | p(s(bool,'const.bool.IN_2'(s(X53,X56),s(fun(X53,bool),X55)))) | s(bool,X57) = s(bool,'const.pred_set.SUBSET_2'(s(fun(X53,bool),X54),s(fun(X53,bool),X55))) | s(bool,'T_0') = s(bool,X57)) )),
  inference(subsumption_resolution,[],[f1127,f76])).
fof(f1127,plain,(
  ( ! [X57,X54,X56,X55,X53] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(X53,X56),s(fun(X53,bool),X54)))) | p(s(bool,'const.bool.IN_2'(s(X53,X56),s(fun(X53,bool),X55)))) | s(bool,X57) = s(bool,'const.pred_set.SUBSET_2'(s(fun(X53,bool),X54),s(fun(X53,bool),X55))) | s(bool,'T_0') = s(bool,X57)) )),
  inference(superposition,[],[f105,f135])).
fof(f1407,plain,(
  ( ! [X92,X90,X91,X89] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X89,bool),X91),s(fun(X89,bool),X90)))) | s(bool,X92) = s(bool,'const.bool.IN_2'(s(X89,sK4(X90,X91,X89)),s(fun(X89,bool),X90))) | s(bool,'T_0') = s(bool,X92)) )),
  inference(subsumption_resolution,[],[f1136,f76])).
fof(f1136,plain,(
  ( ! [X92,X90,X91,X89] : (~p(s(bool,'T_0')) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X89,bool),X91),s(fun(X89,bool),X90)))) | s(bool,X92) = s(bool,'const.bool.IN_2'(s(X89,sK4(X90,X91,X89)),s(fun(X89,bool),X90))) | s(bool,'T_0') = s(bool,X92)) )),
  inference(superposition,[],[f107,f135])).
fof(f2100,plain,(
  ( ! [X17,X18,X16] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X16,sK4(X18,X17,X16)),s(fun(X16,bool),X17))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(X16,bool),X17),s(fun(X16,bool),X18)))) )),
  inference(subsumption_resolution,[],[f2095,f78])).
fof(f2095,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X16,sK4(X18,X17,X16)),s(fun(X16,bool),X17))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(X16,bool),X17),s(fun(X16,bool),X18)))) )),
  inference(superposition,[],[f675,f79])).
fof(f103,plain,(
  ( ! [X4,X2,X0,X1] : (s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) )),
  inference(cnf_transformation,[],[f67])).
fof(f109,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f72])).
fof(f674,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,'const.bool.IN_2'(s(X8,X9),s(fun(X8,bool),X10)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X8,X9),s(fun(X8,bool),'const.pred_set.UNION_2'(s(fun(X8,bool),X10),s(fun(X8,bool),X11)))))) )),
  inference(subsumption_resolution,[],[f425,f78])).
fof(f425,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(X8,X9),s(fun(X8,bool),X10)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X8,X9),s(fun(X8,bool),'const.pred_set.UNION_2'(s(fun(X8,bool),X10),s(fun(X8,bool),X11)))))) )),
  inference(superposition,[],[f109,f79])).
fof(f1406,plain,(
  ( ! [X66,X64,X62,X65,X63] : (p(s(bool,'const.bool.IN_2'(s(X62,X63),s(fun(X62,bool),X64)))) | p(s(bool,'const.bool.IN_2'(s(X62,X63),s(fun(X62,bool),X65)))) | s(bool,X66) = s(bool,'const.bool.IN_2'(s(X62,X63),s(fun(X62,bool),'const.pred_set.UNION_2'(s(fun(X62,bool),X64),s(fun(X62,bool),X65))))) | s(bool,'T_0') = s(bool,X66)) )),
  inference(subsumption_resolution,[],[f1130,f76])).
fof(f1130,plain,(
  ( ! [X66,X64,X62,X65,X63] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(X62,X63),s(fun(X62,bool),X64)))) | p(s(bool,'const.bool.IN_2'(s(X62,X63),s(fun(X62,bool),X65)))) | s(bool,X66) = s(bool,'const.bool.IN_2'(s(X62,X63),s(fun(X62,bool),'const.pred_set.UNION_2'(s(fun(X62,bool),X64),s(fun(X62,bool),X65))))) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f108,f135])).
fof(f1097,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'const.bool.IN_2'(s(bool,X35),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X36),s(fun(bool,bool),X37)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X36)))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X37)))) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f108,f135])).
fof(f1098,plain,(
  ( ! [X39,X41,X42,X40] : (p(s(bool,'const.bool.IN_2'(s(bool,X39),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X40),s(fun(bool,bool),X41)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X41)))) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f110,f135])).
fof(f1099,plain,(
  ( ! [X45,X43,X46,X44] : (p(s(bool,'const.bool.IN_2'(s(bool,X43),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X44),s(fun(bool,bool),X45)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X44)))) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f109,f135])).
fof(f1141,plain,(
  ( ! [X114,X112,X113] : (s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X112))) != s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X113))) | s(fun(bool,bool),X112) = s(fun(bool,bool),X113) | s(bool,X114) = s(bool,sK3(X112,X113,bool)) | s(bool,'T_0') = s(bool,X114)) )),
  inference(superposition,[],[f104,f135])).
fof(f1025,plain,(
  ( ! [X103,X101,X102,X104] : (~p(s(bool,'const.bool.IN_2'(s(bool,X102),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X103),s(fun(bool,bool),X104)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X101),s(fun(bool,bool),X103)))) | p(s(bool,'const.bool.IN_2'(s(bool,X101),s(fun(bool,bool),X104)))) | s(bool,'T_0') = s(bool,X102) | s(bool,'T_0') = s(bool,X101)) )),
  inference(superposition,[],[f108,f135])).
fof(f1026,plain,(
  ( ! [X107,X105,X108,X106] : (p(s(bool,'const.bool.IN_2'(s(bool,X106),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X107),s(fun(bool,bool),X108)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X105),s(fun(bool,bool),X108)))) | s(bool,'T_0') = s(bool,X106) | s(bool,'T_0') = s(bool,X105)) )),
  inference(superposition,[],[f110,f135])).
fof(f1027,plain,(
  ( ! [X111,X109,X112,X110] : (p(s(bool,'const.bool.IN_2'(s(bool,X110),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X111),s(fun(bool,bool),X112)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X109),s(fun(bool,bool),X111)))) | s(bool,'T_0') = s(bool,X110) | s(bool,'T_0') = s(bool,X109)) )),
  inference(superposition,[],[f109,f135])).
fof(f1137,plain,(
  ( ! [X94,X95,X93,X96] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X94),s(fun(bool,bool),X95)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X93),s(fun(bool,bool),X94)))) | p(s(bool,'const.bool.IN_2'(s(bool,X93),s(fun(bool,bool),X95)))) | s(bool,X93) = s(bool,X96) | s(bool,'T_0') = s(bool,X96)) )),
  inference(superposition,[],[f108,f135])).
fof(f1138,plain,(
  ( ! [X99,X97,X100,X98] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X98),s(fun(bool,bool),X99)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X97),s(fun(bool,bool),X99)))) | s(bool,X97) = s(bool,X100) | s(bool,'T_0') = s(bool,X100)) )),
  inference(superposition,[],[f110,f135])).
fof(f1139,plain,(
  ( ! [X103,X101,X102,X104] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X102),s(fun(bool,bool),X103)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X101),s(fun(bool,bool),X102)))) | s(bool,X101) = s(bool,X104) | s(bool,'T_0') = s(bool,X104)) )),
  inference(superposition,[],[f109,f135])).
fof(f1029,plain,(
  ( ! [X123,X121,X122] : (s(bool,'const.bool.IN_2'(s(bool,X123),s(fun(bool,bool),X121))) != s(bool,'const.bool.IN_2'(s(bool,X123),s(fun(bool,bool),X122))) | s(fun(bool,bool),X121) = s(fun(bool,bool),X122) | s(bool,'T_0') = s(bool,X123) | s(bool,'T_0') = s(bool,sK3(X121,X122,bool))) )),
  inference(superposition,[],[f104,f135])).
fof(f1144,plain,(
  ( ! [X123,X121,X124,X122] : (s(X123,f(s(bool,'T_0'))) != s(X123,g(s(bool,'T_0'))) | s(fun(bool,X123),X121) = s(fun(bool,X123),X122) | s(bool,X124) = s(bool,sK5(X121,X122,X123,bool)) | s(bool,'T_0') = s(bool,X124)) )),
  inference(superposition,[],[f115,f135])).
fof(f1134,plain,(
  ( ! [X83,X81,X84,X82] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X81,sK3(X82,X83,X81)),s(fun(X81,bool),X83))) | s(fun(X81,bool),X82) = s(fun(X81,bool),X83) | s(bool,X84) = s(bool,'const.bool.IN_2'(s(X81,sK3(X82,X83,X81)),s(fun(X81,bool),X82))) | s(bool,'T_0') = s(bool,X84)) )),
  inference(superposition,[],[f104,f135])).
fof(f1133,plain,(
  ( ! [X80,X78,X79,X77] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X77,sK3(X78,X79,X77)),s(fun(X77,bool),X78))) | s(fun(X77,bool),X78) = s(fun(X77,bool),X79) | s(bool,X80) = s(bool,'const.bool.IN_2'(s(X77,sK3(X78,X79,X77)),s(fun(X77,bool),X79))) | s(bool,'T_0') = s(bool,X80)) )),
  inference(superposition,[],[f104,f135])).
fof(f1402,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f1107,f76])).
fof(f1107,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f114,f135])).
fof(f1401,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f1103,f76])).
fof(f1103,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f84,f135])).
fof(f1083,plain,(
  ( ! [X4,X2,X3] : (s(bool,X2) != s(bool,X3) | p(s(bool,X3)) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f80,f135])).
fof(f1032,plain,(
  ( ! [X132,X130,X133,X131] : (s(X132,f(s(bool,X133))) != s(X132,g(s(bool,X133))) | s(fun(bool,X132),X130) = s(fun(bool,X132),X131) | s(bool,'T_0') = s(bool,X133) | s(bool,'T_0') = s(bool,sK5(X130,X131,X132,bool))) )),
  inference(superposition,[],[f115,f135])).
fof(f1022,plain,(
  ( ! [X92,X90,X91,X89] : (s(bool,X92) != s(bool,'const.bool.IN_2'(s(X89,sK3(X90,X91,X89)),s(fun(X89,bool),X91))) | s(fun(X89,bool),X90) = s(fun(X89,bool),X91) | s(bool,'T_0') = s(bool,X92) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X89,sK3(X90,X91,X89)),s(fun(X89,bool),X90)))) )),
  inference(superposition,[],[f104,f135])).
fof(f1021,plain,(
  ( ! [X88,X87,X85,X86] : (s(bool,X88) != s(bool,'const.bool.IN_2'(s(X85,sK3(X86,X87,X85)),s(fun(X85,bool),X86))) | s(fun(X85,bool),X86) = s(fun(X85,bool),X87) | s(bool,'T_0') = s(bool,X88) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X85,sK3(X86,X87,X85)),s(fun(X85,bool),X87)))) )),
  inference(superposition,[],[f104,f135])).
fof(f960,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f135,f135])).
fof(f14138,plain,(
  $false | ($spl8 | $spl21 | $spl24)),
  inference(global_subsumption,[],[f14137,f75,f74,f76,f79,f81,f80,f114,f111,f115,f121,f135,f171,f960,f1021,f1022,f1032,f1083,f1401,f1402,f1133,f1134,f1144,f2096,f2099,f2218,f2217,f2216,f2215,f2214,f2445,f2447,f1029,f1139,f1138,f1137,f1027,f1026,f1025,f1141,f1099,f1098,f1097,f1406,f674,f673,f110,f109,f108,f103,f104,f2100,f1407,f1404,f675,f107,f106,f105,f1143,f1142,f1031,f1030,f1028,f1140,f2448,f2312,f723,f722,f623,f620,f618,f500,f499,f413,f412,f345,f343,f288,f247,f85,f84,f78,f5559,f5561,f5564,f5566,f5595,f5738,f5740,f5743,f5745,f5776,f6376,f6379,f6394,f6685,f6851,f7189,f7485,f7749,f7750,f7751,f7762,f7763,f7764,f8056,f8057,f8058,f8070,f8515,f8516,f8517,f8921,f9089,f1631,f1634,f13576,f13702,f14132])).
fof(f14137,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl8 | $spl21 | $spl24)),
  inference(subsumption_resolution,[],[f14130,f80])).
fof(f14130,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | ($spl8 | $spl21 | $spl24)),
  inference(condensation,[],[f13802])).
fof(f13802,plain,(
  ( ! [X6,X5] : (s(bool,'T_0') = s(bool,X6) | p(s(bool,X5)) | p(s(bool,X6))) ) | ($spl8 | $spl21 | $spl24)),
  inference(backward_demodulation,[],[f13702,f3595])).
fof(f14136,plain,(
  $false | ($spl21 | $spl24 | $spl47)),
  inference(subsumption_resolution,[],[f13726,f13702])).
fof(f13726,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl21 | $spl24 | $spl47)),
  inference(backward_demodulation,[],[f13702,f2225])).
fof(f14133,plain,(
  $false | ($spl21 | $spl24)),
  inference(global_subsumption,[],[f75,f74,f76,f79,f81,f80,f114,f111,f115,f121,f135,f171,f960,f1021,f1022,f1032,f1083,f1401,f1402,f1133,f1134,f1144,f2096,f2099,f2218,f2217,f2216,f2215,f2214,f2445,f2447,f1029,f1139,f1138,f1137,f1027,f1026,f1025,f1141,f1099,f1098,f1097,f1406,f674,f673,f110,f109,f108,f103,f104,f2100,f1407,f1404,f675,f107,f106,f105,f1143,f1142,f1031,f1030,f1028,f1140,f2448,f2312,f723,f722,f623,f620,f618,f500,f499,f413,f412,f345,f343,f288,f247,f85,f84,f78,f5559,f5561,f5564,f5566,f5595,f5738,f5740,f5743,f5745,f5776,f6376,f6379,f6394,f6685,f6851,f7189,f7485,f7749,f7750,f7751,f7762,f7763,f7764,f8056,f8057,f8058,f8070,f8515,f8516,f8517,f8921,f9089,f1631,f1634,f13576,f13702,f14132])).
fof(f14131,plain,(
  $false | ($spl21 | $spl24)),
  inference(subsumption_resolution,[],[f13703,f76])).
fof(f13703,plain,(
  ~p(s(bool,'T_0')) | ($spl21 | $spl24)),
  inference(backward_demodulation,[],[f13702,f78])).
fof(f13701,plain,(
  $false | ($spl21 | $spl27)),
  inference(subsumption_resolution,[],[f13485,f1636])).
fof(f13485,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | $spl21),
  inference(unit_resulting_resolution,[],[f1631,f79])).
fof(f13700,plain,(
  $false | ($spl21 | $spl27)),
  inference(subsumption_resolution,[],[f13487,f1636])).
fof(f13487,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | $spl21),
  inference(unit_resulting_resolution,[],[f121,f1631,f135])).
fof(f13699,plain,(
  $false | ($spl8 | $spl21 | $spl27 | $spl47)),
  inference(subsumption_resolution,[],[f13698,f1636])).
fof(f13698,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl8 | $spl21 | $spl47)),
  inference(forward_demodulation,[],[f13488,f366])).
fof(f13488,plain,(
  s(bool,f(s(bool,'F_0'))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl21 | $spl47)),
  inference(unit_resulting_resolution,[],[f2225,f1631,f135])).
fof(f13697,plain,(
  $false | ($spl3 | $spl13 | $spl21 | $spl27)),
  inference(subsumption_resolution,[],[f13696,f1636])).
fof(f13696,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl13 | $spl21)),
  inference(forward_demodulation,[],[f13489,f756])).
fof(f13489,plain,(
  s(bool,g(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl13 | $spl21)),
  inference(unit_resulting_resolution,[],[f750,f1631,f135])).
fof(f13695,plain,(
  $false | ($spl3 | $spl17 | $spl21 | $spl27)),
  inference(subsumption_resolution,[],[f13694,f1636])).
fof(f13694,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl17 | $spl21)),
  inference(forward_demodulation,[],[f13490,f831])).
fof(f13490,plain,(
  s(bool,g(s(sK0,sK5(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl17 | $spl21)),
  inference(unit_resulting_resolution,[],[f825,f1631,f135])).
fof(f13693,plain,(
  $false | ($spl3 | $spl21 | $spl27 | $spl32)),
  inference(subsumption_resolution,[],[f13692,f1636])).
fof(f13692,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl21 | $spl32)),
  inference(forward_demodulation,[],[f13491,f1987])).
fof(f13491,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl21 | $spl32)),
  inference(unit_resulting_resolution,[],[f1971,f1631,f135])).
fof(f13691,plain,(
  $false | ($spl3 | $spl21 | $spl23)),
  inference(subsumption_resolution,[],[f13494,f1632])).
fof(f13494,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl3 | $spl21)),
  inference(unit_resulting_resolution,[],[f4513,f1631,f135])).
fof(f13690,plain,(
  $false | ($spl21 | $spl27)),
  inference(subsumption_resolution,[],[f13495,f1636])).
fof(f13495,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | $spl21),
  inference(unit_resulting_resolution,[],[f121,f1631,f135])).
fof(f13689,plain,(
  $false | ($spl8 | $spl21 | $spl27 | $spl47)),
  inference(subsumption_resolution,[],[f13688,f1636])).
fof(f13688,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl8 | $spl21 | $spl47)),
  inference(forward_demodulation,[],[f13496,f366])).
fof(f13496,plain,(
  s(bool,f(s(bool,'F_0'))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl21 | $spl47)),
  inference(unit_resulting_resolution,[],[f2225,f1631,f135])).
fof(f13687,plain,(
  $false | ($spl3 | $spl13 | $spl21 | $spl27)),
  inference(subsumption_resolution,[],[f13686,f1636])).
fof(f13686,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl13 | $spl21)),
  inference(forward_demodulation,[],[f13497,f756])).
fof(f13497,plain,(
  s(bool,g(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl13 | $spl21)),
  inference(unit_resulting_resolution,[],[f750,f1631,f135])).
fof(f13685,plain,(
  $false | ($spl3 | $spl17 | $spl21 | $spl27)),
  inference(subsumption_resolution,[],[f13684,f1636])).
fof(f13684,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl17 | $spl21)),
  inference(forward_demodulation,[],[f13498,f831])).
fof(f13498,plain,(
  s(bool,g(s(sK0,sK5(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl17 | $spl21)),
  inference(unit_resulting_resolution,[],[f825,f1631,f135])).
fof(f13683,plain,(
  $false | ($spl3 | $spl21 | $spl27 | $spl32)),
  inference(subsumption_resolution,[],[f13682,f1636])).
fof(f13682,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl21 | $spl32)),
  inference(forward_demodulation,[],[f13499,f1987])).
fof(f13499,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl21 | $spl32)),
  inference(unit_resulting_resolution,[],[f1971,f1631,f135])).
fof(f13681,plain,(
  $false | ($spl3 | $spl21 | $spl23)),
  inference(subsumption_resolution,[],[f13502,f1632])).
fof(f13502,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | ($spl3 | $spl21)),
  inference(unit_resulting_resolution,[],[f4513,f1631,f135])).
fof(f13680,plain,(
  $false | ($spl21 | $spl27)),
  inference(subsumption_resolution,[],[f13503,f1636])).
fof(f13503,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | $spl21),
  inference(unit_resulting_resolution,[],[f78,f1631,f171])).
fof(f13679,plain,(
  $false | ($spl3 | $spl21 | $spl27 | $spl32)),
  inference(subsumption_resolution,[],[f13678,f1636])).
fof(f13678,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl21 | $spl32)),
  inference(forward_demodulation,[],[f13506,f1987])).
fof(f13506,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl21 | $spl32)),
  inference(unit_resulting_resolution,[],[f1631,f2026])).
fof(f13677,plain,(
  $false | ($spl8 | $spl21 | $spl27 | $spl50)),
  inference(subsumption_resolution,[],[f13676,f1636])).
fof(f13676,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl8 | $spl21 | $spl50)),
  inference(forward_demodulation,[],[f13509,f13635])).
fof(f13635,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))))) | ($spl8 | $spl21 | $spl50)),
  inference(forward_demodulation,[],[f13634,f13613])).
fof(f13613,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) = s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))))) | ($spl8 | $spl21 | $spl50)),
  inference(forward_demodulation,[],[f13537,f13538])).
fof(f13538,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) = s(bool,f(s(bool,f(s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))))))))) | ($spl8 | $spl21 | $spl50)),
  inference(unit_resulting_resolution,[],[f1631,f5037])).
fof(f5037,plain,(
  ( ! [X12] : (s(bool,X12) = s(bool,f(s(bool,f(s(bool,f(s(bool,X12))))))) | s(bool,'T_0') = s(bool,X12)) ) | ($spl8 | $spl50)),
  inference(subsumption_resolution,[],[f4918,f121])).
fof(f4918,plain,(
  ( ! [X12] : (s(bool,X12) = s(bool,f(s(bool,f(s(bool,f(s(bool,X12))))))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X12)) ) | ($spl8 | $spl50)),
  inference(superposition,[],[f3931,f79])).
fof(f13537,plain,(
  s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))))) = s(bool,f(s(bool,f(s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))))))))) | ($spl8 | $spl21 | $spl50)),
  inference(unit_resulting_resolution,[],[f1631,f4762])).
fof(f4762,plain,(
  ( ! [X15] : (s(bool,f(s(bool,X15))) = s(bool,f(s(bool,f(s(bool,f(s(bool,X15))))))) | s(bool,'T_0') = s(bool,X15)) ) | ($spl8 | $spl50)),
  inference(subsumption_resolution,[],[f4746,f78])).
fof(f4746,plain,(
  ( ! [X15] : (p(s(bool,'F_0')) | s(bool,f(s(bool,X15))) = s(bool,f(s(bool,f(s(bool,f(s(bool,X15))))))) | s(bool,'T_0') = s(bool,X15)) ) | ($spl8 | $spl50)),
  inference(superposition,[],[f4669,f2894])).
fof(f4669,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,X0) = s(bool,f(s(bool,f(s(bool,X0)))))) ) | $spl8),
  inference(factoring,[],[f4044])).
fof(f4044,plain,(
  ( ! [X31,X32] : (p(s(bool,X31)) | p(s(bool,X32)) | s(bool,X32) = s(bool,f(s(bool,f(s(bool,X31)))))) ) | $spl8),
  inference(subsumption_resolution,[],[f4007,f80])).
fof(f4007,plain,(
  ( ! [X31,X32] : (p(s(bool,X31)) | p(s(bool,X32)) | s(bool,X32) = s(bool,f(s(bool,f(s(bool,X31))))) | s(bool,'T_0') = s(bool,X31)) ) | $spl8),
  inference(superposition,[],[f3595,f3373])).
fof(f13634,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))))))) | ($spl8 | $spl21 | $spl50)),
  inference(forward_demodulation,[],[f13633,f13613])).
fof(f13633,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))))))))) | ($spl8 | $spl21 | $spl50)),
  inference(forward_demodulation,[],[f13535,f13613])).
fof(f13535,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))))))))))) | ($spl8 | $spl21 | $spl50)),
  inference(unit_resulting_resolution,[],[f1631,f4321])).
fof(f13509,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) = s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))))) | ($spl8 | $spl21)),
  inference(unit_resulting_resolution,[],[f1631,f3373])).
fof(f13675,plain,(
  $false | ($spl8 | $spl21 | $spl27)),
  inference(subsumption_resolution,[],[f13674,f1636])).
fof(f13674,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl8 | $spl21)),
  inference(forward_demodulation,[],[f13510,f366])).
fof(f13510,plain,(
  s(bool,f(s(bool,'F_0'))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl8 | $spl21)),
  inference(unit_resulting_resolution,[],[f1631,f3393])).
fof(f13667,plain,(
  $false | ($spl8 | $spl21 | $spl27 | $spl50)),
  inference(subsumption_resolution,[],[f13666,f1636])).
fof(f13666,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl8 | $spl21 | $spl50)),
  inference(forward_demodulation,[],[f13519,f13635])).
fof(f13519,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) = s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))))) | ($spl8 | $spl21)),
  inference(unit_resulting_resolution,[],[f1631,f1631,f3584])).
fof(f13665,plain,(
  $false | ($spl8 | $spl21 | $spl27)),
  inference(subsumption_resolution,[],[f13664,f1636])).
fof(f13664,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl8 | $spl21)),
  inference(forward_demodulation,[],[f13520,f366])).
fof(f13520,plain,(
  s(bool,f(s(bool,'F_0'))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl8 | $spl21)),
  inference(unit_resulting_resolution,[],[f121,f1631,f3584])).
fof(f13663,plain,(
  $false | ($spl8 | $spl21 | $spl27 | $spl47)),
  inference(subsumption_resolution,[],[f13662,f1636])).
fof(f13662,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl8 | $spl21 | $spl47)),
  inference(forward_demodulation,[],[f13661,f366])).
fof(f13661,plain,(
  s(bool,f(s(bool,'F_0'))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl8 | $spl21 | $spl47)),
  inference(forward_demodulation,[],[f13521,f366])).
fof(f13521,plain,(
  s(bool,f(s(bool,f(s(bool,'F_0'))))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl8 | $spl21 | $spl47)),
  inference(unit_resulting_resolution,[],[f2225,f1631,f3584])).
fof(f13660,plain,(
  $false | ($spl3 | $spl8 | $spl13 | $spl21 | $spl27)),
  inference(subsumption_resolution,[],[f13659,f1636])).
fof(f13659,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl8 | $spl13 | $spl21)),
  inference(forward_demodulation,[],[f13658,f366])).
fof(f13658,plain,(
  s(bool,f(s(bool,'F_0'))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl8 | $spl13 | $spl21)),
  inference(forward_demodulation,[],[f13522,f756])).
fof(f13522,plain,(
  s(bool,f(s(bool,g(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl8 | $spl13 | $spl21)),
  inference(unit_resulting_resolution,[],[f750,f1631,f3584])).
fof(f13657,plain,(
  $false | ($spl3 | $spl8 | $spl17 | $spl21 | $spl27)),
  inference(subsumption_resolution,[],[f13656,f1636])).
fof(f13656,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl8 | $spl17 | $spl21)),
  inference(forward_demodulation,[],[f13655,f366])).
fof(f13655,plain,(
  s(bool,f(s(bool,'F_0'))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl8 | $spl17 | $spl21)),
  inference(forward_demodulation,[],[f13523,f831])).
fof(f13523,plain,(
  s(bool,f(s(bool,g(s(sK0,sK5(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl8 | $spl17 | $spl21)),
  inference(unit_resulting_resolution,[],[f825,f1631,f3584])).
fof(f13654,plain,(
  $false | ($spl3 | $spl8 | $spl21 | $spl27 | $spl32)),
  inference(subsumption_resolution,[],[f13653,f1636])).
fof(f13653,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl8 | $spl21 | $spl32)),
  inference(forward_demodulation,[],[f13652,f366])).
fof(f13652,plain,(
  s(bool,f(s(bool,'F_0'))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl8 | $spl21 | $spl32)),
  inference(forward_demodulation,[],[f13524,f1987])).
fof(f13524,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) = s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))))) | ($spl3 | $spl8 | $spl21 | $spl32)),
  inference(unit_resulting_resolution,[],[f1971,f1631,f3584])).
fof(f13651,plain,(
  $false | ($spl8 | $spl21 | $spl27 | $spl50)),
  inference(subsumption_resolution,[],[f13650,f1636])).
fof(f13650,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl8 | $spl21 | $spl50)),
  inference(forward_demodulation,[],[f13525,f13635])).
fof(f13525,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) = s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))))) | ($spl8 | $spl21)),
  inference(unit_resulting_resolution,[],[f1631,f1631,f3584])).
fof(f13649,plain,(
  $false | ($spl8 | $spl21 | $spl27)),
  inference(subsumption_resolution,[],[f13648,f1636])).
fof(f13648,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl8 | $spl21)),
  inference(forward_demodulation,[],[f13526,f366])).
fof(f13526,plain,(
  s(bool,f(s(bool,'F_0'))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl8 | $spl21)),
  inference(unit_resulting_resolution,[],[f78,f1631,f3596])).
fof(f13640,plain,(
  $false | ($spl8 | $spl21 | $spl27 | $spl50)),
  inference(subsumption_resolution,[],[f13639,f1636])).
fof(f13639,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl8 | $spl21 | $spl50)),
  inference(forward_demodulation,[],[f13638,f366])).
fof(f13638,plain,(
  s(bool,f(s(bool,'F_0'))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl8 | $spl21 | $spl50)),
  inference(forward_demodulation,[],[f13534,f13635])).
fof(f13534,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) = s(bool,f(s(bool,f(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))))))) | ($spl8 | $spl21 | $spl50)),
  inference(unit_resulting_resolution,[],[f1631,f4310])).
fof(f13637,plain,(
  $false | ($spl8 | $spl21 | $spl27 | $spl50)),
  inference(subsumption_resolution,[],[f13636,f1636])).
fof(f13636,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl8 | $spl21 | $spl50)),
  inference(backward_demodulation,[],[f13635,f13613])).
fof(f13632,plain,(
  $false | ($spl8 | $spl21 | $spl27 | $spl50)),
  inference(subsumption_resolution,[],[f13631,f1636])).
fof(f13631,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl8 | $spl21 | $spl50)),
  inference(forward_demodulation,[],[f13630,f366])).
fof(f13630,plain,(
  s(bool,f(s(bool,'F_0'))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl8 | $spl21 | $spl50)),
  inference(forward_demodulation,[],[f13536,f366])).
fof(f13536,plain,(
  s(bool,f(s(bool,f(s(bool,'F_0'))))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl8 | $spl21 | $spl50)),
  inference(unit_resulting_resolution,[],[f78,f1631,f4508])).
fof(f12758,plain,(
  $false | ($spl0 | $spl3 | $spl22)),
  inference(subsumption_resolution,[],[f12757,f78])).
fof(f12757,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl3 | $spl22)),
  inference(forward_demodulation,[],[f12756,f1813])).
fof(f1813,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl22)),
  inference(unit_resulting_resolution,[],[f78,f1637,f171])).
fof(f1637,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl22)),
  inference(backward_demodulation,[],[f1633,f719])).
fof(f719,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl3),
  inference(unit_resulting_resolution,[],[f120,f104])).
fof(f12756,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2)))) | ($spl0 | $spl22)),
  inference(subsumption_resolution,[],[f12755,f76])).
fof(f12755,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2)))) | ($spl0 | $spl22)),
  inference(subsumption_resolution,[],[f12656,f708])).
fof(f12656,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK1)))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2)))) | $spl22),
  inference(superposition,[],[f108,f1633])).
fof(f6398,plain,(
  ( ! [X4,X2,X3] : (s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,X4) | s(bool,X4) = s(bool,sK5(X2,X3,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X2,X3,bool,bool))))) ) | $spl62),
  inference(cnf_transformation,[],[f6398_D])).
fof(f6398_D,plain,(
  ( ! [X4,X2,X3] : (s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,X4) | s(bool,X4) = s(bool,sK5(X2,X3,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X2,X3,bool,bool))))) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f6400,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | $spl65),
  inference(cnf_transformation,[],[f6400_D])).
fof(f6400_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl65),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl65])])).
fof(f6379,plain,(
  ( ! [X12,X13,X11] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X11) = s(fun(bool,bool),X12) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X11,X12,bool,bool)))) | s(bool,X13) = s(bool,sK5(X11,X12,bool,bool)) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f345,f135])).
fof(f6376,plain,(
  ( ! [X4,X2,X3] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X2,X3,bool,bool)))) | s(bool,X4) = s(bool,sK5(X2,X3,bool,bool)) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f345,f135])).
fof(f3322,plain,(
  $false | ($spl8 | $spl44 | $spl47 | $spl50)),
  inference(subsumption_resolution,[],[f2225,f3237])).
fof(f3237,plain,(
  ( ! [X4] : (s(bool,'T_0') = s(bool,X4)) ) | ($spl8 | $spl44 | $spl50)),
  inference(global_subsumption,[],[f2894,f75,f74,f76,f79,f81,f80,f114,f111,f115,f121,f135,f171,f960,f1021,f1022,f1032,f1083,f1401,f1402,f1133,f1134,f1144,f2096,f2099,f2218,f2217,f2216,f2215,f2214,f2271,f2445,f2447,f2297,f366,f2850,f2860,f2861,f2862,f2893,f2895,f2896,f2899,f1141,f1099,f1098,f1097,f1139,f1138,f1137,f1029,f1027,f1026,f1025,f1406,f674,f673,f110,f109,f108,f103,f104,f2100,f1407,f1404,f107,f106,f105,f1143,f1142,f1031,f1030,f1028,f1140,f2448,f2312,f723,f722,f623,f620,f618,f500,f499,f413,f412,f345,f343,f288,f247,f85,f84,f78,f2223,f675])).
fof(f2899,plain,(
  ( ! [X14] : (p(s(bool,X14)) | s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,X14)))))) ) | $spl8),
  inference(subsumption_resolution,[],[f2889,f80])).
fof(f2889,plain,(
  ( ! [X14] : (p(s(bool,X14)) | s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,X14))))) | s(bool,'T_0') = s(bool,X14)) ) | $spl8),
  inference(superposition,[],[f2860,f2271])).
fof(f2896,plain,(
  ( ! [X8] : (s(bool,'F_0') = s(bool,f(s(bool,X8))) | s(bool,'T_0') = s(bool,X8)) ) | $spl8),
  inference(subsumption_resolution,[],[f2882,f78])).
fof(f2882,plain,(
  ( ! [X8] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,f(s(bool,X8))) | s(bool,'T_0') = s(bool,X8)) ) | $spl8),
  inference(superposition,[],[f2860,f79])).
fof(f2862,plain,(
  ( ! [X19,X20] : (p(s(bool,X19)) | s(bool,X20) = s(bool,f(s(bool,X19))) | s(bool,'T_0') = s(bool,X20)) ) | $spl8),
  inference(subsumption_resolution,[],[f2844,f80])).
fof(f2844,plain,(
  ( ! [X19,X20] : (p(s(bool,X19)) | s(bool,X20) = s(bool,f(s(bool,X19))) | s(bool,'T_0') = s(bool,X20) | s(bool,'T_0') = s(bool,X19)) ) | $spl8),
  inference(superposition,[],[f171,f2271])).
fof(f2861,plain,(
  ( ! [X6,X5] : (p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X6) = s(bool,f(s(bool,X5)))) ) | $spl8),
  inference(subsumption_resolution,[],[f2837,f80])).
fof(f2837,plain,(
  ( ! [X6,X5] : (p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X6) = s(bool,f(s(bool,X5))) | s(bool,'T_0') = s(bool,X5)) ) | $spl8),
  inference(superposition,[],[f111,f2271])).
fof(f2860,plain,(
  ( ! [X4] : (p(s(bool,X4)) | s(bool,'F_0') = s(bool,f(s(bool,X4)))) ) | $spl8),
  inference(subsumption_resolution,[],[f2836,f80])).
fof(f2836,plain,(
  ( ! [X4] : (p(s(bool,X4)) | s(bool,'F_0') = s(bool,f(s(bool,X4))) | s(bool,'T_0') = s(bool,X4)) ) | $spl8),
  inference(superposition,[],[f85,f2271])).
fof(f2850,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') = s(bool,X1) | s(bool,X2) = s(bool,f(s(bool,X1))) | s(bool,'T_0') = s(bool,X2)) ) | $spl8),
  inference(duplicate_literal_removal,[],[f2827])).
fof(f2827,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,X2) = s(bool,f(s(bool,X1))) | s(bool,'T_0') = s(bool,X2)) ) | $spl8),
  inference(superposition,[],[f2271,f135])).
fof(f2297,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl8),
  inference(subsumption_resolution,[],[f2277,f121])).
fof(f2277,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X0) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl8),
  inference(superposition,[],[f366,f135])).
fof(f2271,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl8),
  inference(superposition,[],[f366,f79])).
fof(f3321,plain,(
  $false | ($spl0 | $spl3 | $spl8 | $spl32 | $spl44 | $spl50)),
  inference(subsumption_resolution,[],[f3284,f76])).
fof(f3284,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl3 | $spl8 | $spl32 | $spl44 | $spl50)),
  inference(backward_demodulation,[],[f3237,f2067])).
fof(f3320,plain,(
  $false | ($spl0 | $spl3 | $spl8 | $spl22 | $spl44 | $spl50)),
  inference(subsumption_resolution,[],[f3283,f76])).
fof(f3283,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl3 | $spl8 | $spl22 | $spl44 | $spl50)),
  inference(backward_demodulation,[],[f3237,f1880])).
fof(f1880,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))) | ($spl0 | $spl3 | $spl22)),
  inference(unit_resulting_resolution,[],[f1800,f1800,f108])).
fof(f1800,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK1)))) | ($spl0 | $spl3 | $spl22)),
  inference(unit_resulting_resolution,[],[f1637,f733])).
fof(f3315,plain,(
  $false | ($spl8 | $spl44 | $spl50)),
  inference(global_subsumption,[],[f75,f74,f76,f79,f81,f80,f114,f111,f115,f121,f135,f171,f960,f1021,f1022,f1032,f1083,f1401,f1402,f1133,f1134,f1144,f2096,f2099,f2218,f2217,f2216,f2215,f2214,f2271,f2445,f2447,f2297,f366,f2850,f2860,f2861,f2862,f2893,f2895,f2896,f2899,f1141,f1099,f1098,f1097,f1139,f1138,f1137,f1029,f1027,f1026,f1025,f1406,f674,f673,f110,f109,f108,f103,f104,f2100,f1407,f1404,f107,f106,f105,f1143,f1142,f1031,f1030,f1028,f1140,f2448,f2312,f723,f722,f623,f620,f618,f500,f499,f413,f412,f345,f343,f288,f247,f85,f84,f78,f2223,f675,f2894,f3237,f3308,f3314])).
fof(f3314,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | ($spl8 | $spl44 | $spl50)),
  inference(subsumption_resolution,[],[f3275,f3237])).
fof(f3275,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X0,sK3(X2,X1,X0)),s(fun(X0,bool),X1))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | ($spl8 | $spl44 | $spl50)),
  inference(backward_demodulation,[],[f3237,f104])).
fof(f3308,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl8 | $spl44 | $spl50)),
  inference(subsumption_resolution,[],[f3239,f3237])).
fof(f3239,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | ($spl8 | $spl44 | $spl50)),
  inference(backward_demodulation,[],[f3237,f84])).
fof(f3309,plain,(
  $false | ($spl8 | $spl44 | $spl50)),
  inference(global_subsumption,[],[f75,f74,f76,f79,f81,f80,f114,f111,f115,f121,f135,f171,f960,f1021,f1022,f1032,f1083,f1401,f1402,f1133,f1134,f1144,f2096,f2099,f2218,f2217,f2216,f2215,f2214,f2271,f2445,f2447,f2297,f366,f2850,f2860,f2861,f2862,f2893,f2895,f2896,f2899,f1141,f1099,f1098,f1097,f1139,f1138,f1137,f1029,f1027,f1026,f1025,f1406,f674,f673,f110,f109,f108,f103,f104,f2100,f1407,f1404,f107,f106,f105,f1143,f1142,f1031,f1030,f1028,f1140,f2448,f2312,f723,f722,f623,f620,f618,f500,f499,f413,f412,f345,f343,f288,f247,f85,f84,f78,f2223,f675,f2894,f3237,f3308])).
fof(f3307,plain,(
  $false | ($spl8 | $spl44 | $spl50)),
  inference(subsumption_resolution,[],[f3238,f76])).
fof(f3238,plain,(
  ~p(s(bool,'T_0')) | ($spl8 | $spl44 | $spl50)),
  inference(backward_demodulation,[],[f3237,f78])).
fof(f3230,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl60),
  inference(cnf_transformation,[],[f3230_D])).
fof(f3230_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f3222,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | $spl54),
  inference(cnf_transformation,[],[f3222_D])).
fof(f3222_D,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f3086,plain,(
  ( ! [X9] : (s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X9)) ) | $spl42),
  inference(duplicate_literal_removal,[],[f2990])).
fof(f2990,plain,(
  ( ! [X9] : (s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) ) | $spl42),
  inference(superposition,[],[f2220,f79])).
fof(f3078,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl42),
  inference(duplicate_literal_removal,[],[f3011])).
fof(f3011,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl42),
  inference(superposition,[],[f79,f2220])).
fof(f3225,plain,(
  ~p(s(bool,g(s(bool,'F_0')))) | $spl59),
  inference(cnf_transformation,[],[f3225_D])).
fof(f3225_D,plain,(
  ~p(s(bool,g(s(bool,'F_0')))) <=> ~$spl59),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl59])])).
fof(f3077,plain,(
  ( ! [X2] : (~p(s(bool,g(s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X2)) ) | $spl42),
  inference(duplicate_literal_removal,[],[f3013])).
fof(f3013,plain,(
  ( ! [X2] : (~p(s(bool,g(s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl42),
  inference(superposition,[],[f81,f2220])).
fof(f3231,plain,(
  ( ! [X3] : (~p(s(bool,g(s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X3)) ) | $spl42),
  inference(subsumption_resolution,[],[f3014,f79])).
fof(f3014,plain,(
  ( ! [X3] : (~p(s(bool,g(s(bool,'F_0')))) | s(bool,'F_0') != s(bool,X3) | s(bool,'T_0') = s(bool,X3)) ) | $spl42),
  inference(superposition,[],[f84,f2220])).
fof(f3229,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,g(s(bool,'F_0'))))) ) | $spl42),
  inference(condensation,[],[f3228])).
fof(f3228,plain,(
  ( ! [X10,X9] : (~p(s(bool,g(s(bool,'F_0')))) | s(bool,X9) = s(bool,X10) | s(bool,'T_0') = s(bool,X9)) ) | $spl42),
  inference(subsumption_resolution,[],[f3018,f171])).
fof(f3018,plain,(
  ( ! [X10,X9] : (~p(s(bool,g(s(bool,'F_0')))) | ~p(s(bool,X10)) | s(bool,X9) = s(bool,X10) | s(bool,'T_0') = s(bool,X9)) ) | $spl42),
  inference(superposition,[],[f114,f2220])).
fof(f3223,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X28,bool),X30),s(fun(X28,bool),X29)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X28,sK4(X29,X30,X28)),s(fun(X28,bool),X29)))) ) | $spl56),
  inference(cnf_transformation,[],[f3223_D])).
fof(f3223_D,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X28,bool),X30),s(fun(X28,bool),X29)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X28,sK4(X29,X30,X28)),s(fun(X28,bool),X29)))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f3037,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,g(s(bool,'F_0')))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X28,bool),X30),s(fun(X28,bool),X29)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X28,sK4(X29,X30,X28)),s(fun(X28,bool),X29)))) ) | $spl42),
  inference(superposition,[],[f107,f2220])).
fof(f2098,plain,(
  ( ! [X11] : (p(s(bool,X11))) ) | $spl40),
  inference(cnf_transformation,[],[f2098_D])).
fof(f2098_D,plain,(
  ( ! [X11] : (p(s(bool,X11))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f3208,plain,(
  ( ! [X7] : (p(s(bool,X7)) | s(bool,'F_0') = s(bool,g(s(bool,'F_0')))) ) | $spl42),
  inference(subsumption_resolution,[],[f3049,f80])).
fof(f3049,plain,(
  ( ! [X7] : (p(s(bool,X7)) | s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X7)) ) | $spl42),
  inference(superposition,[],[f85,f2220])).
fof(f3220,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(bool,'F_0')))) ) | $spl52),
  inference(cnf_transformation,[],[f3220_D])).
fof(f3220_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(bool,'F_0')))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f363,plain,(
  s(bool,'T_0') != s(bool,g(s(bool,'F_0'))) | $spl5),
  inference(cnf_transformation,[],[f363_D])).
fof(f363_D,plain,(
  s(bool,'T_0') != s(bool,g(s(bool,'F_0'))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f3060,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,g(s(bool,'F_0'))) | s(bool,X0) = s(bool,g(s(bool,'F_0')))) ) | $spl42),
  inference(equality_factoring,[],[f2220])).
fof(f3219,plain,(
  $false | ($spl4 | $spl42)),
  inference(global_subsumption,[],[f75,f74,f76,f79,f81,f80,f114,f111,f115,f121,f135,f171,f960,f1021,f1022,f1032,f1083,f1401,f1402,f1133,f1134,f1144,f2096,f2099,f2218,f2217,f2216,f2215,f2214,f2227,f2445,f2447,f2354,f1141,f1139,f1138,f1137,f1029,f1099,f1098,f1097,f1027,f1026,f1025,f1406,f674,f673,f110,f109,f108,f103,f104,f1143,f1142,f1031,f1030,f2100,f1407,f1404,f675,f107,f106,f105,f2378,f2362,f2357,f2356,f2355,f1028,f1140,f2448,f2377,f2312,f723,f722,f623,f620,f618,f500,f499,f413,f412,f364,f345,f343,f288,f247,f85,f84,f78,f2220,f2902,f3108,f2996,f3085,f3191,f3197,f3206,f3218])).
fof(f3218,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(bool,'T_0')))) ) | ($spl4 | $spl42)),
  inference(forward_demodulation,[],[f3217,f3085])).
fof(f3217,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(bool,'F_0')))) ) | ($spl4 | $spl42)),
  inference(subsumption_resolution,[],[f3060,f3085])).
fof(f3206,plain,(
  ( ! [X1] : (p(s(bool,X1))) ) | ($spl4 | $spl42)),
  inference(subsumption_resolution,[],[f3205,f3085])).
fof(f3205,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,g(s(bool,'F_0'))) | p(s(bool,X1))) ) | $spl42),
  inference(subsumption_resolution,[],[f3012,f80])).
fof(f3012,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,g(s(bool,'F_0'))) | p(s(bool,X1)) | s(bool,'T_0') = s(bool,X1)) ) | $spl42),
  inference(superposition,[],[f80,f2220])).
fof(f3197,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | ($spl4 | $spl42)),
  inference(subsumption_resolution,[],[f3154,f3085])).
fof(f3154,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X0,sK3(X2,X1,X0)),s(fun(X0,bool),X1))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | ($spl4 | $spl42)),
  inference(backward_demodulation,[],[f3085,f104])).
fof(f3191,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl4 | $spl42)),
  inference(subsumption_resolution,[],[f3110,f3085])).
fof(f3110,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | ($spl4 | $spl42)),
  inference(backward_demodulation,[],[f3085,f84])).
fof(f3085,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,X2)) ) | ($spl4 | $spl42)),
  inference(duplicate_literal_removal,[],[f3003])).
fof(f3003,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | ($spl4 | $spl42)),
  inference(superposition,[],[f2220,f364])).
fof(f2996,plain,(
  ( ! [X0,X1] : (s(bool,X1) = s(bool,g(s(bool,X0))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | $spl42),
  inference(superposition,[],[f2220,f79])).
fof(f3108,plain,(
  ( ! [X12] : (s(bool,X12) = s(bool,g(s(bool,X12))) | s(bool,'T_0') = s(bool,X12)) ) | $spl42),
  inference(subsumption_resolution,[],[f2991,f121])).
fof(f2991,plain,(
  ( ! [X12] : (s(bool,X12) = s(bool,g(s(bool,X12))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X12)) ) | $spl42),
  inference(superposition,[],[f2220,f79])).
fof(f2902,plain,(
  ( ! [X19,X20] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X19) = s(fun(bool,bool),X20) | s(bool,'T_0') = s(bool,sK5(X19,X20,bool,bool))) ) | $spl4),
  inference(subsumption_resolution,[],[f2353,f121])).
fof(f2353,plain,(
  ( ! [X19,X20] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X19) = s(fun(bool,bool),X20) | s(bool,'T_0') = s(bool,sK5(X19,X20,bool,bool)) | s(bool,'T_0') = s(bool,'F_0')) ) | $spl4),
  inference(superposition,[],[f343,f2227])).
fof(f2377,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,g(s(bool,X3))) | s(bool,'T_0') = s(bool,X3)) ) | $spl4),
  inference(subsumption_resolution,[],[f2345,f76])).
fof(f2345,plain,(
  ( ! [X3] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,g(s(bool,X3))) | s(bool,'T_0') = s(bool,X3)) ) | $spl4),
  inference(superposition,[],[f84,f2227])).
fof(f2355,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X24),s(fun(bool,bool),X25)))))) | p(s(bool,'const.bool.IN_2'(s(bool,g(s(bool,X23))),s(fun(bool,bool),X24)))) | p(s(bool,'const.bool.IN_2'(s(bool,g(s(bool,X23))),s(fun(bool,bool),X25)))) | s(bool,'T_0') = s(bool,X23)) ) | $spl4),
  inference(superposition,[],[f108,f2227])).
fof(f2356,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X27),s(fun(bool,bool),X28)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,g(s(bool,X26))),s(fun(bool,bool),X28)))) | s(bool,'T_0') = s(bool,X26)) ) | $spl4),
  inference(superposition,[],[f110,f2227])).
fof(f2357,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X30),s(fun(bool,bool),X31)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,g(s(bool,X29))),s(fun(bool,bool),X30)))) | s(bool,'T_0') = s(bool,X29)) ) | $spl4),
  inference(superposition,[],[f109,f2227])).
fof(f2362,plain,(
  ( ! [X1] : (p(s(bool,g(s(bool,X1)))) | s(bool,'T_0') = s(bool,X1)) ) | $spl4),
  inference(trivial_inequality_removal,[],[f2343])).
fof(f2343,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,g(s(bool,X1)))) | s(bool,'T_0') = s(bool,X1)) ) | $spl4),
  inference(superposition,[],[f80,f2227])).
fof(f2378,plain,(
  ( ! [X10,X9] : (~p(s(bool,X10)) | s(bool,X10) = s(bool,g(s(bool,X9))) | s(bool,'T_0') = s(bool,X9)) ) | $spl4),
  inference(subsumption_resolution,[],[f2349,f76])).
fof(f2349,plain,(
  ( ! [X10,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,X10)) | s(bool,X10) = s(bool,g(s(bool,X9))) | s(bool,'T_0') = s(bool,X9)) ) | $spl4),
  inference(superposition,[],[f114,f2227])).
fof(f2354,plain,(
  ( ! [X21,X22] : (s(bool,'T_0') != s(bool,f(s(bool,sK5(X21,X22,bool,bool)))) | s(fun(bool,bool),X21) = s(fun(bool,bool),X22) | s(bool,'T_0') = s(bool,sK5(X21,X22,bool,bool))) ) | $spl4),
  inference(superposition,[],[f115,f2227])).
fof(f2227,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,g(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl4),
  inference(superposition,[],[f364,f79])).
fof(f3207,plain,(
  $false | ($spl4 | $spl42)),
  inference(global_subsumption,[],[f75,f74,f76,f79,f81,f80,f114,f111,f115,f121,f135,f171,f960,f1021,f1022,f1032,f1083,f1401,f1402,f1133,f1134,f1144,f2096,f2099,f2218,f2217,f2216,f2215,f2214,f2227,f2445,f2447,f2354,f1141,f1139,f1138,f1137,f1029,f1099,f1098,f1097,f1027,f1026,f1025,f1406,f674,f673,f110,f109,f108,f103,f104,f1143,f1142,f1031,f1030,f2100,f1407,f1404,f675,f107,f106,f105,f2378,f2362,f2357,f2356,f2355,f1028,f1140,f2448,f2377,f2312,f723,f722,f623,f620,f618,f500,f499,f413,f412,f364,f345,f343,f288,f247,f85,f84,f78,f2220,f2902,f3108,f2996,f3085,f3191,f3197,f3206])).
fof(f3204,plain,(
  $false | ($spl0 | $spl3 | $spl4 | $spl32 | $spl42)),
  inference(subsumption_resolution,[],[f3163,f76])).
fof(f3163,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl3 | $spl4 | $spl32 | $spl42)),
  inference(backward_demodulation,[],[f3085,f2067])).
fof(f3203,plain,(
  $false | ($spl0 | $spl3 | $spl4 | $spl22 | $spl42)),
  inference(subsumption_resolution,[],[f3162,f76])).
fof(f3162,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl3 | $spl4 | $spl22 | $spl42)),
  inference(backward_demodulation,[],[f3085,f1880])).
fof(f3198,plain,(
  $false | ($spl4 | $spl42)),
  inference(global_subsumption,[],[f75,f74,f76,f79,f81,f80,f114,f111,f115,f121,f135,f171,f960,f1021,f1022,f1032,f1083,f1401,f1402,f1133,f1134,f1144,f2096,f2099,f2218,f2217,f2216,f2215,f2214,f2227,f2445,f2447,f2354,f1141,f1139,f1138,f1137,f1029,f1099,f1098,f1097,f1027,f1026,f1025,f1406,f674,f673,f110,f109,f108,f103,f104,f1143,f1142,f1031,f1030,f2100,f1407,f1404,f675,f107,f106,f105,f2378,f2362,f2357,f2356,f2355,f1028,f1140,f2448,f2377,f2312,f723,f722,f623,f620,f618,f500,f499,f413,f412,f364,f345,f343,f288,f247,f85,f84,f78,f2220,f2902,f3108,f2996,f3085,f3191,f3197])).
fof(f3192,plain,(
  $false | ($spl4 | $spl42)),
  inference(global_subsumption,[],[f75,f74,f76,f79,f81,f80,f114,f111,f115,f121,f135,f171,f960,f1021,f1022,f1032,f1083,f1401,f1402,f1133,f1134,f1144,f2096,f2099,f2218,f2217,f2216,f2215,f2214,f2227,f2445,f2447,f2354,f1141,f1139,f1138,f1137,f1029,f1099,f1098,f1097,f1027,f1026,f1025,f1406,f674,f673,f110,f109,f108,f103,f104,f1143,f1142,f1031,f1030,f2100,f1407,f1404,f675,f107,f106,f105,f2378,f2362,f2357,f2356,f2355,f1028,f1140,f2448,f2377,f2312,f723,f722,f623,f620,f618,f500,f499,f413,f412,f364,f345,f343,f288,f247,f85,f84,f78,f2220,f2902,f3108,f2996,f3085,f3191])).
fof(f3190,plain,(
  $false | ($spl4 | $spl42)),
  inference(subsumption_resolution,[],[f3109,f76])).
fof(f3109,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl42)),
  inference(backward_demodulation,[],[f3085,f78])).
fof(f3107,plain,(
  $false | ($spl4 | $spl42)),
  inference(subsumption_resolution,[],[f3106,f121])).
fof(f3106,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl4 | $spl42)),
  inference(condensation,[],[f3105])).
fof(f3105,plain,(
  ( ! [X9] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X9)) ) | ($spl4 | $spl42)),
  inference(forward_demodulation,[],[f3086,f364])).
fof(f3104,plain,(
  $false | ($spl4 | $spl42)),
  inference(subsumption_resolution,[],[f3103,f121])).
fof(f3103,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl4 | $spl42)),
  inference(forward_demodulation,[],[f2981,f364])).
fof(f2981,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | $spl42),
  inference(unit_resulting_resolution,[],[f121,f2220])).
fof(f3102,plain,(
  $false | ($spl3 | $spl4 | $spl13 | $spl42)),
  inference(subsumption_resolution,[],[f3101,f121])).
fof(f3101,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl3 | $spl4 | $spl13 | $spl42)),
  inference(forward_demodulation,[],[f3100,f364])).
fof(f3100,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | ($spl3 | $spl13 | $spl42)),
  inference(forward_demodulation,[],[f2982,f756])).
fof(f2982,plain,(
  s(bool,g(s(bool,'F_0'))) = s(bool,g(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) | ($spl3 | $spl13 | $spl42)),
  inference(unit_resulting_resolution,[],[f750,f2220])).
fof(f3099,plain,(
  $false | ($spl3 | $spl4 | $spl17 | $spl42)),
  inference(subsumption_resolution,[],[f3098,f121])).
fof(f3098,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl3 | $spl4 | $spl17 | $spl42)),
  inference(forward_demodulation,[],[f3097,f364])).
fof(f3097,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | ($spl3 | $spl17 | $spl42)),
  inference(forward_demodulation,[],[f2983,f831])).
fof(f2983,plain,(
  s(bool,g(s(bool,'F_0'))) = s(bool,g(s(sK0,sK5(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | ($spl3 | $spl17 | $spl42)),
  inference(unit_resulting_resolution,[],[f825,f2220])).
fof(f3096,plain,(
  $false | ($spl3 | $spl4 | $spl22 | $spl42)),
  inference(subsumption_resolution,[],[f3095,f121])).
fof(f3095,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl3 | $spl4 | $spl22 | $spl42)),
  inference(forward_demodulation,[],[f3094,f364])).
fof(f3094,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | ($spl3 | $spl22 | $spl42)),
  inference(forward_demodulation,[],[f2984,f1813])).
fof(f2984,plain,(
  s(bool,g(s(bool,'F_0'))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl22 | $spl42)),
  inference(unit_resulting_resolution,[],[f1637,f2220])).
fof(f3093,plain,(
  $false | ($spl3 | $spl4 | $spl32 | $spl42)),
  inference(subsumption_resolution,[],[f3092,f121])).
fof(f3092,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl3 | $spl4 | $spl32 | $spl42)),
  inference(forward_demodulation,[],[f3091,f364])).
fof(f3091,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | ($spl3 | $spl32 | $spl42)),
  inference(forward_demodulation,[],[f2985,f1987])).
fof(f2985,plain,(
  s(bool,g(s(bool,'F_0'))) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl32 | $spl42)),
  inference(unit_resulting_resolution,[],[f1971,f2220])).
fof(f2901,plain,(
  $false | ($spl9 | $spl47)),
  inference(global_subsumption,[],[f75,f74,f76,f79,f81,f80,f114,f111,f115,f121,f135,f171,f380,f379,f381,f391,f487,f495,f496,f486,f502,f605,f725,f960,f1021,f1022,f1032,f1083,f1401,f1402,f1403,f1133,f1134,f1144,f1579,f1543,f1544,f1545,f1695,f2096,f2099,f2219,f2221,f2218,f2217,f2216,f2215,f2214,f2445,f2447,f1141,f1139,f1138,f1137,f1029,f1099,f1098,f1097,f1027,f1026,f1025,f1406,f674,f673,f110,f109,f108,f103,f104,f1143,f1142,f1031,f1030,f2100,f1407,f1404,f675,f107,f106,f105,f1028,f1140,f2448,f2312,f2225,f723,f722,f623,f620,f618,f500,f499,f413,f412,f345,f343,f288,f247,f85,f84,f78,f367])).
fof(f1695,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl9),
  inference(condensation,[],[f1694])).
fof(f1694,plain,(
  ( ! [X12,X13] : (s(bool,X13) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X12)) ) | $spl9),
  inference(duplicate_literal_removal,[],[f1666])).
fof(f1666,plain,(
  ( ! [X12,X13] : (s(bool,X13) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X12) | s(bool,'T_0') = s(bool,X12)) ) | $spl9),
  inference(superposition,[],[f1579,f79])).
fof(f1545,plain,(
  ( ! [X68,X66,X67,X65] : (p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,X66))),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X67),s(fun(bool,bool),X68)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,X65))),s(fun(bool,bool),X67)))) | s(bool,'T_0') = s(bool,X66) | s(bool,'T_0') = s(bool,X65)) ) | $spl9),
  inference(superposition,[],[f109,f605])).
fof(f1544,plain,(
  ( ! [X61,X64,X62,X63] : (p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,X62))),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X63),s(fun(bool,bool),X64)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,X61))),s(fun(bool,bool),X64)))) | s(bool,'T_0') = s(bool,X62) | s(bool,'T_0') = s(bool,X61)) ) | $spl9),
  inference(superposition,[],[f110,f605])).
fof(f1543,plain,(
  ( ! [X59,X57,X60,X58] : (~p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,X58))),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X59),s(fun(bool,bool),X60)))))) | p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,X57))),s(fun(bool,bool),X59)))) | p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,X57))),s(fun(bool,bool),X60)))) | s(bool,'T_0') = s(bool,X58) | s(bool,'T_0') = s(bool,X57)) ) | $spl9),
  inference(superposition,[],[f108,f605])).
fof(f1579,plain,(
  ( ! [X8,X7] : (s(bool,X7) != s(bool,f(s(bool,X8))) | s(bool,'T_0') = s(bool,X7) | s(bool,'T_0') = s(bool,X8)) ) | $spl9),
  inference(duplicate_literal_removal,[],[f1524])).
fof(f1524,plain,(
  ( ! [X8,X7] : (s(bool,X7) != s(bool,f(s(bool,X8))) | s(bool,'T_0') = s(bool,X7) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X7)) ) | $spl9),
  inference(superposition,[],[f381,f605])).
fof(f1403,plain,(
  ( ! [X30,X31,X32] : (s(bool,X30) = s(bool,f(s(bool,X31))) | s(bool,'T_0') = s(bool,X31) | s(bool,X30) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) ) | $spl9),
  inference(subsumption_resolution,[],[f1113,f76])).
fof(f1113,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,'T_0')) | s(bool,X30) = s(bool,f(s(bool,X31))) | s(bool,'T_0') = s(bool,X31) | s(bool,X30) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) ) | $spl9),
  inference(superposition,[],[f496,f135])).
fof(f725,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X9),s(fun(bool,bool),X10)))))) | p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,X8))),s(fun(bool,bool),X9)))) | p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,X8))),s(fun(bool,bool),X10)))) | s(bool,'T_0') = s(bool,X8)) ) | $spl9),
  inference(superposition,[],[f108,f391])).
fof(f605,plain,(
  ( ! [X8,X7] : (s(bool,f(s(bool,X7))) = s(bool,f(s(bool,X8))) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X7)) ) | $spl9),
  inference(subsumption_resolution,[],[f595,f76])).
fof(f595,plain,(
  ( ! [X8,X7] : (~p(s(bool,'T_0')) | s(bool,f(s(bool,X7))) = s(bool,f(s(bool,X8))) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X7)) ) | $spl9),
  inference(superposition,[],[f496,f391])).
fof(f502,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X14),s(fun(bool,bool),X15)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,X13))),s(fun(bool,bool),X15)))) | s(bool,'T_0') = s(bool,X13)) ) | $spl9),
  inference(superposition,[],[f110,f391])).
fof(f486,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.UNION_2'(s(fun(bool,bool),X19),s(fun(bool,bool),X20)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,X18))),s(fun(bool,bool),X19)))) | s(bool,'T_0') = s(bool,X18)) ) | $spl9),
  inference(superposition,[],[f109,f391])).
fof(f496,plain,(
  ( ! [X10,X11] : (~p(s(bool,X11)) | s(bool,X11) = s(bool,f(s(bool,X10))) | s(bool,'T_0') = s(bool,X10)) ) | $spl9),
  inference(subsumption_resolution,[],[f482,f76])).
fof(f482,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | ~p(s(bool,X11)) | s(bool,X11) = s(bool,f(s(bool,X10))) | s(bool,'T_0') = s(bool,X10)) ) | $spl9),
  inference(superposition,[],[f114,f391])).
fof(f495,plain,(
  ( ! [X4] : (s(bool,'F_0') != s(bool,f(s(bool,X4))) | s(bool,'T_0') = s(bool,X4)) ) | $spl9),
  inference(subsumption_resolution,[],[f478,f76])).
fof(f478,plain,(
  ( ! [X4] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,f(s(bool,X4))) | s(bool,'T_0') = s(bool,X4)) ) | $spl9),
  inference(superposition,[],[f84,f391])).
fof(f487,plain,(
  ( ! [X2] : (p(s(bool,f(s(bool,X2)))) | s(bool,'T_0') = s(bool,X2)) ) | $spl9),
  inference(trivial_inequality_removal,[],[f476])).
fof(f476,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,f(s(bool,X2)))) | s(bool,'T_0') = s(bool,X2)) ) | $spl9),
  inference(superposition,[],[f80,f391])).
fof(f391,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl9),
  inference(superposition,[],[f379,f79])).
fof(f381,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl9),
  inference(superposition,[],[f367,f79])).
fof(f379,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | $spl9),
  inference(unit_resulting_resolution,[],[f78,f367,f171])).
fof(f380,plain,(
  p(s(bool,f(s(bool,'F_0')))) | $spl9),
  inference(unit_resulting_resolution,[],[f78,f367,f111])).
fof(f2900,plain,(
  $false | ($spl9 | $spl42 | $spl47)),
  inference(global_subsumption,[],[f2220,f75,f74,f76,f79,f81,f80,f114,f111,f115,f121,f135,f171,f367,f380,f379,f381,f391,f487,f495,f496,f486,f502,f605,f725,f960,f1021,f1022,f1032,f1083,f1401,f1402,f1403,f1133,f1134,f1144,f1579,f1543,f1544,f1545,f1695,f2096,f2099,f2219,f2221,f2218,f2217,f2216,f2215,f2214,f2445,f2447,f1141,f1139,f1138,f1137,f1029,f1099,f1098,f1097,f1027,f1026,f1025,f1406,f674,f673,f110,f109,f108,f103,f104,f1143,f1142,f1031,f1030,f2100,f1407,f1404,f675,f107,f106,f105,f1028,f1140,f2448,f2312,f2225,f723,f722,f623,f620,f618,f500,f499,f413,f412,f345,f343,f288,f247,f85,f84,f78])).
fof(f2894,plain,(
  ( ! [X4] : (s(bool,'F_0') = s(bool,f(s(bool,X4))) | s(bool,'T_0') = s(bool,X4)) ) | $spl50),
  inference(cnf_transformation,[],[f2894_D])).
fof(f2894_D,plain,(
  ( ! [X4] : (s(bool,'F_0') = s(bool,f(s(bool,X4))) | s(bool,'T_0') = s(bool,X4)) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f2895,plain,(
  ( ! [X6,X7] : (p(s(bool,X7)) | s(bool,'F_0') = s(bool,f(s(bool,X6))) | s(bool,'T_0') = s(bool,X6)) ) | $spl8),
  inference(subsumption_resolution,[],[f2881,f80])).
fof(f2881,plain,(
  ( ! [X6,X7] : (p(s(bool,X7)) | s(bool,'F_0') = s(bool,f(s(bool,X6))) | s(bool,'T_0') = s(bool,X7) | s(bool,'T_0') = s(bool,X6)) ) | $spl8),
  inference(superposition,[],[f2860,f135])).
fof(f2893,plain,(
  ( ! [X4,X5] : (p(s(bool,X5)) | s(bool,'F_0') = s(bool,f(s(bool,X4))) | s(bool,'T_0') = s(bool,X4)) ) | $spl8),
  inference(subsumption_resolution,[],[f2880,f80])).
fof(f2880,plain,(
  ( ! [X4,X5] : (p(s(bool,X5)) | s(bool,'F_0') = s(bool,f(s(bool,X4))) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,X5)) ) | $spl8),
  inference(superposition,[],[f2860,f135])).
fof(f2641,plain,(
  $false | $spl40),
  inference(subsumption_resolution,[],[f2452,f76])).
fof(f2452,plain,(
  ~p(s(bool,'T_0')) | $spl40),
  inference(unit_resulting_resolution,[],[f121,f2098,f114])).
fof(f2640,plain,(
  $false | ($spl3 | $spl13 | $spl40)),
  inference(subsumption_resolution,[],[f2454,f76])).
fof(f2454,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl13 | $spl40)),
  inference(unit_resulting_resolution,[],[f750,f2098,f114])).
fof(f2639,plain,(
  $false | ($spl3 | $spl17 | $spl40)),
  inference(subsumption_resolution,[],[f2456,f76])).
fof(f2456,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl17 | $spl40)),
  inference(unit_resulting_resolution,[],[f825,f2098,f114])).
fof(f2638,plain,(
  $false | ($spl3 | $spl22 | $spl40)),
  inference(subsumption_resolution,[],[f2458,f76])).
fof(f2458,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl22 | $spl40)),
  inference(unit_resulting_resolution,[],[f1637,f2098,f114])).
fof(f2637,plain,(
  $false | ($spl3 | $spl32 | $spl40)),
  inference(subsumption_resolution,[],[f2460,f76])).
fof(f2460,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl32 | $spl40)),
  inference(unit_resulting_resolution,[],[f1971,f2098,f114])).
fof(f2636,plain,(
  $false | $spl40),
  inference(subsumption_resolution,[],[f2462,f2098])).
fof(f2462,plain,(
  ~p(s(bool,'F_0')) | $spl40),
  inference(unit_resulting_resolution,[],[f121,f2098,f114])).
fof(f2635,plain,(
  $false | ($spl3 | $spl13 | $spl40)),
  inference(subsumption_resolution,[],[f2464,f2098])).
fof(f2464,plain,(
  ~p(s(bool,g(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0))))) | ($spl3 | $spl13 | $spl40)),
  inference(unit_resulting_resolution,[],[f750,f2098,f114])).
fof(f2634,plain,(
  $false | ($spl3 | $spl17 | $spl40)),
  inference(subsumption_resolution,[],[f2466,f2098])).
fof(f2466,plain,(
  ~p(s(bool,g(s(sK0,sK5(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0))))) | ($spl3 | $spl17 | $spl40)),
  inference(unit_resulting_resolution,[],[f825,f2098,f114])).
fof(f2633,plain,(
  $false | ($spl3 | $spl22 | $spl40)),
  inference(subsumption_resolution,[],[f2468,f2098])).
fof(f2468,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2)))) | ($spl3 | $spl22 | $spl40)),
  inference(unit_resulting_resolution,[],[f1637,f2098,f114])).
fof(f2632,plain,(
  $false | ($spl3 | $spl32 | $spl40)),
  inference(subsumption_resolution,[],[f2470,f2098])).
fof(f2470,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) | ($spl3 | $spl32 | $spl40)),
  inference(unit_resulting_resolution,[],[f1971,f2098,f114])).
fof(f2631,plain,(
  $false | $spl40),
  inference(subsumption_resolution,[],[f2474,f76])).
fof(f2474,plain,(
  ~p(s(bool,'T_0')) | $spl40),
  inference(unit_resulting_resolution,[],[f121,f2098,f114])).
fof(f2630,plain,(
  $false | ($spl3 | $spl13 | $spl40)),
  inference(subsumption_resolution,[],[f2476,f76])).
fof(f2476,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl13 | $spl40)),
  inference(unit_resulting_resolution,[],[f750,f2098,f114])).
fof(f2629,plain,(
  $false | ($spl3 | $spl17 | $spl40)),
  inference(subsumption_resolution,[],[f2478,f76])).
fof(f2478,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl17 | $spl40)),
  inference(unit_resulting_resolution,[],[f825,f2098,f114])).
fof(f2628,plain,(
  $false | ($spl3 | $spl22 | $spl40)),
  inference(subsumption_resolution,[],[f2480,f76])).
fof(f2480,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl22 | $spl40)),
  inference(unit_resulting_resolution,[],[f1637,f2098,f114])).
fof(f2627,plain,(
  $false | ($spl3 | $spl32 | $spl40)),
  inference(subsumption_resolution,[],[f2482,f76])).
fof(f2482,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl32 | $spl40)),
  inference(unit_resulting_resolution,[],[f1971,f2098,f114])).
fof(f2626,plain,(
  $false | $spl40),
  inference(subsumption_resolution,[],[f2484,f2098])).
fof(f2484,plain,(
  ~p(s(bool,'F_0')) | $spl40),
  inference(unit_resulting_resolution,[],[f121,f2098,f114])).
fof(f2625,plain,(
  $false | ($spl3 | $spl13 | $spl40)),
  inference(subsumption_resolution,[],[f2486,f2098])).
fof(f2486,plain,(
  ~p(s(bool,g(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0))))) | ($spl3 | $spl13 | $spl40)),
  inference(unit_resulting_resolution,[],[f750,f2098,f114])).
fof(f2624,plain,(
  $false | ($spl3 | $spl17 | $spl40)),
  inference(subsumption_resolution,[],[f2488,f2098])).
fof(f2488,plain,(
  ~p(s(bool,g(s(sK0,sK5(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0))))) | ($spl3 | $spl17 | $spl40)),
  inference(unit_resulting_resolution,[],[f825,f2098,f114])).
fof(f2623,plain,(
  $false | ($spl3 | $spl22 | $spl40)),
  inference(subsumption_resolution,[],[f2490,f2098])).
fof(f2490,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2)))) | ($spl3 | $spl22 | $spl40)),
  inference(unit_resulting_resolution,[],[f1637,f2098,f114])).
fof(f2622,plain,(
  $false | ($spl3 | $spl32 | $spl40)),
  inference(subsumption_resolution,[],[f2492,f2098])).
fof(f2492,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) | ($spl3 | $spl32 | $spl40)),
  inference(unit_resulting_resolution,[],[f1971,f2098,f114])).
fof(f2621,plain,(
  $false | ($spl0 | $spl3 | $spl32 | $spl40)),
  inference(subsumption_resolution,[],[f2599,f76])).
fof(f2599,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl3 | $spl32 | $spl40)),
  inference(backward_demodulation,[],[f2496,f2067])).
fof(f2496,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl40),
  inference(unit_resulting_resolution,[],[f2098,f81])).
fof(f2620,plain,(
  $false | ($spl0 | $spl3 | $spl22 | $spl40)),
  inference(subsumption_resolution,[],[f2598,f76])).
fof(f2598,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl3 | $spl22 | $spl40)),
  inference(backward_demodulation,[],[f2496,f1880])).
fof(f2618,plain,(
  $false | ($spl40 | $spl47)),
  inference(subsumption_resolution,[],[f2556,f2496])).
fof(f2556,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl40 | $spl47)),
  inference(backward_demodulation,[],[f2496,f2225])).
fof(f2617,plain,(
  $false | $spl40),
  inference(subsumption_resolution,[],[f2616,f2098])).
fof(f2616,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl40),
  inference(subsumption_resolution,[],[f2535,f2496])).
fof(f2535,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl40),
  inference(backward_demodulation,[],[f2496,f84])).
fof(f2615,plain,(
  $false | $spl40),
  inference(subsumption_resolution,[],[f2534,f76])).
fof(f2534,plain,(
  ~p(s(bool,'T_0')) | $spl40),
  inference(backward_demodulation,[],[f2496,f78])).
fof(f2533,plain,(
  $false | ($spl0 | $spl3 | $spl32 | $spl40)),
  inference(subsumption_resolution,[],[f2531,f121])).
fof(f2531,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl0 | $spl3 | $spl32 | $spl40)),
  inference(backward_demodulation,[],[f2502,f1987])).
fof(f2502,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK2)))) ) | ($spl0 | $spl40)),
  inference(unit_resulting_resolution,[],[f2098,f733])).
fof(f2532,plain,(
  $false | ($spl0 | $spl3 | $spl22 | $spl40)),
  inference(subsumption_resolution,[],[f2530,f121])).
fof(f2530,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl0 | $spl3 | $spl22 | $spl40)),
  inference(backward_demodulation,[],[f2502,f1813])).
fof(f2503,plain,(
  $false | ($spl0 | $spl3 | $spl22 | $spl40)),
  inference(unit_resulting_resolution,[],[f1637,f2098,f733])).
fof(f2497,plain,(
  $false | $spl40),
  inference(unit_resulting_resolution,[],[f121,f2098,f81])).
fof(f2495,plain,(
  $false | ($spl3 | $spl13 | $spl40)),
  inference(unit_resulting_resolution,[],[f756,f2098,f84])).
fof(f2493,plain,(
  $false | ($spl3 | $spl32 | $spl40)),
  inference(unit_resulting_resolution,[],[f2098,f1971,f2098,f114])).
fof(f2491,plain,(
  $false | ($spl3 | $spl22 | $spl40)),
  inference(unit_resulting_resolution,[],[f2098,f1637,f2098,f114])).
fof(f2489,plain,(
  $false | ($spl3 | $spl17 | $spl40)),
  inference(unit_resulting_resolution,[],[f2098,f825,f2098,f114])).
fof(f2487,plain,(
  $false | ($spl3 | $spl13 | $spl40)),
  inference(unit_resulting_resolution,[],[f2098,f750,f2098,f114])).
fof(f2485,plain,(
  $false | $spl40),
  inference(unit_resulting_resolution,[],[f2098,f121,f2098,f114])).
fof(f2483,plain,(
  $false | ($spl3 | $spl32 | $spl40)),
  inference(unit_resulting_resolution,[],[f76,f1971,f2098,f114])).
fof(f2481,plain,(
  $false | ($spl3 | $spl22 | $spl40)),
  inference(unit_resulting_resolution,[],[f76,f1637,f2098,f114])).
fof(f2479,plain,(
  $false | ($spl3 | $spl17 | $spl40)),
  inference(unit_resulting_resolution,[],[f76,f825,f2098,f114])).
fof(f2477,plain,(
  $false | ($spl3 | $spl13 | $spl40)),
  inference(unit_resulting_resolution,[],[f76,f750,f2098,f114])).
fof(f2475,plain,(
  $false | $spl40),
  inference(unit_resulting_resolution,[],[f76,f121,f2098,f114])).
fof(f2471,plain,(
  $false | ($spl3 | $spl32 | $spl40)),
  inference(unit_resulting_resolution,[],[f2098,f1971,f2098,f114])).
fof(f2469,plain,(
  $false | ($spl3 | $spl22 | $spl40)),
  inference(unit_resulting_resolution,[],[f2098,f1637,f2098,f114])).
fof(f2467,plain,(
  $false | ($spl3 | $spl17 | $spl40)),
  inference(unit_resulting_resolution,[],[f2098,f825,f2098,f114])).
fof(f2465,plain,(
  $false | ($spl3 | $spl13 | $spl40)),
  inference(unit_resulting_resolution,[],[f2098,f750,f2098,f114])).
fof(f2463,plain,(
  $false | $spl40),
  inference(unit_resulting_resolution,[],[f2098,f121,f2098,f114])).
fof(f2461,plain,(
  $false | ($spl3 | $spl32 | $spl40)),
  inference(unit_resulting_resolution,[],[f76,f1971,f2098,f114])).
fof(f2459,plain,(
  $false | ($spl3 | $spl22 | $spl40)),
  inference(unit_resulting_resolution,[],[f76,f1637,f2098,f114])).
fof(f2457,plain,(
  $false | ($spl3 | $spl17 | $spl40)),
  inference(unit_resulting_resolution,[],[f76,f825,f2098,f114])).
fof(f2455,plain,(
  $false | ($spl3 | $spl13 | $spl40)),
  inference(unit_resulting_resolution,[],[f76,f750,f2098,f114])).
fof(f2453,plain,(
  $false | $spl40),
  inference(unit_resulting_resolution,[],[f76,f121,f2098,f114])).
fof(f2449,plain,(
  $false | $spl40),
  inference(unit_resulting_resolution,[],[f78,f2098])).
fof(f2446,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X7)))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(bool,bool),X6),s(fun(bool,bool),X7))) | s(bool,'T_0') = s(bool,sK4(X7,X6,bool))) ) | $spl48),
  inference(cnf_transformation,[],[f2446_D])).
fof(f2446_D,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X7)))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(bool,bool),X6),s(fun(bool,bool),X7))) | s(bool,'T_0') = s(bool,sK4(X7,X6,bool))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f2447,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X11)) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X10)))) | s(bool,'T_0') = s(bool,sK4(X10,X9,bool)) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(bool,bool),X9),s(fun(bool,bool),X10)))) )),
  inference(subsumption_resolution,[],[f2443,f80])).
fof(f2443,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X11)) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X10)))) | s(bool,'T_0') = s(bool,sK4(X10,X9,bool)) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(bool,bool),X9),s(fun(bool,bool),X10)))) )),
  inference(superposition,[],[f288,f135])).
fof(f2445,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X8)) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X7)))) | s(bool,'T_0') = s(bool,sK4(X7,X6,bool)) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(bool,bool),X6),s(fun(bool,bool),X7)))) )),
  inference(subsumption_resolution,[],[f2442,f80])).
fof(f2442,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X8)) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X7)))) | s(bool,'T_0') = s(bool,sK4(X7,X6,bool)) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(bool,bool),X6),s(fun(bool,bool),X7))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f288,f135])).
fof(f2220,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl42),
  inference(cnf_transformation,[],[f2220_D])).
fof(f2220_D,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f365,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK5(X0,X1,bool,bool))) ) | $spl6),
  inference(cnf_transformation,[],[f365_D])).
fof(f365_D,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK5(X0,X1,bool,bool))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f2225,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | $spl47),
  inference(cnf_transformation,[],[f2225_D])).
fof(f2225_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl47),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl47])])).
fof(f2214,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK5(X0,X1,bool,bool)) | s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f343,f135])).
fof(f2215,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,sK5(X3,X4,bool,bool)) | s(bool,X5) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f343,f135])).
fof(f364,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl4),
  inference(cnf_transformation,[],[f364_D])).
fof(f364_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f2223,plain,(
  ( ! [X9] : (s(bool,X9) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X9)) ) | $spl44),
  inference(cnf_transformation,[],[f2223_D])).
fof(f2223_D,plain,(
  ( ! [X9] : (s(bool,X9) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X9)) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f2216,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X7) = s(fun(bool,bool),X8) | s(bool,'T_0') = s(bool,sK5(X7,X8,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f343,f135])).
fof(f2217,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X10) = s(fun(bool,bool),X11) | s(bool,'T_0') = s(bool,sK5(X10,X11,bool,bool)) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f343,f135])).
fof(f367,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl9),
  inference(cnf_transformation,[],[f367_D])).
fof(f367_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f2218,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X12) = s(fun(bool,bool),X13) | s(bool,'T_0') = s(bool,sK5(X12,X13,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f343,f79])).
fof(f2221,plain,(
  ( ! [X4,X5,X3] : (s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,sK5(X3,X4,bool,bool)) | s(bool,X5) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) ) | $spl9),
  inference(subsumption_resolution,[],[f2215,f379])).
fof(f2219,plain,(
  ( ! [X2,X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK5(X0,X1,bool,bool)) | s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl9),
  inference(subsumption_resolution,[],[f2214,f379])).
fof(f2097,plain,(
  ( ! [X10,X8,X9] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X8,sK4(X10,X9,X8)),s(fun(X8,bool),X9))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(X8,bool),X9),s(fun(X8,bool),X10)))) ) | $spl38),
  inference(cnf_transformation,[],[f2097_D])).
fof(f2097_D,plain,(
  ( ! [X10,X8,X9] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X8,sK4(X10,X9,X8)),s(fun(X8,bool),X9))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(X8,bool),X9),s(fun(X8,bool),X10)))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f2099,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X15)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X12,sK4(X14,X13,X12)),s(fun(X12,bool),X13))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(X12,bool),X13),s(fun(X12,bool),X14)))) )),
  inference(subsumption_resolution,[],[f2094,f80])).
fof(f2094,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X15)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X12,sK4(X14,X13,X12)),s(fun(X12,bool),X13))) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(X12,bool),X13),s(fun(X12,bool),X14)))) )),
  inference(superposition,[],[f675,f135])).
fof(f2096,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X11)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X8,sK4(X10,X9,X8)),s(fun(X8,bool),X9))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(X8,bool),X9),s(fun(X8,bool),X10)))) )),
  inference(subsumption_resolution,[],[f2093,f80])).
fof(f2093,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X11)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X8,sK4(X10,X9,X8)),s(fun(X8,bool),X9))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(X8,bool),X9),s(fun(X8,bool),X10))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f675,f135])).
fof(f1967,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl32),
  inference(cnf_transformation,[],[f1967_D])).
fof(f1967_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f1970,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | $spl37),
  inference(cnf_transformation,[],[f1970_D])).
fof(f1970_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) <=> ~$spl37),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl37])])).
fof(f1962,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl3),
  inference(superposition,[],[f721,f79])).
fof(f1968,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X2)) ) | $spl34),
  inference(cnf_transformation,[],[f1968_D])).
fof(f1968_D,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f1961,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) ) | $spl3),
  inference(superposition,[],[f721,f135])).
fof(f1960,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,X2)) ) | $spl3),
  inference(superposition,[],[f721,f135])).
fof(f1963,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl28),
  inference(cnf_transformation,[],[f1963_D])).
fof(f1963_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f1965,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | $spl31),
  inference(cnf_transformation,[],[f1965_D])).
fof(f1965_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) <=> ~$spl31),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl31])])).
fof(f1959,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | s(bool,X1) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,X1)) ) | $spl3),
  inference(superposition,[],[f721,f135])).
fof(f1958,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(superposition,[],[f721,f135])).
fof(f1633,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl22),
  inference(cnf_transformation,[],[f1633_D])).
fof(f1633_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f1636,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | $spl27),
  inference(cnf_transformation,[],[f1636_D])).
fof(f1636_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) <=> ~$spl27),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl27])])).
fof(f1628,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl3),
  inference(superposition,[],[f719,f79])).
fof(f1634,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X2)) ) | $spl24),
  inference(cnf_transformation,[],[f1634_D])).
fof(f1634_D,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f1627,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) ) | $spl3),
  inference(superposition,[],[f719,f135])).
fof(f1626,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,X2)) ) | $spl3),
  inference(superposition,[],[f719,f135])).
fof(f1625,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | s(bool,X1) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,X1)) ) | $spl3),
  inference(superposition,[],[f719,f135])).
fof(f1624,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(superposition,[],[f719,f135])).
fof(f775,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK5(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl14),
  inference(cnf_transformation,[],[f775_D])).
fof(f775_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK5(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f777,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK5(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl17),
  inference(cnf_transformation,[],[f777_D])).
fof(f777_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK5(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f773,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK5(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK5(sK2,'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl3),
  inference(superposition,[],[f720,f79])).
fof(f737,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) | $spl10),
  inference(cnf_transformation,[],[f737_D])).
fof(f737_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f739,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) | $spl13),
  inference(cnf_transformation,[],[f739_D])).
fof(f739_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f735,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) | $spl3),
  inference(superposition,[],[f718,f79])).
fof(f663,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f662,f78])).
fof(f662,plain,(
  p(s(bool,'F_0')) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f661,f129])).
fof(f129,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | $spl1),
  inference(unit_resulting_resolution,[],[f118,f85])).
fof(f661,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f640,f76])).
fof(f640,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) | ($spl1 | $spl2)),
  inference(superposition,[],[f107,f634])).
fof(f634,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f630,f76])).
fof(f630,plain,(
  ~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl1 | $spl2)),
  inference(superposition,[],[f615,f263])).
fof(f263,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl1),
  inference(condensation,[],[f262])).
fof(f262,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X0)) ) | $spl1),
  inference(subsumption_resolution,[],[f257,f81])).
fof(f257,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X0)) ) | $spl1),
  inference(superposition,[],[f254,f197])).
fof(f197,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X4)) ) | $spl1),
  inference(subsumption_resolution,[],[f195,f78])).
fof(f195,plain,(
  ( ! [X4] : (p(s(bool,'F_0')) | s(bool,X4) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X4)) ) | $spl1),
  inference(superposition,[],[f171,f129])).
fof(f254,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,sK4(X1,X2,X0)),s(fun(X0,bool),X2)))) ) | $spl1),
  inference(subsumption_resolution,[],[f253,f78])).
fof(f253,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,sK4(X1,X2,X0)),s(fun(X0,bool),X2)))) ) | $spl1),
  inference(forward_demodulation,[],[f249,f129])).
fof(f249,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,sK4(X1,X2,X0)),s(fun(X0,bool),X2)))) ) | $spl1),
  inference(superposition,[],[f106,f197])).
fof(f615,plain,(
  ( ! [X1] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK2)))) ) | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f614,f78])).
fof(f614,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK2)))) ) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f612,f129])).
fof(f612,plain,(
  ( ! [X1] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK2)))) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f422,f197])).
fof(f422,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK2)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK1))))) ) | $spl2),
  inference(superposition,[],[f109,f119])).
fof(f362,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK5(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) ) | $spl1),
  inference(superposition,[],[f346,f79])).
fof(f346,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK5(X0,X1,X2,bool))) ) | $spl1),
  inference(forward_demodulation,[],[f342,f129])).
fof(f342,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) != s(X2,g(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK5(X0,X1,X2,bool))) ) | $spl1),
  inference(superposition,[],[f115,f197])).
fof(f117,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) | $spl0),
  inference(cnf_transformation,[],[f117_D])).
fof(f117_D,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f119,plain,(
  s(fun(sK0,bool),sK2) = s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | $spl2),
  inference(cnf_transformation,[],[f119_D])).
fof(f119_D,plain,(
  s(fun(sK0,bool),sK2) = s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f74,plain,(
  s(fun(sK0,bool),sK2) = s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),
  inference(cnf_transformation,[],[f50])).
fof(f50,plain,(
  (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) | s(fun(sK0,bool),sK2) = s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) & (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) | s(fun(sK0,bool),sK2) != s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f49])).
fof(f49,plain,(
  ? [X0,X1,X2] : ((p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | s(fun(X0,bool),X2) = s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) & (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | s(fun(X0,bool),X2) != s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))),
  inference(nnf_transformation,[],[f44])).
fof(f44,plain,(
  ? [X0,X1,X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <~> s(fun(X0,bool),X2) = s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))),
  inference(ennf_transformation,[],[f22])).
fof(f22,plain,(
  ~! [X0,X1,X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> s(fun(X0,bool),X2) = s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))),
  inference(flattening,[],[f21])).
fof(f21,plain,(
  ~! [X0] : ! [X1,X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> s(fun(X0,bool),X2) = s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))) <=> s(fun(X6,bool),X8) = s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))) <=> s(fun(X6,bool),X8) = s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8))))),
  file('thm.pred_set.SUBSET_UNION_ABSORPTION','thm.pred_set.SUBSET_UNION_ABSORPTION')).
fof(f118,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) | $spl1),
  inference(cnf_transformation,[],[f118_D])).
fof(f118_D,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f120,plain,(
  s(fun(sK0,bool),sK2) != s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | $spl3),
  inference(cnf_transformation,[],[f120_D])).
fof(f120_D,plain,(
  s(fun(sK0,bool),sK2) != s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f75,plain,(
  s(fun(sK0,bool),sK2) != s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),
  inference(cnf_transformation,[],[f50])).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 281
% Passive clauses: 1076
% Generated clauses: 23214
% Final active clauses: 114
% Final passive clauses: 264
% Input formulas: 19
% Initial clauses: 42
% 
% Duplicate literals: 1149
% Trivial inequalities: 92
% Fw subsumption resolutions: 2087
% Fw demodulations: 5495
% Bw demodulations: 587
% Condensations: 1012
% Global subsumptions: 14
% 
% Simple tautologies: 191
% Equational tautologies: 144
% Forward subsumptions: 11726
% Fw demodulations to eq. taut.: 22
% Bw demodulations to eq. taut.: 14
% 
% Unit resulting resolution: 1268
% Factoring: 11
% Forward superposition: 7028
% Backward superposition: 3648
% Self superposition: 89
% Equality factoring: 40
% Equality resolution: 15
% 
% Split clauses: 41
% Split components: 88
% SAT solver clauses: 1339
% SAT solver unit clauses: 19
% SAT solver binary clauses: 163
% 
% Sat splits: 41
% Sat splitting refutations: 165
% 
% Memory used [KB]: 8187
% Time elapsed: 0.562 s
% ------------------------------
% ------------------------------
% Success in time 1.873 s
