Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 44
dis+11_7_41 on thm.numeral
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.numeral
% SZS output start Proof for thm.numeral
fof(f149566,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f120399,f2031,f2831,f2829,f1866,f2830,f1845,f2828,f2826,f149564,f1844,f2827,f2825,f149237,f119493,f119497,f119495])).
fof(f119495,plain,(
  s('type.num.num',sK16) = s('type.num.num',sK17) | $spl604),
  inference(cnf_transformation,[],[f119495_D])).
fof(f119495_D,plain,(
  s('type.num.num',sK16) = s('type.num.num',sK17) <=> ~$spl604),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl604])])).
fof(f119497,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK17)))) | $spl606),
  inference(cnf_transformation,[],[f119497_D])).
fof(f119497_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK17)))) <=> ~$spl606),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl606])])).
fof(f119493,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK17)))) | s('type.num.num',sK16) = s('type.num.num',sK17) | ($spl3 | $spl5)),
  inference(subsumption_resolution,[],[f119012,f2830])).
fof(f119012,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK17)))) | s('type.num.num',sK16) = s('type.num.num',sK17) | $spl3),
  inference(superposition,[],[f2297,f3099])).
fof(f3099,plain,(
  s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK17),s('type.num.num',sK16))) | $spl3),
  inference(resolution,[],[f1975,f2827])).
fof(f1975,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f1443])).
fof(f1443,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f625])).
fof(f625,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f624])).
fof(f624,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X8] : (s(bool,'F_0') = s(bool,X8) <=> ~p(s(bool,X8)))),
  file('thm.numeral.numeral_|49||48|_prove_ltb|split|2','thm.bool.EQ_CLAUSES|split|2')).
fof(f2297,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num',X0) = s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f892])).
fof(f892,plain,(
  ! [X0,X1] : (s('type.num.num',X0) = s('type.num.num',X1) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(rectify,[],[f201])).
fof(f201,axiom,(
  ! [X9,X10] : (s('type.num.num',X9) = s('type.num.num',X10) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num',X10)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X10),s('type.num.num',X9)))))),
  file('thm.numeral.numeral_|49||48|_prove_ltb|split|2','thm.arithmetic.LESS_LESS_CASES')).
fof(f149237,plain,(
  $false | ($spl1 | $spl606)),
  inference(subsumption_resolution,[],[f148651,f121200])).
fof(f121200,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK17)))) | $spl606),
  inference(resolution,[],[f119497,f2378])).
fof(f2378,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f1649])).
fof(f1649,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num',X0) = s('type.num.num',X1)) & ((~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) & s('type.num.num',X0) != s('type.num.num',X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(flattening,[],[f1648])).
fof(f1648,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num',X0) = s('type.num.num',X1))) & ((~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) & s('type.num.num',X0) != s('type.num.num',X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(nnf_transformation,[],[f962])).
fof(f962,plain,(
  ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num',X0) = s('type.num.num',X1)))),
  inference(rectify,[],[f170])).
fof(f170,axiom,(
  ! [X9,X10] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X9),s('type.num.num',X10)))) <=> (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num',X10)))) | s('type.num.num',X9) = s('type.num.num',X10)))),
  file('thm.numeral.numeral_|49||48|_prove_ltb|split|2','thm.arithmetic.LESS_OR_EQ')).
fof(f148651,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK17)))) | $spl1),
  inference(backward_demodulation,[],[f2271,f80189])).
fof(f80189,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK16),s('type.num.num',sK16))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK17),s('type.num.num',sK17)))))) | $spl1),
  inference(forward_demodulation,[],[f2826,f2250])).
fof(f2250,plain,(
  ( ! [X0,X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f845])).
fof(f845,plain,(
  ! [X0,X1] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f70])).
fof(f70,axiom,(
  ! [X10,X9] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X10),s('type.num.num',X9))) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X10),s('type.num.num','const.num.SUC_1'(s('type.num.num',X9)))))),
  file('thm.numeral.numeral_|49||48|_prove_ltb|split|2','thm.arithmetic.LESS_EQ_IFF_LESS_SUC')).
fof(f2271,plain,(
  ( ! [X0,X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X0))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f866])).
fof(f866,plain,(
  ! [X0,X1] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X0))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X1)))))),
  inference(rectify,[],[f504])).
fof(f504,axiom,(
  ! [X10,X9] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X10),s('type.num.num',X9))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X10),s('type.num.num',X10))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X9),s('type.num.num',X9)))))),
  file('thm.numeral.numeral_|49||48|_prove_ltb|split|2','thm.numeral.numeral_9_prove_ltb|split|1')).
fof(f2825,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK16),s('type.num.num',sK16))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK17),s('type.num.num',sK17)))))))) | $spl0),
  inference(cnf_transformation,[],[f2825_D])).
fof(f2825_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK16),s('type.num.num',sK16))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK17),s('type.num.num',sK17)))))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f2827,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK17),s('type.num.num',sK16)))) | $spl3),
  inference(cnf_transformation,[],[f2827_D])).
fof(f2827_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK17),s('type.num.num',sK16)))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f1844,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK17),s('type.num.num',sK16)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK16),s('type.num.num',sK16))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK17),s('type.num.num',sK17))))))))),
  inference(cnf_transformation,[],[f1420])).
fof(f1420,plain,(
  (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK16),s('type.num.num',sK16))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK17),s('type.num.num',sK17)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK17),s('type.num.num',sK16))))) & (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK16),s('type.num.num',sK16))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK17),s('type.num.num',sK17)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK17),s('type.num.num',sK16)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK16,sK17])],[f1419])).
fof(f1419,plain,(
  ? [X0,X1] : ((p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X1)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))))) & (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X1)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))))))),
  inference(nnf_transformation,[],[f1182])).
fof(f1182,plain,(
  ? [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X1)))))))) <~> ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(ennf_transformation,[],[f519])).
fof(f519,plain,(
  ~! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X1)))))))) <=> ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(flattening,[],[f518])).
fof(f518,plain,(
  ~! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X1)))))))) <=> ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X6))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num',X7)))))))) <=> ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num',X6)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X6))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num',X7)))))))) <=> ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num',X6)))))),
  file('thm.numeral.numeral_|49||48|_prove_ltb|split|2','thm.numeral.numeral_|49||48|_prove_ltb|split|2')).
fof(f149564,plain,(
  $false | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f149485,f2828])).
fof(f149485,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK17),s('type.num.num',sK16)))) | $spl0),
  inference(resolution,[],[f149384,f2472])).
fof(f2472,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f1704])).
fof(f1704,plain,(
  ! [X0,X1] : ((p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))))) & (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(nnf_transformation,[],[f996])).
fof(f996,plain,(
  ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(flattening,[],[f995])).
fof(f995,plain,(
  ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(rectify,[],[f301])).
fof(f301,axiom,(
  ! [X9,X10] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X9),s('type.num.num',X10)))) <=> p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X10),s('type.num.num',X9)))))),
  file('thm.numeral.numeral_|49||48|_prove_ltb|split|2','thm.arithmetic.NOT_LESS_EQUAL')).
fof(f149384,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK17)))) | $spl0),
  inference(forward_demodulation,[],[f149383,f2271])).
fof(f149383,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK16),s('type.num.num',sK16))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK17),s('type.num.num',sK17)))))) | $spl0),
  inference(forward_demodulation,[],[f2825,f2250])).
fof(f2826,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK16),s('type.num.num',sK16))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK17),s('type.num.num',sK17)))))))) | $spl1),
  inference(cnf_transformation,[],[f2826_D])).
fof(f2826_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK16),s('type.num.num',sK16))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK17),s('type.num.num',sK17)))))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f2828,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK17),s('type.num.num',sK16)))) | $spl2),
  inference(cnf_transformation,[],[f2828_D])).
fof(f2828_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK17),s('type.num.num',sK16)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f1845,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK17),s('type.num.num',sK16)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK16),s('type.num.num',sK16))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK17),s('type.num.num',sK17))))))))),
  inference(cnf_transformation,[],[f1420])).
fof(f2830,plain,(
  ~p(s(bool,'F_0')) | $spl5),
  inference(cnf_transformation,[],[f2830_D])).
fof(f2830_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f1866,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f522])).
fof(f522,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.numeral.numeral_|49||48|_prove_ltb|split|2',aHL_FALSITY)).
fof(f2829,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(cnf_transformation,[],[f2829_D])).
fof(f2829_D,plain,(
  p(s(bool,'F_0')) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f2831,plain,(
  p(s(bool,'T_0')) | $spl6),
  inference(cnf_transformation,[],[f2831_D])).
fof(f2831_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f2031,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f1482])).
fof(f1482,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f1481])).
fof(f1481,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f1214])).
fof(f1214,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f652])).
fof(f652,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X8] : ((p(s(bool,'F_0')) => p(s(bool,X8))) <=> p(s(bool,'T_0')))),
  file('thm.numeral.numeral_|49||48|_prove_ltb|split|2','thm.bool.IMP_CLAUSES|split|2')).
fof(f120399,plain,(
  $false | ($spl1 | $spl6 | $spl604)),
  inference(subsumption_resolution,[],[f120398,f2831])).
fof(f120398,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl604)),
  inference(forward_demodulation,[],[f119821,f3063])).
fof(f3063,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X0)))) )),
  inference(resolution,[],[f1971,f1873])).
fof(f1873,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f529])).
fof(f529,plain,(
  ! [X0] : p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X0))))),
  inference(rectify,[],[f407])).
fof(f407,axiom,(
  ! [X9] : p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X9),s('type.num.num',X9))))),
  file('thm.numeral.numeral_|49||48|_prove_ltb|split|2','thm.arithmetic.LESS_EQ_REFL')).
fof(f1971,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f1441])).
fof(f1441,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f622])).
fof(f622,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X8] : (s(bool,'T_0') = s(bool,X8) <=> p(s(bool,X8)))),
  file('thm.numeral.numeral_|49||48|_prove_ltb|split|2','thm.bool.EQ_CLAUSES|split|1')).
fof(f119821,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK16),s('type.num.num',sK16))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK16),s('type.num.num',sK16)))))) | ($spl1 | $spl604)),
  inference(backward_demodulation,[],[f119495,f80189])).
% SZS output end Proof for thm.numeral
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 2564
% Passive clauses: 33678
% Generated clauses: 147602
% Final active clauses: 2419
% Final passive clauses: 29712
% Input formulas: 516
% Initial clauses: 969
% 
% Duplicate literals: 274
% Trivial inequalities: 283
% Fw subsumption resolutions: 5038
% Fw demodulations: 26640
% Bw demodulations: 119
% 
% Simple tautologies: 463
% Equational tautologies: 35
% Forward subsumptions: 80607
% Fw demodulations to eq. taut.: 49
% 
% Binary resolution: 24553
% Factoring: 947
% Forward superposition: 46007
% Backward superposition: 42184
% Self superposition: 119
% Equality factoring: 2
% Equality resolution: 34
% 
% Split clauses: 1138
% Split components: 2278
% SAT solver clauses: 1474
% SAT solver unit clauses: 35
% SAT solver binary clauses: 233
% SAT solver learnt clauses: 3
% SAT solver learnt literals: 3
% 
% Sat splits: 1138
% Sat splitting refutations: 18
% 
% TWLsolver clauses: 1471
% TWLsolver calls for satisfiability: 716
% 
% Memory used [KB]: 91725
% Time elapsed: 2.868 s
% ------------------------------
% ------------------------------
% Success in time 2.901 s
