Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.pred_set
% (75055)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 140
% Passive clauses: 3075
% Generated clauses: 20139
% Final active clauses: 118
% Final passive clauses: 2632
% Input formulas: 24
% Initial clauses: 56
% 
% Duplicate literals: 777
% Trivial inequalities: 245
% Fw subsumption resolutions: 2849
% Fw demodulations: 2898
% Bw demodulations: 19
% 
% Simple tautologies: 397
% Equational tautologies: 47
% Forward subsumptions: 7576
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 877
% Factoring: 6
% Forward superposition: 6267
% Backward superposition: 5690
% Self superposition: 174
% Equality factoring: 114
% Equality resolution: 30
% 
% Split clauses: 528
% Split components: 1079
% SAT solver clauses: 1210
% SAT solver unit clauses: 22
% SAT solver binary clauses: 200
% 
% Sat splits: 528
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1200
% TWLsolver calls for satisfiability: 200
% 
% Memory used [KB]: 9210
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.pred_set
% (75087)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 93
% Passive clauses: 797
% Generated clauses: 9959
% Final active clauses: 81
% Final passive clauses: 637
% Input formulas: 24
% Initial clauses: 56
% 
% Duplicate literals: 784
% Trivial inequalities: 108
% Fw subsumption resolutions: 888
% Fw demodulations: 481
% Bw demodulations: 59
% Condensations: 1236
% Global subsumptions: 8
% 
% Simple tautologies: 219
% Equational tautologies: 50
% Forward subsumptions: 3935
% Bw demodulations to eq. taut.: 12
% 
% Binary resolution: 438
% Factoring: 3
% Forward superposition: 3031
% Backward superposition: 2496
% Self superposition: 144
% Equality factoring: 122
% Equality resolution: 17
% 
% Split clauses: 83
% Split components: 186
% SAT solver clauses: 1284
% SAT solver unit clauses: 60
% SAT solver binary clauses: 52
% SAT solver learnt clauses: 16
% 
% Sat splits: 83
% Sat splitting refutations: 16
% 
% TWLsolver clauses: 1252
% TWLsolver calls for satisfiability: 261
% 
% Memory used [KB]: 8443
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.pred_set
% (75124)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 549
% Passive clauses: 5845
% Generated clauses: 169173
% Final active clauses: 156
% Final passive clauses: 1127
% Input formulas: 24
% Initial clauses: 56
% Discarded non-redundant clauses: 832
% 
% Duplicate literals: 8229
% Trivial inequalities: 1532
% Fw subsumption resolutions: 18443
% Fw demodulations: 23231
% Bw demodulations: 3582
% Condensations: 18032
% Global subsumptions: 692
% 
% Simple tautologies: 1798
% Equational tautologies: 977
% Forward subsumptions: 88879
% Fw demodulations to eq. taut.: 35
% Bw demodulations to eq. taut.: 103
% 
% Unit resulting resolution: 4903
% Factoring: 7
% Forward superposition: 41062
% Backward superposition: 32000
% Self superposition: 760
% Equality factoring: 416
% Equality resolution: 125
% 
% Split clauses: 482
% Split components: 969
% SAT solver clauses: 8666
% SAT solver unit clauses: 76
% SAT solver binary clauses: 692
% 
% Sat splits: 482
% Sat splitting refutations: 1163
% 
% Memory used [KB]: 66011
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.pred_set
% (75398)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1866
% Passive clauses: 45528
% Generated clauses: 159198
% Final active clauses: 139
% Final passive clauses: 24936
% Input formulas: 24
% Initial clauses: 56
% 
% Selected by SInE selection: 24
% SInE iterations: 1
% 
% Duplicate literals: 4021
% Trivial inequalities: 2936
% Fw subsumption resolutions: 15578
% Fw demodulations: 25359
% Bw demodulations: 925
% 
% Simple tautologies: 5144
% Equational tautologies: 119
% Forward subsumptions: 58827
% Bw demodulations to eq. taut.: 20
% 
% Binary resolution: 13189
% Factoring: 16
% Forward superposition: 56707
% Backward superposition: 34369
% Self superposition: 356
% Equality factoring: 158
% Equality resolution: 147
% 
% Split clauses: 3274
% Split components: 6913
% SAT solver clauses: 3553
% SAT solver unit clauses: 8
% SAT solver binary clauses: 840
% 
% Sat splits: 3274
% Sat splitting refutations: 242
% 
% Memory used [KB]: 42984
% Time elapsed: 4.0000 s
% ------------------------------
% ------------------------------
% remaining time: 142 next slice time: 6
lrs+1002_2:3_br=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:sas=minisat:stl=30:sac=on:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none:sp=reverse_arity:urr=on_4 on thm.pred_set
% (75523)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 83
% Passive clauses: 1952
% Generated clauses: 11021
% Final active clauses: 60
% Final passive clauses: 1212
% Input formulas: 24
% Initial clauses: 56
% 
% Duplicate literals: 459
% Trivial inequalities: 326
% Fw subsumption resolutions: 1599
% Bw demodulations: 572
% Global subsumptions: 16
% 
% Simple tautologies: 504
% Equational tautologies: 61
% Forward subsumptions: 6018
% Bw demodulations to eq. taut.: 6
% 
% Unit resulting resolution: 160
% Forward superposition: 3679
% Backward superposition: 3295
% Self superposition: 112
% Equality factoring: 62
% Equality resolution: 14
% 
% Split clauses: 12
% Split components: 24
% SAT solver clauses: 2057
% SAT solver unit clauses: 56
% SAT solver binary clauses: 87
% 
% Sat splits: 12
% Sat splitting refutations: 75
% 
% Memory used [KB]: 6780
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 135 next slice time: 15
lrs+1011_12_bs=on:bsr=unit_only:cond=on:gs=on:gsaa=from_current:gsssp=full:nwc=1.1:sas=minisat:stl=60:sos=all:sac=on:sdd=off:sser=off:sfr=on:ssfp=100000:ssfq=1.2:ssnc=none:sp=reverse_arity:updr=off_14 on thm.pred_set
% (75556)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 348
% Passive clauses: 4757
% Generated clauses: 7838
% Final active clauses: 348
% Final passive clauses: 4438
% Input formulas: 24
% Initial clauses: 56
% 
% Duplicate literals: 308
% Trivial inequalities: 194
% Fw subsumption resolutions: 133
% Condensations: 152
% 
% Simple tautologies: 172
% Forward subsumptions: 2146
% 
% Binary resolution: 1087
% Factoring: 6
% Forward superposition: 5954
% Equality resolution: 4
% 
% Split clauses: 1
% Split components: 2
% SAT solver clauses: 4760
% SAT solver unit clauses: 5
% SAT solver binary clauses: 26
% 
% Sat splits: 1
% 
% Memory used [KB]: 18293
% Time elapsed: 1.600 s
% ------------------------------
% ------------------------------
% remaining time: 119 next slice time: 9
lrs+10_8:1_bd=preordered:bs=on:ccuc=first:fsr=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:nicw=on:sas=minisat:stl=120:sos=on:sscc=on:sser=off:sfr=on:ssfp=4000:ssfq=1.0:ssnc=none:sp=reverse_arity:urr=on_7 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f2981,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f91,f150,f154,f156,f153,f158,f152,f160,f151,f162,f163,f167,f165,f168,f172,f170,f175,f182,f186,f191,f190,f189,f188,f187,f195,f193,f197,f198,f200,f201,f203,f249,f196,f301,f300,f299,f298,f297,f305,f303,f307,f308,f310,f345,f306,f397,f396,f395,f394,f393,f401,f399,f403,f404,f406,f411,f439,f412,f443,f460,f459,f458,f457,f456,f464,f462,f466,f467,f469,f474,f502,f475,f504,f550,f402,f521,f552,f522,f556,f601,f465,f640,f639,f612,f614,f616,f618,f620,f645,f646,f624,f626,f634,f638,f647,f573,f652,f572,f654,f655,f176,f658,f183,f676,f675,f674,f673,f672,f680,f678,f177,f682,f689,f688,f687,f686,f685,f693,f691,f184,f864,f831,f833,f835,f837,f839,f841,f843,f874,f875,f854,f859,f876,f699,f651,f653,f909,f907,f902,f901,f900,f899,f919,f896,f895,f921,f913,f914,f917,f922,f924,f932,f998,f997,f995,f994,f989,f988,f916,f987,f986,f985,f984,f1009,f1000,f1001,f1010,f918,f1026,f1024,f1019,f1018,f1017,f1016,f1036,f1013,f1012,f1038,f1030,f1031,f1034,f1055,f1054,f1052,f1051,f1046,f1045,f1033,f1044,f1043,f1042,f1041,f1066,f1057,f1058,f1067,f1035,f1083,f1081,f1076,f1075,f1074,f1073,f1094,f1070,f1069,f1096,f1085,f1098,f1088,f1089,f1092,f1119,f1118,f1116,f1115,f1110,f1109,f1091,f1108,f1107,f1106,f1105,f1132,f1102,f1101,f1121,f1122,f1099,f1123,f1124,f1133,f1093,f1149,f1147,f1142,f1141,f1140,f1139,f1160,f1136,f1135,f1162,f1151,f1166,f1154,f1155,f1158,f1181,f1179,f1174,f1173,f1172,f1171,f1191,f1168,f1167,f1193,f1185,f1186,f1212,f1211,f1209,f1208,f1203,f1202,f1157,f1201,f1200,f1199,f1198,f1225,f1195,f1194,f1214,f1215,f1216,f1217,f1226,f1159,f1244,f1242,f1237,f1236,f1235,f1234,f1256,f1231,f1230,f1258,f1250,f1251,f1254,f1277,f1276,f1274,f1273,f1268,f1267,f1253,f1266,f1265,f1264,f1263,f1289,f1290,f1097,f1280,f1281,f1291,f1255,f1309,f1307,f1302,f1301,f1300,f1299,f1321,f1296,f1295,f1323,f1294,f1312,f1315,f1316,f1319,f1342,f1341,f1339,f1338,f1333,f1332,f1318,f1331,f1330,f1329,f1328,f1354,f1345,f1346,f1355,f1320,f1374,f1372,f1367,f1366,f1365,f1364,f1386,f1361,f1360,f1388,f1380,f1381,f1384,f1407,f1406,f1404,f1403,f1398,f1397,f1383,f1396,f1395,f1394,f1393,f1419,f1420,f1165,f1410,f1411,f1421,f1385,f1439,f1437,f1432,f1431,f1430,f1429,f1451,f1426,f1425,f1453,f1424,f1442,f1164,f1445,f1446,f1449,f1470,f1468,f1463,f1462,f1461,f1460,f1480,f1457,f1456,f1482,f1474,f1475,f1478,f1499,f1498,f1496,f1495,f1490,f1489,f1448,f1488,f1487,f1486,f1485,f1511,f1512,f1163,f1502,f1503,f1513,f1450,f1549,f1704,f1716,f501,f1735,f1734,f1732,f1731,f1726,f1725,f1477,f1724,f1723,f1722,f1721,f1746,f1737,f1738,f1747,f1479,f1753,f179,f1816,f1814,f1809,f1808,f1807,f1806,f1826,f1803,f1802,f1828,f1820,f1821,f1824,f1900,f1898,f1893,f1892,f1891,f1890,f1910,f1887,f1886,f1912,f1904,f1905,f1908,f2011,f2009,f2004,f2003,f2002,f2001,f2023,f1998,f1997,f2025,f2026,f2027,f2028,f2029,f2017,f2018,f2021,f2046,f2045,f2043,f2042,f2037,f2036,f2020,f2035,f2034,f2033,f2032,f2059,f2049,f2050,f2061,f2022,f2083,f2082,f2080,f2079,f2074,f2073,f2098,f2072,f2071,f2070,f2069,f2100,f2087,f2088,f2105,f2096,f2164,f2162,f2157,f2156,f2155,f2154,f2176,f2151,f2150,f2178,f2179,f2180,f2181,f2182,f2170,f2171,f2174,f2199,f2198,f2196,f2195,f2190,f2189,f2173,f2188,f2187,f2186,f2185,f2212,f2202,f2203,f2214,f2175,f2236,f2235,f2233,f2232,f2227,f2226,f2251,f2225,f2224,f2223,f2222,f2253,f2240,f2241,f2258,f2249,f2392,f1761,f555,f551,f2494,f2508,f2086,f2085,f2065,f2066,f1845,f438,f442,f2409,f2410,f2411,f1189,f2412,f2413,f1100,f2518,f2585,f2583,f2578,f2577,f2572,f2571,f2595,f2590,f2591,f2097,f2675,f2969,f441,f554])).
fof(f554,plain,(
  ~p(s(bool,'T_0')) | $spl65),
  inference(cnf_transformation,[],[f554_D])).
fof(f554_D,plain,(
  ~p(s(bool,'T_0')) <=> ~$spl65),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl65])])).
fof(f441,plain,(
  p(s(bool,'F_0')) | $spl48),
  inference(cnf_transformation,[],[f441_D])).
fof(f441_D,plain,(
  p(s(bool,'F_0')) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f2969,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'T_0')) | ($spl76 | $spl114 | $spl132)),
  inference(duplicate_literal_removal,[],[f2968])).
fof(f2968,plain,(
  p(s(bool,'F_0')) | p(s(bool,'F_0')) | ~p(s(bool,'T_0')) | ($spl76 | $spl114 | $spl132)),
  inference(forward_demodulation,[],[f2967,f1189])).
fof(f2967,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3)))) | ($spl76 | $spl132)),
  inference(forward_demodulation,[],[f2938,f917])).
fof(f2938,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1)))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3)))) | $spl132),
  inference(superposition,[],[f139,f1320])).
fof(f139,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f89])).
fof(f89,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) & ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))))),
  inference(flattening,[],[f88])).
fof(f88,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))) & ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))))),
  inference(nnf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(flattening,[],[f49])).
fof(f49,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X6] : ! [X7,X8,X12] : (p(s(bool,'const.bool.IN_2'(s(X6,X12),s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X6,X12),s(fun(X6,bool),X7)))) | p(s(bool,'const.bool.IN_2'(s(X6,X12),s(fun(X6,bool),X8))))))),
  file('thm.pred_set.INTER_OVER_UNION','thm.pred_set.IN_UNION')).
fof(f2675,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'F_0')) | ($spl122 | $spl154)),
  inference(forward_demodulation,[],[f2648,f1479])).
fof(f2648,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2)))) | $spl122),
  inference(superposition,[],[f141,f1254])).
fof(f141,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f89])).
fof(f2097,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl184),
  inference(cnf_transformation,[],[f2097_D])).
fof(f2097_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f2591,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl83),
  inference(resolution,[],[f1008,f101])).
fof(f101,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f61])).
fof(f61,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f32])).
fof(f32,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f31])).
fof(f31,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X8] : (s(bool,'F_0') = s(bool,X8) <=> ~p(s(bool,X8)))),
  file('thm.pred_set.INTER_OVER_UNION','thm.bool.EQ_CLAUSES|split|3')).
fof(f1008,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))) | $spl83),
  inference(cnf_transformation,[],[f1008_D])).
fof(f1008_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))) <=> ~$spl83),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl83])])).
fof(f2590,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl83),
  inference(resolution,[],[f1008,f103])).
fof(f103,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f62])).
fof(f62,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f34])).
fof(f34,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f33])).
fof(f33,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X8] : (s(bool,'F_0') = s(bool,X8) <=> ~p(s(bool,X8)))),
  file('thm.pred_set.INTER_OVER_UNION','thm.bool.EQ_CLAUSES|split|2')).
fof(f2595,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | $spl203),
  inference(cnf_transformation,[],[f2595_D])).
fof(f2595_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) <=> ~$spl203),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl203])])).
fof(f2571,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | $spl83),
  inference(unit_resulting_resolution,[],[f1008,f1008,f139])).
fof(f2572,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | $spl83),
  inference(unit_resulting_resolution,[],[f1008,f1008,f139])).
fof(f2577,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl83),
  inference(unit_resulting_resolution,[],[f1008,f101])).
fof(f2578,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl83),
  inference(unit_resulting_resolution,[],[f1008,f103])).
fof(f2583,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl83),
  inference(unit_resulting_resolution,[],[f94,f1008,f142])).
fof(f142,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f56])).
fof(f56,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f55])).
fof(f55,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X10,X11] : ((p(s(bool,X10)) => p(s(bool,X11))) => ((p(s(bool,X11)) => p(s(bool,X10))) => s(bool,X10) = s(bool,X11)))),
  file('thm.pred_set.INTER_OVER_UNION','thm.bool.IMP_ANTISYM_AX')).
fof(f94,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f28])).
fof(f28,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.pred_set.INTER_OVER_UNION',aHL_FALSITY)).
fof(f2585,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl83),
  inference(unit_resulting_resolution,[],[f94,f1008,f142])).
fof(f2518,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))))),
  introduced(theory_axiom,[])).
fof(f1100,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3)))) | $spl101),
  inference(cnf_transformation,[],[f1100_D])).
fof(f1100_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3)))) <=> ~$spl101),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl101])])).
fof(f2413,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3)))) | $spl117),
  inference(unit_resulting_resolution,[],[f93,f1191,f145])).
fof(f145,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f56])).
fof(f93,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.INTER_OVER_UNION','thm.bool.TRUTH')).
fof(f2412,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3)))) | $spl117),
  inference(unit_resulting_resolution,[],[f93,f1191,f145])).
fof(f1189,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3))) | $spl114),
  inference(cnf_transformation,[],[f1189_D])).
fof(f1189_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3))) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f2411,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3))) | $spl117),
  inference(unit_resulting_resolution,[],[f1191,f95])).
fof(f95,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.pred_set.INTER_OVER_UNION',aHL_BOOL_CASES)).
fof(f2410,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3)))) | $spl117),
  inference(unit_resulting_resolution,[],[f1191,f97])).
fof(f97,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f29])).
fof(f29,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X8] : (s(bool,'T_0') = s(bool,X8) <=> p(s(bool,X8)))),
  file('thm.pred_set.INTER_OVER_UNION','thm.bool.EQ_CLAUSES|split|1')).
fof(f2409,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3)))) | $spl117),
  inference(unit_resulting_resolution,[],[f1191,f99])).
fof(f99,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f30])).
fof(f30,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X8] : (s(bool,'T_0') = s(bool,X8) <=> p(s(bool,X8)))),
  file('thm.pred_set.INTER_OVER_UNION','thm.bool.EQ_CLAUSES|split|0')).
fof(f442,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))) | $spl50),
  inference(cnf_transformation,[],[f442_D])).
fof(f442_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f438,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1)))) | $spl46),
  inference(cnf_transformation,[],[f438_D])).
fof(f438_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1)))) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f1845,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1)))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))) | $spl42),
  inference(superposition,[],[f139,f403])).
fof(f2066,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1)))) | $spl82),
  inference(unit_resulting_resolution,[],[f1009,f137])).
fof(f137,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f87])).
fof(f87,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))) & (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))))),
  inference(flattening,[],[f86])).
fof(f86,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))) & ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))))),
  inference(nnf_transformation,[],[f48])).
fof(f48,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(flattening,[],[f47])).
fof(f47,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X6] : ! [X7,X8,X12] : (p(s(bool,'const.bool.IN_2'(s(X6,X12),s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X6,X12),s(fun(X6,bool),X7)))) & p(s(bool,'const.bool.IN_2'(s(X6,X12),s(fun(X6,bool),X8))))))),
  file('thm.pred_set.INTER_OVER_UNION','thm.pred_set.IN_INTER')).
fof(f2065,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1)))) | $spl82),
  inference(unit_resulting_resolution,[],[f1009,f136])).
fof(f136,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f87])).
fof(f2085,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1)))) | $spl82),
  inference(resolution,[],[f1009,f136])).
fof(f2086,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1)))) | $spl82),
  inference(resolution,[],[f1009,f137])).
fof(f2508,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3))) | p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3))))),
  introduced(theory_axiom,[])).
fof(f2494,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'F_0')) | ($spl78 | $spl130)),
  inference(forward_demodulation,[],[f2467,f918])).
fof(f2467,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1)))) | $spl130),
  inference(superposition,[],[f140,f1319])).
fof(f140,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f89])).
fof(f551,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) | $spl63),
  inference(cnf_transformation,[],[f551_D])).
fof(f551_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) <=> ~$spl63),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl63])])).
fof(f555,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))))) | $spl67),
  inference(cnf_transformation,[],[f555_D])).
fof(f555_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))))) <=> ~$spl67),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl67])])).
fof(f1761,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) | $spl44),
  inference(superposition,[],[f138,f405])).
fof(f405,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))))) | $spl44),
  inference(cnf_transformation,[],[f405_D])).
fof(f405_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))))) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f138,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f87])).
fof(f2392,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'F_0')) | ($spl78 | $spl122)),
  inference(forward_demodulation,[],[f2365,f918])).
fof(f2365,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1)))) | $spl122),
  inference(superposition,[],[f140,f1254])).
fof(f2249,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl196),
  inference(cnf_transformation,[],[f2249_D])).
fof(f2249_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f2258,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl90),
  inference(global_subsumption,[],[f91,f1066,f2236,f2235,f2233,f2232,f2227,f2226,f2225,f2224,f2223,f2222,f2221,f2220,f2240,f2241,f2246,f2247])).
fof(f2247,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl90),
  inference(superposition,[],[f1066,f95])).
fof(f2246,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))) ) | $spl90),
  inference(resolution,[],[f1066,f145])).
fof(f2220,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),X0))))))) ) | $spl90),
  inference(unit_resulting_resolution,[],[f1066,f140])).
fof(f2221,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),X0),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))))))) ) | $spl90),
  inference(unit_resulting_resolution,[],[f1066,f141])).
fof(f2241,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl90),
  inference(resolution,[],[f1066,f99])).
fof(f2240,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl90),
  inference(resolution,[],[f1066,f97])).
fof(f2253,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | $spl200),
  inference(cnf_transformation,[],[f2253_D])).
fof(f2253_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f2222,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | $spl90),
  inference(unit_resulting_resolution,[],[f1066,f1066,f138])).
fof(f2223,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | $spl90),
  inference(unit_resulting_resolution,[],[f1066,f1066,f138])).
fof(f2224,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl90),
  inference(unit_resulting_resolution,[],[f1066,f97])).
fof(f2225,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl90),
  inference(unit_resulting_resolution,[],[f1066,f99])).
fof(f2251,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl199),
  inference(cnf_transformation,[],[f2251_D])).
fof(f2251_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) <=> ~$spl199),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl199])])).
fof(f2226,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl90),
  inference(unit_resulting_resolution,[],[f1066,f100])).
fof(f100,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f61])).
fof(f2227,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl90),
  inference(unit_resulting_resolution,[],[f1066,f102])).
fof(f102,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f62])).
fof(f2232,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl90),
  inference(unit_resulting_resolution,[],[f93,f1066,f145])).
fof(f2233,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl90),
  inference(unit_resulting_resolution,[],[f92,f1066,f145])).
fof(f92,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.INTER_OVER_UNION',aHL_TRUTH)).
fof(f2235,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl90),
  inference(unit_resulting_resolution,[],[f93,f1066,f145])).
fof(f2236,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl90),
  inference(unit_resulting_resolution,[],[f92,f1066,f145])).
fof(f2175,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl190),
  inference(cnf_transformation,[],[f2175_D])).
fof(f2175_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f2214,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl88),
  inference(global_subsumption,[],[f91,f1037,f2199,f2198,f2196,f2195,f2190,f2189,f2188,f2187,f2186,f2185,f2184,f2183,f2202,f2203,f2208,f2209])).
fof(f2209,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl88),
  inference(superposition,[],[f1037,f95])).
fof(f2208,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))) ) | $spl88),
  inference(resolution,[],[f1037,f145])).
fof(f2183,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),X0))))))) ) | $spl88),
  inference(unit_resulting_resolution,[],[f1037,f140])).
fof(f2184,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),X0),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))))))) ) | $spl88),
  inference(unit_resulting_resolution,[],[f1037,f141])).
fof(f1037,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))) | $spl88),
  inference(cnf_transformation,[],[f1037_D])).
fof(f1037_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f2203,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl88),
  inference(resolution,[],[f1037,f99])).
fof(f2202,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl88),
  inference(resolution,[],[f1037,f97])).
fof(f2212,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | $spl194),
  inference(cnf_transformation,[],[f2212_D])).
fof(f2212_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f2185,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | $spl88),
  inference(unit_resulting_resolution,[],[f1037,f1037,f138])).
fof(f2186,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | $spl88),
  inference(unit_resulting_resolution,[],[f1037,f1037,f138])).
fof(f2187,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl88),
  inference(unit_resulting_resolution,[],[f1037,f97])).
fof(f2188,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl88),
  inference(unit_resulting_resolution,[],[f1037,f99])).
fof(f2173,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl189),
  inference(cnf_transformation,[],[f2173_D])).
fof(f2173_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) <=> ~$spl189),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl189])])).
fof(f2189,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl88),
  inference(unit_resulting_resolution,[],[f1037,f100])).
fof(f2190,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl88),
  inference(unit_resulting_resolution,[],[f1037,f102])).
fof(f2195,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl88),
  inference(unit_resulting_resolution,[],[f93,f1037,f145])).
fof(f2196,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl88),
  inference(unit_resulting_resolution,[],[f92,f1037,f145])).
fof(f2198,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl88),
  inference(unit_resulting_resolution,[],[f93,f1037,f145])).
fof(f2199,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl88),
  inference(unit_resulting_resolution,[],[f92,f1037,f145])).
fof(f2174,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl188),
  inference(cnf_transformation,[],[f2174_D])).
fof(f2174_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f2171,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl89),
  inference(resolution,[],[f1038,f101])).
fof(f2170,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl89),
  inference(resolution,[],[f1038,f103])).
fof(f2182,plain,(
  ~p(s(bool,'T_0')) | ($spl86 | $spl89)),
  inference(forward_demodulation,[],[f2167,f1035])).
fof(f2167,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1)))) | $spl89),
  inference(resolution,[],[f1038,f141])).
fof(f2181,plain,(
  ~p(s(bool,'T_0')) | ($spl86 | $spl89)),
  inference(forward_demodulation,[],[f2166,f1035])).
fof(f2166,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1)))) | $spl89),
  inference(resolution,[],[f1038,f140])).
fof(f2180,plain,(
  ~p(s(bool,'T_0')) | ($spl86 | $spl89)),
  inference(forward_demodulation,[],[f2148,f1035])).
fof(f2148,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1)))) | $spl89),
  inference(unit_resulting_resolution,[],[f1038,f140])).
fof(f2179,plain,(
  ~p(s(bool,'T_0')) | ($spl86 | $spl89)),
  inference(forward_demodulation,[],[f2149,f1035])).
fof(f2149,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1)))) | $spl89),
  inference(unit_resulting_resolution,[],[f1038,f141])).
fof(f2178,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | $spl193),
  inference(cnf_transformation,[],[f2178_D])).
fof(f2178_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) <=> ~$spl193),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl193])])).
fof(f2150,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | $spl89),
  inference(unit_resulting_resolution,[],[f1038,f1038,f139])).
fof(f2151,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | $spl89),
  inference(unit_resulting_resolution,[],[f1038,f1038,f139])).
fof(f2176,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl191),
  inference(cnf_transformation,[],[f2176_D])).
fof(f2176_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) <=> ~$spl191),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl191])])).
fof(f2154,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl89),
  inference(unit_resulting_resolution,[],[f1038,f96])).
fof(f96,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f59])).
fof(f2155,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl89),
  inference(unit_resulting_resolution,[],[f1038,f98])).
fof(f98,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f60])).
fof(f2156,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl89),
  inference(unit_resulting_resolution,[],[f1038,f101])).
fof(f2157,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl89),
  inference(unit_resulting_resolution,[],[f1038,f103])).
fof(f2162,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl89),
  inference(unit_resulting_resolution,[],[f94,f1038,f142])).
fof(f2164,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl89),
  inference(unit_resulting_resolution,[],[f94,f1038,f142])).
fof(f2096,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl182),
  inference(cnf_transformation,[],[f2096_D])).
fof(f2096_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f2105,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl82),
  inference(global_subsumption,[],[f91,f1009,f2083,f2082,f2080,f2079,f2074,f2073,f2072,f2071,f2070,f2069,f2068,f2067,f2087,f2088,f2093,f2094])).
fof(f2094,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl82),
  inference(superposition,[],[f1009,f95])).
fof(f2093,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))) ) | $spl82),
  inference(resolution,[],[f1009,f145])).
fof(f2067,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),X0))))))) ) | $spl82),
  inference(unit_resulting_resolution,[],[f1009,f140])).
fof(f2068,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),X0),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))))))) ) | $spl82),
  inference(unit_resulting_resolution,[],[f1009,f141])).
fof(f2088,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl82),
  inference(resolution,[],[f1009,f99])).
fof(f2087,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl82),
  inference(resolution,[],[f1009,f97])).
fof(f2100,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | $spl186),
  inference(cnf_transformation,[],[f2100_D])).
fof(f2100_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f2069,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | $spl82),
  inference(unit_resulting_resolution,[],[f1009,f1009,f138])).
fof(f2070,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | $spl82),
  inference(unit_resulting_resolution,[],[f1009,f1009,f138])).
fof(f2071,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl82),
  inference(unit_resulting_resolution,[],[f1009,f97])).
fof(f2072,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl82),
  inference(unit_resulting_resolution,[],[f1009,f99])).
fof(f2098,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl185),
  inference(cnf_transformation,[],[f2098_D])).
fof(f2098_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) <=> ~$spl185),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl185])])).
fof(f2073,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl82),
  inference(unit_resulting_resolution,[],[f1009,f100])).
fof(f2074,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl82),
  inference(unit_resulting_resolution,[],[f1009,f102])).
fof(f2079,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl82),
  inference(unit_resulting_resolution,[],[f93,f1009,f145])).
fof(f2080,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl82),
  inference(unit_resulting_resolution,[],[f92,f1009,f145])).
fof(f2082,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl82),
  inference(unit_resulting_resolution,[],[f93,f1009,f145])).
fof(f2083,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl82),
  inference(unit_resulting_resolution,[],[f92,f1009,f145])).
fof(f2022,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl176),
  inference(cnf_transformation,[],[f2022_D])).
fof(f2022_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f2061,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl80),
  inference(global_subsumption,[],[f91,f920,f2046,f2045,f2043,f2042,f2037,f2036,f2035,f2034,f2033,f2032,f2031,f2030,f2049,f2050,f2055,f2056])).
fof(f2056,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl80),
  inference(superposition,[],[f920,f95])).
fof(f2055,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))) ) | $spl80),
  inference(resolution,[],[f920,f145])).
fof(f2030,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),X0))))))) ) | $spl80),
  inference(unit_resulting_resolution,[],[f920,f140])).
fof(f2031,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),X0),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))))))) ) | $spl80),
  inference(unit_resulting_resolution,[],[f920,f141])).
fof(f920,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))) | $spl80),
  inference(cnf_transformation,[],[f920_D])).
fof(f920_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f2050,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl80),
  inference(resolution,[],[f920,f99])).
fof(f2049,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl80),
  inference(resolution,[],[f920,f97])).
fof(f2059,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | $spl180),
  inference(cnf_transformation,[],[f2059_D])).
fof(f2059_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f2032,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | $spl80),
  inference(unit_resulting_resolution,[],[f920,f920,f138])).
fof(f2033,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | $spl80),
  inference(unit_resulting_resolution,[],[f920,f920,f138])).
fof(f2034,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl80),
  inference(unit_resulting_resolution,[],[f920,f97])).
fof(f2035,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl80),
  inference(unit_resulting_resolution,[],[f920,f99])).
fof(f2020,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl175),
  inference(cnf_transformation,[],[f2020_D])).
fof(f2020_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) <=> ~$spl175),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl175])])).
fof(f2036,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl80),
  inference(unit_resulting_resolution,[],[f920,f100])).
fof(f2037,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl80),
  inference(unit_resulting_resolution,[],[f920,f102])).
fof(f2042,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl80),
  inference(unit_resulting_resolution,[],[f93,f920,f145])).
fof(f2043,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl80),
  inference(unit_resulting_resolution,[],[f92,f920,f145])).
fof(f2045,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl80),
  inference(unit_resulting_resolution,[],[f93,f920,f145])).
fof(f2046,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl80),
  inference(unit_resulting_resolution,[],[f92,f920,f145])).
fof(f2021,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl174),
  inference(cnf_transformation,[],[f2021_D])).
fof(f2021_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f2018,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl81),
  inference(resolution,[],[f921,f101])).
fof(f2017,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl81),
  inference(resolution,[],[f921,f103])).
fof(f2029,plain,(
  ~p(s(bool,'T_0')) | ($spl78 | $spl81)),
  inference(forward_demodulation,[],[f2014,f918])).
fof(f2014,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1)))) | $spl81),
  inference(resolution,[],[f921,f141])).
fof(f2028,plain,(
  ~p(s(bool,'T_0')) | ($spl78 | $spl81)),
  inference(forward_demodulation,[],[f2013,f918])).
fof(f2013,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1)))) | $spl81),
  inference(resolution,[],[f921,f140])).
fof(f2027,plain,(
  ~p(s(bool,'T_0')) | ($spl78 | $spl81)),
  inference(forward_demodulation,[],[f1995,f918])).
fof(f1995,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1)))) | $spl81),
  inference(unit_resulting_resolution,[],[f921,f140])).
fof(f2026,plain,(
  ~p(s(bool,'T_0')) | ($spl78 | $spl81)),
  inference(forward_demodulation,[],[f1996,f918])).
fof(f1996,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1)))) | $spl81),
  inference(unit_resulting_resolution,[],[f921,f141])).
fof(f2025,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | $spl179),
  inference(cnf_transformation,[],[f2025_D])).
fof(f2025_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) <=> ~$spl179),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl179])])).
fof(f1997,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | $spl81),
  inference(unit_resulting_resolution,[],[f921,f921,f139])).
fof(f1998,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))))) | $spl81),
  inference(unit_resulting_resolution,[],[f921,f921,f139])).
fof(f2023,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl177),
  inference(cnf_transformation,[],[f2023_D])).
fof(f2023_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) <=> ~$spl177),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl177])])).
fof(f2001,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl81),
  inference(unit_resulting_resolution,[],[f921,f96])).
fof(f2002,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl81),
  inference(unit_resulting_resolution,[],[f921,f98])).
fof(f2003,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl81),
  inference(unit_resulting_resolution,[],[f921,f101])).
fof(f2004,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl81),
  inference(unit_resulting_resolution,[],[f921,f103])).
fof(f2009,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl81),
  inference(unit_resulting_resolution,[],[f94,f921,f142])).
fof(f2011,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1))))) | $spl81),
  inference(unit_resulting_resolution,[],[f94,f921,f142])).
fof(f1908,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3))) | $spl168),
  inference(cnf_transformation,[],[f1908_D])).
fof(f1908_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3))) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f1905,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3))) | $spl113),
  inference(resolution,[],[f1166,f101])).
fof(f1904,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3))) | $spl113),
  inference(resolution,[],[f1166,f103])).
fof(f1912,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),sK3)))))) | $spl173),
  inference(cnf_transformation,[],[f1912_D])).
fof(f1912_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),sK3)))))) <=> ~$spl173),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl173])])).
fof(f1886,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),sK3)))))) | $spl113),
  inference(unit_resulting_resolution,[],[f1166,f1166,f139])).
fof(f1887,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),sK3)))))) | $spl113),
  inference(unit_resulting_resolution,[],[f1166,f1166,f139])).
fof(f1910,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3))) | $spl171),
  inference(cnf_transformation,[],[f1910_D])).
fof(f1910_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3))) <=> ~$spl171),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl171])])).
fof(f1890,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3))) | $spl113),
  inference(unit_resulting_resolution,[],[f1166,f96])).
fof(f1891,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3))) | $spl113),
  inference(unit_resulting_resolution,[],[f1166,f98])).
fof(f1892,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3))) | $spl113),
  inference(unit_resulting_resolution,[],[f1166,f101])).
fof(f1893,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3))) | $spl113),
  inference(unit_resulting_resolution,[],[f1166,f103])).
fof(f1898,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3))) | $spl113),
  inference(unit_resulting_resolution,[],[f94,f1166,f142])).
fof(f1900,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3))) | $spl113),
  inference(unit_resulting_resolution,[],[f94,f1166,f142])).
fof(f1824,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2))) | $spl162),
  inference(cnf_transformation,[],[f1824_D])).
fof(f1824_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2))) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f1821,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2))) | $spl111),
  inference(resolution,[],[f1164,f101])).
fof(f1820,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2))) | $spl111),
  inference(resolution,[],[f1164,f103])).
fof(f1828,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))) | $spl167),
  inference(cnf_transformation,[],[f1828_D])).
fof(f1828_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))) <=> ~$spl167),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl167])])).
fof(f1802,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))) | $spl111),
  inference(unit_resulting_resolution,[],[f1164,f1164,f139])).
fof(f1803,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))) | $spl111),
  inference(unit_resulting_resolution,[],[f1164,f1164,f139])).
fof(f1826,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2))) | $spl165),
  inference(cnf_transformation,[],[f1826_D])).
fof(f1826_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2))) <=> ~$spl165),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl165])])).
fof(f1806,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2))) | $spl111),
  inference(unit_resulting_resolution,[],[f1164,f96])).
fof(f1807,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2))) | $spl111),
  inference(unit_resulting_resolution,[],[f1164,f98])).
fof(f1808,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2))) | $spl111),
  inference(unit_resulting_resolution,[],[f1164,f101])).
fof(f1809,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2))) | $spl111),
  inference(unit_resulting_resolution,[],[f1164,f103])).
fof(f1814,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2))) | $spl111),
  inference(unit_resulting_resolution,[],[f94,f1164,f142])).
fof(f1816,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2))) | $spl111),
  inference(unit_resulting_resolution,[],[f94,f1164,f142])).
fof(f179,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) | $spl21),
  inference(cnf_transformation,[],[f179_D])).
fof(f179_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f1753,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) | ($spl5 | $spl44)),
  inference(backward_demodulation,[],[f405,f158])).
fof(f1479,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2))) | $spl154),
  inference(cnf_transformation,[],[f1479_D])).
fof(f1479_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2))) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f1747,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2))) | $spl98),
  inference(global_subsumption,[],[f91,f1097,f1735,f1734,f1732,f1731,f1726,f1725,f1724,f1723,f1722,f1721,f1720,f1719,f1737,f1738,f1743,f1744])).
fof(f1744,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2))) | $spl98),
  inference(superposition,[],[f1097,f95])).
fof(f1743,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2)))) ) | $spl98),
  inference(resolution,[],[f1097,f145])).
fof(f1719,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X0))))))) ) | $spl98),
  inference(unit_resulting_resolution,[],[f1097,f140])).
fof(f1720,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),X0),s(fun(sK0,bool),sK2))))))) ) | $spl98),
  inference(unit_resulting_resolution,[],[f1097,f141])).
fof(f1738,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2))) | $spl98),
  inference(resolution,[],[f1097,f99])).
fof(f1737,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2))) | $spl98),
  inference(resolution,[],[f1097,f97])).
fof(f1746,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))) | $spl160),
  inference(cnf_transformation,[],[f1746_D])).
fof(f1746_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f1721,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))) | $spl98),
  inference(unit_resulting_resolution,[],[f1097,f1097,f138])).
fof(f1722,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))) | $spl98),
  inference(unit_resulting_resolution,[],[f1097,f1097,f138])).
fof(f1723,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2))) | $spl98),
  inference(unit_resulting_resolution,[],[f1097,f97])).
fof(f1724,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2))) | $spl98),
  inference(unit_resulting_resolution,[],[f1097,f99])).
fof(f1477,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2))) | $spl153),
  inference(cnf_transformation,[],[f1477_D])).
fof(f1477_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2))) <=> ~$spl153),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl153])])).
fof(f1725,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2))) | $spl98),
  inference(unit_resulting_resolution,[],[f1097,f100])).
fof(f1726,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2))) | $spl98),
  inference(unit_resulting_resolution,[],[f1097,f102])).
fof(f1731,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2))) | $spl98),
  inference(unit_resulting_resolution,[],[f93,f1097,f145])).
fof(f1732,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2))) | $spl98),
  inference(unit_resulting_resolution,[],[f92,f1097,f145])).
fof(f1734,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2))) | $spl98),
  inference(unit_resulting_resolution,[],[f93,f1097,f145])).
fof(f1735,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2))) | $spl98),
  inference(unit_resulting_resolution,[],[f92,f1097,f145])).
fof(f501,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1)))) | $spl58),
  inference(cnf_transformation,[],[f501_D])).
fof(f501_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1)))) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f1716,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1)))) | ($spl54 | $spl104)),
  inference(forward_demodulation,[],[f1673,f1158])).
fof(f1673,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1)))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))) | $spl54),
  inference(superposition,[],[f139,f466])).
fof(f1704,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'T_0')) | ($spl54 | $spl84 | $spl104)),
  inference(duplicate_literal_removal,[],[f1703])).
fof(f1703,plain,(
  p(s(bool,'F_0')) | p(s(bool,'F_0')) | ~p(s(bool,'T_0')) | ($spl54 | $spl84 | $spl104)),
  inference(forward_demodulation,[],[f1702,f1158])).
fof(f1702,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))) | ($spl54 | $spl84)),
  inference(forward_demodulation,[],[f1673,f1034])).
fof(f1549,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'T_0')) | ($spl42 | $spl76 | $spl92)),
  inference(duplicate_literal_removal,[],[f1548])).
fof(f1548,plain,(
  p(s(bool,'F_0')) | p(s(bool,'F_0')) | ~p(s(bool,'T_0')) | ($spl42 | $spl76 | $spl92)),
  inference(forward_demodulation,[],[f1547,f1092])).
fof(f1547,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))) | ($spl42 | $spl76)),
  inference(forward_demodulation,[],[f1518,f917])).
fof(f1518,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1)))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))) | $spl42),
  inference(superposition,[],[f139,f403])).
fof(f1450,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl148),
  inference(cnf_transformation,[],[f1450_D])).
fof(f1450_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f1513,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl70),
  inference(global_subsumption,[],[f91,f654,f1499,f1498,f1496,f1495,f1490,f1489,f1488,f1487,f1486,f1485,f1484,f1483,f1502,f1503,f1508,f1509])).
fof(f1509,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl70),
  inference(superposition,[],[f654,f95])).
fof(f1508,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) ) | $spl70),
  inference(resolution,[],[f654,f145])).
fof(f1483,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),X0))))))) ) | $spl70),
  inference(unit_resulting_resolution,[],[f654,f140])).
fof(f1484,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),X0),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))) ) | $spl70),
  inference(unit_resulting_resolution,[],[f654,f141])).
fof(f1503,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl70),
  inference(resolution,[],[f654,f99])).
fof(f1502,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl70),
  inference(resolution,[],[f654,f97])).
fof(f1163,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2)))) | $spl110),
  inference(cnf_transformation,[],[f1163_D])).
fof(f1163_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2)))) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f1512,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2)))) | ($spl70 | $spl84)),
  inference(forward_demodulation,[],[f1501,f1034])).
fof(f1501,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1)))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2)))) | $spl70),
  inference(resolution,[],[f654,f139])).
fof(f1511,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | $spl158),
  inference(cnf_transformation,[],[f1511_D])).
fof(f1511_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f1485,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | $spl70),
  inference(unit_resulting_resolution,[],[f654,f654,f138])).
fof(f1486,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | $spl70),
  inference(unit_resulting_resolution,[],[f654,f654,f138])).
fof(f1487,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl70),
  inference(unit_resulting_resolution,[],[f654,f97])).
fof(f1488,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl70),
  inference(unit_resulting_resolution,[],[f654,f99])).
fof(f1448,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl147),
  inference(cnf_transformation,[],[f1448_D])).
fof(f1448_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl147),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl147])])).
fof(f1489,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl70),
  inference(unit_resulting_resolution,[],[f654,f100])).
fof(f1490,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl70),
  inference(unit_resulting_resolution,[],[f654,f102])).
fof(f1495,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl70),
  inference(unit_resulting_resolution,[],[f93,f654,f145])).
fof(f1496,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl70),
  inference(unit_resulting_resolution,[],[f92,f654,f145])).
fof(f1498,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl70),
  inference(unit_resulting_resolution,[],[f93,f654,f145])).
fof(f1499,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl70),
  inference(unit_resulting_resolution,[],[f92,f654,f145])).
fof(f1478,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2))) | $spl152),
  inference(cnf_transformation,[],[f1478_D])).
fof(f1478_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2))) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f1475,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2))) | $spl99),
  inference(resolution,[],[f1098,f101])).
fof(f1474,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2))) | $spl99),
  inference(resolution,[],[f1098,f103])).
fof(f1482,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))) | $spl157),
  inference(cnf_transformation,[],[f1482_D])).
fof(f1482_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))) <=> ~$spl157),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl157])])).
fof(f1456,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))) | $spl99),
  inference(unit_resulting_resolution,[],[f1098,f1098,f139])).
fof(f1457,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK2)))))) | $spl99),
  inference(unit_resulting_resolution,[],[f1098,f1098,f139])).
fof(f1480,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2))) | $spl155),
  inference(cnf_transformation,[],[f1480_D])).
fof(f1480_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2))) <=> ~$spl155),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl155])])).
fof(f1460,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2))) | $spl99),
  inference(unit_resulting_resolution,[],[f1098,f96])).
fof(f1461,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2))) | $spl99),
  inference(unit_resulting_resolution,[],[f1098,f98])).
fof(f1462,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2))) | $spl99),
  inference(unit_resulting_resolution,[],[f1098,f101])).
fof(f1463,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2))) | $spl99),
  inference(unit_resulting_resolution,[],[f1098,f103])).
fof(f1468,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2))) | $spl99),
  inference(unit_resulting_resolution,[],[f94,f1098,f142])).
fof(f1470,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2))) | $spl99),
  inference(unit_resulting_resolution,[],[f94,f1098,f142])).
fof(f1449,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl146),
  inference(cnf_transformation,[],[f1449_D])).
fof(f1449_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f1446,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl71),
  inference(resolution,[],[f653,f101])).
fof(f1445,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl71),
  inference(resolution,[],[f653,f103])).
fof(f1164,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2)))) | $spl111),
  inference(cnf_transformation,[],[f1164_D])).
fof(f1164_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2)))) <=> ~$spl111),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl111])])).
fof(f1442,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2)))) | $spl71),
  inference(resolution,[],[f653,f141])).
fof(f1424,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2)))) | $spl71),
  inference(unit_resulting_resolution,[],[f653,f141])).
fof(f1453,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | $spl151),
  inference(cnf_transformation,[],[f1453_D])).
fof(f1453_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) <=> ~$spl151),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl151])])).
fof(f1425,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | $spl71),
  inference(unit_resulting_resolution,[],[f653,f653,f139])).
fof(f1426,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | $spl71),
  inference(unit_resulting_resolution,[],[f653,f653,f139])).
fof(f1451,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl149),
  inference(cnf_transformation,[],[f1451_D])).
fof(f1451_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl149),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl149])])).
fof(f1429,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl71),
  inference(unit_resulting_resolution,[],[f653,f96])).
fof(f1430,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl71),
  inference(unit_resulting_resolution,[],[f653,f98])).
fof(f1431,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl71),
  inference(unit_resulting_resolution,[],[f653,f101])).
fof(f1432,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl71),
  inference(unit_resulting_resolution,[],[f653,f103])).
fof(f1437,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl71),
  inference(unit_resulting_resolution,[],[f94,f653,f142])).
fof(f1439,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl71),
  inference(unit_resulting_resolution,[],[f94,f653,f142])).
fof(f1385,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))) | $spl140),
  inference(cnf_transformation,[],[f1385_D])).
fof(f1385_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f1421,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))) | $spl68),
  inference(global_subsumption,[],[f91,f652,f1407,f1406,f1404,f1403,f1398,f1397,f1396,f1395,f1394,f1393,f1392,f1391,f1410,f1411,f1416,f1417])).
fof(f1417,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))) | $spl68),
  inference(superposition,[],[f652,f95])).
fof(f1416,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))))) ) | $spl68),
  inference(resolution,[],[f652,f145])).
fof(f1391,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))),s(fun(sK0,bool),X0))))))) ) | $spl68),
  inference(unit_resulting_resolution,[],[f652,f140])).
fof(f1392,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),X0),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))))))) ) | $spl68),
  inference(unit_resulting_resolution,[],[f652,f141])).
fof(f1411,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))) | $spl68),
  inference(resolution,[],[f652,f99])).
fof(f1410,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))) | $spl68),
  inference(resolution,[],[f652,f97])).
fof(f1165,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3)))) | $spl112),
  inference(cnf_transformation,[],[f1165_D])).
fof(f1165_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3)))) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f1420,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3)))) | ($spl68 | $spl84)),
  inference(forward_demodulation,[],[f1409,f1034])).
fof(f1409,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1)))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3)))) | $spl68),
  inference(resolution,[],[f652,f139])).
fof(f1419,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))))))) | $spl144),
  inference(cnf_transformation,[],[f1419_D])).
fof(f1419_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))))))) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f1393,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))))))) | $spl68),
  inference(unit_resulting_resolution,[],[f652,f652,f138])).
fof(f1394,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))))))) | $spl68),
  inference(unit_resulting_resolution,[],[f652,f652,f138])).
fof(f1395,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))) | $spl68),
  inference(unit_resulting_resolution,[],[f652,f97])).
fof(f1396,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))) | $spl68),
  inference(unit_resulting_resolution,[],[f652,f99])).
fof(f1383,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))) | $spl139),
  inference(cnf_transformation,[],[f1383_D])).
fof(f1383_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))) <=> ~$spl139),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl139])])).
fof(f1397,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))) | $spl68),
  inference(unit_resulting_resolution,[],[f652,f100])).
fof(f1398,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))) | $spl68),
  inference(unit_resulting_resolution,[],[f652,f102])).
fof(f1403,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))) | $spl68),
  inference(unit_resulting_resolution,[],[f93,f652,f145])).
fof(f1404,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))) | $spl68),
  inference(unit_resulting_resolution,[],[f92,f652,f145])).
fof(f1406,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))) | $spl68),
  inference(unit_resulting_resolution,[],[f93,f652,f145])).
fof(f1407,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))) | $spl68),
  inference(unit_resulting_resolution,[],[f92,f652,f145])).
fof(f1384,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))) | $spl138),
  inference(cnf_transformation,[],[f1384_D])).
fof(f1384_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f1381,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))) | $spl69),
  inference(resolution,[],[f651,f101])).
fof(f1380,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))) | $spl69),
  inference(resolution,[],[f651,f103])).
fof(f1388,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))))))) | $spl143),
  inference(cnf_transformation,[],[f1388_D])).
fof(f1388_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))))))) <=> ~$spl143),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl143])])).
fof(f1360,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))))))) | $spl69),
  inference(unit_resulting_resolution,[],[f651,f651,f139])).
fof(f1361,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))))))) | $spl69),
  inference(unit_resulting_resolution,[],[f651,f651,f139])).
fof(f1386,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))) | $spl141),
  inference(cnf_transformation,[],[f1386_D])).
fof(f1386_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))) <=> ~$spl141),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl141])])).
fof(f1364,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))) | $spl69),
  inference(unit_resulting_resolution,[],[f651,f96])).
fof(f1365,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))) | $spl69),
  inference(unit_resulting_resolution,[],[f651,f98])).
fof(f1366,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))) | $spl69),
  inference(unit_resulting_resolution,[],[f651,f101])).
fof(f1367,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))) | $spl69),
  inference(unit_resulting_resolution,[],[f651,f103])).
fof(f1372,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))) | $spl69),
  inference(unit_resulting_resolution,[],[f94,f651,f142])).
fof(f1374,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))) | $spl69),
  inference(unit_resulting_resolution,[],[f94,f651,f142])).
fof(f1320,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))) | $spl132),
  inference(cnf_transformation,[],[f1320_D])).
fof(f1320_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f1355,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))) | $spl66),
  inference(global_subsumption,[],[f91,f556,f1342,f1341,f1339,f1338,f1333,f1332,f1331,f1330,f1329,f1328,f1327,f1326,f1345,f1346,f1351,f1352])).
fof(f1352,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))) | $spl66),
  inference(superposition,[],[f556,f95])).
fof(f1351,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))))) ) | $spl66),
  inference(resolution,[],[f556,f145])).
fof(f1326,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))),s(fun(sK0,bool),X0))))))) ) | $spl66),
  inference(unit_resulting_resolution,[],[f556,f140])).
fof(f1327,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),X0),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))))))) ) | $spl66),
  inference(unit_resulting_resolution,[],[f556,f141])).
fof(f1346,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))) | $spl66),
  inference(resolution,[],[f556,f99])).
fof(f1345,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))) | $spl66),
  inference(resolution,[],[f556,f97])).
fof(f1354,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))))))) | $spl136),
  inference(cnf_transformation,[],[f1354_D])).
fof(f1354_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))))))) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f1328,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))))))) | $spl66),
  inference(unit_resulting_resolution,[],[f556,f556,f138])).
fof(f1329,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))))))) | $spl66),
  inference(unit_resulting_resolution,[],[f556,f556,f138])).
fof(f1330,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))) | $spl66),
  inference(unit_resulting_resolution,[],[f556,f97])).
fof(f1331,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))) | $spl66),
  inference(unit_resulting_resolution,[],[f556,f99])).
fof(f1318,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))) | $spl131),
  inference(cnf_transformation,[],[f1318_D])).
fof(f1318_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))) <=> ~$spl131),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl131])])).
fof(f1332,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))) | $spl66),
  inference(unit_resulting_resolution,[],[f556,f100])).
fof(f1333,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))) | $spl66),
  inference(unit_resulting_resolution,[],[f556,f102])).
fof(f1338,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))) | $spl66),
  inference(unit_resulting_resolution,[],[f93,f556,f145])).
fof(f1339,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))) | $spl66),
  inference(unit_resulting_resolution,[],[f92,f556,f145])).
fof(f1341,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))) | $spl66),
  inference(unit_resulting_resolution,[],[f93,f556,f145])).
fof(f1342,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))) | $spl66),
  inference(unit_resulting_resolution,[],[f92,f556,f145])).
fof(f1319,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))) | $spl130),
  inference(cnf_transformation,[],[f1319_D])).
fof(f1319_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f1316,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))) | $spl67),
  inference(resolution,[],[f555,f101])).
fof(f1315,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))) | $spl67),
  inference(resolution,[],[f555,f103])).
fof(f1312,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3)))) | $spl67),
  inference(resolution,[],[f555,f141])).
fof(f1294,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3)))) | $spl67),
  inference(unit_resulting_resolution,[],[f555,f141])).
fof(f1323,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))))))) | $spl135),
  inference(cnf_transformation,[],[f1323_D])).
fof(f1323_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))))))) <=> ~$spl135),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl135])])).
fof(f1295,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))))))) | $spl67),
  inference(unit_resulting_resolution,[],[f555,f555,f139])).
fof(f1296,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))))))) | $spl67),
  inference(unit_resulting_resolution,[],[f555,f555,f139])).
fof(f1321,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))) | $spl133),
  inference(cnf_transformation,[],[f1321_D])).
fof(f1321_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))) <=> ~$spl133),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl133])])).
fof(f1299,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))) | $spl67),
  inference(unit_resulting_resolution,[],[f555,f96])).
fof(f1300,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))) | $spl67),
  inference(unit_resulting_resolution,[],[f555,f98])).
fof(f1301,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))) | $spl67),
  inference(unit_resulting_resolution,[],[f555,f101])).
fof(f1302,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))) | $spl67),
  inference(unit_resulting_resolution,[],[f555,f103])).
fof(f1307,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))) | $spl67),
  inference(unit_resulting_resolution,[],[f94,f555,f142])).
fof(f1309,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))) | $spl67),
  inference(unit_resulting_resolution,[],[f94,f555,f142])).
fof(f1255,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl124),
  inference(cnf_transformation,[],[f1255_D])).
fof(f1255_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f1291,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl62),
  inference(global_subsumption,[],[f91,f552,f1277,f1276,f1274,f1273,f1268,f1267,f1266,f1265,f1264,f1263,f1262,f1261,f1280,f1281,f1286,f1287])).
fof(f1287,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl62),
  inference(superposition,[],[f552,f95])).
fof(f1286,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) ) | $spl62),
  inference(resolution,[],[f552,f145])).
fof(f1261,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),X0))))))) ) | $spl62),
  inference(unit_resulting_resolution,[],[f552,f140])).
fof(f1262,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),X0),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))) ) | $spl62),
  inference(unit_resulting_resolution,[],[f552,f141])).
fof(f1281,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl62),
  inference(resolution,[],[f552,f99])).
fof(f1280,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl62),
  inference(resolution,[],[f552,f97])).
fof(f1097,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2)))) | $spl98),
  inference(cnf_transformation,[],[f1097_D])).
fof(f1097_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2)))) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f1290,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2)))) | ($spl62 | $spl76)),
  inference(forward_demodulation,[],[f1279,f917])).
fof(f1279,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1)))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2)))) | $spl62),
  inference(resolution,[],[f552,f139])).
fof(f1289,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | $spl128),
  inference(cnf_transformation,[],[f1289_D])).
fof(f1289_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f1263,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | $spl62),
  inference(unit_resulting_resolution,[],[f552,f552,f138])).
fof(f1264,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | $spl62),
  inference(unit_resulting_resolution,[],[f552,f552,f138])).
fof(f1265,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl62),
  inference(unit_resulting_resolution,[],[f552,f97])).
fof(f1266,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl62),
  inference(unit_resulting_resolution,[],[f552,f99])).
fof(f1253,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl123),
  inference(cnf_transformation,[],[f1253_D])).
fof(f1253_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl123),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl123])])).
fof(f1267,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl62),
  inference(unit_resulting_resolution,[],[f552,f100])).
fof(f1268,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl62),
  inference(unit_resulting_resolution,[],[f552,f102])).
fof(f1273,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl62),
  inference(unit_resulting_resolution,[],[f93,f552,f145])).
fof(f1274,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl62),
  inference(unit_resulting_resolution,[],[f92,f552,f145])).
fof(f1276,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl62),
  inference(unit_resulting_resolution,[],[f93,f552,f145])).
fof(f1277,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl62),
  inference(unit_resulting_resolution,[],[f92,f552,f145])).
fof(f1254,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl122),
  inference(cnf_transformation,[],[f1254_D])).
fof(f1254_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f1251,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl63),
  inference(resolution,[],[f551,f101])).
fof(f1250,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl63),
  inference(resolution,[],[f551,f103])).
fof(f1258,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | $spl127),
  inference(cnf_transformation,[],[f1258_D])).
fof(f1258_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) <=> ~$spl127),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl127])])).
fof(f1230,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | $spl63),
  inference(unit_resulting_resolution,[],[f551,f551,f139])).
fof(f1231,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))))) | $spl63),
  inference(unit_resulting_resolution,[],[f551,f551,f139])).
fof(f1256,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl125),
  inference(cnf_transformation,[],[f1256_D])).
fof(f1256_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl125),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl125])])).
fof(f1234,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl63),
  inference(unit_resulting_resolution,[],[f551,f96])).
fof(f1235,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl63),
  inference(unit_resulting_resolution,[],[f551,f98])).
fof(f1236,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl63),
  inference(unit_resulting_resolution,[],[f551,f101])).
fof(f1237,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl63),
  inference(unit_resulting_resolution,[],[f551,f103])).
fof(f1242,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl63),
  inference(unit_resulting_resolution,[],[f94,f551,f142])).
fof(f1244,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl63),
  inference(unit_resulting_resolution,[],[f94,f551,f142])).
fof(f1159,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | $spl106),
  inference(cnf_transformation,[],[f1159_D])).
fof(f1159_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f1226,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | $spl60),
  inference(global_subsumption,[],[f91,f503,f1212,f1211,f1209,f1208,f1203,f1202,f1201,f1200,f1199,f1198,f1197,f1196,f1195,f1194,f1214,f1215,f1216,f1217,f1222,f1223])).
fof(f1223,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | $spl60),
  inference(superposition,[],[f503,f95])).
fof(f1222,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))) ) | $spl60),
  inference(resolution,[],[f503,f145])).
fof(f1196,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))),s(fun(sK0,bool),X0))))))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f503,f140])).
fof(f1197,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),X0),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f503,f141])).
fof(f503,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))) | $spl60),
  inference(cnf_transformation,[],[f503_D])).
fof(f503_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f1217,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | $spl60),
  inference(resolution,[],[f503,f99])).
fof(f1216,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | $spl60),
  inference(resolution,[],[f503,f97])).
fof(f1215,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3)))) | $spl60),
  inference(resolution,[],[f503,f137])).
fof(f1214,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2)))) | $spl60),
  inference(resolution,[],[f503,f136])).
fof(f1194,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2)))) | $spl60),
  inference(unit_resulting_resolution,[],[f503,f136])).
fof(f1195,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3)))) | $spl60),
  inference(unit_resulting_resolution,[],[f503,f137])).
fof(f1225,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))))) | $spl120),
  inference(cnf_transformation,[],[f1225_D])).
fof(f1225_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))))) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f1198,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f503,f503,f138])).
fof(f1199,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))))) | $spl60),
  inference(unit_resulting_resolution,[],[f503,f503,f138])).
fof(f1200,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | $spl60),
  inference(unit_resulting_resolution,[],[f503,f97])).
fof(f1201,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | $spl60),
  inference(unit_resulting_resolution,[],[f503,f99])).
fof(f1157,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | $spl105),
  inference(cnf_transformation,[],[f1157_D])).
fof(f1157_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) <=> ~$spl105),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl105])])).
fof(f1202,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | $spl60),
  inference(unit_resulting_resolution,[],[f503,f100])).
fof(f1203,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | $spl60),
  inference(unit_resulting_resolution,[],[f503,f102])).
fof(f1208,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | $spl60),
  inference(unit_resulting_resolution,[],[f93,f503,f145])).
fof(f1209,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | $spl60),
  inference(unit_resulting_resolution,[],[f92,f503,f145])).
fof(f1211,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | $spl60),
  inference(unit_resulting_resolution,[],[f93,f503,f145])).
fof(f1212,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | $spl60),
  inference(unit_resulting_resolution,[],[f92,f503,f145])).
fof(f1186,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3))) | $spl101),
  inference(resolution,[],[f1100,f101])).
fof(f1185,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3))) | $spl101),
  inference(resolution,[],[f1100,f103])).
fof(f1193,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),sK3)))))) | $spl119),
  inference(cnf_transformation,[],[f1193_D])).
fof(f1193_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),sK3)))))) <=> ~$spl119),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl119])])).
fof(f1167,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),sK3)))))) | $spl101),
  inference(unit_resulting_resolution,[],[f1100,f1100,f139])).
fof(f1168,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),sK3)))))) | $spl101),
  inference(unit_resulting_resolution,[],[f1100,f1100,f139])).
fof(f1191,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3))) | $spl117),
  inference(cnf_transformation,[],[f1191_D])).
fof(f1191_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3))) <=> ~$spl117),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl117])])).
fof(f1171,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3))) | $spl101),
  inference(unit_resulting_resolution,[],[f1100,f96])).
fof(f1172,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3))) | $spl101),
  inference(unit_resulting_resolution,[],[f1100,f98])).
fof(f1173,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3))) | $spl101),
  inference(unit_resulting_resolution,[],[f1100,f101])).
fof(f1174,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3))) | $spl101),
  inference(unit_resulting_resolution,[],[f1100,f103])).
fof(f1179,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3))) | $spl101),
  inference(unit_resulting_resolution,[],[f94,f1100,f142])).
fof(f1181,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3))) | $spl101),
  inference(unit_resulting_resolution,[],[f94,f1100,f142])).
fof(f1158,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | $spl104),
  inference(cnf_transformation,[],[f1158_D])).
fof(f1158_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f1155,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | $spl61),
  inference(resolution,[],[f504,f101])).
fof(f1154,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | $spl61),
  inference(resolution,[],[f504,f103])).
fof(f1166,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3)))) | $spl113),
  inference(cnf_transformation,[],[f1166_D])).
fof(f1166_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3)))) <=> ~$spl113),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl113])])).
fof(f1151,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2)))) | $spl61),
  inference(resolution,[],[f504,f138])).
fof(f1162,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))))) | $spl109),
  inference(cnf_transformation,[],[f1162_D])).
fof(f1162_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))))) <=> ~$spl109),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl109])])).
fof(f1135,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))))) | $spl61),
  inference(unit_resulting_resolution,[],[f504,f504,f139])).
fof(f1136,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))))) | $spl61),
  inference(unit_resulting_resolution,[],[f504,f504,f139])).
fof(f1160,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | $spl107),
  inference(cnf_transformation,[],[f1160_D])).
fof(f1160_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) <=> ~$spl107),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl107])])).
fof(f1139,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | $spl61),
  inference(unit_resulting_resolution,[],[f504,f96])).
fof(f1140,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | $spl61),
  inference(unit_resulting_resolution,[],[f504,f98])).
fof(f1141,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | $spl61),
  inference(unit_resulting_resolution,[],[f504,f101])).
fof(f1142,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | $spl61),
  inference(unit_resulting_resolution,[],[f504,f103])).
fof(f1147,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | $spl61),
  inference(unit_resulting_resolution,[],[f94,f504,f142])).
fof(f1149,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | $spl61),
  inference(unit_resulting_resolution,[],[f94,f504,f142])).
fof(f1093,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | $spl94),
  inference(cnf_transformation,[],[f1093_D])).
fof(f1093_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f1133,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | $spl50),
  inference(global_subsumption,[],[f91,f442,f1119,f1118,f1116,f1115,f1110,f1109,f1108,f1107,f1106,f1105,f1104,f1103,f1102,f1101,f1121,f1122,f1123,f1124,f1129,f1130])).
fof(f1130,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | $spl50),
  inference(superposition,[],[f442,f95])).
fof(f1129,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))) ) | $spl50),
  inference(resolution,[],[f442,f145])).
fof(f1103,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))),s(fun(sK0,bool),X0))))))) ) | $spl50),
  inference(unit_resulting_resolution,[],[f442,f140])).
fof(f1104,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),X0),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))))) ) | $spl50),
  inference(unit_resulting_resolution,[],[f442,f141])).
fof(f1124,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | $spl50),
  inference(resolution,[],[f442,f99])).
fof(f1123,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | $spl50),
  inference(resolution,[],[f442,f97])).
fof(f1099,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3)))) | $spl100),
  inference(cnf_transformation,[],[f1099_D])).
fof(f1099_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3)))) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f1122,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3)))) | $spl50),
  inference(resolution,[],[f442,f137])).
fof(f1121,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2)))) | $spl50),
  inference(resolution,[],[f442,f136])).
fof(f1101,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2)))) | $spl50),
  inference(unit_resulting_resolution,[],[f442,f136])).
fof(f1102,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3)))) | $spl50),
  inference(unit_resulting_resolution,[],[f442,f137])).
fof(f1132,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))))) | $spl102),
  inference(cnf_transformation,[],[f1132_D])).
fof(f1132_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))))) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f1105,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))))) | $spl50),
  inference(unit_resulting_resolution,[],[f442,f442,f138])).
fof(f1106,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))))) | $spl50),
  inference(unit_resulting_resolution,[],[f442,f442,f138])).
fof(f1107,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | $spl50),
  inference(unit_resulting_resolution,[],[f442,f97])).
fof(f1108,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | $spl50),
  inference(unit_resulting_resolution,[],[f442,f99])).
fof(f1091,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | $spl93),
  inference(cnf_transformation,[],[f1091_D])).
fof(f1091_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) <=> ~$spl93),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl93])])).
fof(f1109,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | $spl50),
  inference(unit_resulting_resolution,[],[f442,f100])).
fof(f1110,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | $spl50),
  inference(unit_resulting_resolution,[],[f442,f102])).
fof(f1115,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | $spl50),
  inference(unit_resulting_resolution,[],[f93,f442,f145])).
fof(f1116,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | $spl50),
  inference(unit_resulting_resolution,[],[f92,f442,f145])).
fof(f1118,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | $spl50),
  inference(unit_resulting_resolution,[],[f93,f442,f145])).
fof(f1119,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | $spl50),
  inference(unit_resulting_resolution,[],[f92,f442,f145])).
fof(f1092,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | $spl92),
  inference(cnf_transformation,[],[f1092_D])).
fof(f1092_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f1089,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | $spl51),
  inference(resolution,[],[f443,f101])).
fof(f1088,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | $spl51),
  inference(resolution,[],[f443,f103])).
fof(f1098,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2)))) | $spl99),
  inference(cnf_transformation,[],[f1098_D])).
fof(f1098_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2)))) <=> ~$spl99),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl99])])).
fof(f1085,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK3)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK2)))) | $spl51),
  inference(resolution,[],[f443,f138])).
fof(f1096,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))))) | $spl97),
  inference(cnf_transformation,[],[f1096_D])).
fof(f1096_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))))) <=> ~$spl97),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl97])])).
fof(f1069,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))))) | $spl51),
  inference(unit_resulting_resolution,[],[f443,f443,f139])).
fof(f1070,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))))) | $spl51),
  inference(unit_resulting_resolution,[],[f443,f443,f139])).
fof(f1094,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | $spl95),
  inference(cnf_transformation,[],[f1094_D])).
fof(f1094_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) <=> ~$spl95),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl95])])).
fof(f1073,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | $spl51),
  inference(unit_resulting_resolution,[],[f443,f96])).
fof(f1074,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | $spl51),
  inference(unit_resulting_resolution,[],[f443,f98])).
fof(f1075,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | $spl51),
  inference(unit_resulting_resolution,[],[f443,f101])).
fof(f1076,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | $spl51),
  inference(unit_resulting_resolution,[],[f443,f103])).
fof(f1081,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | $spl51),
  inference(unit_resulting_resolution,[],[f94,f443,f142])).
fof(f1083,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) | $spl51),
  inference(unit_resulting_resolution,[],[f94,f443,f142])).
fof(f1035,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1))) | $spl86),
  inference(cnf_transformation,[],[f1035_D])).
fof(f1035_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f1067,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1))) | $spl58),
  inference(global_subsumption,[],[f91,f501,f1055,f1054,f1052,f1051,f1046,f1045,f1044,f1043,f1042,f1041,f1040,f1039,f1057,f1058,f1063,f1064])).
fof(f1064,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1))) | $spl58),
  inference(superposition,[],[f501,f95])).
fof(f1063,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1)))) ) | $spl58),
  inference(resolution,[],[f501,f145])).
fof(f1039,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X0))))))) ) | $spl58),
  inference(unit_resulting_resolution,[],[f501,f140])).
fof(f1040,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),X0),s(fun(sK0,bool),sK1))))))) ) | $spl58),
  inference(unit_resulting_resolution,[],[f501,f141])).
fof(f1058,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1))) | $spl58),
  inference(resolution,[],[f501,f99])).
fof(f1057,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1))) | $spl58),
  inference(resolution,[],[f501,f97])).
fof(f1066,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))) | $spl90),
  inference(cnf_transformation,[],[f1066_D])).
fof(f1066_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f1041,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))) | $spl58),
  inference(unit_resulting_resolution,[],[f501,f501,f138])).
fof(f1042,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))) | $spl58),
  inference(unit_resulting_resolution,[],[f501,f501,f138])).
fof(f1043,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1))) | $spl58),
  inference(unit_resulting_resolution,[],[f501,f97])).
fof(f1044,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1))) | $spl58),
  inference(unit_resulting_resolution,[],[f501,f99])).
fof(f1033,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1))) | $spl85),
  inference(cnf_transformation,[],[f1033_D])).
fof(f1033_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl85),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl85])])).
fof(f1045,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1))) | $spl58),
  inference(unit_resulting_resolution,[],[f501,f100])).
fof(f1046,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1))) | $spl58),
  inference(unit_resulting_resolution,[],[f501,f102])).
fof(f1051,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1))) | $spl58),
  inference(unit_resulting_resolution,[],[f93,f501,f145])).
fof(f1052,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1))) | $spl58),
  inference(unit_resulting_resolution,[],[f92,f501,f145])).
fof(f1054,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1))) | $spl58),
  inference(unit_resulting_resolution,[],[f93,f501,f145])).
fof(f1055,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1))) | $spl58),
  inference(unit_resulting_resolution,[],[f92,f501,f145])).
fof(f1034,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1))) | $spl84),
  inference(cnf_transformation,[],[f1034_D])).
fof(f1034_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f1031,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1))) | $spl59),
  inference(resolution,[],[f502,f101])).
fof(f1030,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1))) | $spl59),
  inference(resolution,[],[f502,f103])).
fof(f1038,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))) | $spl89),
  inference(cnf_transformation,[],[f1038_D])).
fof(f1038_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))) <=> ~$spl89),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl89])])).
fof(f1012,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))) | $spl59),
  inference(unit_resulting_resolution,[],[f502,f502,f139])).
fof(f1013,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))) | $spl59),
  inference(unit_resulting_resolution,[],[f502,f502,f139])).
fof(f1036,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1))) | $spl87),
  inference(cnf_transformation,[],[f1036_D])).
fof(f1036_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl87),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl87])])).
fof(f1016,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1))) | $spl59),
  inference(unit_resulting_resolution,[],[f502,f96])).
fof(f1017,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1))) | $spl59),
  inference(unit_resulting_resolution,[],[f502,f98])).
fof(f1018,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1))) | $spl59),
  inference(unit_resulting_resolution,[],[f502,f101])).
fof(f1019,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1))) | $spl59),
  inference(unit_resulting_resolution,[],[f502,f103])).
fof(f1024,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1))) | $spl59),
  inference(unit_resulting_resolution,[],[f94,f502,f142])).
fof(f1026,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1))) | $spl59),
  inference(unit_resulting_resolution,[],[f94,f502,f142])).
fof(f918,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1))) | $spl78),
  inference(cnf_transformation,[],[f918_D])).
fof(f918_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f1010,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1))) | $spl46),
  inference(global_subsumption,[],[f91,f438,f998,f997,f995,f994,f989,f988,f987,f986,f985,f984,f983,f982,f1000,f1001,f1006,f1007])).
fof(f1007,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1))) | $spl46),
  inference(superposition,[],[f438,f95])).
fof(f1006,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1)))) ) | $spl46),
  inference(resolution,[],[f438,f145])).
fof(f982,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X0))))))) ) | $spl46),
  inference(unit_resulting_resolution,[],[f438,f140])).
fof(f983,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),X0),s(fun(sK0,bool),sK1))))))) ) | $spl46),
  inference(unit_resulting_resolution,[],[f438,f141])).
fof(f1001,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1))) | $spl46),
  inference(resolution,[],[f438,f99])).
fof(f1000,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1))) | $spl46),
  inference(resolution,[],[f438,f97])).
fof(f1009,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))) | $spl82),
  inference(cnf_transformation,[],[f1009_D])).
fof(f1009_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f984,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))) | $spl46),
  inference(unit_resulting_resolution,[],[f438,f438,f138])).
fof(f985,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))) | $spl46),
  inference(unit_resulting_resolution,[],[f438,f438,f138])).
fof(f986,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1))) | $spl46),
  inference(unit_resulting_resolution,[],[f438,f97])).
fof(f987,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1))) | $spl46),
  inference(unit_resulting_resolution,[],[f438,f99])).
fof(f916,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1))) | $spl77),
  inference(cnf_transformation,[],[f916_D])).
fof(f916_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl77),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl77])])).
fof(f988,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1))) | $spl46),
  inference(unit_resulting_resolution,[],[f438,f100])).
fof(f989,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1))) | $spl46),
  inference(unit_resulting_resolution,[],[f438,f102])).
fof(f994,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1))) | $spl46),
  inference(unit_resulting_resolution,[],[f93,f438,f145])).
fof(f995,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1))) | $spl46),
  inference(unit_resulting_resolution,[],[f92,f438,f145])).
fof(f997,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1))) | $spl46),
  inference(unit_resulting_resolution,[],[f93,f438,f145])).
fof(f998,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1))) | $spl46),
  inference(unit_resulting_resolution,[],[f92,f438,f145])).
fof(f932,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) | $spl44),
  inference(superposition,[],[f138,f405])).
fof(f924,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) | ($spl5 | $spl44)),
  inference(backward_demodulation,[],[f405,f158])).
fof(f922,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))))) | p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))))))),
  introduced(theory_axiom,[])).
fof(f917,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1))) | $spl76),
  inference(cnf_transformation,[],[f917_D])).
fof(f917_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f914,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1))) | $spl47),
  inference(resolution,[],[f439,f101])).
fof(f913,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1))) | $spl47),
  inference(resolution,[],[f439,f103])).
fof(f921,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))) | $spl81),
  inference(cnf_transformation,[],[f921_D])).
fof(f921_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))) <=> ~$spl81),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl81])])).
fof(f895,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))) | $spl47),
  inference(unit_resulting_resolution,[],[f439,f439,f139])).
fof(f896,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK1)))))) | $spl47),
  inference(unit_resulting_resolution,[],[f439,f439,f139])).
fof(f919,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1))) | $spl79),
  inference(cnf_transformation,[],[f919_D])).
fof(f919_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl79),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl79])])).
fof(f899,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1))) | $spl47),
  inference(unit_resulting_resolution,[],[f439,f96])).
fof(f900,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1))) | $spl47),
  inference(unit_resulting_resolution,[],[f439,f98])).
fof(f901,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1))) | $spl47),
  inference(unit_resulting_resolution,[],[f439,f101])).
fof(f902,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1))) | $spl47),
  inference(unit_resulting_resolution,[],[f439,f103])).
fof(f907,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1))) | $spl47),
  inference(unit_resulting_resolution,[],[f94,f439,f142])).
fof(f909,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1))) | $spl47),
  inference(unit_resulting_resolution,[],[f94,f439,f142])).
fof(f653,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) | $spl71),
  inference(cnf_transformation,[],[f653_D])).
fof(f653_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) <=> ~$spl71),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl71])])).
fof(f651,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))))) | $spl69),
  inference(cnf_transformation,[],[f651_D])).
fof(f651_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))))) <=> ~$spl69),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl69])])).
fof(f699,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) | $spl56),
  inference(superposition,[],[f138,f468])).
fof(f468,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))))) | $spl56),
  inference(cnf_transformation,[],[f468_D])).
fof(f468_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))))) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f876,plain,(
  $false | ($spl33 | $spl48)),
  inference(global_subsumption,[],[f91,f203,f441,f858,f856,f855,f853])).
fof(f853,plain,(
  ~p(s(bool,'T_0')) | ($spl33 | $spl48)),
  inference(unit_resulting_resolution,[],[f203,f441,f145])).
fof(f855,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl48),
  inference(unit_resulting_resolution,[],[f93,f441,f145])).
fof(f856,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl48),
  inference(unit_resulting_resolution,[],[f92,f441,f145])).
fof(f858,plain,(
  ~p(s(bool,'T_0')) | ($spl33 | $spl48)),
  inference(unit_resulting_resolution,[],[f203,f441,f145])).
fof(f859,plain,(
  $false | ($spl33 | $spl48)),
  inference(unit_resulting_resolution,[],[f93,f203,f441,f145])).
fof(f854,plain,(
  $false | ($spl33 | $spl48)),
  inference(unit_resulting_resolution,[],[f93,f203,f441,f145])).
fof(f875,plain,(
  $false | $spl48),
  inference(trivial_inequality_removal,[],[f845])).
fof(f845,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl48),
  inference(unit_resulting_resolution,[],[f441,f102])).
fof(f874,plain,(
  $false | $spl48),
  inference(trivial_inequality_removal,[],[f844])).
fof(f844,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl48),
  inference(unit_resulting_resolution,[],[f441,f100])).
fof(f843,plain,(
  $false | ($spl33 | $spl48)),
  inference(unit_resulting_resolution,[],[f203,f441,f99])).
fof(f841,plain,(
  $false | ($spl33 | $spl48)),
  inference(unit_resulting_resolution,[],[f203,f441,f97])).
fof(f839,plain,(
  $false | $spl48),
  inference(unit_resulting_resolution,[],[f94,f441,f106])).
fof(f106,plain,(
  ( ! [X0] : (~p(s(bool,'F_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f64])).
fof(f64,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | ~p(s(bool,X0)) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,'F_0')) & p(s(bool,X0)))))),
  inference(flattening,[],[f63])).
fof(f63,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) | ~p(s(bool,X0))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,'F_0')) & p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0] : ((p(s(bool,'F_0')) & p(s(bool,X0))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X8] : ((p(s(bool,'F_0')) & p(s(bool,X8))) <=> p(s(bool,'F_0')))),
  file('thm.pred_set.INTER_OVER_UNION','thm.bool.AND_CLAUSES|split|2')).
fof(f837,plain,(
  $false | $spl48),
  inference(unit_resulting_resolution,[],[f94,f441,f108])).
fof(f108,plain,(
  ( ! [X0] : (~p(s(bool,'F_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f66])).
fof(f66,plain,(
  ! [X0] : ((~p(s(bool,X0)) | ~p(s(bool,'F_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,X0)) & p(s(bool,'F_0')))))),
  inference(flattening,[],[f65])).
fof(f65,plain,(
  ! [X0] : (((~p(s(bool,X0)) | ~p(s(bool,'F_0'))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,X0)) & p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f36])).
fof(f36,plain,(
  ! [X0] : ((p(s(bool,X0)) & p(s(bool,'F_0'))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X8] : ((p(s(bool,X8)) & p(s(bool,'F_0'))) <=> p(s(bool,'F_0')))),
  file('thm.pred_set.INTER_OVER_UNION','thm.bool.AND_CLAUSES|split|3')).
fof(f835,plain,(
  $false | $spl48),
  inference(unit_resulting_resolution,[],[f94,f441,f119])).
fof(f119,plain,(
  ( ! [X0] : (~p(s(bool,'F_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f73])).
fof(f73,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f40])).
fof(f40,plain,(
  ! [X0] : ((p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X8] : ((p(s(bool,'F_0')) | p(s(bool,X8))) <=> p(s(bool,X8)))),
  file('thm.pred_set.INTER_OVER_UNION','thm.bool.OR_CLAUSES|split|2')).
fof(f833,plain,(
  $false | $spl48),
  inference(unit_resulting_resolution,[],[f94,f441,f123])).
fof(f123,plain,(
  ( ! [X0] : (~p(s(bool,'F_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0] : (((~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | p(s(bool,X0))) & (~p(s(bool,X0)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f75])).
fof(f75,plain,(
  ! [X0] : (((~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f41])).
fof(f41,plain,(
  ! [X0] : ((p(s(bool,X0)) | p(s(bool,'F_0'))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X8] : ((p(s(bool,X8)) | p(s(bool,'F_0'))) <=> p(s(bool,X8)))),
  file('thm.pred_set.INTER_OVER_UNION','thm.bool.OR_CLAUSES|split|3')).
fof(f831,plain,(
  $false | $spl48),
  inference(unit_resulting_resolution,[],[f94,f441])).
fof(f864,plain,(
  $false | $spl48),
  inference(resolution,[],[f441,f94])).
fof(f184,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))))) | $spl22),
  inference(cnf_transformation,[],[f184_D])).
fof(f184_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f691,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))))) | $spl57),
  inference(trivial_inequality_removal,[],[f690])).
fof(f690,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))))) | $spl57),
  inference(superposition,[],[f469,f95])).
fof(f693,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))))))) | $spl74),
  inference(cnf_transformation,[],[f693_D])).
fof(f693_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))))))) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f685,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))))))) | $spl57),
  inference(unit_resulting_resolution,[],[f469,f103])).
fof(f686,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))))))) | $spl57),
  inference(unit_resulting_resolution,[],[f469,f101])).
fof(f687,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))))) | $spl57),
  inference(unit_resulting_resolution,[],[f469,f95])).
fof(f688,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))))))) | $spl57),
  inference(unit_resulting_resolution,[],[f94,f469,f142])).
fof(f689,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))))))) | $spl57),
  inference(unit_resulting_resolution,[],[f94,f469,f142])).
fof(f682,plain,(
  p(s(bool,'F_0')) | ($spl20 | $spl40)),
  inference(backward_demodulation,[],[f178,f401])).
fof(f178,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) | $spl20),
  inference(cnf_transformation,[],[f178_D])).
fof(f178_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f177,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))))) | $spl18),
  inference(cnf_transformation,[],[f177_D])).
fof(f177_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))))) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f678,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))))) | $spl45),
  inference(trivial_inequality_removal,[],[f677])).
fof(f677,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))))) | $spl45),
  inference(superposition,[],[f406,f95])).
fof(f680,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))))))) | $spl72),
  inference(cnf_transformation,[],[f680_D])).
fof(f680_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))))))) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f672,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))))))) | $spl45),
  inference(unit_resulting_resolution,[],[f406,f103])).
fof(f673,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))))))) | $spl45),
  inference(unit_resulting_resolution,[],[f406,f101])).
fof(f674,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))))) | $spl45),
  inference(unit_resulting_resolution,[],[f406,f95])).
fof(f675,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))))))) | $spl45),
  inference(unit_resulting_resolution,[],[f94,f406,f142])).
fof(f676,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))))))) | $spl45),
  inference(unit_resulting_resolution,[],[f94,f406,f142])).
fof(f183,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))))) | $spl23),
  inference(cnf_transformation,[],[f183_D])).
fof(f183_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))))) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f658,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))))) | ($spl9 | $spl54)),
  inference(backward_demodulation,[],[f466,f162])).
fof(f176,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))))) | $spl19),
  inference(cnf_transformation,[],[f176_D])).
fof(f176_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))))) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f655,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))))) | ($spl5 | $spl42)),
  inference(backward_demodulation,[],[f403,f158])).
fof(f654,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) | $spl70),
  inference(cnf_transformation,[],[f654_D])).
fof(f654_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f572,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) | $spl22),
  inference(superposition,[],[f136,f184])).
fof(f652,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))))) | $spl68),
  inference(cnf_transformation,[],[f652_D])).
fof(f652_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))))) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f573,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))))) | $spl22),
  inference(superposition,[],[f137,f184])).
fof(f647,plain,(
  $false | ($spl33 | $spl65)),
  inference(global_subsumption,[],[f91,f203,f554,f637,f635,f633])).
fof(f633,plain,(
  p(s(bool,'F_0')) | ($spl33 | $spl65)),
  inference(unit_resulting_resolution,[],[f203,f554,f142])).
fof(f635,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl65),
  inference(unit_resulting_resolution,[],[f94,f554,f142])).
fof(f637,plain,(
  p(s(bool,'F_0')) | ($spl33 | $spl65)),
  inference(unit_resulting_resolution,[],[f203,f554,f142])).
fof(f638,plain,(
  $false | ($spl33 | $spl65)),
  inference(unit_resulting_resolution,[],[f94,f203,f554,f142])).
fof(f634,plain,(
  $false | ($spl33 | $spl65)),
  inference(unit_resulting_resolution,[],[f94,f203,f554,f142])).
fof(f626,plain,(
  $false | ($spl33 | $spl65)),
  inference(unit_resulting_resolution,[],[f203,f554,f103])).
fof(f624,plain,(
  $false | ($spl33 | $spl65)),
  inference(unit_resulting_resolution,[],[f203,f554,f101])).
fof(f646,plain,(
  $false | $spl65),
  inference(trivial_inequality_removal,[],[f622])).
fof(f622,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl65),
  inference(unit_resulting_resolution,[],[f554,f98])).
fof(f645,plain,(
  $false | $spl65),
  inference(trivial_inequality_removal,[],[f621])).
fof(f621,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl65),
  inference(unit_resulting_resolution,[],[f554,f96])).
fof(f620,plain,(
  $false | $spl65),
  inference(unit_resulting_resolution,[],[f93,f554,f112])).
fof(f112,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f68])).
fof(f68,plain,(
  ! [X0] : ((~p(s(bool,X0)) | ~p(s(bool,'T_0')) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,X0)) & p(s(bool,'T_0')))))),
  inference(flattening,[],[f67])).
fof(f67,plain,(
  ! [X0] : (((~p(s(bool,X0)) | ~p(s(bool,'T_0'))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,X0)) & p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f37])).
fof(f37,plain,(
  ! [X0] : ((p(s(bool,X0)) & p(s(bool,'T_0'))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X8] : ((p(s(bool,X8)) & p(s(bool,'T_0'))) <=> p(s(bool,X8)))),
  file('thm.pred_set.INTER_OVER_UNION','thm.bool.AND_CLAUSES|split|1')).
fof(f618,plain,(
  $false | $spl65),
  inference(unit_resulting_resolution,[],[f93,f554,f114])).
fof(f114,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f70])).
fof(f70,plain,(
  ! [X0] : ((~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,'T_0')) & p(s(bool,X0)))))),
  inference(flattening,[],[f69])).
fof(f69,plain,(
  ! [X0] : (((~p(s(bool,'T_0')) | ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,'T_0')) & p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0] : ((p(s(bool,'T_0')) & p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X8] : ((p(s(bool,'T_0')) & p(s(bool,X8))) <=> p(s(bool,X8)))),
  file('thm.pred_set.INTER_OVER_UNION','thm.bool.AND_CLAUSES|split|0')).
fof(f616,plain,(
  $false | $spl65),
  inference(unit_resulting_resolution,[],[f93,f554,f128])).
fof(f128,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0] : (((~p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | p(s(bool,X0)) | p(s(bool,'T_0'))))),
  inference(flattening,[],[f79])).
fof(f79,plain,(
  ! [X0] : (((~p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (p(s(bool,X0)) | p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f43])).
fof(f43,plain,(
  ! [X0] : ((p(s(bool,X0)) | p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X8] : ((p(s(bool,X8)) | p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  file('thm.pred_set.INTER_OVER_UNION','thm.bool.OR_CLAUSES|split|1')).
fof(f614,plain,(
  $false | $spl65),
  inference(unit_resulting_resolution,[],[f93,f554,f132])).
fof(f132,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0] : (((~p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | p(s(bool,'T_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f81])).
fof(f81,plain,(
  ! [X0] : (((~p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (p(s(bool,'T_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f44])).
fof(f44,plain,(
  ! [X0] : ((p(s(bool,'T_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X8] : ((p(s(bool,'T_0')) | p(s(bool,X8))) <=> p(s(bool,'T_0')))),
  file('thm.pred_set.INTER_OVER_UNION','thm.bool.OR_CLAUSES|split|0')).
fof(f612,plain,(
  $false | $spl65),
  inference(unit_resulting_resolution,[],[f93,f554])).
fof(f639,plain,(
  $false | $spl65),
  inference(resolution,[],[f554,f93])).
fof(f640,plain,(
  $false | $spl65),
  inference(resolution,[],[f554,f92])).
fof(f465,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) | $spl55),
  inference(cnf_transformation,[],[f465_D])).
fof(f465_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) <=> ~$spl55),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl55])])).
fof(f601,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) | $spl22),
  inference(global_subsumption,[],[f91,f184,f571])).
fof(f571,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) | s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) = s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))) | $spl22),
  inference(superposition,[],[f135,f184])).
fof(f135,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) )),
  inference(cnf_transformation,[],[f85])).
fof(f85,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f84])).
fof(f84,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(rectify,[],[f83])).
fof(f83,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(nnf_transformation,[],[f46])).
fof(f46,plain,(
  ! [X0,X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(flattening,[],[f45])).
fof(f45,plain,(
  ! [X0] : ! [X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X6] : ! [X7,X8] : (s(fun(X6,bool),X7) = s(fun(X6,bool),X8) <=> ! [X12] : s(bool,'const.bool.IN_2'(s(X6,X12),s(fun(X6,bool),X7))) = s(bool,'const.bool.IN_2'(s(X6,X12),s(fun(X6,bool),X8))))),
  file('thm.pred_set.INTER_OVER_UNION','thm.pred_set.EXTENSION')).
fof(f556,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))))) | $spl66),
  inference(cnf_transformation,[],[f556_D])).
fof(f556_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))))) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f522,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))))) | $spl18),
  inference(superposition,[],[f137,f177])).
fof(f552,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) | $spl62),
  inference(cnf_transformation,[],[f552_D])).
fof(f552_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f521,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))))) | $spl18),
  inference(superposition,[],[f136,f177])).
fof(f402,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) | $spl43),
  inference(cnf_transformation,[],[f402_D])).
fof(f402_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) <=> ~$spl43),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl43])])).
fof(f550,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) | $spl18),
  inference(global_subsumption,[],[f91,f177,f520])).
fof(f520,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) | s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) = s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))) | $spl18),
  inference(superposition,[],[f135,f177])).
fof(f504,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))) | $spl61),
  inference(cnf_transformation,[],[f504_D])).
fof(f504_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))) <=> ~$spl61),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl61])])).
fof(f475,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))) | $spl24),
  inference(superposition,[],[f141,f185])).
fof(f185,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) | $spl24),
  inference(cnf_transformation,[],[f185_D])).
fof(f185_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f502,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1)))) | $spl59),
  inference(cnf_transformation,[],[f502_D])).
fof(f502_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1)))) <=> ~$spl59),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl59])])).
fof(f474,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1)))) | $spl24),
  inference(superposition,[],[f140,f185])).
fof(f469,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))))) | $spl57),
  inference(cnf_transformation,[],[f469_D])).
fof(f469_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))))) <=> ~$spl57),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl57])])).
fof(f467,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))))) | ($spl9 | $spl24)),
  inference(backward_demodulation,[],[f185,f162])).
fof(f466,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) | $spl54),
  inference(cnf_transformation,[],[f466_D])).
fof(f466_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f462,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) | $spl25),
  inference(trivial_inequality_removal,[],[f461])).
fof(f461,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) | $spl25),
  inference(superposition,[],[f186,f95])).
fof(f464,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))))) | $spl52),
  inference(cnf_transformation,[],[f464_D])).
fof(f464_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))))) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f456,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))))) | $spl25),
  inference(unit_resulting_resolution,[],[f186,f103])).
fof(f457,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))))) | $spl25),
  inference(unit_resulting_resolution,[],[f186,f101])).
fof(f458,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) | $spl25),
  inference(unit_resulting_resolution,[],[f186,f95])).
fof(f459,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))))) | $spl25),
  inference(unit_resulting_resolution,[],[f94,f186,f142])).
fof(f460,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))))) | $spl25),
  inference(unit_resulting_resolution,[],[f94,f186,f142])).
fof(f443,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))) | $spl51),
  inference(cnf_transformation,[],[f443_D])).
fof(f443_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))) <=> ~$spl51),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl51])])).
fof(f412,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))) | $spl20),
  inference(superposition,[],[f141,f178])).
fof(f439,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1)))) | $spl47),
  inference(cnf_transformation,[],[f439_D])).
fof(f439_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1)))) <=> ~$spl47),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl47])])).
fof(f411,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),sK1)))) | $spl20),
  inference(superposition,[],[f140,f178])).
fof(f406,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))))) | $spl45),
  inference(cnf_transformation,[],[f406_D])).
fof(f406_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))))) <=> ~$spl45),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl45])])).
fof(f404,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))))) | ($spl5 | $spl20)),
  inference(backward_demodulation,[],[f178,f158])).
fof(f403,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) | $spl42),
  inference(cnf_transformation,[],[f403_D])).
fof(f403_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f399,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) | $spl21),
  inference(trivial_inequality_removal,[],[f398])).
fof(f398,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) | $spl21),
  inference(superposition,[],[f179,f95])).
fof(f401,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))))) | $spl40),
  inference(cnf_transformation,[],[f401_D])).
fof(f401_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))))) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f393,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))))) | $spl21),
  inference(unit_resulting_resolution,[],[f179,f103])).
fof(f394,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))))) | $spl21),
  inference(unit_resulting_resolution,[],[f179,f101])).
fof(f395,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) | $spl21),
  inference(unit_resulting_resolution,[],[f179,f95])).
fof(f396,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))))) | $spl21),
  inference(unit_resulting_resolution,[],[f94,f179,f142])).
fof(f397,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))))))) | $spl21),
  inference(unit_resulting_resolution,[],[f94,f179,f142])).
fof(f306,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),bool,sK0)))) | $spl37),
  inference(cnf_transformation,[],[f306_D])).
fof(f306_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),bool,sK0)))) <=> ~$spl37),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl37])])).
fof(f345,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),bool,sK0)))) | $spl14),
  inference(global_subsumption,[],[f91,f170,f331])).
fof(f331,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),bool,sK0)))) | s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) = s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))) | $spl14),
  inference(superposition,[],[f146,f170])).
fof(f146,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK5(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK5(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f90])).
fof(f90,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK5(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK5(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5])],[f57])).
fof(f57,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f52])).
fof(f52,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.pred_set.INTER_OVER_UNION',aHL_EXT)).
fof(f310,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),bool,sK0)))) | $spl39),
  inference(cnf_transformation,[],[f310_D])).
fof(f310_D,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),bool,sK0)))) <=> ~$spl39),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl39])])).
fof(f308,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),bool,sK0)))) | ($spl7 | $spl16)),
  inference(backward_demodulation,[],[f171,f160])).
fof(f171,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),bool,sK0)))) | $spl16),
  inference(cnf_transformation,[],[f171_D])).
fof(f171_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),bool,sK0)))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f307,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),bool,sK0)))) | $spl36),
  inference(cnf_transformation,[],[f307_D])).
fof(f307_D,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),bool,sK0)))) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f303,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),bool,sK0)))) | $spl17),
  inference(trivial_inequality_removal,[],[f302])).
fof(f302,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),bool,sK0)))) | $spl17),
  inference(superposition,[],[f172,f95])).
fof(f305,plain,(
  p(s(bool,f(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),bool,sK0))))) | $spl34),
  inference(cnf_transformation,[],[f305_D])).
fof(f305_D,plain,(
  p(s(bool,f(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),bool,sK0))))) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f297,plain,(
  p(s(bool,f(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),bool,sK0))))) | $spl17),
  inference(unit_resulting_resolution,[],[f172,f103])).
fof(f298,plain,(
  p(s(bool,f(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),bool,sK0))))) | $spl17),
  inference(unit_resulting_resolution,[],[f172,f101])).
fof(f299,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),bool,sK0)))) | $spl17),
  inference(unit_resulting_resolution,[],[f172,f95])).
fof(f300,plain,(
  p(s(bool,f(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),bool,sK0))))) | $spl17),
  inference(unit_resulting_resolution,[],[f94,f172,f142])).
fof(f301,plain,(
  p(s(bool,f(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),bool,sK0))))) | $spl17),
  inference(unit_resulting_resolution,[],[f94,f172,f142])).
fof(f196,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),bool,sK0)))) | $spl29),
  inference(cnf_transformation,[],[f196_D])).
fof(f196_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),bool,sK0)))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f249,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),bool,sK0)))) | $spl10),
  inference(global_subsumption,[],[f91,f165,f235])).
fof(f235,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),bool,sK0)))) | s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) = s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))) | $spl10),
  inference(superposition,[],[f146,f165])).
fof(f203,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl33),
  inference(cnf_transformation,[],[f203_D])).
fof(f203_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f201,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ($spl10 | $spl31)),
  inference(forward_demodulation,[],[f200,f165])).
fof(f200,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),bool,sK0)))) | $spl31),
  inference(cnf_transformation,[],[f200_D])).
fof(f200_D,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),bool,sK0)))) <=> ~$spl31),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl31])])).
fof(f198,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),bool,sK0)))) | ($spl3 | $spl12)),
  inference(backward_demodulation,[],[f166,f156])).
fof(f166,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),bool,sK0)))) | $spl12),
  inference(cnf_transformation,[],[f166_D])).
fof(f166_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),bool,sK0)))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f197,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),bool,sK0)))) | $spl28),
  inference(cnf_transformation,[],[f197_D])).
fof(f197_D,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),bool,sK0)))) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f193,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),bool,sK0)))) | $spl13),
  inference(trivial_inequality_removal,[],[f192])).
fof(f192,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),bool,sK0)))) | $spl13),
  inference(superposition,[],[f167,f95])).
fof(f195,plain,(
  p(s(bool,f(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),bool,sK0))))) | $spl26),
  inference(cnf_transformation,[],[f195_D])).
fof(f195_D,plain,(
  p(s(bool,f(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),bool,sK0))))) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f187,plain,(
  p(s(bool,f(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),bool,sK0))))) | $spl13),
  inference(unit_resulting_resolution,[],[f167,f103])).
fof(f188,plain,(
  p(s(bool,f(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),bool,sK0))))) | $spl13),
  inference(unit_resulting_resolution,[],[f167,f101])).
fof(f189,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),bool,sK0)))) | $spl13),
  inference(unit_resulting_resolution,[],[f167,f95])).
fof(f190,plain,(
  p(s(bool,f(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),bool,sK0))))) | $spl13),
  inference(unit_resulting_resolution,[],[f94,f167,f142])).
fof(f191,plain,(
  p(s(bool,f(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),bool,sK0))))) | $spl13),
  inference(unit_resulting_resolution,[],[f94,f167,f142])).
fof(f186,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) | $spl25),
  inference(cnf_transformation,[],[f186_D])).
fof(f186_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f182,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))))) | $spl9),
  inference(superposition,[],[f162,f95])).
fof(f175,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))))) | $spl5),
  inference(superposition,[],[f158,f95])).
fof(f170,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),bool,sK0)))) | $spl14),
  inference(cnf_transformation,[],[f170_D])).
fof(f170_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),bool,sK0)))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f172,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),bool,sK0)))) | $spl17),
  inference(cnf_transformation,[],[f172_D])).
fof(f172_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),bool,sK0)))) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f168,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),bool,sK0)))) | $spl7),
  inference(superposition,[],[f160,f95])).
fof(f165,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),bool,sK0)))) | $spl10),
  inference(cnf_transformation,[],[f165_D])).
fof(f165_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),bool,sK0)))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f167,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),bool,sK0)))) | $spl13),
  inference(cnf_transformation,[],[f167_D])).
fof(f167_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),bool,sK0)))) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f163,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),bool,sK0)))) | $spl3),
  inference(superposition,[],[f156,f95])).
fof(f162,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))))) | $spl9),
  inference(cnf_transformation,[],[f162_D])).
fof(f162_D,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f151,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))))) | $spl1),
  inference(unit_resulting_resolution,[],[f150,f135])).
fof(f160,plain,(
  s(bool,f(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),bool,sK0)))) != s(bool,g(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),bool,sK0)))) | $spl7),
  inference(cnf_transformation,[],[f160_D])).
fof(f160_D,plain,(
  s(bool,f(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),bool,sK0)))) != s(bool,g(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),bool,sK0)))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f152,plain,(
  s(bool,f(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),bool,sK0)))) != s(bool,g(s(sK0,sK5('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),bool,sK0)))) | $spl1),
  inference(unit_resulting_resolution,[],[f150,f146])).
fof(f158,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))))) | $spl5),
  inference(cnf_transformation,[],[f158_D])).
fof(f158_D,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f153,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))))) | $spl1),
  inference(unit_resulting_resolution,[],[f150,f135])).
fof(f156,plain,(
  s(bool,f(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),bool,sK0)))) != s(bool,g(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),bool,sK0)))) | $spl3),
  inference(cnf_transformation,[],[f156_D])).
fof(f156_D,plain,(
  s(bool,f(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),bool,sK0)))) != s(bool,g(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),bool,sK0)))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f154,plain,(
  s(bool,f(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),bool,sK0)))) != s(bool,g(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),bool,sK0)))) | $spl1),
  inference(unit_resulting_resolution,[],[f150,f146])).
fof(f150,plain,(
  s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))) | $spl1),
  inference(cnf_transformation,[],[f150_D])).
fof(f150_D,plain,(
  s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f91,plain,(
  s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))))),
  inference(cnf_transformation,[],[f58])).
fof(f58,plain,(
  s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f54])).
fof(f54,plain,(
  ? [X0,X1,X2,X3] : s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))) != s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X3)))))),
  inference(ennf_transformation,[],[f27])).
fof(f27,plain,(
  ~! [X0,X1,X2,X3] : s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))) = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X3)))))),
  inference(flattening,[],[f26])).
fof(f26,plain,(
  ~! [X0] : ! [X1,X2,X3] : s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))) = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X3)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),X7),s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X9))))) = s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8))),s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X9)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),X7),s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X9))))) = s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8))),s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X9)))))),
  file('thm.pred_set.INTER_OVER_UNION','thm.pred_set.INTER_OVER_UNION')).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 120
% Passive clauses: 258
% Generated clauses: 2880
% Final active clauses: 55
% Final passive clauses: 46
% Input formulas: 24
% Initial clauses: 56
% 
% Duplicate literals: 8
% Trivial inequalities: 80
% Fw demodulations: 629
% Bw demodulations: 40
% Global subsumptions: 23
% 
% Simple tautologies: 110
% Equational tautologies: 62
% Forward subsumptions: 1248
% Backward subsumptions: 1
% 
% Binary resolution: 198
% Unit resulting resolution: 728
% Forward superposition: 82
% Backward superposition: 839
% 
% Split clauses: 28
% Split components: 61
% SAT solver clauses: 1203
% SAT solver unit clauses: 19
% SAT solver binary clauses: 1034
% 
% Sat splits: 28
% Sat splitting refutations: 27
% 
% Memory used [KB]: 2430
% Time elapsed: 0.096 s
% ------------------------------
% ------------------------------
% Success in time 18.133 s
