Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.integral
% (47527)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 995
% Passive clauses: 9649
% Generated clauses: 28546
% Final active clauses: 889
% Final passive clauses: 8051
% Input formulas: 48
% Initial clauses: 200
% 
% Duplicate literals: 366
% Trivial inequalities: 211
% Fw subsumption resolutions: 3214
% Fw demodulations: 1925
% Bw demodulations: 104
% 
% Simple tautologies: 246
% Equational tautologies: 40
% Forward subsumptions: 11887
% 
% Binary resolution: 9891
% Factoring: 952
% Forward superposition: 2913
% Backward superposition: 8317
% Self superposition: 29
% Equality factoring: 26
% Equality resolution: 12
% 
% Split clauses: 979
% Split components: 1966
% SAT solver clauses: 1419
% SAT solver unit clauses: 10
% SAT solver binary clauses: 187
% 
% Sat splits: 979
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 1417
% TWLsolver calls for satisfiability: 506
% 
% Memory used [KB]: 13048
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 294 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.integral
% (47672)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 524
% Passive clauses: 3616
% Generated clauses: 15177
% Final active clauses: 481
% Final passive clauses: 3081
% Input formulas: 48
% Initial clauses: 200
% 
% Duplicate literals: 202
% Trivial inequalities: 48
% Fw subsumption resolutions: 1038
% Fw demodulations: 731
% Bw demodulations: 19
% Condensations: 383
% Global subsumptions: 13
% 
% Simple tautologies: 260
% Equational tautologies: 21
% Forward subsumptions: 4582
% 
% Binary resolution: 5276
% Factoring: 240
% Forward superposition: 2025
% Backward superposition: 1283
% Self superposition: 26
% Equality factoring: 26
% Equality resolution: 6
% 
% Split clauses: 115
% Split components: 232
% SAT solver clauses: 4388
% SAT solver unit clauses: 193
% SAT solver binary clauses: 83
% SAT solver learnt clauses: 26
% SAT solver learnt literals: 10
% 
% Sat splits: 115
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 4336
% TWLsolver calls for satisfiability: 1924
% 
% Memory used [KB]: 16758
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.integral
% (47743)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 4717
% Passive clauses: 21525
% Generated clauses: 140431
% Final active clauses: 2975
% Final passive clauses: 6782
% Input formulas: 48
% Initial clauses: 200
% 
% Duplicate literals: 3234
% Trivial inequalities: 79
% Fw subsumption resolutions: 1692
% Fw demodulations: 31277
% Bw demodulations: 1181
% Condensations: 1039
% Global subsumptions: 41
% 
% Simple tautologies: 139
% Equational tautologies: 433
% Forward subsumptions: 79818
% 
% Unit resulting resolution: 72165
% Factoring: 1
% Forward superposition: 21114
% Backward superposition: 7759
% Self superposition: 49
% Equality factoring: 42
% Equality resolution: 2
% 
% Split clauses: 178
% Split components: 363
% SAT solver clauses: 22886
% SAT solver unit clauses: 2534
% SAT solver binary clauses: 9089
% 
% Sat splits: 178
% Sat splitting refutations: 961
% 
% Memory used [KB]: 52707
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.integral
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.integral
% SZS output start Proof for thm.integral
fof(f22546,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f261,f469,f467,f465,f463,f851,f1396,f2343,f2345,f2473,f2469,f2441,f2520,f2521,f2522,f2523,f2996,f2997,f3003,f3129,f3137,f3130,f3138,f3131,f3139,f3132,f3305,f3346,f3347,f3348,f3611,f3612,f3613,f3749,f3747,f3746,f3745,f3721,f3713,f3712,f3711,f3811,f3812,f3813,f3814,f3815,f3816,f4092,f3310,f4131,f3354,f4132,f3357,f4133,f3356,f4317,f4320,f4318,f4559,f4566,f4561,f6704,f8960,f8961,f9265,f9275,f9418,f9419,f9687,f9699,f9727,f9745,f9873,f10131,f10200,f10202,f4567,f10501,f10513,f10505,f10135,f10535,f3141,f10559,f10543,f10565,f11772,f11774,f11773,f11836,f11839,f11840,f11841,f11837,f11842,f11843,f11844,f11847,f11848,f11845,f11849,f11850,f11851,f11852,f11919,f11918,f11922,f11923,f12090,f12175,f12176,f12177,f12178,f12222,f12223,f12224,f12225,f12226,f12227,f12228,f12279,f12280,f12281,f12388,f12389,f12390,f12391,f12392,f12387,f12393,f12733,f12735,f12736,f12737,f12734,f12738,f12739,f12740,f11846,f12742,f12743,f11838,f12741,f12744,f12745,f12746,f12747,f12730,f12729,f12750,f12748,f13643,f13644,f13645,f13648,f13649,f13646,f13652,f13653,f13654,f13655,f13656,f13657,f13658,f13659,f13660,f13661,f13634,f13633,f13663,f13662,f13664,f13985,f13989,f14369,f14370,f14371,f14373,f14374,f13647,f14372,f14377,f14378,f14379,f14380,f14382,f14383,f14384,f14385,f14483,f14482,f14486,f14487,f19958,f19961,f19962,f19963,f19977,f19964,f20130,f20131,f20132,f20133,f20145,f19980,f20146,f20140,f20147,f20139,f20148,f19979,f20149,f20138,f20151,f19976,f20967,f20968,f20969,f20970,f21125,f21126,f21375,f21882,f22340,f22345,f10210,f22348,f22372,f22349,f22350,f22353,f22354,f22355,f3618,f22356,f22357,f22408,f22417,f22411,f22418,f22414,f22423,f22373,f22424,f22374,f22425,f3309,f22426,f22433,f3002,f22474,f22477,f22478,f4319,f22479,f22489,f22480,f3000,f22524,f22525,f22526,f22527,f22537,f22490,f22538,f22532,f22539,f22531,f22540,f3355,f22541,f22530,f22542,f2344,f3353,f22545])).
fof(f22545,plain,(
  $false | ($spl0 | $spl5)),
  inference(subsumption_resolution,[],[f19568,f467])).
fof(f19568,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16)))))) | $spl0),
  inference(resolution,[],[f852,f580])).
fof(f580,plain,(
  ( ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK12),s('type.num.num',X0))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',X0))))))) )),
  inference(resolution,[],[f260,f263])).
fof(f263,plain,(
  ( ! [X2,X0,X5,X3,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num',X5))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X3),s('type.num.num',X5)))))) | ~p(s(bool,'const.transc.tdiv_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.pair.prod'(fun('type.num.num','type.realax.real'),fun('type.num.num','type.realax.real')),'const.pair.,_2'(s(fun('type.num.num','type.realax.real'),X2),s(fun('type.num.num','type.realax.real'),X3))))))) )),
  inference(cnf_transformation,[],[f165])).
fof(f165,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.transc.tdiv_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.pair.prod'(fun('type.num.num','type.realax.real'),fun('type.num.num','type.realax.real')),'const.pair.,_2'(s(fun('type.num.num','type.realax.real'),X2),s(fun('type.num.num','type.realax.real'),X3)))))) | (p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s(fun('type.num.num','type.realax.real'),X2)))) & ! [X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num',X5))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X3),s('type.num.num',X5)))))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X3),s('type.num.num',X5))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))))))))))) & (~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s(fun('type.num.num','type.realax.real'),X2)))) | (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num',sK17(X3,X2,X1,X0)))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X3),s('type.num.num',sK17(X3,X2,X1,X0))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X3),s('type.num.num',sK17(X3,X2,X1,X0)))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK17(X3,X2,X1,X0)))))))))) | p(s(bool,'const.transc.tdiv_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.pair.prod'(fun('type.num.num','type.realax.real'),fun('type.num.num','type.realax.real')),'const.pair.,_2'(s(fun('type.num.num','type.realax.real'),X2),s(fun('type.num.num','type.realax.real'),X3))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK17])],[f164])).
fof(f164,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.transc.tdiv_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.pair.prod'(fun('type.num.num','type.realax.real'),fun('type.num.num','type.realax.real')),'const.pair.,_2'(s(fun('type.num.num','type.realax.real'),X2),s(fun('type.num.num','type.realax.real'),X3)))))) | (p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s(fun('type.num.num','type.realax.real'),X2)))) & ! [X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num',X5))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X3),s('type.num.num',X5)))))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X3),s('type.num.num',X5))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))))))))))) & (~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s(fun('type.num.num','type.realax.real'),X2)))) | ? [X4] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num',X4))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X3),s('type.num.num',X4)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X3),s('type.num.num',X4))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))))))))) | p(s(bool,'const.transc.tdiv_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.pair.prod'(fun('type.num.num','type.realax.real'),fun('type.num.num','type.realax.real')),'const.pair.,_2'(s(fun('type.num.num','type.realax.real'),X2),s(fun('type.num.num','type.realax.real'),X3))))))))),
  inference(rectify,[],[f163])).
fof(f163,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.transc.tdiv_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.pair.prod'(fun('type.num.num','type.realax.real'),fun('type.num.num','type.realax.real')),'const.pair.,_2'(s(fun('type.num.num','type.realax.real'),X2),s(fun('type.num.num','type.realax.real'),X3)))))) | (p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s(fun('type.num.num','type.realax.real'),X2)))) & ! [X4] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num',X4))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X3),s('type.num.num',X4)))))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X3),s('type.num.num',X4))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))))))))))) & (~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s(fun('type.num.num','type.realax.real'),X2)))) | ? [X4] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num',X4))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X3),s('type.num.num',X4)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X3),s('type.num.num',X4))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))))))))) | p(s(bool,'const.transc.tdiv_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.pair.prod'(fun('type.num.num','type.realax.real'),fun('type.num.num','type.realax.real')),'const.pair.,_2'(s(fun('type.num.num','type.realax.real'),X2),s(fun('type.num.num','type.realax.real'),X3))))))))),
  inference(flattening,[],[f162])).
fof(f162,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.transc.tdiv_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.pair.prod'(fun('type.num.num','type.realax.real'),fun('type.num.num','type.realax.real')),'const.pair.,_2'(s(fun('type.num.num','type.realax.real'),X2),s(fun('type.num.num','type.realax.real'),X3)))))) | (p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s(fun('type.num.num','type.realax.real'),X2)))) & ! [X4] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num',X4))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X3),s('type.num.num',X4)))))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X3),s('type.num.num',X4))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))))))))))) & ((~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s(fun('type.num.num','type.realax.real'),X2)))) | ? [X4] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num',X4))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X3),s('type.num.num',X4)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X3),s('type.num.num',X4))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X4)))))))))) | p(s(bool,'const.transc.tdiv_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.pair.prod'(fun('type.num.num','type.realax.real'),fun('type.num.num','type.realax.real')),'const.pair.,_2'(s(fun('type.num.num','type.realax.real'),X2),s(fun('type.num.num','type.realax.real'),X3))))))))),
  inference(nnf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.transc.tdiv_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.pair.prod'(fun('type.num.num','type.realax.real'),fun('type.num.num','type.realax.real')),'const.pair.,_2'(s(fun('type.num.num','type.realax.real'),X2),s(fun('type.num.num','type.realax.real'),X3)))))) <=> (p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s(fun('type.num.num','type.realax.real'),X2)))) & ! [X4] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num',X4))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X3),s('type.num.num',X4)))))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X3),s('type.num.num',X4))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X4)))))))))))),
  inference(rectify,[],[f48])).
fof(f48,axiom,(
  ! [X8,X9,X24,X7] : (p(s(bool,'const.transc.tdiv_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X8),s('type.realax.real',X9))),s('type.pair.prod'(fun('type.num.num','type.realax.real'),fun('type.num.num','type.realax.real')),'const.pair.,_2'(s(fun('type.num.num','type.realax.real'),X24),s(fun('type.num.num','type.realax.real'),X7)))))) <=> (p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X8),s('type.realax.real',X9))),s(fun('type.num.num','type.realax.real'),X24)))) & ! [X10] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X24),s('type.num.num',X10))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X7),s('type.num.num',X10)))))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X7),s('type.num.num',X10))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X24),s('type.num.num','const.num.SUC_1'(s('type.num.num',X10)))))))))))),
  file('thm.integral.TDIV_BOUNDS','thm.transc.tdiv')).
fof(f260,plain,(
  p(s(bool,'const.transc.tdiv_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK14),s('type.realax.real',sK15))),s('type.pair.prod'(fun('type.num.num','type.realax.real'),fun('type.num.num','type.realax.real')),'const.pair.,_2'(s(fun('type.num.num','type.realax.real'),sK12),s(fun('type.num.num','type.realax.real'),sK13))))))),
  inference(cnf_transformation,[],[f161])).
fof(f161,plain,(
  p(s(bool,'const.transc.tdiv_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK14),s('type.realax.real',sK15))),s('type.pair.prod'(fun('type.num.num','type.realax.real'),fun('type.num.num','type.realax.real')),'const.pair.,_2'(s(fun('type.num.num','type.realax.real'),sK12),s(fun('type.num.num','type.realax.real'),sK13)))))) & (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK12),s('type.num.num',sK16)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK12),s('type.num.num',sK16))),s('type.realax.real',sK15)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13,sK14,sK15,sK16])],[f117])).
fof(f117,plain,(
  ? [X0,X1,X2,X3] : (p(s(bool,'const.transc.tdiv_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X2),s('type.realax.real',X3))),s('type.pair.prod'(fun('type.num.num','type.realax.real'),fun('type.num.num','type.realax.real')),'const.pair.,_2'(s(fun('type.num.num','type.realax.real'),X0),s(fun('type.num.num','type.realax.real'),X1)))))) & ? [X4] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X4)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X4))),s('type.realax.real',X3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X1),s('type.num.num',X4)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X1),s('type.num.num',X4))),s('type.realax.real',X3))))))),
  inference(ennf_transformation,[],[f50])).
fof(f50,plain,(
  ~! [X0,X1,X2,X3] : (p(s(bool,'const.transc.tdiv_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X2),s('type.realax.real',X3))),s('type.pair.prod'(fun('type.num.num','type.realax.real'),fun('type.num.num','type.realax.real')),'const.pair.,_2'(s(fun('type.num.num','type.realax.real'),X0),s(fun('type.num.num','type.realax.real'),X1)))))) => ! [X4] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X4)))))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X4))),s('type.realax.real',X3)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X1),s('type.num.num',X4)))))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X1),s('type.num.num',X4))),s('type.realax.real',X3))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8,X9] : (p(s(bool,'const.transc.tdiv_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X8),s('type.realax.real',X9))),s('type.pair.prod'(fun('type.num.num','type.realax.real'),fun('type.num.num','type.realax.real')),'const.pair.,_2'(s(fun('type.num.num','type.realax.real'),X6),s(fun('type.num.num','type.realax.real'),X7)))))) => ! [X10] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X8),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X6),s('type.num.num',X10)))))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X6),s('type.num.num',X10))),s('type.realax.real',X9)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X8),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X7),s('type.num.num',X10)))))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X7),s('type.num.num',X10))),s('type.realax.real',X9))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8,X9] : (p(s(bool,'const.transc.tdiv_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X8),s('type.realax.real',X9))),s('type.pair.prod'(fun('type.num.num','type.realax.real'),fun('type.num.num','type.realax.real')),'const.pair.,_2'(s(fun('type.num.num','type.realax.real'),X6),s(fun('type.num.num','type.realax.real'),X7)))))) => ! [X10] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X8),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X6),s('type.num.num',X10)))))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X6),s('type.num.num',X10))),s('type.realax.real',X9)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X8),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X7),s('type.num.num',X10)))))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X7),s('type.num.num',X10))),s('type.realax.real',X9))))))),
  file('thm.integral.TDIV_BOUNDS','thm.integral.TDIV_BOUNDS')).
fof(f852,plain,(
  ( ! [X0] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK12),s('type.num.num',sK16))),s('type.realax.real',X0)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',X0))))) ) | $spl0),
  inference(resolution,[],[f462,f266])).
fof(f266,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X2)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1))))) )),
  inference(cnf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X2)))))),
  inference(flattening,[],[f118])).
fof(f118,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X2)))))),
  inference(ennf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0,X1,X2] : ((p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2))))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X2)))))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X15,X16,X21] : ((p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X15),s('type.realax.real',X16)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X16),s('type.realax.real',X21))))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X15),s('type.realax.real',X21)))))),
  file('thm.integral.TDIV_BOUNDS','thm.real.REAL_LE_TRANS')).
fof(f462,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK12),s('type.num.num',sK16)))))) | $spl0),
  inference(cnf_transformation,[],[f462_D])).
fof(f462_D,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK12),s('type.num.num',sK16)))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f3353,plain,(
  s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15))) | $spl29),
  inference(cnf_transformation,[],[f3353_D])).
fof(f3353_D,plain,(
  s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f2344,plain,(
  ( ! [X3] : (p(s(bool,X3))) ) | $spl8),
  inference(cnf_transformation,[],[f2344_D])).
fof(f2344_D,plain,(
  ( ! [X3] : (p(s(bool,X3))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f22542,plain,(
  ( ! [X18] : (p(s(bool,X18)) | s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15)))) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f3527,f2950])).
fof(f2950,plain,(
  ( ! [X4,X5] : (~sP6(X5,'F_0',X4) | p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5)) ) | $spl0),
  inference(subsumption_resolution,[],[f2944,f348])).
fof(f348,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f213])).
fof(f213,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f212])).
fof(f212,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f211])).
fof(f211,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f154])).
fof(f154,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f2944,plain,(
  ( ! [X4,X5] : (~sP8(X4,X5,'F_0') | ~sP6(X5,'F_0',X4) | p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5)) ) | $spl0),
  inference(resolution,[],[f2921,f360])).
fof(f360,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | s(bool,X0) != s(bool,X1) | ~sP9(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f219])).
fof(f219,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP9(X2,X1,X0)) & (~sP9(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f156])).
fof(f156,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP9(X2,X1,X0))),
  inference(definition_folding,[],[f71,f155,f154,f153,f152])).
fof(f152,plain,(
  ! [X0,X1,X2] : (sP6(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f153,plain,(
  ! [X1,X0,X2] : (sP7(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f155,plain,(
  ! [X2,X1,X0] : (sP9(X2,X1,X0) <=> (sP6(X0,X1,X2) & sP7(X1,X0,X2) & sP8(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f71,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X22,X23,X7] : ((p(s(bool,X7)) <=> s(bool,X22) = s(bool,X23)) <=> ((p(s(bool,X7)) | p(s(bool,X23)) | p(s(bool,X22))) & (p(s(bool,X7)) | ~p(s(bool,X22)) | ~p(s(bool,X23))) & (p(s(bool,X23)) | ~p(s(bool,X22)) | ~p(s(bool,X7))) & (p(s(bool,X22)) | ~p(s(bool,X23)) | ~p(s(bool,X7)))))),
  file('thm.integral.TDIV_BOUNDS','thm.sat.sat_2|49|_prove_ltb')).
fof(f2921,plain,(
  ( ! [X4,X5] : (sP9(X4,'F_0',X5) | ~sP8(X4,X5,'F_0') | ~sP6(X5,'F_0',X4)) ) | $spl0),
  inference(subsumption_resolution,[],[f2919,f447])).
fof(f447,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f112])).
fof(f112,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.integral.TDIV_BOUNDS',aHL_FALSITY)).
fof(f2919,plain,(
  ( ! [X4,X5] : (sP9(X4,'F_0',X5) | p(s(bool,'F_0')) | ~sP8(X4,X5,'F_0') | ~sP6(X5,'F_0',X4)) ) | $spl0),
  inference(resolution,[],[f2916,f343])).
fof(f343,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | p(s(bool,X1)) | ~sP8(X0,X2,X1) | ~sP7(X1,X2,X0) | ~sP6(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f210])).
fof(f210,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | (sP6(X2,X1,X0) & sP7(X1,X2,X0) & sP8(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP6(X2,X1,X0) | ~sP7(X1,X2,X0) | ~sP8(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f209])).
fof(f209,plain,(
  ! [X2,X1,X0] : ((~sP9(X2,X1,X0) | (sP6(X0,X1,X2) & sP7(X1,X0,X2) & sP8(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP6(X0,X1,X2) | ~sP7(X1,X0,X2) | ~sP8(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP9(X2,X1,X0)))),
  inference(flattening,[],[f208])).
fof(f208,plain,(
  ! [X2,X1,X0] : ((~sP9(X2,X1,X0) | (sP6(X0,X1,X2) & sP7(X1,X0,X2) & sP8(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP6(X0,X1,X2) | ~sP7(X1,X0,X2) | ~sP8(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP9(X2,X1,X0)))),
  inference(nnf_transformation,[],[f155])).
fof(f2916,plain,(
  ( ! [X4,X5] : (sP7('F_0',X4,X5)) ) | $spl0),
  inference(subsumption_resolution,[],[f2910,f350])).
fof(f350,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f216])).
fof(f216,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f215])).
fof(f215,plain,(
  ! [X1,X0,X2] : ((~sP7(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP7(X1,X0,X2)))),
  inference(flattening,[],[f214])).
fof(f214,plain,(
  ! [X1,X0,X2] : ((~sP7(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP7(X1,X0,X2)))),
  inference(nnf_transformation,[],[f153])).
fof(f2910,plain,(
  ( ! [X4,X5] : (sP7('F_0',X4,X5) | p(s(bool,X5))) ) | $spl0),
  inference(resolution,[],[f2907,f348])).
fof(f2907,plain,(
  ( ! [X6,X7] : (~sP8(X7,X6,'F_0') | sP7('F_0',X6,X7)) ) | $spl0),
  inference(subsumption_resolution,[],[f2903,f351])).
fof(f351,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f216])).
fof(f2903,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | ~sP8(X7,X6,'F_0') | sP7('F_0',X6,X7)) ) | $spl0),
  inference(resolution,[],[f2897,f339])).
fof(f339,plain,(
  ( ! [X2,X0,X1] : (sP7(X1,X2,X0) | ~sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f210])).
fof(f2897,plain,(
  ( ! [X2,X3] : (sP9(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP8(X3,X2,'F_0')) ) | $spl0),
  inference(subsumption_resolution,[],[f2896,f436])).
fof(f436,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f255])).
fof(f255,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X14] : (s(bool,'T_0') = s(bool,X14) <=> p(s(bool,X14)))),
  file('thm.integral.TDIV_BOUNDS','thm.bool.EQ_CLAUSES|split|0')).
fof(f2896,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | sP9(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP8(X3,X2,'F_0')) ) | $spl0),
  inference(subsumption_resolution,[],[f2893,f2586])).
fof(f2586,plain,(
  ( ! [X4,X5] : (sP6(X4,'F_0',X5) | s(bool,'T_0') != s(bool,X4)) ) | $spl0),
  inference(subsumption_resolution,[],[f2581,f354])).
fof(f354,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f218])).
fof(f218,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(flattening,[],[f217])).
fof(f217,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(nnf_transformation,[],[f152])).
fof(f2581,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | p(s(bool,X5)) | sP6(X4,'F_0',X5)) ) | $spl0),
  inference(resolution,[],[f2322,f338])).
fof(f338,plain,(
  ( ! [X2,X0,X1] : (sP6(X2,X1,X0) | ~sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f210])).
fof(f2322,plain,(
  ( ! [X6,X5] : (sP9(X6,'F_0',X5) | s(bool,'T_0') != s(bool,X5) | p(s(bool,X6))) ) | $spl0),
  inference(superposition,[],[f2315,f357])).
fof(f357,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f219])).
fof(f2315,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl0),
  inference(subsumption_resolution,[],[f2308,f462])).
fof(f2308,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK12),s('type.num.num',sK16)))))) | $spl0),
  inference(superposition,[],[f932,f436])).
fof(f932,plain,(
  s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK12),s('type.num.num',sK16))))) | $spl0),
  inference(resolution,[],[f462,f431])).
fof(f431,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f253])).
fof(f253,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f99])).
fof(f99,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f98])).
fof(f98,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X14] : (s(bool,'F_0') = s(bool,X14) <=> ~p(s(bool,X14)))),
  file('thm.integral.TDIV_BOUNDS','thm.bool.EQ_CLAUSES|split|3')).
fof(f2893,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | sP9(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP8(X3,X2,'F_0') | ~sP6(X2,'F_0',X3)) ) | $spl0),
  inference(resolution,[],[f2587,f342])).
fof(f342,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~p(s(bool,X2)) | ~sP8(X0,X2,X1) | ~sP7(X1,X2,X0) | ~sP6(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f210])).
fof(f2587,plain,(
  ( ! [X6,X7] : (sP7('F_0',X6,X7) | s(bool,'T_0') != s(bool,X6)) ) | $spl0),
  inference(subsumption_resolution,[],[f2582,f350])).
fof(f2582,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X6) | p(s(bool,X7)) | sP7('F_0',X6,X7)) ) | $spl0),
  inference(resolution,[],[f2322,f339])).
fof(f3527,plain,(
  ( ! [X121,X120] : (sP6('const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15)),X120,X121)) ) | $spl6),
  inference(resolution,[],[f468,f356])).
fof(f356,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f218])).
fof(f468,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15)))) | $spl6),
  inference(cnf_transformation,[],[f468_D])).
fof(f468_D,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15)))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f22530,plain,(
  ( ! [X11] : (s(bool,X11) != s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15)))) ) | $spl162),
  inference(cnf_transformation,[],[f22530_D])).
fof(f22530_D,plain,(
  ( ! [X11] : (s(bool,X11) != s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15)))) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f22541,plain,(
  ( ! [X17,X16] : (p(s(bool,X16)) | s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15))) | s(bool,X17) != s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15)))) ) | ($spl6 | $spl92)),
  inference(resolution,[],[f3527,f13137])).
fof(f13137,plain,(
  ( ! [X17,X15,X16] : (~sP6(X15,X17,X16) | p(s(bool,X16)) | s(bool,'F_0') != s(bool,X15) | s(bool,X15) != s(bool,X17)) ) | $spl92),
  inference(duplicate_literal_removal,[],[f13115])).
fof(f13115,plain,(
  ( ! [X17,X15,X16] : (s(bool,'F_0') != s(bool,X15) | p(s(bool,X16)) | ~sP6(X15,X17,X16) | p(s(bool,X16)) | s(bool,X15) != s(bool,X17)) ) | $spl92),
  inference(resolution,[],[f13065,f360])).
fof(f13065,plain,(
  ( ! [X17,X15,X16] : (sP9(X16,X17,X15) | s(bool,'F_0') != s(bool,X15) | p(s(bool,X16)) | ~sP6(X15,X17,X16)) ) | $spl92),
  inference(subsumption_resolution,[],[f13061,f348])).
fof(f13061,plain,(
  ( ! [X17,X15,X16] : (s(bool,'F_0') != s(bool,X15) | sP9(X16,X17,X15) | p(s(bool,X16)) | ~sP8(X16,X15,X17) | ~sP6(X15,X17,X16)) ) | $spl92),
  inference(resolution,[],[f13042,f344])).
fof(f344,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | p(s(bool,X0)) | ~sP8(X0,X2,X1) | ~sP7(X1,X2,X0) | ~sP6(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f210])).
fof(f13042,plain,(
  ( ! [X4,X5,X3] : (sP7(X5,X4,X3) | s(bool,'F_0') != s(bool,X4)) ) | $spl92),
  inference(subsumption_resolution,[],[f13041,f350])).
fof(f13041,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | s(bool,'F_0') != s(bool,X4) | sP7(X5,X4,X3)) ) | $spl92),
  inference(subsumption_resolution,[],[f13011,f352])).
fof(f352,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f216])).
fof(f13011,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | s(bool,'F_0') != s(bool,X4) | ~p(s(bool,X5)) | sP7(X5,X4,X3)) ) | $spl92),
  inference(resolution,[],[f12748,f339])).
fof(f3355,plain,(
  s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15))) | $spl30),
  inference(cnf_transformation,[],[f3355_D])).
fof(f3355_D,plain,(
  s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15))) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f22540,plain,(
  ( ! [X14,X15] : (p(s(bool,X14)) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15))) | s(bool,X15) != s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15)))) ) | ($spl6 | $spl84)),
  inference(resolution,[],[f3527,f13093])).
fof(f13093,plain,(
  ( ! [X17,X15,X16] : (~sP6(X15,X17,X16) | p(s(bool,X16)) | s(bool,'T_0') = s(bool,X15) | s(bool,X15) != s(bool,X17)) ) | $spl84),
  inference(duplicate_literal_removal,[],[f13072])).
fof(f13072,plain,(
  ( ! [X17,X15,X16] : (s(bool,'T_0') = s(bool,X15) | p(s(bool,X16)) | ~sP6(X15,X17,X16) | p(s(bool,X16)) | s(bool,X15) != s(bool,X17)) ) | $spl84),
  inference(resolution,[],[f12972,f360])).
fof(f12972,plain,(
  ( ! [X17,X15,X16] : (sP9(X16,X17,X15) | s(bool,'T_0') = s(bool,X15) | p(s(bool,X16)) | ~sP6(X15,X17,X16)) ) | $spl84),
  inference(subsumption_resolution,[],[f12969,f348])).
fof(f12969,plain,(
  ( ! [X17,X15,X16] : (s(bool,'T_0') = s(bool,X15) | sP9(X16,X17,X15) | p(s(bool,X16)) | ~sP8(X16,X15,X17) | ~sP6(X15,X17,X16)) ) | $spl84),
  inference(resolution,[],[f12953,f344])).
fof(f12953,plain,(
  ( ! [X4,X5,X3] : (sP7(X5,X3,X4) | s(bool,'T_0') = s(bool,X3)) ) | $spl84),
  inference(subsumption_resolution,[],[f12952,f350])).
fof(f12952,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') = s(bool,X3) | p(s(bool,X4)) | sP7(X5,X3,X4)) ) | $spl84),
  inference(subsumption_resolution,[],[f12923,f352])).
fof(f12923,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') = s(bool,X3) | p(s(bool,X4)) | ~p(s(bool,X5)) | sP7(X5,X3,X4)) ) | $spl84),
  inference(resolution,[],[f12739,f339])).
fof(f22531,plain,(
  ( ! [X15] : (s(bool,'T_0') = s(bool,X15) | s(bool,X15) != s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15)))) ) | $spl164),
  inference(cnf_transformation,[],[f22531_D])).
fof(f22531_D,plain,(
  ( ! [X15] : (s(bool,'T_0') = s(bool,X15) | s(bool,X15) != s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15)))) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f22539,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | s(bool,'T_0') = s(bool,X13) | s(bool,X13) != s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15)))) ) | ($spl6 | $spl70)),
  inference(resolution,[],[f3527,f12607])).
fof(f12607,plain,(
  ( ! [X17,X15,X16] : (~sP6(X17,X15,X16) | p(s(bool,X16)) | s(bool,'T_0') = s(bool,X15) | s(bool,X15) != s(bool,X17)) ) | $spl70),
  inference(duplicate_literal_removal,[],[f12586])).
fof(f12586,plain,(
  ( ! [X17,X15,X16] : (s(bool,'T_0') = s(bool,X15) | p(s(bool,X16)) | ~sP6(X17,X15,X16) | p(s(bool,X16)) | s(bool,X15) != s(bool,X17)) ) | $spl70),
  inference(resolution,[],[f12541,f360])).
fof(f12541,plain,(
  ( ! [X17,X15,X16] : (sP9(X16,X15,X17) | s(bool,'T_0') = s(bool,X15) | p(s(bool,X16)) | ~sP6(X17,X15,X16)) ) | $spl70),
  inference(subsumption_resolution,[],[f12537,f348])).
fof(f12537,plain,(
  ( ! [X17,X15,X16] : (s(bool,'T_0') = s(bool,X15) | sP9(X16,X15,X17) | p(s(bool,X16)) | ~sP8(X16,X17,X15) | ~sP6(X17,X15,X16)) ) | $spl70),
  inference(resolution,[],[f12519,f344])).
fof(f12519,plain,(
  ( ! [X4,X5,X3] : (sP7(X3,X5,X4) | s(bool,'T_0') = s(bool,X3)) ) | $spl70),
  inference(subsumption_resolution,[],[f12518,f350])).
fof(f12518,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') = s(bool,X3) | p(s(bool,X4)) | sP7(X3,X5,X4)) ) | $spl70),
  inference(subsumption_resolution,[],[f12487,f351])).
fof(f12487,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') = s(bool,X3) | p(s(bool,X4)) | ~p(s(bool,X5)) | sP7(X3,X5,X4)) ) | $spl70),
  inference(resolution,[],[f11843,f339])).
fof(f22532,plain,(
  ( ! [X17] : (s(bool,'F_0') != s(bool,X17) | s(bool,X17) != s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15)))) ) | $spl166),
  inference(cnf_transformation,[],[f22532_D])).
fof(f22532_D,plain,(
  ( ! [X17] : (s(bool,'F_0') != s(bool,X17) | s(bool,X17) != s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15)))) ) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f22538,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | s(bool,'F_0') != s(bool,X11) | s(bool,X11) != s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15)))) ) | ($spl6 | $spl10)),
  inference(resolution,[],[f3527,f9656])).
fof(f9656,plain,(
  ( ! [X4,X5,X3] : (~sP6(X5,X3,X4) | p(s(bool,X4)) | s(bool,'F_0') != s(bool,X3) | s(bool,X3) != s(bool,X5)) ) | $spl10),
  inference(duplicate_literal_removal,[],[f9639])).
fof(f9639,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X3) | p(s(bool,X4)) | ~sP6(X5,X3,X4) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5)) ) | $spl10),
  inference(resolution,[],[f9434,f360])).
fof(f9434,plain,(
  ( ! [X10,X11,X9] : (sP9(X10,X9,X11) | s(bool,'F_0') != s(bool,X9) | p(s(bool,X10)) | ~sP6(X11,X9,X10)) ) | $spl10),
  inference(subsumption_resolution,[],[f9430,f348])).
fof(f9430,plain,(
  ( ! [X10,X11,X9] : (s(bool,'F_0') != s(bool,X9) | sP9(X10,X9,X11) | p(s(bool,X10)) | ~sP8(X10,X11,X9) | ~sP6(X11,X9,X10)) ) | $spl10),
  inference(resolution,[],[f9421,f344])).
fof(f9421,plain,(
  ( ! [X10,X11,X9] : (sP7(X10,X9,X11) | s(bool,'F_0') != s(bool,X10)) ) | $spl10),
  inference(subsumption_resolution,[],[f9420,f351])).
fof(f9420,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | s(bool,'F_0') != s(bool,X10) | sP7(X10,X9,X11)) ) | $spl10),
  inference(subsumption_resolution,[],[f9403,f350])).
fof(f9403,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | s(bool,'F_0') != s(bool,X10) | p(s(bool,X11)) | sP7(X10,X9,X11)) ) | $spl10),
  inference(resolution,[],[f9258,f339])).
fof(f9258,plain,(
  ( ! [X6,X8,X7] : (sP9(X6,X7,X8) | ~p(s(bool,X8)) | s(bool,'F_0') != s(bool,X7) | p(s(bool,X6))) ) | $spl10),
  inference(resolution,[],[f9208,f348])).
fof(f9208,plain,(
  ( ! [X4,X5,X3] : (~sP8(X5,X4,X3) | sP9(X5,X3,X4) | ~p(s(bool,X4)) | s(bool,'F_0') != s(bool,X3)) ) | $spl10),
  inference(subsumption_resolution,[],[f9207,f436])).
fof(f9207,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X3) | s(bool,'T_0') != s(bool,X4) | sP9(X5,X3,X4) | ~p(s(bool,X4)) | ~sP8(X5,X4,X3)) ) | $spl10),
  inference(subsumption_resolution,[],[f9202,f9184])).
fof(f9184,plain,(
  ( ! [X6,X8,X7] : (sP6(X6,X7,X8) | s(bool,'F_0') != s(bool,X7) | s(bool,'T_0') != s(bool,X6)) ) | $spl10),
  inference(subsumption_resolution,[],[f9167,f354])).
fof(f9167,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X6) | s(bool,'F_0') != s(bool,X7) | p(s(bool,X8)) | sP6(X6,X7,X8)) ) | $spl10),
  inference(resolution,[],[f2552,f338])).
fof(f2552,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X1,X0) | s(bool,'T_0') != s(bool,X0) | s(bool,'F_0') != s(bool,X1) | p(s(bool,X2))) ) | $spl10),
  inference(superposition,[],[f2345,f357])).
fof(f9202,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X3) | s(bool,'T_0') != s(bool,X4) | sP9(X5,X3,X4) | ~p(s(bool,X4)) | ~sP8(X5,X4,X3) | ~sP6(X4,X3,X5)) ) | $spl10),
  inference(resolution,[],[f9185,f342])).
fof(f9185,plain,(
  ( ! [X10,X11,X9] : (sP7(X10,X9,X11) | s(bool,'F_0') != s(bool,X10) | s(bool,'T_0') != s(bool,X9)) ) | $spl10),
  inference(subsumption_resolution,[],[f9168,f350])).
fof(f9168,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') != s(bool,X9) | s(bool,'F_0') != s(bool,X10) | p(s(bool,X11)) | sP7(X10,X9,X11)) ) | $spl10),
  inference(resolution,[],[f2552,f339])).
fof(f22490,plain,(
  ( ! [X24] : (s(bool,X24) != s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15))) | p(s(bool,X24))) ) | $spl160),
  inference(cnf_transformation,[],[f22490_D])).
fof(f22490_D,plain,(
  ( ! [X24] : (s(bool,X24) != s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15))) | p(s(bool,X24))) ) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f22537,plain,(
  ( ! [X8,X9] : (p(s(bool,X8)) | p(s(bool,X9)) | s(bool,X8) != s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15)))) ) | ($spl6 | $spl10)),
  inference(resolution,[],[f3527,f9548])).
fof(f9548,plain,(
  ( ! [X4,X5,X3] : (~sP6(X4,X3,X5) | p(s(bool,X3)) | p(s(bool,X5)) | s(bool,X3) != s(bool,X4)) ) | $spl10),
  inference(duplicate_literal_removal,[],[f9530])).
fof(f9530,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~sP6(X4,X3,X5) | p(s(bool,X5)) | p(s(bool,X5)) | s(bool,X3) != s(bool,X4)) ) | $spl10),
  inference(resolution,[],[f9438,f360])).
fof(f9438,plain,(
  ( ! [X6,X8,X7] : (sP9(X7,X6,X8) | p(s(bool,X6)) | ~sP6(X8,X6,X7) | p(s(bool,X7))) ) | $spl10),
  inference(resolution,[],[f9433,f348])).
fof(f9433,plain,(
  ( ! [X6,X8,X7] : (~sP8(X7,X8,X6) | p(s(bool,X6)) | sP9(X7,X6,X8) | ~sP6(X8,X6,X7)) ) | $spl10),
  inference(subsumption_resolution,[],[f9429,f432])).
fof(f432,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f253])).
fof(f9429,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') != s(bool,X6) | sP9(X7,X6,X8) | p(s(bool,X6)) | ~sP8(X7,X8,X6) | ~sP6(X8,X6,X7)) ) | $spl10),
  inference(resolution,[],[f9421,f343])).
fof(f22527,plain,(
  ( ! [X17,X16] : (p(s(bool,X16)) | s(bool,'F_0') != s(bool,X17) | s(bool,X17) != s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15)))) ) | ($spl6 | $spl92)),
  inference(resolution,[],[f3526,f13137])).
fof(f3526,plain,(
  ( ! [X118,X119] : (sP6(X118,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15)),X119)) ) | $spl6),
  inference(resolution,[],[f468,f355])).
fof(f355,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f218])).
fof(f22526,plain,(
  ( ! [X14,X15] : (p(s(bool,X14)) | s(bool,'T_0') = s(bool,X15) | s(bool,X15) != s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15)))) ) | ($spl6 | $spl84)),
  inference(resolution,[],[f3526,f13093])).
fof(f22525,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15))) | s(bool,X13) != s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15)))) ) | ($spl6 | $spl70)),
  inference(resolution,[],[f3526,f12607])).
fof(f22524,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15))) | s(bool,X11) != s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15)))) ) | ($spl6 | $spl10)),
  inference(resolution,[],[f3526,f9656])).
fof(f3000,plain,(
  ( ! [X8] : (~p(s(bool,X8))) ) | $spl12),
  inference(cnf_transformation,[],[f3000_D])).
fof(f3000_D,plain,(
  ( ! [X8] : (~p(s(bool,X8))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f22480,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X24) != s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15))) | p(s(bool,X24))) ) | ($spl6 | $spl10)),
  inference(resolution,[],[f3521,f13515])).
fof(f13515,plain,(
  ( ! [X2,X0,X3,X1] : (~sP8(X2,X1,X3) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X0) != s(bool,X3) | p(s(bool,X0))) ) | $spl10),
  inference(subsumption_resolution,[],[f13514,f356])).
fof(f13514,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X0) != s(bool,X3) | ~sP8(X2,X1,X3) | ~sP6(X1,X3,X2)) ) | $spl10),
  inference(subsumption_resolution,[],[f13513,f350])).
fof(f13513,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X0) != s(bool,X3) | ~sP8(X2,X1,X3) | ~sP7(X3,X1,X2) | ~sP6(X1,X3,X2)) ) | $spl10),
  inference(duplicate_literal_removal,[],[f13461])).
fof(f13461,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X0) != s(bool,X3) | ~p(s(bool,X1)) | ~sP8(X2,X1,X3) | ~sP7(X3,X1,X2) | ~sP6(X1,X3,X2)) ) | $spl10),
  inference(resolution,[],[f4324,f342])).
fof(f4324,plain,(
  ( ! [X10,X8,X11,X9] : (~sP9(X11,X9,X8) | p(s(bool,X10)) | ~p(s(bool,X8)) | ~p(s(bool,X11)) | s(bool,X9) != s(bool,X10)) ) | $spl10),
  inference(superposition,[],[f4313,f359])).
fof(f359,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X2)) | ~sP9(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f219])).
fof(f4313,plain,(
  ( ! [X35,X34] : (s(bool,X34) != s(bool,X35) | p(s(bool,X35)) | ~p(s(bool,X34))) ) | $spl10),
  inference(trivial_inequality_removal,[],[f4297])).
fof(f4297,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X34) != s(bool,X35) | p(s(bool,X35)) | ~p(s(bool,X34))) ) | $spl10),
  inference(superposition,[],[f2536,f438])).
fof(f438,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f256])).
fof(f256,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X14] : (s(bool,'T_0') = s(bool,X14) <=> p(s(bool,X14)))),
  file('thm.integral.TDIV_BOUNDS','thm.bool.EQ_CLAUSES|split|1')).
fof(f2536,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X3) | s(bool,X2) != s(bool,X3) | p(s(bool,X2))) ) | $spl10),
  inference(superposition,[],[f2345,f434])).
fof(f434,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f254])).
fof(f254,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f100])).
fof(f100,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X14] : (s(bool,'F_0') = s(bool,X14) <=> ~p(s(bool,X14)))),
  file('thm.integral.TDIV_BOUNDS','thm.bool.EQ_CLAUSES|split|2')).
fof(f3521,plain,(
  ( ! [X109,X108] : (sP8(X108,X109,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15)))) ) | $spl6),
  inference(resolution,[],[f468,f346])).
fof(f346,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f213])).
fof(f22489,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15)),X11) | ~p(s(bool,X11))) ) | $spl158),
  inference(cnf_transformation,[],[f22489_D])).
fof(f22489_D,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15)),X11) | ~p(s(bool,X11))) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f22479,plain,(
  ( ! [X21,X20] : (sP9(X20,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15)),X21) | ~p(s(bool,X21)) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15)))) ) | ($spl6 | $spl70)),
  inference(resolution,[],[f3521,f12540])).
fof(f12540,plain,(
  ( ! [X10,X11,X9] : (~sP8(X10,X11,X9) | sP9(X10,X9,X11) | ~p(s(bool,X11)) | s(bool,'T_0') = s(bool,X9)) ) | $spl70),
  inference(subsumption_resolution,[],[f12535,f356])).
fof(f12535,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') = s(bool,X9) | sP9(X10,X9,X11) | ~p(s(bool,X11)) | ~sP8(X10,X11,X9) | ~sP6(X11,X9,X10)) ) | $spl70),
  inference(resolution,[],[f12519,f342])).
fof(f4319,plain,(
  ( ! [X23] : (s(bool,'T_0') != s(bool,X23)) ) | $spl38),
  inference(cnf_transformation,[],[f4319_D])).
fof(f4319_D,plain,(
  ( ! [X23] : (s(bool,'T_0') != s(bool,X23)) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f22478,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15))) | s(bool,'T_0') != s(bool,X19)) ) | ($spl6 | $spl10)),
  inference(resolution,[],[f3521,f9858])).
fof(f9858,plain,(
  ( ! [X4,X5,X3] : (~sP8(X4,X3,X5) | ~p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5) | s(bool,'T_0') != s(bool,X3)) ) | $spl10),
  inference(subsumption_resolution,[],[f9857,f431])).
fof(f9857,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5) | p(s(bool,X5)) | ~sP8(X4,X3,X5)) ) | $spl10),
  inference(subsumption_resolution,[],[f9856,f9184])).
fof(f9856,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5) | p(s(bool,X5)) | ~sP8(X4,X3,X5) | ~sP6(X3,X5,X4)) ) | $spl10),
  inference(subsumption_resolution,[],[f9817,f9185])).
fof(f9817,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5) | p(s(bool,X5)) | ~sP8(X4,X3,X5) | ~sP7(X5,X3,X4) | ~sP6(X3,X5,X4)) ) | $spl10),
  inference(resolution,[],[f2554,f343])).
fof(f2554,plain,(
  ( ! [X6,X8,X7] : (~sP9(X8,X7,X6) | s(bool,'T_0') != s(bool,X6) | ~p(s(bool,X8)) | s(bool,'F_0') != s(bool,X7)) ) | $spl10),
  inference(superposition,[],[f2345,f359])).
fof(f22477,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15))) | ~p(s(bool,X16)) | ~p(s(bool,X17))) ) | ($spl6 | $spl10)),
  inference(resolution,[],[f3521,f9855])).
fof(f9855,plain,(
  ( ! [X2,X0,X1] : (~sP8(X1,X0,X2) | s(bool,'F_0') != s(bool,X2) | ~p(s(bool,X0)) | ~p(s(bool,X1))) ) | $spl10),
  inference(subsumption_resolution,[],[f9854,f436])).
fof(f9854,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X1)) | s(bool,'F_0') != s(bool,X2) | ~p(s(bool,X0)) | ~sP8(X1,X0,X2)) ) | $spl10),
  inference(subsumption_resolution,[],[f9853,f9184])).
fof(f9853,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X1)) | s(bool,'F_0') != s(bool,X2) | ~p(s(bool,X0)) | ~sP8(X1,X0,X2) | ~sP6(X0,X2,X1)) ) | $spl10),
  inference(subsumption_resolution,[],[f9816,f9185])).
fof(f9816,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X1)) | s(bool,'F_0') != s(bool,X2) | ~p(s(bool,X0)) | ~sP8(X1,X0,X2) | ~sP7(X2,X0,X1) | ~sP6(X0,X2,X1)) ) | $spl10),
  inference(resolution,[],[f2554,f342])).
fof(f22474,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15)),X11) | ~p(s(bool,X11)) | s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15)))) ) | ($spl6 | $spl10)),
  inference(resolution,[],[f3521,f9208])).
fof(f3002,plain,(
  s(bool,'T_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))))) | $spl15),
  inference(cnf_transformation,[],[f3002_D])).
fof(f3002_D,plain,(
  s(bool,'T_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))))) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f22433,plain,(
  ( ! [X12,X13] : (sP9(X12,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16)))),X13) | p(s(bool,X12))) ) | $spl156),
  inference(cnf_transformation,[],[f22433_D])).
fof(f22433_D,plain,(
  ( ! [X12,X13] : (sP9(X12,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16)))),X13) | p(s(bool,X12))) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f22426,plain,(
  ( ! [X14,X15] : (sP9(X14,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16)))),X15) | p(s(bool,X14)) | s(bool,'T_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16)))))) ) | ($spl5 | $spl110)),
  inference(resolution,[],[f1670,f14050])).
fof(f14050,plain,(
  ( ! [X6,X8,X7] : (~sP7(X6,X8,X7) | sP9(X7,X6,X8) | p(s(bool,X7)) | s(bool,'T_0') != s(bool,X6)) ) | $spl110),
  inference(subsumption_resolution,[],[f14041,f348])).
fof(f14041,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X6) | sP9(X7,X6,X8) | p(s(bool,X7)) | ~sP8(X7,X8,X6) | ~sP7(X6,X8,X7)) ) | $spl110),
  inference(resolution,[],[f14028,f344])).
fof(f14028,plain,(
  ( ! [X2,X0,X1] : (sP6(X2,X1,X0) | s(bool,'T_0') != s(bool,X1)) ) | $spl110),
  inference(subsumption_resolution,[],[f14027,f354])).
fof(f14027,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X1) | sP6(X2,X1,X0)) ) | $spl110),
  inference(subsumption_resolution,[],[f13990,f356])).
fof(f13990,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X1) | p(s(bool,X2)) | sP6(X2,X1,X0)) ) | $spl110),
  inference(resolution,[],[f13662,f338])).
fof(f1670,plain,(
  ( ! [X114,X115] : (sP7('const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16)))),X114,X115)) ) | $spl5),
  inference(resolution,[],[f467,f352])).
fof(f3309,plain,(
  s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))))) | $spl26),
  inference(cnf_transformation,[],[f3309_D])).
fof(f3309_D,plain,(
  s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))))) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f22425,plain,(
  ( ! [X12,X13] : (sP9(X12,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16)))),X13) | p(s(bool,X12)) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16)))))) ) | ($spl5 | $spl94)),
  inference(resolution,[],[f1670,f13988])).
fof(f13988,plain,(
  ( ! [X6,X8,X7] : (~sP7(X6,X8,X7) | sP9(X7,X6,X8) | p(s(bool,X7)) | s(bool,'F_0') = s(bool,X6)) ) | $spl94),
  inference(subsumption_resolution,[],[f13979,f348])).
fof(f13979,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') = s(bool,X6) | sP9(X7,X6,X8) | p(s(bool,X7)) | ~sP8(X7,X8,X6) | ~sP7(X6,X8,X7)) ) | $spl94),
  inference(resolution,[],[f13907,f344])).
fof(f13907,plain,(
  ( ! [X2,X0,X1] : (sP6(X2,X0,X1) | s(bool,'F_0') = s(bool,X0)) ) | $spl94),
  inference(subsumption_resolution,[],[f13906,f354])).
fof(f13906,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X1)) | sP6(X2,X0,X1)) ) | $spl94),
  inference(subsumption_resolution,[],[f13872,f356])).
fof(f13872,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | sP6(X2,X0,X1)) ) | $spl94),
  inference(resolution,[],[f13644,f338])).
fof(f22374,plain,(
  ( ! [X35] : (s(bool,'T_0') != s(bool,X35) | s(bool,X35) != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16)))))) ) | $spl150),
  inference(cnf_transformation,[],[f22374_D])).
fof(f22374_D,plain,(
  ( ! [X35] : (s(bool,'T_0') != s(bool,X35) | s(bool,X35) != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16)))))) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f22424,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | s(bool,'T_0') != s(bool,X11) | s(bool,X11) != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16)))))) ) | ($spl5 | $spl10)),
  inference(resolution,[],[f1670,f9629])).
fof(f9629,plain,(
  ( ! [X4,X5,X3] : (~sP7(X5,X3,X4) | p(s(bool,X4)) | s(bool,'T_0') != s(bool,X3) | s(bool,X3) != s(bool,X5)) ) | $spl10),
  inference(duplicate_literal_removal,[],[f9611])).
fof(f9611,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,X3) | p(s(bool,X4)) | ~sP7(X5,X3,X4) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5)) ) | $spl10),
  inference(resolution,[],[f9310,f360])).
fof(f9310,plain,(
  ( ! [X10,X11,X9] : (sP9(X10,X11,X9) | s(bool,'T_0') != s(bool,X9) | p(s(bool,X10)) | ~sP7(X11,X9,X10)) ) | $spl10),
  inference(subsumption_resolution,[],[f9307,f348])).
fof(f9307,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') != s(bool,X9) | sP9(X10,X11,X9) | p(s(bool,X10)) | ~sP8(X10,X9,X11) | ~sP7(X11,X9,X10)) ) | $spl10),
  inference(resolution,[],[f9297,f344])).
fof(f9297,plain,(
  ( ! [X6,X8,X7] : (sP6(X7,X6,X8) | s(bool,'T_0') != s(bool,X7)) ) | $spl10),
  inference(subsumption_resolution,[],[f9296,f354])).
fof(f9296,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X7) | p(s(bool,X8)) | sP6(X7,X6,X8)) ) | $spl10),
  inference(subsumption_resolution,[],[f9279,f355])).
fof(f9279,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | s(bool,'T_0') != s(bool,X7) | p(s(bool,X8)) | sP6(X7,X6,X8)) ) | $spl10),
  inference(resolution,[],[f9215,f338])).
fof(f9215,plain,(
  ( ! [X6,X8,X7] : (sP9(X6,X7,X8) | p(s(bool,X7)) | s(bool,'T_0') != s(bool,X8) | p(s(bool,X6))) ) | $spl10),
  inference(resolution,[],[f9200,f348])).
fof(f9200,plain,(
  ( ! [X6,X8,X7] : (~sP8(X8,X7,X6) | sP9(X8,X6,X7) | p(s(bool,X6)) | s(bool,'T_0') != s(bool,X7)) ) | $spl10),
  inference(subsumption_resolution,[],[f9199,f352])).
fof(f9199,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X7) | sP9(X8,X6,X7) | p(s(bool,X6)) | ~sP8(X8,X7,X6) | ~sP7(X6,X7,X8)) ) | $spl10),
  inference(subsumption_resolution,[],[f9194,f432])).
fof(f9194,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,'T_0') != s(bool,X7) | sP9(X8,X6,X7) | p(s(bool,X6)) | ~sP8(X8,X7,X6) | ~sP7(X6,X7,X8)) ) | $spl10),
  inference(resolution,[],[f9184,f343])).
fof(f22373,plain,(
  ( ! [X27] : (~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16)))))) ) | $spl148),
  inference(cnf_transformation,[],[f22373_D])).
fof(f22373_D,plain,(
  ( ! [X27] : (~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16)))))) ) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f22423,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | p(s(bool,X9)) | s(bool,X8) != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16)))))) ) | ($spl5 | $spl10)),
  inference(resolution,[],[f1670,f9391])).
fof(f9391,plain,(
  ( ! [X4,X5,X3] : (~sP7(X4,X3,X5) | ~p(s(bool,X3)) | p(s(bool,X5)) | s(bool,X3) != s(bool,X4)) ) | $spl10),
  inference(duplicate_literal_removal,[],[f9374])).
fof(f9374,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | ~sP7(X4,X3,X5) | p(s(bool,X5)) | p(s(bool,X5)) | s(bool,X3) != s(bool,X4)) ) | $spl10),
  inference(resolution,[],[f9313,f360])).
fof(f9313,plain,(
  ( ! [X6,X8,X7] : (sP9(X7,X8,X6) | ~p(s(bool,X6)) | ~sP7(X8,X6,X7) | p(s(bool,X7))) ) | $spl10),
  inference(resolution,[],[f9309,f348])).
fof(f9309,plain,(
  ( ! [X4,X5,X3] : (~sP8(X4,X3,X5) | ~p(s(bool,X3)) | sP9(X4,X5,X3) | ~sP7(X5,X3,X4)) ) | $spl10),
  inference(subsumption_resolution,[],[f9305,f436])).
fof(f9305,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,X3) | sP9(X4,X5,X3) | ~p(s(bool,X3)) | ~sP8(X4,X3,X5) | ~sP7(X5,X3,X4)) ) | $spl10),
  inference(resolution,[],[f9297,f342])).
fof(f22414,plain,(
  ( ! [X20] : (p(s(bool,X20)) | s(bool,'T_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16)))))) ) | ($spl5 | $spl96)),
  inference(resolution,[],[f1669,f13824])).
fof(f13824,plain,(
  ( ! [X10,X11] : (~sP7('T_0',X11,X10) | p(s(bool,X10)) | s(bool,'T_0') != s(bool,X11)) ) | $spl96),
  inference(duplicate_literal_removal,[],[f13808])).
fof(f13808,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | ~sP7('T_0',X11,X10) | p(s(bool,X10)) | s(bool,'T_0') != s(bool,X11)) ) | $spl96),
  inference(resolution,[],[f13768,f360])).
fof(f13768,plain,(
  ( ! [X4,X5] : (sP9(X4,'T_0',X5) | p(s(bool,X4)) | ~sP7('T_0',X5,X4)) ) | $spl96),
  inference(subsumption_resolution,[],[f13759,f348])).
fof(f13759,plain,(
  ( ! [X4,X5] : (sP9(X4,'T_0',X5) | p(s(bool,X4)) | ~sP8(X4,X5,'T_0') | ~sP7('T_0',X5,X4)) ) | $spl96),
  inference(resolution,[],[f13713,f344])).
fof(f13713,plain,(
  ( ! [X0,X1] : (sP6(X1,'T_0',X0)) ) | $spl96),
  inference(subsumption_resolution,[],[f13712,f354])).
fof(f13712,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | sP6(X1,'T_0',X0)) ) | $spl96),
  inference(subsumption_resolution,[],[f13687,f356])).
fof(f13687,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | sP6(X1,'T_0',X0)) ) | $spl96),
  inference(resolution,[],[f13646,f338])).
fof(f1669,plain,(
  ( ! [X112,X113] : (sP7(X112,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16)))),X113)) ) | $spl5),
  inference(resolution,[],[f467,f351])).
fof(f22418,plain,(
  ( ! [X17,X16] : (sP9(X16,X17,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))))) | p(s(bool,X16))) ) | $spl154),
  inference(cnf_transformation,[],[f22418_D])).
fof(f22418_D,plain,(
  ( ! [X17,X16] : (sP9(X16,X17,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))))) | p(s(bool,X16))) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f22411,plain,(
  ( ! [X17,X16] : (sP9(X16,X17,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))))) | p(s(bool,X16)) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16)))))) ) | ($spl5 | $spl116)),
  inference(resolution,[],[f1669,f14811])).
fof(f14811,plain,(
  ( ! [X6,X8,X7] : (~sP7(X8,X6,X7) | sP9(X7,X8,X6) | p(s(bool,X7)) | s(bool,'F_0') = s(bool,X6)) ) | $spl116),
  inference(subsumption_resolution,[],[f14801,f348])).
fof(f14801,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') = s(bool,X6) | sP9(X7,X8,X6) | p(s(bool,X7)) | ~sP8(X7,X6,X8) | ~sP7(X8,X6,X7)) ) | $spl116),
  inference(resolution,[],[f14744,f344])).
fof(f14744,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X2,X1) | s(bool,'F_0') = s(bool,X0)) ) | $spl116),
  inference(subsumption_resolution,[],[f14743,f354])).
fof(f14743,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X1)) | sP6(X0,X2,X1)) ) | $spl116),
  inference(subsumption_resolution,[],[f14711,f355])).
fof(f14711,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | sP6(X0,X2,X1)) ) | $spl116),
  inference(resolution,[],[f14370,f338])).
fof(f22417,plain,(
  ( ! [X11] : (s(bool,X11) != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16)))))) ) | $spl152),
  inference(cnf_transformation,[],[f22417_D])).
fof(f22417_D,plain,(
  ( ! [X11] : (s(bool,X11) != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16)))))) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f22408,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | s(bool,'T_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))))) | s(bool,X11) != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16)))))) ) | ($spl5 | $spl10)),
  inference(resolution,[],[f1669,f9629])).
fof(f22357,plain,(
  ( ! [X39,X40] : (~p(s(bool,X39)) | s(bool,X40) != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))))) | ~p(s(bool,X40))) ) | ($spl0 | $spl5 | $spl10)),
  inference(resolution,[],[f1666,f15224])).
fof(f15224,plain,(
  ( ! [X165,X166,X164] : (~sP8(X165,X166,'F_0') | ~p(s(bool,X165)) | s(bool,X164) != s(bool,X166) | ~p(s(bool,X164))) ) | ($spl0 | $spl10)),
  inference(subsumption_resolution,[],[f15223,f354])).
fof(f15223,plain,(
  ( ! [X165,X166,X164] : (~p(s(bool,X164)) | ~p(s(bool,X165)) | s(bool,X164) != s(bool,X166) | ~sP8(X165,X166,'F_0') | ~sP6(X166,'F_0',X165)) ) | ($spl0 | $spl10)),
  inference(subsumption_resolution,[],[f15192,f417])).
fof(f417,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f244])).
fof(f244,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f243])).
fof(f243,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f93])).
fof(f93,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f92])).
fof(f92,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X14] : ((p(s(bool,X14)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X14)))),
  file('thm.integral.TDIV_BOUNDS','thm.bool.IMP_CLAUSES|split|4')).
fof(f15192,plain,(
  ( ! [X165,X166,X164] : (p(s(bool,'F_0')) | ~p(s(bool,X164)) | ~p(s(bool,X165)) | s(bool,X164) != s(bool,X166) | ~sP8(X165,X166,'F_0') | ~sP6(X166,'F_0',X165)) ) | ($spl0 | $spl10)),
  inference(resolution,[],[f4346,f2921])).
fof(f4346,plain,(
  ( ! [X14,X12,X15,X13] : (~sP9(X15,X12,X13) | p(s(bool,X12)) | ~p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X13) != s(bool,X14)) ) | $spl10),
  inference(superposition,[],[f4313,f359])).
fof(f1666,plain,(
  ( ! [X107,X106] : (sP8(X106,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16)))),X107)) ) | $spl5),
  inference(resolution,[],[f467,f347])).
fof(f347,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f213])).
fof(f22356,plain,(
  ( ! [X37,X38] : (~p(s(bool,X37)) | s(bool,X38) != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))))) | s(bool,'T_0') != s(bool,X38)) ) | ($spl0 | $spl5 | $spl10)),
  inference(resolution,[],[f1666,f9156])).
fof(f9156,plain,(
  ( ! [X12,X10,X11] : (~sP8(X11,X12,'F_0') | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,'T_0') != s(bool,X10)) ) | ($spl0 | $spl10)),
  inference(subsumption_resolution,[],[f9141,f2536])).
fof(f9141,plain,(
  ( ! [X12,X10,X11] : (s(bool,'T_0') != s(bool,X10) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | ~p(s(bool,X12)) | ~sP8(X11,X12,'F_0')) ) | ($spl0 | $spl10)),
  inference(resolution,[],[f2541,f2897])).
fof(f2541,plain,(
  ( ! [X17,X15,X16] : (~sP9(X17,'F_0',X15) | s(bool,'T_0') != s(bool,X16) | ~p(s(bool,X17)) | s(bool,X15) != s(bool,X16)) ) | $spl10),
  inference(superposition,[],[f2345,f359])).
fof(f3618,plain,(
  ( ! [X9] : (sP9(X9,'F_0','const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16)))))) ) | $spl34),
  inference(cnf_transformation,[],[f3618_D])).
fof(f3618_D,plain,(
  ( ! [X9] : (sP9(X9,'F_0','const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16)))))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f22355,plain,(
  ( ! [X35,X36] : (s(bool,'T_0') != s(bool,X35) | sP9(X36,'F_0','const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))))) | s(bool,X35) != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16)))))) ) | ($spl0 | $spl5 | $spl10)),
  inference(resolution,[],[f1666,f9092])).
fof(f9092,plain,(
  ( ! [X6,X8,X7] : (~sP8(X8,X7,'F_0') | s(bool,'T_0') != s(bool,X6) | sP9(X8,'F_0',X7) | s(bool,X6) != s(bool,X7)) ) | ($spl0 | $spl10)),
  inference(subsumption_resolution,[],[f9091,f2916])).
fof(f9091,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,X7) | s(bool,'T_0') != s(bool,X6) | sP9(X8,'F_0',X7) | ~sP8(X8,X7,'F_0') | ~sP7('F_0',X7,X8)) ) | $spl10),
  inference(subsumption_resolution,[],[f9088,f2536])).
fof(f9088,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,X7) | s(bool,'T_0') != s(bool,X6) | sP9(X8,'F_0',X7) | ~p(s(bool,X7)) | ~sP8(X8,X7,'F_0') | ~sP7('F_0',X7,X8)) ) | $spl10),
  inference(resolution,[],[f9077,f342])).
fof(f9077,plain,(
  ( ! [X14,X12,X13] : (sP6(X13,'F_0',X14) | s(bool,X12) != s(bool,X13) | s(bool,'T_0') != s(bool,X12)) ) | $spl10),
  inference(subsumption_resolution,[],[f9066,f354])).
fof(f9066,plain,(
  ( ! [X14,X12,X13] : (s(bool,'T_0') != s(bool,X12) | s(bool,X12) != s(bool,X13) | p(s(bool,X14)) | sP6(X13,'F_0',X14)) ) | $spl10),
  inference(resolution,[],[f2539,f338])).
fof(f2539,plain,(
  ( ! [X10,X11,X9] : (sP9(X11,'F_0',X9) | s(bool,'T_0') != s(bool,X10) | s(bool,X9) != s(bool,X10) | p(s(bool,X11))) ) | $spl10),
  inference(superposition,[],[f2345,f357])).
fof(f22354,plain,(
  ( ! [X34] : (s(bool,'T_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))))) | ~p(s(bool,X34))) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f1666,f2991])).
fof(f2991,plain,(
  ( ! [X14,X15] : (~sP8(X14,X15,'F_0') | s(bool,'T_0') != s(bool,X15) | ~p(s(bool,X14))) ) | $spl0),
  inference(subsumption_resolution,[],[f2990,f354])).
fof(f2990,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | ~sP8(X14,X15,'F_0') | ~sP6(X15,'F_0',X14)) ) | $spl0),
  inference(subsumption_resolution,[],[f2989,f417])).
fof(f2989,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | p(s(bool,'F_0')) | ~sP8(X14,X15,'F_0') | ~sP6(X15,'F_0',X14)) ) | $spl0),
  inference(subsumption_resolution,[],[f2982,f352])).
fof(f2982,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | p(s(bool,'F_0')) | ~sP8(X14,X15,'F_0') | ~sP7('F_0',X15,X14) | ~sP6(X15,'F_0',X14)) ) | $spl0),
  inference(resolution,[],[f2324,f343])).
fof(f2324,plain,(
  ( ! [X10,X9] : (~sP9(X10,'F_0',X9) | ~p(s(bool,X10)) | s(bool,'T_0') != s(bool,X9)) ) | $spl0),
  inference(superposition,[],[f2315,f359])).
fof(f22353,plain,(
  ( ! [X33] : (s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))))) | ~p(s(bool,X33))) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f1666,f2949])).
fof(f2949,plain,(
  ( ! [X2,X3] : (~sP8(X2,X3,'F_0') | s(bool,'F_0') = s(bool,X3) | ~p(s(bool,X2))) ) | $spl0),
  inference(subsumption_resolution,[],[f2943,f354])).
fof(f2943,plain,(
  ( ! [X2,X3] : (~sP8(X2,X3,'F_0') | ~sP6(X3,'F_0',X2) | s(bool,'F_0') = s(bool,X3) | ~p(s(bool,X2))) ) | $spl0),
  inference(resolution,[],[f2921,f359])).
fof(f22350,plain,(
  ( ! [X30] : (sP9(X30,'F_0','const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))))) | s(bool,'T_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16)))))) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f1666,f2891])).
fof(f2891,plain,(
  ( ! [X4,X5] : (~sP8(X5,X4,'F_0') | sP9(X5,'F_0',X4) | s(bool,'T_0') != s(bool,X4)) ) | $spl0),
  inference(subsumption_resolution,[],[f2890,f447])).
fof(f2890,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | sP9(X5,'F_0',X4) | p(s(bool,'F_0')) | ~sP8(X5,X4,'F_0')) ) | $spl0),
  inference(subsumption_resolution,[],[f2887,f352])).
fof(f2887,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | sP9(X5,'F_0',X4) | p(s(bool,'F_0')) | ~sP8(X5,X4,'F_0') | ~sP7('F_0',X4,X5)) ) | $spl0),
  inference(resolution,[],[f2586,f343])).
fof(f22349,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))))) | p(s(bool,X29))) ) | ($spl5 | $spl10)),
  inference(resolution,[],[f1666,f15221])).
fof(f15221,plain,(
  ( ! [X6,X4,X7,X5] : (~sP8(X6,X7,X4) | ~p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | p(s(bool,X4))) ) | $spl10),
  inference(subsumption_resolution,[],[f15220,f354])).
fof(f15220,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | ~sP8(X6,X7,X4) | ~sP6(X7,X4,X6)) ) | $spl10),
  inference(subsumption_resolution,[],[f15219,f352])).
fof(f15219,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | ~sP8(X6,X7,X4) | ~sP7(X4,X7,X6) | ~sP6(X7,X4,X6)) ) | $spl10),
  inference(duplicate_literal_removal,[],[f15149])).
fof(f15149,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | p(s(bool,X4)) | ~sP8(X6,X7,X4) | ~sP7(X4,X7,X6) | ~sP6(X7,X4,X6)) ) | $spl10),
  inference(resolution,[],[f4346,f343])).
fof(f22372,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))))) | p(s(bool,X9))) ) | $spl146),
  inference(cnf_transformation,[],[f22372_D])).
fof(f22372_D,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))))) | p(s(bool,X9))) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f22348,plain,(
  ( ! [X26,X25] : (sP9(X25,X26,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))))) | p(s(bool,X26)) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16)))))) ) | ($spl5 | $spl116)),
  inference(resolution,[],[f1666,f14810])).
fof(f14810,plain,(
  ( ! [X4,X5,X3] : (~sP8(X4,X3,X5) | sP9(X4,X5,X3) | p(s(bool,X5)) | s(bool,'F_0') = s(bool,X3)) ) | $spl116),
  inference(subsumption_resolution,[],[f14800,f352])).
fof(f14800,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') = s(bool,X3) | sP9(X4,X5,X3) | p(s(bool,X5)) | ~sP8(X4,X3,X5) | ~sP7(X5,X3,X4)) ) | $spl116),
  inference(resolution,[],[f14744,f343])).
fof(f10210,plain,(
  ( ! [X29] : (s(bool,'F_0') != s(bool,X29)) ) | $spl56),
  inference(cnf_transformation,[],[f10210_D])).
fof(f10210_D,plain,(
  ( ! [X29] : (s(bool,'F_0') != s(bool,X29)) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f22345,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'F_0') != s(bool,X19) | s(bool,'T_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16)))))) ) | ($spl5 | $spl10)),
  inference(resolution,[],[f1666,f9858])).
fof(f22340,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))))) | p(s(bool,X9)) | s(bool,'T_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16)))))) ) | ($spl5 | $spl10)),
  inference(resolution,[],[f1666,f9200])).
fof(f21882,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | p(s(bool,X11))) ) | ($spl0 | $spl2 | $spl130)),
  inference(subsumption_resolution,[],[f21881,f436])).
fof(f21881,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') != s(bool,X10) | ~p(s(bool,X10)) | p(s(bool,X11))) ) | ($spl0 | $spl2 | $spl130)),
  inference(forward_demodulation,[],[f21859,f19767])).
fof(f19767,plain,(
  s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK15))) | ($spl0 | $spl2)),
  inference(resolution,[],[f19677,f436])).
fof(f19677,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK15)))) | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f19654,f460])).
fof(f460,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f410])).
fof(f410,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f240])).
fof(f240,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f87])).
fof(f87,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f86])).
fof(f86,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X13] : ! [X15] : (s(X13,X15) = s(X13,X15) <=> p(s(bool,'T_0')))),
  file('thm.integral.TDIV_BOUNDS','thm.bool.REFL_CLAUSE')).
fof(f19654,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK15)))) | ($spl0 | $spl2)),
  inference(superposition,[],[f852,f1586])).
fof(f1586,plain,(
  s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK12),s('type.num.num',sK16))),s('type.realax.real',sK15))) | $spl2),
  inference(resolution,[],[f464,f436])).
fof(f464,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK12),s('type.num.num',sK16))),s('type.realax.real',sK15)))) | $spl2),
  inference(cnf_transformation,[],[f464_D])).
fof(f464_D,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK12),s('type.num.num',sK16))),s('type.realax.real',sK15)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f21859,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK15)))) ) | $spl130),
  inference(resolution,[],[f19977,f360])).
fof(f21375,plain,(
  $false | $spl7),
  inference(subsumption_resolution,[],[f21374,f453])).
fof(f453,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  p(s(bool,'T_0'))),
  file('thm.integral.TDIV_BOUNDS','thm.bool.TRUTH')).
fof(f21374,plain,(
  ~p(s(bool,'T_0')) | $spl7),
  inference(forward_demodulation,[],[f21268,f1914])).
fof(f1914,plain,(
  ( ! [X219] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK12),s('type.num.num',X219))),s('type.realax.real',sK15)))) )),
  inference(resolution,[],[f1181,f436])).
fof(f1181,plain,(
  ( ! [X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK12),s('type.num.num',X1))),s('type.realax.real',sK15))))) )),
  inference(resolution,[],[f579,f268])).
fof(f268,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X3))),s('type.realax.real',X2)))) | ~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s(fun('type.num.num','type.realax.real'),X0))))) )),
  inference(cnf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s(fun('type.num.num','type.realax.real'),X0)))) | ! [X3] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X3)))))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X3))),s('type.realax.real',X2))))))),
  inference(ennf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s(fun('type.num.num','type.realax.real'),X0)))) => ! [X3] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X3)))))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X3))),s('type.realax.real',X2))))))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X6,X8,X9] : (p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X8),s('type.realax.real',X9))),s(fun('type.num.num','type.realax.real'),X6)))) => ! [X10] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X8),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X6),s('type.num.num',X10)))))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X6),s('type.num.num',X10))),s('type.realax.real',X9))))))),
  file('thm.integral.TDIV_BOUNDS','thm.integral.DIVISION_BOUNDS')).
fof(f579,plain,(
  p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK14),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK12))))),
  inference(resolution,[],[f260,f262])).
fof(f262,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s(fun('type.num.num','type.realax.real'),X2)))) | ~p(s(bool,'const.transc.tdiv_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.pair.prod'(fun('type.num.num','type.realax.real'),fun('type.num.num','type.realax.real')),'const.pair.,_2'(s(fun('type.num.num','type.realax.real'),X2),s(fun('type.num.num','type.realax.real'),X3))))))) )),
  inference(cnf_transformation,[],[f165])).
fof(f21268,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK12),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16))))),s('type.realax.real',sK15)))) | $spl7),
  inference(resolution,[],[f581,f3933])).
fof(f3933,plain,(
  ( ! [X0] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',X0)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',sK15))))) ) | $spl7),
  inference(resolution,[],[f469,f266])).
fof(f581,plain,(
  ( ! [X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',X1))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK12),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))))))) )),
  inference(resolution,[],[f260,f264])).
fof(f264,plain,(
  ( ! [X2,X0,X5,X3,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X3),s('type.num.num',X5))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X5)))))))) | ~p(s(bool,'const.transc.tdiv_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.pair.prod'(fun('type.num.num','type.realax.real'),fun('type.num.num','type.realax.real')),'const.pair.,_2'(s(fun('type.num.num','type.realax.real'),X2),s(fun('type.num.num','type.realax.real'),X3))))))) )),
  inference(cnf_transformation,[],[f165])).
fof(f21126,plain,(
  $false | ($spl0 | $spl2 | $spl136)),
  inference(trivial_inequality_removal,[],[f21104])).
fof(f21104,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl0 | $spl2 | $spl136)),
  inference(superposition,[],[f20138,f19767])).
fof(f21125,plain,(
  $false | $spl136),
  inference(equality_resolution,[],[f20138])).
fof(f20970,plain,(
  ( ! [X339,X338] : (p(s(bool,X338)) | p(s(bool,X339)) | sP9(X339,X338,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK15)))) ) | $spl144),
  inference(cnf_transformation,[],[f20970_D])).
fof(f20970_D,plain,(
  ( ! [X339,X338] : (p(s(bool,X338)) | p(s(bool,X339)) | sP9(X339,X338,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK15)))) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f20969,plain,(
  ( ! [X337,X339,X338] : (s(bool,'T_0') != s(bool,X337) | p(s(bool,X338)) | sP9(X339,X338,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK15))) | p(s(bool,X339))) ) | ($spl0 | $spl2 | $spl10)),
  inference(subsumption_resolution,[],[f20831,f435])).
fof(f435,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f255])).
fof(f20831,plain,(
  ( ! [X337,X339,X338] : (s(bool,'T_0') != s(bool,X337) | p(s(bool,X338)) | ~p(s(bool,X337)) | sP9(X339,X338,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK15))) | p(s(bool,X339))) ) | ($spl0 | $spl2 | $spl10)),
  inference(superposition,[],[f4344,f19767])).
fof(f4344,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X5) != s(bool,X6) | p(s(bool,X4)) | ~p(s(bool,X6)) | sP9(X7,X4,X5) | p(s(bool,X7))) ) | $spl10),
  inference(superposition,[],[f4313,f357])).
fof(f20968,plain,(
  ( ! [X333,X332] : (p(s(bool,X332)) | p(s(bool,X333)) | sP9(X333,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK15)),X332)) ) | $spl142),
  inference(cnf_transformation,[],[f20968_D])).
fof(f20968_D,plain,(
  ( ! [X333,X332] : (p(s(bool,X332)) | p(s(bool,X333)) | sP9(X333,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK15)),X332)) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f20967,plain,(
  ( ! [X333,X331,X332] : (s(bool,'T_0') != s(bool,X331) | p(s(bool,X332)) | sP9(X333,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK15)),X332) | p(s(bool,X333))) ) | ($spl0 | $spl2 | $spl10)),
  inference(subsumption_resolution,[],[f20829,f435])).
fof(f20829,plain,(
  ( ! [X333,X331,X332] : (s(bool,'T_0') != s(bool,X331) | p(s(bool,X332)) | ~p(s(bool,X331)) | sP9(X333,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK15)),X332) | p(s(bool,X333))) ) | ($spl0 | $spl2 | $spl10)),
  inference(superposition,[],[f4343,f19767])).
fof(f4343,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X1) != s(bool,X2) | p(s(bool,X0)) | ~p(s(bool,X2)) | sP9(X3,X1,X0) | p(s(bool,X3))) ) | $spl10),
  inference(superposition,[],[f4313,f357])).
fof(f19976,plain,(
  s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK15))) | $spl129),
  inference(cnf_transformation,[],[f19976_D])).
fof(f19976_D,plain,(
  s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK15))) <=> ~$spl129),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl129])])).
fof(f20151,plain,(
  ( ! [X19] : (p(s(bool,X19)) | s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK15)))) ) | ($spl0 | $spl2)),
  inference(resolution,[],[f19745,f2950])).
fof(f19745,plain,(
  ( ! [X121,X120] : (sP6('const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK15)),X120,X121)) ) | ($spl0 | $spl2)),
  inference(resolution,[],[f19677,f356])).
fof(f20138,plain,(
  ( ! [X11] : (s(bool,X11) != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK15)))) ) | $spl136),
  inference(cnf_transformation,[],[f20138_D])).
fof(f20138_D,plain,(
  ( ! [X11] : (s(bool,X11) != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK15)))) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f20149,plain,(
  ( ! [X17,X16] : (p(s(bool,X16)) | s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK15))) | s(bool,X17) != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK15)))) ) | ($spl0 | $spl2 | $spl92)),
  inference(resolution,[],[f19745,f13137])).
fof(f19979,plain,(
  s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK15))) | $spl132),
  inference(cnf_transformation,[],[f19979_D])).
fof(f19979_D,plain,(
  s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK15))) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f20148,plain,(
  ( ! [X14,X15] : (p(s(bool,X14)) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK15))) | s(bool,X15) != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK15)))) ) | ($spl0 | $spl2 | $spl84)),
  inference(resolution,[],[f19745,f13093])).
fof(f20139,plain,(
  ( ! [X15] : (s(bool,'T_0') = s(bool,X15) | s(bool,X15) != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK15)))) ) | $spl138),
  inference(cnf_transformation,[],[f20139_D])).
fof(f20139_D,plain,(
  ( ! [X15] : (s(bool,'T_0') = s(bool,X15) | s(bool,X15) != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK15)))) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f20147,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | s(bool,'T_0') = s(bool,X13) | s(bool,X13) != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK15)))) ) | ($spl0 | $spl2 | $spl70)),
  inference(resolution,[],[f19745,f12607])).
fof(f20140,plain,(
  ( ! [X17] : (s(bool,'F_0') != s(bool,X17) | s(bool,X17) != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK15)))) ) | $spl140),
  inference(cnf_transformation,[],[f20140_D])).
fof(f20140_D,plain,(
  ( ! [X17] : (s(bool,'F_0') != s(bool,X17) | s(bool,X17) != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK15)))) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f20146,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | s(bool,'F_0') != s(bool,X11) | s(bool,X11) != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK15)))) ) | ($spl0 | $spl2 | $spl10)),
  inference(resolution,[],[f19745,f9656])).
fof(f19980,plain,(
  ( ! [X24] : (s(bool,X24) != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK15))) | p(s(bool,X24))) ) | $spl134),
  inference(cnf_transformation,[],[f19980_D])).
fof(f19980_D,plain,(
  ( ! [X24] : (s(bool,X24) != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK15))) | p(s(bool,X24))) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f20145,plain,(
  ( ! [X8,X9] : (p(s(bool,X8)) | p(s(bool,X9)) | s(bool,X8) != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK15)))) ) | ($spl0 | $spl2 | $spl10)),
  inference(resolution,[],[f19745,f9548])).
fof(f20133,plain,(
  ( ! [X17,X16] : (p(s(bool,X16)) | s(bool,'F_0') != s(bool,X17) | s(bool,X17) != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK15)))) ) | ($spl0 | $spl2 | $spl92)),
  inference(resolution,[],[f19744,f13137])).
fof(f19744,plain,(
  ( ! [X118,X119] : (sP6(X118,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK15)),X119)) ) | ($spl0 | $spl2)),
  inference(resolution,[],[f19677,f355])).
fof(f20132,plain,(
  ( ! [X14,X15] : (p(s(bool,X14)) | s(bool,'T_0') = s(bool,X15) | s(bool,X15) != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK15)))) ) | ($spl0 | $spl2 | $spl84)),
  inference(resolution,[],[f19744,f13093])).
fof(f20131,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK15))) | s(bool,X13) != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK15)))) ) | ($spl0 | $spl2 | $spl70)),
  inference(resolution,[],[f19744,f12607])).
fof(f20130,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK15))) | s(bool,X11) != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK15)))) ) | ($spl0 | $spl2 | $spl10)),
  inference(resolution,[],[f19744,f9656])).
fof(f19964,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X24) != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK15))) | p(s(bool,X24))) ) | ($spl0 | $spl2 | $spl10)),
  inference(resolution,[],[f19739,f13515])).
fof(f19739,plain,(
  ( ! [X109,X108] : (sP8(X108,X109,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK15)))) ) | ($spl0 | $spl2)),
  inference(resolution,[],[f19677,f346])).
fof(f19977,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK15)),X11) | ~p(s(bool,X11))) ) | $spl130),
  inference(cnf_transformation,[],[f19977_D])).
fof(f19977_D,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK15)),X11) | ~p(s(bool,X11))) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f19963,plain,(
  ( ! [X21,X20] : (sP9(X20,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK15)),X21) | ~p(s(bool,X21)) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK15)))) ) | ($spl0 | $spl2 | $spl70)),
  inference(resolution,[],[f19739,f12540])).
fof(f19962,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK15))) | s(bool,'T_0') != s(bool,X19)) ) | ($spl0 | $spl2 | $spl10)),
  inference(resolution,[],[f19739,f9858])).
fof(f19961,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK15))) | ~p(s(bool,X16)) | ~p(s(bool,X17))) ) | ($spl0 | $spl2 | $spl10)),
  inference(resolution,[],[f19739,f9855])).
fof(f19958,plain,(
  ( ! [X10,X11] : (sP9(X10,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK15)),X11) | ~p(s(bool,X11)) | s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK15)))) ) | ($spl0 | $spl2 | $spl10)),
  inference(resolution,[],[f19739,f9208])).
fof(f14487,plain,(
  ( ! [X39,X38] : (s(bool,'F_0') != s(bool,X39) | ~p(s(bool,X38))) ) | $spl98),
  inference(subsumption_resolution,[],[f14466,f13647])).
fof(f14466,plain,(
  ( ! [X39,X38] : (s(bool,'F_0') != s(bool,X39) | ~p(s(bool,X38)) | s(bool,'T_0') = s(bool,X38)) ) | $spl98),
  inference(superposition,[],[f13647,f456])).
fof(f456,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.integral.TDIV_BOUNDS',aHL_BOOL_CASES)).
fof(f14486,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,X37) | ~p(s(bool,X36))) ) | $spl98),
  inference(subsumption_resolution,[],[f14465,f13647])).
fof(f14465,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,X37) | ~p(s(bool,X36)) | s(bool,'F_0') = s(bool,X36)) ) | $spl98),
  inference(superposition,[],[f13647,f456])).
fof(f14482,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') != s(bool,X35) | ~p(s(bool,X34))) ) | $spl98),
  inference(duplicate_literal_removal,[],[f14464])).
fof(f14464,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') != s(bool,X35) | ~p(s(bool,X34)) | ~p(s(bool,X34))) ) | $spl98),
  inference(superposition,[],[f13647,f438])).
fof(f14483,plain,(
  ( ! [X33,X32] : (s(bool,'T_0') != s(bool,X33) | ~p(s(bool,X32))) ) | $spl98),
  inference(duplicate_literal_removal,[],[f14463])).
fof(f14463,plain,(
  ( ! [X33,X32] : (s(bool,'T_0') != s(bool,X33) | ~p(s(bool,X32)) | ~p(s(bool,X32))) ) | $spl98),
  inference(superposition,[],[f13647,f436])).
fof(f14385,plain,(
  ( ! [X114,X113] : (p(s(bool,X113)) | p(s(bool,X114)) | sP9(X114,X113,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK14),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK12)))) ) | $spl126),
  inference(cnf_transformation,[],[f14385_D])).
fof(f14385_D,plain,(
  ( ! [X114,X113] : (p(s(bool,X113)) | p(s(bool,X114)) | sP9(X114,X113,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK14),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK12)))) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f14384,plain,(
  ( ! [X114,X112,X113] : (s(bool,'T_0') != s(bool,X112) | p(s(bool,X113)) | sP9(X114,X113,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK14),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK12))) | p(s(bool,X114))) ) | $spl10),
  inference(subsumption_resolution,[],[f14328,f435])).
fof(f14328,plain,(
  ( ! [X114,X112,X113] : (s(bool,'T_0') != s(bool,X112) | p(s(bool,X113)) | ~p(s(bool,X112)) | sP9(X114,X113,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK14),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK12))) | p(s(bool,X114))) ) | $spl10),
  inference(superposition,[],[f4344,f1263])).
fof(f1263,plain,(
  s(bool,'T_0') = s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK14),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK12)))),
  inference(resolution,[],[f579,f436])).
fof(f14383,plain,(
  ( ! [X111,X110,X108] : (p(s(bool,X110)) | p(s(bool,X111)) | sP9(X111,X110,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK12),s('type.num.num',X108))),s('type.realax.real',sK15)))) ) | $spl124),
  inference(cnf_transformation,[],[f14383_D])).
fof(f14383_D,plain,(
  ( ! [X111,X110,X108] : (p(s(bool,X110)) | p(s(bool,X111)) | sP9(X111,X110,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK12),s('type.num.num',X108))),s('type.realax.real',sK15)))) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f14382,plain,(
  ( ! [X111,X109,X110,X108] : (s(bool,'T_0') != s(bool,X109) | p(s(bool,X110)) | sP9(X111,X110,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK12),s('type.num.num',X108))),s('type.realax.real',sK15))) | p(s(bool,X111))) ) | $spl10),
  inference(subsumption_resolution,[],[f14327,f435])).
fof(f14327,plain,(
  ( ! [X111,X109,X110,X108] : (s(bool,'T_0') != s(bool,X109) | p(s(bool,X110)) | ~p(s(bool,X109)) | sP9(X111,X110,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK12),s('type.num.num',X108))),s('type.realax.real',sK15))) | p(s(bool,X111))) ) | $spl10),
  inference(superposition,[],[f4344,f1914])).
fof(f14380,plain,(
  ( ! [X101,X100,X98] : (p(s(bool,X100)) | p(s(bool,X101)) | sP9(X101,X100,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK12),s('type.num.num',X98)))))) ) | $spl122),
  inference(cnf_transformation,[],[f14380_D])).
fof(f14380_D,plain,(
  ( ! [X101,X100,X98] : (p(s(bool,X100)) | p(s(bool,X101)) | sP9(X101,X100,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK12),s('type.num.num',X98)))))) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f14379,plain,(
  ( ! [X101,X99,X100,X98] : (s(bool,'T_0') != s(bool,X99) | p(s(bool,X100)) | sP9(X101,X100,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK12),s('type.num.num',X98))))) | p(s(bool,X101))) ) | $spl10),
  inference(subsumption_resolution,[],[f14324,f435])).
fof(f14324,plain,(
  ( ! [X101,X99,X100,X98] : (s(bool,'T_0') != s(bool,X99) | p(s(bool,X100)) | ~p(s(bool,X99)) | sP9(X101,X100,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK12),s('type.num.num',X98))))) | p(s(bool,X101))) ) | $spl10),
  inference(superposition,[],[f4344,f1368])).
fof(f1368,plain,(
  ( ! [X219] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK12),s('type.num.num',X219)))))) )),
  inference(resolution,[],[f1180,f436])).
fof(f1180,plain,(
  ( ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK12),s('type.num.num',X0))))))) )),
  inference(resolution,[],[f579,f267])).
fof(f267,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X3)))))) | ~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s(fun('type.num.num','type.realax.real'),X0))))) )),
  inference(cnf_transformation,[],[f120])).
fof(f14378,plain,(
  ( ! [X97,X96] : (p(s(bool,X96)) | p(s(bool,X97)) | sP9(X97,X96,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16)))))) ) | $spl120),
  inference(cnf_transformation,[],[f14378_D])).
fof(f14378_D,plain,(
  ( ! [X97,X96] : (p(s(bool,X96)) | p(s(bool,X97)) | sP9(X97,X96,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16)))))) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f14377,plain,(
  ( ! [X97,X95,X96] : (s(bool,'T_0') != s(bool,X95) | p(s(bool,X96)) | sP9(X97,X96,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))))) | p(s(bool,X97))) ) | ($spl4 | $spl10)),
  inference(subsumption_resolution,[],[f14323,f435])).
fof(f14323,plain,(
  ( ! [X97,X95,X96] : (s(bool,'T_0') != s(bool,X95) | p(s(bool,X96)) | ~p(s(bool,X95)) | sP9(X97,X96,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))))) | p(s(bool,X97))) ) | ($spl4 | $spl10)),
  inference(superposition,[],[f4344,f3911])).
fof(f3911,plain,(
  s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))))) | $spl4),
  inference(resolution,[],[f466,f436])).
fof(f466,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16)))))) | $spl4),
  inference(cnf_transformation,[],[f466_D])).
fof(f466_D,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16)))))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f14372,plain,(
  ( ! [X70,X71] : (sP9(X71,X70,'T_0') | p(s(bool,X71)) | p(s(bool,X70))) ) | $spl118),
  inference(cnf_transformation,[],[f14372_D])).
fof(f14372_D,plain,(
  ( ! [X70,X71] : (sP9(X71,X70,'T_0') | p(s(bool,X71)) | p(s(bool,X70))) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f13647,plain,(
  ( ! [X68,X69] : (s(bool,X68) != s(bool,X69) | ~p(s(bool,X69))) ) | $spl98),
  inference(cnf_transformation,[],[f13647_D])).
fof(f13647_D,plain,(
  ( ! [X68,X69] : (s(bool,X68) != s(bool,X69) | ~p(s(bool,X69))) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f14374,plain,(
  ( ! [X78,X76,X79,X77] : (s(bool,X76) != s(bool,X77) | p(s(bool,X78)) | sP9(X79,X78,'T_0') | p(s(bool,X79)) | ~p(s(bool,X76))) ) | $spl10),
  inference(subsumption_resolution,[],[f14318,f4313])).
fof(f14318,plain,(
  ( ! [X78,X76,X79,X77] : (s(bool,X76) != s(bool,X77) | p(s(bool,X78)) | ~p(s(bool,X77)) | sP9(X79,X78,'T_0') | p(s(bool,X79)) | ~p(s(bool,X76))) ) | $spl10),
  inference(superposition,[],[f4344,f436])).
fof(f14373,plain,(
  ( ! [X74,X72,X75,X73] : (s(bool,X72) != s(bool,X73) | p(s(bool,X74)) | sP9(X75,X74,'T_0') | p(s(bool,X75)) | ~p(s(bool,X72))) ) | $spl10),
  inference(subsumption_resolution,[],[f14317,f4313])).
fof(f14317,plain,(
  ( ! [X74,X72,X75,X73] : (s(bool,X72) != s(bool,X73) | p(s(bool,X74)) | ~p(s(bool,X73)) | sP9(X75,X74,'T_0') | p(s(bool,X75)) | ~p(s(bool,X72))) ) | $spl10),
  inference(superposition,[],[f4344,f438])).
fof(f14371,plain,(
  ( ! [X70,X68,X71,X69] : (s(bool,X68) != s(bool,X69) | p(s(bool,X70)) | ~p(s(bool,X69)) | sP9(X71,X70,'T_0') | p(s(bool,X71))) ) | ($spl10 | $spl48)),
  inference(subsumption_resolution,[],[f14316,f9419])).
fof(f14316,plain,(
  ( ! [X70,X68,X71,X69] : (s(bool,X68) != s(bool,X69) | p(s(bool,X70)) | ~p(s(bool,X69)) | sP9(X71,X70,'T_0') | p(s(bool,X71)) | s(bool,'F_0') = s(bool,X68)) ) | $spl10),
  inference(superposition,[],[f4344,f456])).
fof(f14370,plain,(
  ( ! [X62,X60,X63] : (sP9(X63,X62,X60) | s(bool,'F_0') = s(bool,X60) | p(s(bool,X63)) | p(s(bool,X62))) ) | $spl116),
  inference(cnf_transformation,[],[f14370_D])).
fof(f14370_D,plain,(
  ( ! [X62,X60,X63] : (sP9(X63,X62,X60) | s(bool,'F_0') = s(bool,X60) | p(s(bool,X63)) | p(s(bool,X62))) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f14369,plain,(
  ( ! [X61,X62,X60,X63] : (s(bool,'T_0') != s(bool,X61) | p(s(bool,X62)) | sP9(X63,X62,X60) | p(s(bool,X63)) | s(bool,'F_0') = s(bool,X60)) ) | $spl10),
  inference(subsumption_resolution,[],[f14314,f435])).
fof(f14314,plain,(
  ( ! [X61,X62,X60,X63] : (s(bool,'T_0') != s(bool,X61) | p(s(bool,X62)) | ~p(s(bool,X61)) | sP9(X63,X62,X60) | p(s(bool,X63)) | s(bool,'F_0') = s(bool,X60)) ) | $spl10),
  inference(superposition,[],[f4344,f456])).
fof(f13989,plain,(
  ( ! [X26,X24] : (s(bool,'F_0') != s(bool,X26) | s(bool,X24) != s(bool,X26) | s(bool,'F_0') = s(bool,X24)) ) | $spl114),
  inference(cnf_transformation,[],[f13989_D])).
fof(f13989_D,plain,(
  ( ! [X26,X24] : (s(bool,'F_0') != s(bool,X26) | s(bool,X24) != s(bool,X26) | s(bool,'F_0') = s(bool,X24)) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f13985,plain,(
  ( ! [X26,X24,X25] : (s(bool,'F_0') = s(bool,X24) | p(s(bool,X25)) | s(bool,'F_0') != s(bool,X26) | s(bool,X24) != s(bool,X26)) ) | ($spl92 | $spl94)),
  inference(resolution,[],[f13907,f13137])).
fof(f13664,plain,(
  ( ! [X60] : (s(bool,'F_0') = s(bool,X60)) ) | $spl112),
  inference(cnf_transformation,[],[f13664_D])).
fof(f13664_D,plain,(
  ( ! [X60] : (s(bool,'F_0') = s(bool,X60)) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f13662,plain,(
  ( ! [X54,X55,X53] : (sP9(X55,X53,X54) | p(s(bool,X55)) | s(bool,'T_0') != s(bool,X53) | p(s(bool,X54))) ) | $spl110),
  inference(cnf_transformation,[],[f13662_D])).
fof(f13662_D,plain,(
  ( ! [X54,X55,X53] : (sP9(X55,X53,X54) | p(s(bool,X55)) | s(bool,'T_0') != s(bool,X53) | p(s(bool,X54))) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f13663,plain,(
  ( ! [X61,X62,X60,X63] : (s(bool,'T_0') != s(bool,X61) | p(s(bool,X62)) | sP9(X63,X61,X62) | p(s(bool,X63)) | s(bool,'F_0') = s(bool,X60)) ) | $spl10),
  inference(subsumption_resolution,[],[f13616,f432])).
fof(f13616,plain,(
  ( ! [X61,X62,X60,X63] : (s(bool,'T_0') != s(bool,X61) | p(s(bool,X62)) | ~p(s(bool,X60)) | sP9(X63,X61,X62) | p(s(bool,X63)) | s(bool,'F_0') = s(bool,X60)) ) | $spl10),
  inference(superposition,[],[f4343,f456])).
fof(f13633,plain,(
  ( ! [X59,X57,X58,X56] : (s(bool,'T_0') != s(bool,X57) | p(s(bool,X58)) | ~p(s(bool,X56)) | sP9(X59,X57,X58) | p(s(bool,X59))) ) | $spl10),
  inference(duplicate_literal_removal,[],[f13615])).
fof(f13615,plain,(
  ( ! [X59,X57,X58,X56] : (s(bool,'T_0') != s(bool,X57) | p(s(bool,X58)) | ~p(s(bool,X56)) | sP9(X59,X57,X58) | p(s(bool,X59)) | ~p(s(bool,X56))) ) | $spl10),
  inference(superposition,[],[f4343,f438])).
fof(f13634,plain,(
  ( ! [X54,X52,X55,X53] : (s(bool,'T_0') != s(bool,X53) | p(s(bool,X54)) | ~p(s(bool,X52)) | sP9(X55,X53,X54) | p(s(bool,X55))) ) | $spl10),
  inference(duplicate_literal_removal,[],[f13614])).
fof(f13614,plain,(
  ( ! [X54,X52,X55,X53] : (s(bool,'T_0') != s(bool,X53) | p(s(bool,X54)) | ~p(s(bool,X52)) | sP9(X55,X53,X54) | p(s(bool,X55)) | ~p(s(bool,X52))) ) | $spl10),
  inference(superposition,[],[f4343,f436])).
fof(f13661,plain,(
  ( ! [X114,X113] : (p(s(bool,X113)) | p(s(bool,X114)) | sP9(X114,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK14),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK12)),X113)) ) | $spl108),
  inference(cnf_transformation,[],[f13661_D])).
fof(f13661_D,plain,(
  ( ! [X114,X113] : (p(s(bool,X113)) | p(s(bool,X114)) | sP9(X114,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK14),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK12)),X113)) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f13660,plain,(
  ( ! [X114,X112,X113] : (s(bool,'T_0') != s(bool,X112) | p(s(bool,X113)) | sP9(X114,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK14),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK12)),X113) | p(s(bool,X114))) ) | $spl10),
  inference(subsumption_resolution,[],[f13602,f435])).
fof(f13602,plain,(
  ( ! [X114,X112,X113] : (s(bool,'T_0') != s(bool,X112) | p(s(bool,X113)) | ~p(s(bool,X112)) | sP9(X114,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK14),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK12)),X113) | p(s(bool,X114))) ) | $spl10),
  inference(superposition,[],[f4343,f1263])).
fof(f13659,plain,(
  ( ! [X111,X110,X108] : (p(s(bool,X110)) | p(s(bool,X111)) | sP9(X111,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK12),s('type.num.num',X108))),s('type.realax.real',sK15)),X110)) ) | $spl106),
  inference(cnf_transformation,[],[f13659_D])).
fof(f13659_D,plain,(
  ( ! [X111,X110,X108] : (p(s(bool,X110)) | p(s(bool,X111)) | sP9(X111,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK12),s('type.num.num',X108))),s('type.realax.real',sK15)),X110)) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f13658,plain,(
  ( ! [X111,X109,X110,X108] : (s(bool,'T_0') != s(bool,X109) | p(s(bool,X110)) | sP9(X111,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK12),s('type.num.num',X108))),s('type.realax.real',sK15)),X110) | p(s(bool,X111))) ) | $spl10),
  inference(subsumption_resolution,[],[f13601,f435])).
fof(f13601,plain,(
  ( ! [X111,X109,X110,X108] : (s(bool,'T_0') != s(bool,X109) | p(s(bool,X110)) | ~p(s(bool,X109)) | sP9(X111,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK12),s('type.num.num',X108))),s('type.realax.real',sK15)),X110) | p(s(bool,X111))) ) | $spl10),
  inference(superposition,[],[f4343,f1914])).
fof(f13657,plain,(
  ( ! [X103,X104] : (p(s(bool,X103)) | p(s(bool,X104)) | sP9(X104,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK12),s('type.num.num',sK16))),s('type.realax.real',sK15)),X103)) ) | $spl104),
  inference(cnf_transformation,[],[f13657_D])).
fof(f13657_D,plain,(
  ( ! [X103,X104] : (p(s(bool,X103)) | p(s(bool,X104)) | sP9(X104,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK12),s('type.num.num',sK16))),s('type.realax.real',sK15)),X103)) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f13656,plain,(
  ( ! [X103,X102,X104] : (s(bool,'T_0') != s(bool,X102) | p(s(bool,X103)) | sP9(X104,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK12),s('type.num.num',sK16))),s('type.realax.real',sK15)),X103) | p(s(bool,X104))) ) | ($spl2 | $spl10)),
  inference(subsumption_resolution,[],[f13599,f435])).
fof(f13599,plain,(
  ( ! [X103,X102,X104] : (s(bool,'T_0') != s(bool,X102) | p(s(bool,X103)) | ~p(s(bool,X102)) | sP9(X104,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK12),s('type.num.num',sK16))),s('type.realax.real',sK15)),X103) | p(s(bool,X104))) ) | ($spl2 | $spl10)),
  inference(superposition,[],[f4343,f1586])).
fof(f13655,plain,(
  ( ! [X101,X100,X98] : (p(s(bool,X100)) | p(s(bool,X101)) | sP9(X101,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK12),s('type.num.num',X98)))),X100)) ) | $spl102),
  inference(cnf_transformation,[],[f13655_D])).
fof(f13655_D,plain,(
  ( ! [X101,X100,X98] : (p(s(bool,X100)) | p(s(bool,X101)) | sP9(X101,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK12),s('type.num.num',X98)))),X100)) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f13654,plain,(
  ( ! [X101,X99,X100,X98] : (s(bool,'T_0') != s(bool,X99) | p(s(bool,X100)) | sP9(X101,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK12),s('type.num.num',X98)))),X100) | p(s(bool,X101))) ) | $spl10),
  inference(subsumption_resolution,[],[f13598,f435])).
fof(f13598,plain,(
  ( ! [X101,X99,X100,X98] : (s(bool,'T_0') != s(bool,X99) | p(s(bool,X100)) | ~p(s(bool,X99)) | sP9(X101,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK12),s('type.num.num',X98)))),X100) | p(s(bool,X101))) ) | $spl10),
  inference(superposition,[],[f4343,f1368])).
fof(f13653,plain,(
  ( ! [X97,X96] : (p(s(bool,X96)) | p(s(bool,X97)) | sP9(X97,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16)))),X96)) ) | $spl100),
  inference(cnf_transformation,[],[f13653_D])).
fof(f13653_D,plain,(
  ( ! [X97,X96] : (p(s(bool,X96)) | p(s(bool,X97)) | sP9(X97,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16)))),X96)) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f13652,plain,(
  ( ! [X97,X95,X96] : (s(bool,'T_0') != s(bool,X95) | p(s(bool,X96)) | sP9(X97,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16)))),X96) | p(s(bool,X97))) ) | ($spl4 | $spl10)),
  inference(subsumption_resolution,[],[f13597,f435])).
fof(f13597,plain,(
  ( ! [X97,X95,X96] : (s(bool,'T_0') != s(bool,X95) | p(s(bool,X96)) | ~p(s(bool,X95)) | sP9(X97,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16)))),X96) | p(s(bool,X97))) ) | ($spl4 | $spl10)),
  inference(superposition,[],[f4343,f3911])).
fof(f13646,plain,(
  ( ! [X70,X71] : (sP9(X71,'T_0',X70) | p(s(bool,X71)) | p(s(bool,X70))) ) | $spl96),
  inference(cnf_transformation,[],[f13646_D])).
fof(f13646_D,plain,(
  ( ! [X70,X71] : (sP9(X71,'T_0',X70) | p(s(bool,X71)) | p(s(bool,X70))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f13649,plain,(
  ( ! [X78,X76,X79,X77] : (s(bool,X76) != s(bool,X77) | p(s(bool,X78)) | sP9(X79,'T_0',X78) | p(s(bool,X79)) | ~p(s(bool,X76))) ) | $spl10),
  inference(subsumption_resolution,[],[f13592,f4313])).
fof(f13592,plain,(
  ( ! [X78,X76,X79,X77] : (s(bool,X76) != s(bool,X77) | p(s(bool,X78)) | ~p(s(bool,X77)) | sP9(X79,'T_0',X78) | p(s(bool,X79)) | ~p(s(bool,X76))) ) | $spl10),
  inference(superposition,[],[f4343,f436])).
fof(f13648,plain,(
  ( ! [X74,X72,X75,X73] : (s(bool,X72) != s(bool,X73) | p(s(bool,X74)) | sP9(X75,'T_0',X74) | p(s(bool,X75)) | ~p(s(bool,X72))) ) | $spl10),
  inference(subsumption_resolution,[],[f13591,f4313])).
fof(f13591,plain,(
  ( ! [X74,X72,X75,X73] : (s(bool,X72) != s(bool,X73) | p(s(bool,X74)) | ~p(s(bool,X73)) | sP9(X75,'T_0',X74) | p(s(bool,X75)) | ~p(s(bool,X72))) ) | $spl10),
  inference(superposition,[],[f4343,f438])).
fof(f13645,plain,(
  ( ! [X70,X68,X71,X69] : (s(bool,X68) != s(bool,X69) | p(s(bool,X70)) | ~p(s(bool,X69)) | sP9(X71,'T_0',X70) | p(s(bool,X71))) ) | ($spl10 | $spl48)),
  inference(subsumption_resolution,[],[f13590,f9419])).
fof(f13590,plain,(
  ( ! [X70,X68,X71,X69] : (s(bool,X68) != s(bool,X69) | p(s(bool,X70)) | ~p(s(bool,X69)) | sP9(X71,'T_0',X70) | p(s(bool,X71)) | s(bool,'F_0') = s(bool,X68)) ) | $spl10),
  inference(superposition,[],[f4343,f456])).
fof(f13644,plain,(
  ( ! [X62,X60,X63] : (sP9(X63,X60,X62) | s(bool,'F_0') = s(bool,X60) | p(s(bool,X63)) | p(s(bool,X62))) ) | $spl94),
  inference(cnf_transformation,[],[f13644_D])).
fof(f13644_D,plain,(
  ( ! [X62,X60,X63] : (sP9(X63,X60,X62) | s(bool,'F_0') = s(bool,X60) | p(s(bool,X63)) | p(s(bool,X62))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f13643,plain,(
  ( ! [X61,X62,X60,X63] : (s(bool,'T_0') != s(bool,X61) | p(s(bool,X62)) | sP9(X63,X60,X62) | p(s(bool,X63)) | s(bool,'F_0') = s(bool,X60)) ) | $spl10),
  inference(subsumption_resolution,[],[f13588,f435])).
fof(f13588,plain,(
  ( ! [X61,X62,X60,X63] : (s(bool,'T_0') != s(bool,X61) | p(s(bool,X62)) | ~p(s(bool,X61)) | sP9(X63,X60,X62) | p(s(bool,X63)) | s(bool,'F_0') = s(bool,X60)) ) | $spl10),
  inference(superposition,[],[f4343,f456])).
fof(f12748,plain,(
  ( ! [X47,X45,X46] : (sP9(X47,X46,X45) | p(s(bool,X47)) | s(bool,'F_0') != s(bool,X45) | ~p(s(bool,X46))) ) | $spl92),
  inference(cnf_transformation,[],[f12748_D])).
fof(f12748_D,plain,(
  ( ! [X47,X45,X46] : (sP9(X47,X46,X45) | p(s(bool,X47)) | s(bool,'F_0') != s(bool,X45) | ~p(s(bool,X46))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f12750,plain,(
  ( ! [X66,X64,X67,X65] : (s(bool,'F_0') != s(bool,X65) | p(s(bool,X64)) | ~p(s(bool,X66)) | sP9(X67,X66,X65) | p(s(bool,X67))) ) | $spl10),
  inference(subsumption_resolution,[],[f12713,f435])).
fof(f12713,plain,(
  ( ! [X66,X64,X67,X65] : (s(bool,'F_0') != s(bool,X65) | p(s(bool,X64)) | ~p(s(bool,X66)) | sP9(X67,X66,X65) | p(s(bool,X67)) | s(bool,'T_0') = s(bool,X64)) ) | $spl10),
  inference(superposition,[],[f4323,f456])).
fof(f4323,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X5) != s(bool,X6) | p(s(bool,X6)) | ~p(s(bool,X4)) | sP9(X7,X4,X5) | p(s(bool,X7))) ) | $spl10),
  inference(superposition,[],[f4313,f357])).
fof(f12729,plain,(
  ( ! [X50,X48,X51,X49] : (s(bool,'F_0') != s(bool,X49) | p(s(bool,X48)) | ~p(s(bool,X50)) | sP9(X51,X50,X49) | p(s(bool,X51))) ) | $spl10),
  inference(duplicate_literal_removal,[],[f12709])).
fof(f12709,plain,(
  ( ! [X50,X48,X51,X49] : (s(bool,'F_0') != s(bool,X49) | p(s(bool,X48)) | ~p(s(bool,X50)) | sP9(X51,X50,X49) | p(s(bool,X51)) | p(s(bool,X48))) ) | $spl10),
  inference(superposition,[],[f4323,f434])).
fof(f12730,plain,(
  ( ! [X47,X45,X46,X44] : (s(bool,'F_0') != s(bool,X45) | p(s(bool,X44)) | ~p(s(bool,X46)) | sP9(X47,X46,X45) | p(s(bool,X47))) ) | $spl10),
  inference(duplicate_literal_removal,[],[f12708])).
fof(f12708,plain,(
  ( ! [X47,X45,X46,X44] : (s(bool,'F_0') != s(bool,X45) | p(s(bool,X44)) | ~p(s(bool,X46)) | sP9(X47,X46,X45) | p(s(bool,X47)) | p(s(bool,X44))) ) | $spl10),
  inference(superposition,[],[f4323,f432])).
fof(f12747,plain,(
  ( ! [X118,X117,X115] : (~p(s(bool,X117)) | p(s(bool,X118)) | sP9(X118,X117,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X115),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK13)))) ) | $spl90),
  inference(cnf_transformation,[],[f12747_D])).
fof(f12747_D,plain,(
  ( ! [X118,X117,X115] : (~p(s(bool,X117)) | p(s(bool,X118)) | sP9(X118,X117,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X115),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK13)))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f12746,plain,(
  ( ! [X118,X116,X117,X115] : (p(s(bool,X116)) | ~p(s(bool,X117)) | sP9(X118,X117,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X115),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK13))) | p(s(bool,X118))) ) | ($spl7 | $spl10)),
  inference(subsumption_resolution,[],[f12699,f432])).
fof(f12699,plain,(
  ( ! [X118,X116,X117,X115] : (s(bool,'F_0') != s(bool,X116) | p(s(bool,X116)) | ~p(s(bool,X117)) | sP9(X118,X117,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X115),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK13))) | p(s(bool,X118))) ) | ($spl7 | $spl10)),
  inference(superposition,[],[f4323,f4244])).
fof(f4244,plain,(
  ( ! [X215] : (s(bool,'F_0') = s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X215),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK13)))) ) | $spl7),
  inference(resolution,[],[f3934,f432])).
fof(f3934,plain,(
  ( ! [X1] : (~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK13))))) ) | $spl7),
  inference(resolution,[],[f469,f268])).
fof(f12745,plain,(
  ( ! [X107,X106] : (~p(s(bool,X106)) | p(s(bool,X107)) | sP9(X107,X106,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15)))) ) | $spl88),
  inference(cnf_transformation,[],[f12745_D])).
fof(f12745_D,plain,(
  ( ! [X107,X106] : (~p(s(bool,X106)) | p(s(bool,X107)) | sP9(X107,X106,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15)))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f12744,plain,(
  ( ! [X107,X105,X106] : (p(s(bool,X105)) | ~p(s(bool,X106)) | sP9(X107,X106,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15))) | p(s(bool,X107))) ) | ($spl7 | $spl10)),
  inference(subsumption_resolution,[],[f12696,f432])).
fof(f12696,plain,(
  ( ! [X107,X105,X106] : (s(bool,'F_0') != s(bool,X105) | p(s(bool,X105)) | ~p(s(bool,X106)) | sP9(X107,X106,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15))) | p(s(bool,X107))) ) | ($spl7 | $spl10)),
  inference(superposition,[],[f4323,f4014])).
fof(f4014,plain,(
  s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15))) | $spl7),
  inference(resolution,[],[f469,f432])).
fof(f12741,plain,(
  ( ! [X83,X82] : (sP9(X83,X82,'F_0') | p(s(bool,X83)) | ~p(s(bool,X82))) ) | $spl86),
  inference(cnf_transformation,[],[f12741_D])).
fof(f12741_D,plain,(
  ( ! [X83,X82] : (sP9(X83,X82,'F_0') | p(s(bool,X83)) | ~p(s(bool,X82))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f11838,plain,(
  ( ! [X26,X25] : (s(bool,X25) != s(bool,X26) | p(s(bool,X26))) ) | $spl68),
  inference(cnf_transformation,[],[f11838_D])).
fof(f11838_D,plain,(
  ( ! [X26,X25] : (s(bool,X25) != s(bool,X26) | p(s(bool,X26))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f12743,plain,(
  ( ! [X90,X88,X91,X89] : (s(bool,X88) != s(bool,X89) | p(s(bool,X89)) | ~p(s(bool,X90)) | sP9(X91,X90,'F_0') | p(s(bool,X91))) ) | $spl10),
  inference(subsumption_resolution,[],[f12691,f4313])).
fof(f12691,plain,(
  ( ! [X90,X88,X91,X89] : (s(bool,X88) != s(bool,X89) | p(s(bool,X89)) | ~p(s(bool,X90)) | sP9(X91,X90,'F_0') | p(s(bool,X91)) | p(s(bool,X88))) ) | $spl10),
  inference(superposition,[],[f4323,f432])).
fof(f12742,plain,(
  ( ! [X87,X85,X86,X84] : (s(bool,X84) != s(bool,X85) | p(s(bool,X85)) | ~p(s(bool,X86)) | sP9(X87,X86,'F_0') | p(s(bool,X87))) ) | $spl10),
  inference(subsumption_resolution,[],[f12690,f4313])).
fof(f12690,plain,(
  ( ! [X87,X85,X86,X84] : (s(bool,X84) != s(bool,X85) | p(s(bool,X85)) | ~p(s(bool,X86)) | sP9(X87,X86,'F_0') | p(s(bool,X87)) | p(s(bool,X84))) ) | $spl10),
  inference(superposition,[],[f4323,f434])).
fof(f11846,plain,(
  ( ! [X80,X81] : (s(bool,X80) != s(bool,X81) | s(bool,'T_0') = s(bool,X80)) ) | $spl74),
  inference(cnf_transformation,[],[f11846_D])).
fof(f11846_D,plain,(
  ( ! [X80,X81] : (s(bool,X80) != s(bool,X81) | s(bool,'T_0') = s(bool,X80)) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f12740,plain,(
  ( ! [X80,X83,X81,X82] : (s(bool,X80) != s(bool,X81) | ~p(s(bool,X82)) | sP9(X83,X82,'F_0') | p(s(bool,X83)) | s(bool,'T_0') = s(bool,X80)) ) | $spl10),
  inference(subsumption_resolution,[],[f12689,f4384])).
fof(f4384,plain,(
  ( ! [X47,X46] : (s(bool,X46) != s(bool,X47) | ~p(s(bool,X47)) | s(bool,'T_0') = s(bool,X46)) ) | $spl10),
  inference(subsumption_resolution,[],[f4360,f417])).
fof(f4360,plain,(
  ( ! [X47,X46] : (s(bool,X46) != s(bool,X47) | p(s(bool,'F_0')) | ~p(s(bool,X47)) | s(bool,'T_0') = s(bool,X46)) ) | $spl10),
  inference(superposition,[],[f4313,f456])).
fof(f12689,plain,(
  ( ! [X80,X83,X81,X82] : (s(bool,X80) != s(bool,X81) | p(s(bool,X81)) | ~p(s(bool,X82)) | sP9(X83,X82,'F_0') | p(s(bool,X83)) | s(bool,'T_0') = s(bool,X80)) ) | $spl10),
  inference(superposition,[],[f4323,f456])).
fof(f12739,plain,(
  ( ! [X66,X64,X67] : (sP9(X67,X66,X64) | s(bool,'T_0') = s(bool,X64) | p(s(bool,X67)) | ~p(s(bool,X66))) ) | $spl84),
  inference(cnf_transformation,[],[f12739_D])).
fof(f12739_D,plain,(
  ( ! [X66,X64,X67] : (sP9(X67,X66,X64) | s(bool,'T_0') = s(bool,X64) | p(s(bool,X67)) | ~p(s(bool,X66))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f12738,plain,(
  ( ! [X66,X64,X67,X65] : (p(s(bool,X65)) | ~p(s(bool,X66)) | sP9(X67,X66,X64) | p(s(bool,X67)) | s(bool,'T_0') = s(bool,X64)) ) | $spl10),
  inference(subsumption_resolution,[],[f12685,f432])).
fof(f12685,plain,(
  ( ! [X66,X64,X67,X65] : (s(bool,'F_0') != s(bool,X65) | p(s(bool,X65)) | ~p(s(bool,X66)) | sP9(X67,X66,X64) | p(s(bool,X67)) | s(bool,'T_0') = s(bool,X64)) ) | $spl10),
  inference(superposition,[],[f4323,f456])).
fof(f12734,plain,(
  ( ! [X28,X24,X27] : (~p(s(bool,X27)) | p(s(bool,X24)) | p(s(bool,X28)) | sP9(X28,X27,X24)) ) | $spl82),
  inference(cnf_transformation,[],[f12734_D])).
fof(f12734_D,plain,(
  ( ! [X28,X24,X27] : (~p(s(bool,X27)) | p(s(bool,X24)) | p(s(bool,X28)) | sP9(X28,X27,X24)) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f12737,plain,(
  ( ! [X50,X48,X51,X49] : (p(s(bool,X49)) | ~p(s(bool,X50)) | sP9(X51,X50,X48) | p(s(bool,X51)) | p(s(bool,X48))) ) | $spl10),
  inference(subsumption_resolution,[],[f12681,f432])).
fof(f12681,plain,(
  ( ! [X50,X48,X51,X49] : (s(bool,'F_0') != s(bool,X49) | p(s(bool,X49)) | ~p(s(bool,X50)) | sP9(X51,X50,X48) | p(s(bool,X51)) | p(s(bool,X48))) ) | $spl10),
  inference(superposition,[],[f4323,f434])).
fof(f12736,plain,(
  ( ! [X47,X45,X46,X44] : (p(s(bool,X45)) | ~p(s(bool,X46)) | sP9(X47,X46,X44) | p(s(bool,X47)) | p(s(bool,X44))) ) | $spl10),
  inference(subsumption_resolution,[],[f12680,f432])).
fof(f12680,plain,(
  ( ! [X47,X45,X46,X44] : (s(bool,'F_0') != s(bool,X45) | p(s(bool,X45)) | ~p(s(bool,X46)) | sP9(X47,X46,X44) | p(s(bool,X47)) | p(s(bool,X44))) ) | $spl10),
  inference(superposition,[],[f4323,f432])).
fof(f12735,plain,(
  ( ! [X30,X33,X31,X29,X32] : (s(bool,X30) != s(bool,X31) | p(s(bool,X31)) | ~p(s(bool,X32)) | sP9(X33,X32,X29) | p(s(bool,X33)) | p(s(bool,X29))) ) | $spl10),
  inference(subsumption_resolution,[],[f12677,f4313])).
fof(f12677,plain,(
  ( ! [X30,X33,X31,X29,X32] : (s(bool,X30) != s(bool,X31) | p(s(bool,X31)) | ~p(s(bool,X32)) | sP9(X33,X32,X29) | p(s(bool,X33)) | p(s(bool,X29)) | p(s(bool,X30))) ) | $spl10),
  inference(superposition,[],[f4323,f404])).
fof(f404,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f130])).
fof(f130,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f129])).
fof(f129,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X11,X12] : ((p(s(bool,X11)) => p(s(bool,X12))) => ((p(s(bool,X12)) => p(s(bool,X11))) => s(bool,X11) = s(bool,X12)))),
  file('thm.integral.TDIV_BOUNDS','thm.bool.IMP_ANTISYM_AX')).
fof(f12733,plain,(
  ( ! [X28,X26,X24,X27,X25] : (s(bool,X25) != s(bool,X26) | p(s(bool,X26)) | ~p(s(bool,X27)) | sP9(X28,X27,X24) | p(s(bool,X28)) | p(s(bool,X24))) ) | $spl10),
  inference(subsumption_resolution,[],[f12676,f4313])).
fof(f12676,plain,(
  ( ! [X28,X26,X24,X27,X25] : (s(bool,X25) != s(bool,X26) | p(s(bool,X26)) | ~p(s(bool,X27)) | sP9(X28,X27,X24) | p(s(bool,X28)) | p(s(bool,X25)) | p(s(bool,X24))) ) | $spl10),
  inference(superposition,[],[f4323,f404])).
fof(f12393,plain,(
  ( ! [X36] : (s(bool,'T_0') = s(bool,X36)) ) | $spl80),
  inference(cnf_transformation,[],[f12393_D])).
fof(f12393_D,plain,(
  ( ! [X36] : (s(bool,'T_0') = s(bool,X36)) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f12387,plain,(
  ( ! [X39,X38] : (s(bool,'F_0') != s(bool,X39) | s(bool,'T_0') = s(bool,X38)) ) | $spl74),
  inference(duplicate_literal_removal,[],[f12343])).
fof(f12343,plain,(
  ( ! [X39,X38] : (s(bool,'F_0') != s(bool,X39) | s(bool,'T_0') = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) ) | $spl74),
  inference(superposition,[],[f11846,f456])).
fof(f12392,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,X37) | s(bool,'T_0') = s(bool,X36)) ) | $spl74),
  inference(subsumption_resolution,[],[f12342,f11846])).
fof(f12342,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,X37) | s(bool,'T_0') = s(bool,X36) | s(bool,'F_0') = s(bool,X36)) ) | $spl74),
  inference(superposition,[],[f11846,f456])).
fof(f12391,plain,(
  ( ! [X30,X31] : (s(bool,'F_0') != s(bool,X31) | p(s(bool,X30))) ) | $spl74),
  inference(subsumption_resolution,[],[f12339,f435])).
fof(f12339,plain,(
  ( ! [X30,X31] : (s(bool,'F_0') != s(bool,X31) | s(bool,'T_0') = s(bool,X30) | p(s(bool,X30))) ) | $spl74),
  inference(superposition,[],[f11846,f434])).
fof(f12390,plain,(
  ( ! [X28,X29] : (s(bool,'F_0') != s(bool,X29) | p(s(bool,X28))) ) | $spl74),
  inference(subsumption_resolution,[],[f12338,f435])).
fof(f12338,plain,(
  ( ! [X28,X29] : (s(bool,'F_0') != s(bool,X29) | s(bool,'T_0') = s(bool,X28) | p(s(bool,X28))) ) | $spl74),
  inference(superposition,[],[f11846,f432])).
fof(f12389,plain,(
  ( ! [X21,X19,X20] : (s(bool,X20) != s(bool,X21) | p(s(bool,X19)) | p(s(bool,X20))) ) | $spl74),
  inference(subsumption_resolution,[],[f12335,f435])).
fof(f12335,plain,(
  ( ! [X21,X19,X20] : (s(bool,X20) != s(bool,X21) | s(bool,'T_0') = s(bool,X19) | p(s(bool,X19)) | p(s(bool,X20))) ) | $spl74),
  inference(superposition,[],[f11846,f404])).
fof(f12388,plain,(
  ( ! [X17,X18,X16] : (s(bool,X17) != s(bool,X18) | p(s(bool,X17)) | p(s(bool,X16))) ) | $spl74),
  inference(subsumption_resolution,[],[f12334,f435])).
fof(f12334,plain,(
  ( ! [X17,X18,X16] : (s(bool,X17) != s(bool,X18) | s(bool,'T_0') = s(bool,X16) | p(s(bool,X17)) | p(s(bool,X16))) ) | $spl74),
  inference(superposition,[],[f11846,f404])).
fof(f12281,plain,(
  $false | ($spl7 | $spl56)),
  inference(trivial_inequality_removal,[],[f12275])).
fof(f12275,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl7 | $spl56)),
  inference(superposition,[],[f10210,f4014])).
fof(f12280,plain,(
  $false | ($spl7 | $spl56)),
  inference(trivial_inequality_removal,[],[f12278])).
fof(f12278,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl7 | $spl56)),
  inference(superposition,[],[f10210,f4244])).
fof(f12279,plain,(
  $false | $spl56),
  inference(equality_resolution,[],[f10210])).
fof(f12228,plain,(
  $false | ($spl0 | $spl38)),
  inference(trivial_inequality_removal,[],[f12214])).
fof(f12214,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl0 | $spl38)),
  inference(superposition,[],[f4319,f934])).
fof(f934,plain,(
  s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK12),s('type.num.num',sK16))))) | $spl0),
  inference(resolution,[],[f462,f436])).
fof(f12227,plain,(
  $false | ($spl4 | $spl38)),
  inference(trivial_inequality_removal,[],[f12215])).
fof(f12215,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl4 | $spl38)),
  inference(superposition,[],[f4319,f3911])).
fof(f12226,plain,(
  $false | $spl38),
  inference(trivial_inequality_removal,[],[f12216])).
fof(f12216,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl38),
  inference(superposition,[],[f4319,f1368])).
fof(f12225,plain,(
  $false | ($spl2 | $spl38)),
  inference(trivial_inequality_removal,[],[f12217])).
fof(f12217,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl2 | $spl38)),
  inference(superposition,[],[f4319,f1586])).
fof(f12224,plain,(
  $false | $spl38),
  inference(trivial_inequality_removal,[],[f12219])).
fof(f12219,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl38),
  inference(superposition,[],[f4319,f1914])).
fof(f12223,plain,(
  $false | $spl38),
  inference(trivial_inequality_removal,[],[f12220])).
fof(f12220,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl38),
  inference(superposition,[],[f4319,f1263])).
fof(f12222,plain,(
  $false | $spl38),
  inference(equality_resolution,[],[f4319])).
fof(f12178,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f12092,f11917])).
fof(f11917,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl68),
  inference(equality_resolution,[],[f11838])).
fof(f12092,plain,(
  ~p(s(bool,'T_0')) | $spl68),
  inference(resolution,[],[f11917,f451])).
fof(f451,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f258])).
fof(f258,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f114])).
fof(f114,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f19])).
fof(f19,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.integral.TDIV_BOUNDS','thm.bool.NOT_CLAUSES|split|2')).
fof(f12177,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f12091,f11917])).
fof(f12091,plain,(
  ~p(s(bool,'T_0')) | $spl68),
  inference(resolution,[],[f11917,f449])).
fof(f449,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f257])).
fof(f257,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f113])).
fof(f113,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f18])).
fof(f18,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.integral.TDIV_BOUNDS','thm.bool.NOT_CLAUSES|split|1')).
fof(f12176,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f12089,f11917])).
fof(f12089,plain,(
  ( ! [X230] : (~p(s(bool,X230))) ) | $spl68),
  inference(resolution,[],[f11917,f445])).
fof(f445,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f140])).
fof(f140,plain,(
  ! [X0] : ((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) => ~p(s(bool,X0)))),
  inference(flattening,[],[f108])).
fof(f108,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) => ~p(s(bool,X0)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X14] : ((p(s(bool,X14)) => p(s(bool,'F_0'))) => ~p(s(bool,X14)))),
  file('thm.integral.TDIV_BOUNDS','thm.bool.IMP_F')).
fof(f12175,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f12087,f11917])).
fof(f12087,plain,(
  ( ! [X228] : (~p(s(bool,X228))) ) | $spl68),
  inference(resolution,[],[f11917,f417])).
fof(f12090,plain,(
  $false | $spl68),
  inference(resolution,[],[f11917,f447])).
fof(f11923,plain,(
  ( ! [X39,X38] : (s(bool,'F_0') != s(bool,X39) | p(s(bool,X38))) ) | $spl68),
  inference(subsumption_resolution,[],[f11902,f11838])).
fof(f11902,plain,(
  ( ! [X39,X38] : (s(bool,'F_0') != s(bool,X39) | p(s(bool,X38)) | s(bool,'T_0') = s(bool,X38)) ) | $spl68),
  inference(superposition,[],[f11838,f456])).
fof(f11922,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,X37) | p(s(bool,X36))) ) | $spl68),
  inference(subsumption_resolution,[],[f11901,f11838])).
fof(f11901,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,X37) | p(s(bool,X36)) | s(bool,'F_0') = s(bool,X36)) ) | $spl68),
  inference(superposition,[],[f11838,f456])).
fof(f11918,plain,(
  ( ! [X30,X31] : (s(bool,'F_0') != s(bool,X31) | p(s(bool,X30))) ) | $spl68),
  inference(duplicate_literal_removal,[],[f11898])).
fof(f11898,plain,(
  ( ! [X30,X31] : (s(bool,'F_0') != s(bool,X31) | p(s(bool,X30)) | p(s(bool,X30))) ) | $spl68),
  inference(superposition,[],[f11838,f434])).
fof(f11919,plain,(
  ( ! [X28,X29] : (s(bool,'F_0') != s(bool,X29) | p(s(bool,X28))) ) | $spl68),
  inference(duplicate_literal_removal,[],[f11897])).
fof(f11897,plain,(
  ( ! [X28,X29] : (s(bool,'F_0') != s(bool,X29) | p(s(bool,X28)) | p(s(bool,X28))) ) | $spl68),
  inference(superposition,[],[f11838,f432])).
fof(f11852,plain,(
  ( ! [X118,X117,X115] : (~p(s(bool,X117)) | p(s(bool,X118)) | sP9(X118,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X115),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK13)),X117)) ) | $spl78),
  inference(cnf_transformation,[],[f11852_D])).
fof(f11852_D,plain,(
  ( ! [X118,X117,X115] : (~p(s(bool,X117)) | p(s(bool,X118)) | sP9(X118,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X115),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK13)),X117)) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f11851,plain,(
  ( ! [X118,X116,X117,X115] : (p(s(bool,X116)) | ~p(s(bool,X117)) | sP9(X118,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X115),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK13)),X117) | p(s(bool,X118))) ) | ($spl7 | $spl10)),
  inference(subsumption_resolution,[],[f11802,f432])).
fof(f11802,plain,(
  ( ! [X118,X116,X117,X115] : (s(bool,'F_0') != s(bool,X116) | p(s(bool,X116)) | ~p(s(bool,X117)) | sP9(X118,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X115),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK13)),X117) | p(s(bool,X118))) ) | ($spl7 | $spl10)),
  inference(superposition,[],[f4322,f4244])).
fof(f4322,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X1) != s(bool,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | sP9(X3,X1,X0) | p(s(bool,X3))) ) | $spl10),
  inference(superposition,[],[f4313,f357])).
fof(f11850,plain,(
  ( ! [X107,X106] : (~p(s(bool,X106)) | p(s(bool,X107)) | sP9(X107,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15)),X106)) ) | $spl76),
  inference(cnf_transformation,[],[f11850_D])).
fof(f11850_D,plain,(
  ( ! [X107,X106] : (~p(s(bool,X106)) | p(s(bool,X107)) | sP9(X107,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15)),X106)) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f11849,plain,(
  ( ! [X107,X105,X106] : (p(s(bool,X105)) | ~p(s(bool,X106)) | sP9(X107,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15)),X106) | p(s(bool,X107))) ) | ($spl7 | $spl10)),
  inference(subsumption_resolution,[],[f11799,f432])).
fof(f11799,plain,(
  ( ! [X107,X105,X106] : (s(bool,'F_0') != s(bool,X105) | p(s(bool,X105)) | ~p(s(bool,X106)) | sP9(X107,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15)),X106) | p(s(bool,X107))) ) | ($spl7 | $spl10)),
  inference(superposition,[],[f4322,f4014])).
fof(f11845,plain,(
  ( ! [X83,X82] : (sP9(X83,'F_0',X82) | p(s(bool,X83)) | ~p(s(bool,X82))) ) | $spl72),
  inference(cnf_transformation,[],[f11845_D])).
fof(f11845_D,plain,(
  ( ! [X83,X82] : (sP9(X83,'F_0',X82) | p(s(bool,X83)) | ~p(s(bool,X82))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f11848,plain,(
  ( ! [X90,X88,X91,X89] : (s(bool,X88) != s(bool,X89) | p(s(bool,X89)) | ~p(s(bool,X90)) | sP9(X91,'F_0',X90) | p(s(bool,X91))) ) | $spl10),
  inference(subsumption_resolution,[],[f11794,f4313])).
fof(f11794,plain,(
  ( ! [X90,X88,X91,X89] : (s(bool,X88) != s(bool,X89) | p(s(bool,X89)) | ~p(s(bool,X90)) | sP9(X91,'F_0',X90) | p(s(bool,X91)) | p(s(bool,X88))) ) | $spl10),
  inference(superposition,[],[f4322,f432])).
fof(f11847,plain,(
  ( ! [X87,X85,X86,X84] : (s(bool,X84) != s(bool,X85) | p(s(bool,X85)) | ~p(s(bool,X86)) | sP9(X87,'F_0',X86) | p(s(bool,X87))) ) | $spl10),
  inference(subsumption_resolution,[],[f11793,f4313])).
fof(f11793,plain,(
  ( ! [X87,X85,X86,X84] : (s(bool,X84) != s(bool,X85) | p(s(bool,X85)) | ~p(s(bool,X86)) | sP9(X87,'F_0',X86) | p(s(bool,X87)) | p(s(bool,X84))) ) | $spl10),
  inference(superposition,[],[f4322,f434])).
fof(f11844,plain,(
  ( ! [X80,X83,X81,X82] : (s(bool,X80) != s(bool,X81) | ~p(s(bool,X82)) | sP9(X83,'F_0',X82) | p(s(bool,X83)) | s(bool,'T_0') = s(bool,X80)) ) | $spl10),
  inference(subsumption_resolution,[],[f11792,f4384])).
fof(f11792,plain,(
  ( ! [X80,X83,X81,X82] : (s(bool,X80) != s(bool,X81) | p(s(bool,X81)) | ~p(s(bool,X82)) | sP9(X83,'F_0',X82) | p(s(bool,X83)) | s(bool,'T_0') = s(bool,X80)) ) | $spl10),
  inference(superposition,[],[f4322,f456])).
fof(f11843,plain,(
  ( ! [X66,X64,X67] : (sP9(X67,X64,X66) | s(bool,'T_0') = s(bool,X64) | p(s(bool,X67)) | ~p(s(bool,X66))) ) | $spl70),
  inference(cnf_transformation,[],[f11843_D])).
fof(f11843_D,plain,(
  ( ! [X66,X64,X67] : (sP9(X67,X64,X66) | s(bool,'T_0') = s(bool,X64) | p(s(bool,X67)) | ~p(s(bool,X66))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f11842,plain,(
  ( ! [X66,X64,X67,X65] : (p(s(bool,X65)) | ~p(s(bool,X66)) | sP9(X67,X64,X66) | p(s(bool,X67)) | s(bool,'T_0') = s(bool,X64)) ) | $spl10),
  inference(subsumption_resolution,[],[f11788,f432])).
fof(f11788,plain,(
  ( ! [X66,X64,X67,X65] : (s(bool,'F_0') != s(bool,X65) | p(s(bool,X65)) | ~p(s(bool,X66)) | sP9(X67,X64,X66) | p(s(bool,X67)) | s(bool,'T_0') = s(bool,X64)) ) | $spl10),
  inference(superposition,[],[f4322,f456])).
fof(f11837,plain,(
  ( ! [X28,X24,X27] : (~p(s(bool,X27)) | p(s(bool,X24)) | p(s(bool,X28)) | sP9(X28,X24,X27)) ) | $spl66),
  inference(cnf_transformation,[],[f11837_D])).
fof(f11837_D,plain,(
  ( ! [X28,X24,X27] : (~p(s(bool,X27)) | p(s(bool,X24)) | p(s(bool,X28)) | sP9(X28,X24,X27)) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f11841,plain,(
  ( ! [X50,X48,X51,X49] : (p(s(bool,X49)) | ~p(s(bool,X50)) | sP9(X51,X48,X50) | p(s(bool,X51)) | p(s(bool,X48))) ) | $spl10),
  inference(subsumption_resolution,[],[f11784,f432])).
fof(f11784,plain,(
  ( ! [X50,X48,X51,X49] : (s(bool,'F_0') != s(bool,X49) | p(s(bool,X49)) | ~p(s(bool,X50)) | sP9(X51,X48,X50) | p(s(bool,X51)) | p(s(bool,X48))) ) | $spl10),
  inference(superposition,[],[f4322,f434])).
fof(f11840,plain,(
  ( ! [X47,X45,X46,X44] : (p(s(bool,X45)) | ~p(s(bool,X46)) | sP9(X47,X44,X46) | p(s(bool,X47)) | p(s(bool,X44))) ) | $spl10),
  inference(subsumption_resolution,[],[f11783,f432])).
fof(f11783,plain,(
  ( ! [X47,X45,X46,X44] : (s(bool,'F_0') != s(bool,X45) | p(s(bool,X45)) | ~p(s(bool,X46)) | sP9(X47,X44,X46) | p(s(bool,X47)) | p(s(bool,X44))) ) | $spl10),
  inference(superposition,[],[f4322,f432])).
fof(f11839,plain,(
  ( ! [X30,X33,X31,X29,X32] : (s(bool,X30) != s(bool,X31) | p(s(bool,X31)) | ~p(s(bool,X32)) | sP9(X33,X29,X32) | p(s(bool,X33)) | p(s(bool,X29))) ) | $spl10),
  inference(subsumption_resolution,[],[f11780,f4313])).
fof(f11780,plain,(
  ( ! [X30,X33,X31,X29,X32] : (s(bool,X30) != s(bool,X31) | p(s(bool,X31)) | ~p(s(bool,X32)) | sP9(X33,X29,X32) | p(s(bool,X33)) | p(s(bool,X29)) | p(s(bool,X30))) ) | $spl10),
  inference(superposition,[],[f4322,f404])).
fof(f11836,plain,(
  ( ! [X28,X26,X24,X27,X25] : (s(bool,X25) != s(bool,X26) | p(s(bool,X26)) | ~p(s(bool,X27)) | sP9(X28,X24,X27) | p(s(bool,X28)) | p(s(bool,X24))) ) | $spl10),
  inference(subsumption_resolution,[],[f11779,f4313])).
fof(f11779,plain,(
  ( ! [X28,X26,X24,X27,X25] : (s(bool,X25) != s(bool,X26) | p(s(bool,X26)) | ~p(s(bool,X27)) | sP9(X28,X24,X27) | p(s(bool,X28)) | p(s(bool,X25)) | p(s(bool,X24))) ) | $spl10),
  inference(superposition,[],[f4322,f404])).
fof(f11773,plain,(
  ( ! [X16] : (~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK14),s('type.realax.real',X16))),s(fun('type.num.num','type.realax.real'),sK13))))) ) | $spl64),
  inference(cnf_transformation,[],[f11773_D])).
fof(f11773_D,plain,(
  ( ! [X16] : (~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK14),s('type.realax.real',X16))),s(fun('type.num.num','type.realax.real'),sK13))))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f11774,plain,(
  ( ! [X19,X18] : (~p(s(bool,X19)) | ~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK14),s('type.realax.real',X18))),s(fun('type.num.num','type.realax.real'),sK13))))) ) | $spl16),
  inference(subsumption_resolution,[],[f11762,f436])).
fof(f11762,plain,(
  ( ! [X19,X18] : (s(bool,'T_0') != s(bool,X19) | ~p(s(bool,X19)) | ~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK14),s('type.realax.real',X18))),s(fun('type.num.num','type.realax.real'),sK13))))) ) | $spl16),
  inference(superposition,[],[f3003,f407])).
fof(f407,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f130])).
fof(f11772,plain,(
  ( ! [X17,X16] : (~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK14),s('type.realax.real',X16))),s(fun('type.num.num','type.realax.real'),sK13)))) | ~p(s(bool,X17))) ) | $spl16),
  inference(subsumption_resolution,[],[f11761,f436])).
fof(f11761,plain,(
  ( ! [X17,X16] : (s(bool,'T_0') != s(bool,X17) | ~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK14),s('type.realax.real',X16))),s(fun('type.num.num','type.realax.real'),sK13)))) | ~p(s(bool,X17))) ) | $spl16),
  inference(superposition,[],[f3003,f407])).
fof(f10565,plain,(
  ( ! [X72] : (s(bool,'F_0') != s(bool,X72) | s(bool,X72) != s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK14),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK12)))) ) | $spl62),
  inference(cnf_transformation,[],[f10565_D])).
fof(f10565_D,plain,(
  ( ! [X72] : (s(bool,'F_0') != s(bool,X72) | s(bool,X72) != s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK14),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK12)))) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f10543,plain,(
  ( ! [X72,X71] : (p(s(bool,X71)) | s(bool,'F_0') != s(bool,X72) | s(bool,X72) != s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK14),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK12)))) ) | $spl10),
  inference(resolution,[],[f9656,f1241])).
fof(f1241,plain,(
  ( ! [X121,X122] : (sP6('const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK14),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK12)),X121,X122)) )),
  inference(resolution,[],[f579,f356])).
fof(f10559,plain,(
  ( ! [X53] : (s(bool,X53) != s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK14),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK12)))) ) | $spl60),
  inference(cnf_transformation,[],[f10559_D])).
fof(f10559_D,plain,(
  ( ! [X53] : (s(bool,X53) != s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK14),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK12)))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f3141,plain,(
  s(bool,'F_0') != s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK14),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK12))) | $spl25),
  inference(cnf_transformation,[],[f3141_D])).
fof(f3141_D,plain,(
  s(bool,'F_0') != s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK14),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK12))) <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f10535,plain,(
  ( ! [X52,X53] : (p(s(bool,X52)) | s(bool,'F_0') != s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK14),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK12))) | s(bool,X53) != s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK14),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK12)))) ) | $spl10),
  inference(resolution,[],[f9656,f1240])).
fof(f1240,plain,(
  ( ! [X120,X119] : (sP6(X119,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK14),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK12)),X120)) )),
  inference(resolution,[],[f579,f355])).
fof(f10135,plain,(
  ( ! [X15,X16] : (s(bool,X15) != s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X16),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK13))) | s(bool,'T_0') != s(bool,X15)) ) | $spl54),
  inference(cnf_transformation,[],[f10135_D])).
fof(f10135_D,plain,(
  ( ! [X15,X16] : (s(bool,X15) != s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X16),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK13))) | s(bool,'T_0') != s(bool,X15)) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f10505,plain,(
  ( ! [X52,X53,X51] : (p(s(bool,X51)) | s(bool,'T_0') != s(bool,X52) | s(bool,X52) != s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X53),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK13)))) ) | ($spl7 | $spl10)),
  inference(resolution,[],[f9629,f4220])).
fof(f4220,plain,(
  ( ! [X171,X169,X170] : (sP7('const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X169),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK13)),X170,X171)) ) | $spl7),
  inference(resolution,[],[f3934,f352])).
fof(f10513,plain,(
  ( ! [X43,X44] : (s(bool,'T_0') != s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X43),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK13))) | s(bool,X44) != s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X43),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK13)))) ) | $spl58),
  inference(cnf_transformation,[],[f10513_D])).
fof(f10513_D,plain,(
  ( ! [X43,X44] : (s(bool,'T_0') != s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X43),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK13))) | s(bool,X44) != s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X43),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK13)))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f10501,plain,(
  ( ! [X43,X44,X42] : (p(s(bool,X42)) | s(bool,'T_0') != s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X43),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK13))) | s(bool,X44) != s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X43),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK13)))) ) | ($spl7 | $spl10)),
  inference(resolution,[],[f9629,f4219])).
fof(f4219,plain,(
  ( ! [X167,X166,X168] : (sP7(X166,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X167),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK13)),X168)) ) | $spl7),
  inference(resolution,[],[f3934,f351])).
fof(f4567,plain,(
  ( ! [X16] : (s(bool,'T_0') != s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X16),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK13)))) ) | $spl42),
  inference(cnf_transformation,[],[f4567_D])).
fof(f4567_D,plain,(
  ( ! [X16] : (s(bool,'T_0') != s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X16),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK13)))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f10202,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,X28)) | s(bool,'F_0') != s(bool,X29) | s(bool,'T_0') != s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X30),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK13)))) ) | ($spl7 | $spl10)),
  inference(resolution,[],[f9858,f4216])).
fof(f4216,plain,(
  ( ! [X158,X159,X157] : (sP8(X157,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X158),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK13)),X159)) ) | $spl7),
  inference(resolution,[],[f3934,f347])).
fof(f10200,plain,(
  ( ! [X24,X25] : (~p(s(bool,X24)) | s(bool,'F_0') != s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK14),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK12))) | s(bool,'T_0') != s(bool,X25)) ) | $spl10),
  inference(resolution,[],[f9858,f1235])).
fof(f1235,plain,(
  ( ! [X109,X110] : (sP8(X109,X110,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK14),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK12)))) )),
  inference(resolution,[],[f579,f346])).
fof(f10131,plain,(
  ( ! [X14,X15,X16] : (~p(s(bool,X14)) | s(bool,X15) != s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X16),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK13))) | s(bool,'T_0') != s(bool,X15)) ) | ($spl0 | $spl7 | $spl10)),
  inference(resolution,[],[f9156,f4216])).
fof(f9873,plain,(
  ( ! [X24,X25] : (s(bool,'F_0') != s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK14),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK12))) | ~p(s(bool,X24)) | ~p(s(bool,X25))) ) | $spl10),
  inference(resolution,[],[f9855,f1235])).
fof(f9745,plain,(
  ( ! [X68] : (p(s(bool,X68)) | s(bool,X68) != s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK14),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK12)))) ) | $spl52),
  inference(cnf_transformation,[],[f9745_D])).
fof(f9745_D,plain,(
  ( ! [X68] : (p(s(bool,X68)) | s(bool,X68) != s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK14),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK12)))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f9727,plain,(
  ( ! [X68,X69] : (p(s(bool,X68)) | p(s(bool,X69)) | s(bool,X68) != s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK14),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK12)))) ) | $spl10),
  inference(resolution,[],[f9548,f1241])).
fof(f9699,plain,(
  ( ! [X50,X48] : (~p(s(bool,X48)) | s(bool,X48) != s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X50),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK13)))) ) | $spl50),
  inference(cnf_transformation,[],[f9699_D])).
fof(f9699_D,plain,(
  ( ! [X50,X48] : (~p(s(bool,X48)) | s(bool,X48) != s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X50),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK13)))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f9687,plain,(
  ( ! [X50,X48,X49] : (~p(s(bool,X48)) | p(s(bool,X49)) | s(bool,X48) != s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X50),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK13)))) ) | ($spl7 | $spl10)),
  inference(resolution,[],[f9391,f4220])).
fof(f9419,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') != s(bool,X4) | s(bool,X3) != s(bool,X4) | ~p(s(bool,X3))) ) | $spl48),
  inference(cnf_transformation,[],[f9419_D])).
fof(f9419_D,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') != s(bool,X4) | s(bool,X3) != s(bool,X4) | ~p(s(bool,X3))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f9418,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | s(bool,'F_0') != s(bool,X4) | p(s(bool,X5)) | s(bool,X3) != s(bool,X4)) ) | $spl10),
  inference(duplicate_literal_removal,[],[f9401])).
fof(f9401,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | s(bool,'F_0') != s(bool,X4) | p(s(bool,X5)) | p(s(bool,X5)) | s(bool,X3) != s(bool,X4)) ) | $spl10),
  inference(resolution,[],[f9258,f360])).
fof(f9275,plain,(
  ( ! [X24,X25] : (sP9(X24,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK14),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK12)),X25) | ~p(s(bool,X25))) ) | $spl46),
  inference(cnf_transformation,[],[f9275_D])).
fof(f9275_D,plain,(
  ( ! [X24,X25] : (sP9(X24,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK14),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK12)),X25) | ~p(s(bool,X25))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f9265,plain,(
  ( ! [X24,X25] : (sP9(X24,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK14),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK12)),X25) | ~p(s(bool,X25)) | s(bool,'F_0') != s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK14),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK12)))) ) | $spl10),
  inference(resolution,[],[f9208,f1235])).
fof(f8961,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X7) | s(bool,'T_0') != s(bool,X6) | s(bool,X6) != s(bool,X7)) ) | $spl44),
  inference(cnf_transformation,[],[f8961_D])).
fof(f8961_D,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X7) | s(bool,'T_0') != s(bool,X6) | s(bool,X6) != s(bool,X7)) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f8960,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X6) | s(bool,X6) != s(bool,X7) | p(s(bool,X8)) | s(bool,'F_0') != s(bool,X7)) ) | $spl10),
  inference(duplicate_literal_removal,[],[f8947])).
fof(f8947,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X6) | s(bool,X6) != s(bool,X7) | p(s(bool,X8)) | p(s(bool,X8)) | s(bool,'F_0') != s(bool,X7)) ) | $spl10),
  inference(resolution,[],[f2538,f360])).
fof(f2538,plain,(
  ( ! [X6,X8,X7] : (sP9(X8,X6,'F_0') | s(bool,'T_0') != s(bool,X7) | s(bool,X6) != s(bool,X7) | p(s(bool,X8))) ) | $spl10),
  inference(superposition,[],[f2345,f357])).
fof(f6704,plain,(
  $false | ($spl4 | $spl15)),
  inference(subsumption_resolution,[],[f3002,f3911])).
fof(f4561,plain,(
  ( ! [X17,X16] : (s(bool,'T_0') != s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X16),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK13))) | ~p(s(bool,X17))) ) | ($spl0 | $spl7)),
  inference(resolution,[],[f4216,f2991])).
fof(f4566,plain,(
  ( ! [X12] : (s(bool,'F_0') = s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X12),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK13)))) ) | $spl40),
  inference(cnf_transformation,[],[f4566_D])).
fof(f4566_D,plain,(
  ( ! [X12] : (s(bool,'F_0') = s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X12),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK13)))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f4559,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') = s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X12),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK13))) | ~p(s(bool,X13))) ) | ($spl0 | $spl7)),
  inference(resolution,[],[f4216,f2949])).
fof(f4318,plain,(
  ( ! [X24,X22] : (s(bool,X22) != s(bool,X24) | ~p(s(bool,X22)) | p(s(bool,X24))) ) | $spl36),
  inference(cnf_transformation,[],[f4318_D])).
fof(f4318_D,plain,(
  ( ! [X24,X22] : (s(bool,X22) != s(bool,X24) | ~p(s(bool,X22)) | p(s(bool,X24))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f4320,plain,(
  ( ! [X26,X27,X25] : (s(bool,'T_0') != s(bool,X26) | s(bool,X25) != s(bool,X27) | p(s(bool,X27)) | ~p(s(bool,X25))) ) | $spl10),
  inference(subsumption_resolution,[],[f4293,f435])).
fof(f4293,plain,(
  ( ! [X26,X27,X25] : (s(bool,'T_0') != s(bool,X26) | s(bool,X25) != s(bool,X27) | p(s(bool,X27)) | ~p(s(bool,X26)) | ~p(s(bool,X25))) ) | $spl10),
  inference(superposition,[],[f2536,f407])).
fof(f4317,plain,(
  ( ! [X24,X23,X22] : (s(bool,'T_0') != s(bool,X23) | s(bool,X22) != s(bool,X24) | p(s(bool,X24)) | ~p(s(bool,X22))) ) | $spl10),
  inference(subsumption_resolution,[],[f4292,f435])).
fof(f4292,plain,(
  ( ! [X24,X23,X22] : (s(bool,'T_0') != s(bool,X23) | s(bool,X22) != s(bool,X24) | p(s(bool,X24)) | ~p(s(bool,X22)) | ~p(s(bool,X23))) ) | $spl10),
  inference(superposition,[],[f2536,f407])).
fof(f3356,plain,(
  s(bool,'T_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15))) | $spl31),
  inference(cnf_transformation,[],[f3356_D])).
fof(f3356_D,plain,(
  s(bool,'T_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15))) <=> ~$spl31),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl31])])).
fof(f4133,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15))) | ~p(s(bool,X10))) ) | ($spl0 | $spl7)),
  inference(resolution,[],[f3197,f2991])).
fof(f3197,plain,(
  ( ! [X107,X106] : (sP8(X106,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15)),X107)) ) | $spl7),
  inference(resolution,[],[f469,f347])).
fof(f3357,plain,(
  ( ! [X9] : (sP9(X9,'F_0','const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15)))) ) | $spl32),
  inference(cnf_transformation,[],[f3357_D])).
fof(f3357_D,plain,(
  ( ! [X9] : (sP9(X9,'F_0','const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15)))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f4132,plain,(
  ( ! [X9] : (sP9(X9,'F_0','const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15))) | s(bool,'T_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15)))) ) | ($spl0 | $spl7)),
  inference(resolution,[],[f3197,f2891])).
fof(f3354,plain,(
  s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15))) | $spl28),
  inference(cnf_transformation,[],[f3354_D])).
fof(f3354_D,plain,(
  s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15))) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f4131,plain,(
  ( ! [X8] : (s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15))) | ~p(s(bool,X8))) ) | ($spl0 | $spl7)),
  inference(resolution,[],[f3197,f2949])).
fof(f3310,plain,(
  s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))))) | $spl27),
  inference(cnf_transformation,[],[f3310_D])).
fof(f3310_D,plain,(
  s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))))) <=> ~$spl27),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl27])])).
fof(f4092,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16)))))) ) | ($spl0 | $spl4)),
  inference(resolution,[],[f3072,f2950])).
fof(f3072,plain,(
  ( ! [X121,X120] : (sP6('const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16)))),X120,X121)) ) | $spl4),
  inference(resolution,[],[f466,f356])).
fof(f3816,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f3710,f3000])).
fof(f3710,plain,(
  p(s(bool,'F_0')) | $spl12),
  inference(resolution,[],[f3000,f450])).
fof(f450,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f258])).
fof(f3815,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f3709,f3000])).
fof(f3709,plain,(
  p(s(bool,'F_0')) | $spl12),
  inference(resolution,[],[f3000,f448])).
fof(f448,plain,(
  p(s(bool,'F_0')) | p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f257])).
fof(f3814,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f3708,f3000])).
fof(f3708,plain,(
  ( ! [X224] : (p(s(bool,X224))) ) | $spl12),
  inference(resolution,[],[f3000,f428])).
fof(f428,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f252])).
fof(f252,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,'T_0'))))),
  inference(flattening,[],[f251])).
fof(f251,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f135])).
fof(f135,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X14] : ((p(s(bool,X14)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  file('thm.integral.TDIV_BOUNDS','thm.bool.IMP_CLAUSES|split|1')).
fof(f3813,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f3706,f3000])).
fof(f3706,plain,(
  p(s(bool,'F_0')) | $spl12),
  inference(resolution,[],[f3000,f425])).
fof(f425,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f250])).
fof(f250,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f249])).
fof(f249,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f134])).
fof(f134,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f96])).
fof(f96,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X14] : ((p(s(bool,'F_0')) => p(s(bool,X14))) <=> p(s(bool,'T_0')))),
  file('thm.integral.TDIV_BOUNDS','thm.bool.IMP_CLAUSES|split|2')).
fof(f3812,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f3704,f3000])).
fof(f3704,plain,(
  ( ! [X221] : (p(s(bool,X221))) ) | $spl12),
  inference(resolution,[],[f3000,f422])).
fof(f422,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f248])).
fof(f248,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))))),
  inference(flattening,[],[f247])).
fof(f247,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f133])).
fof(f133,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f95])).
fof(f95,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X14] : ((p(s(bool,X14)) => p(s(bool,X14))) <=> p(s(bool,'T_0')))),
  file('thm.integral.TDIV_BOUNDS','thm.bool.IMP_CLAUSES|split|3')).
fof(f3811,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f3703,f3000])).
fof(f3703,plain,(
  ( ! [X220] : (p(s(bool,X220))) ) | $spl12),
  inference(resolution,[],[f3000,f419])).
fof(f419,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f246])).
fof(f246,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,'T_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f245])).
fof(f245,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (~p(s(bool,'T_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f132])).
fof(f132,plain,(
  ! [X0] : ((~p(s(bool,'T_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(ennf_transformation,[],[f94])).
fof(f94,plain,(
  ! [X0] : ((p(s(bool,'T_0')) => p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X14] : ((p(s(bool,'T_0')) => p(s(bool,X14))) <=> p(s(bool,X14)))),
  file('thm.integral.TDIV_BOUNDS','thm.bool.IMP_CLAUSES|split|0')).
fof(f3711,plain,(
  $false | $spl12),
  inference(resolution,[],[f3000,f452])).
fof(f452,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.integral.TDIV_BOUNDS',aHL_TRUTH)).
fof(f3712,plain,(
  $false | $spl12),
  inference(resolution,[],[f3000,f453])).
fof(f3713,plain,(
  $false | $spl12),
  inference(resolution,[],[f3000,f460])).
fof(f3721,plain,(
  $false | $spl12),
  inference(resolution,[],[f3000,f260])).
fof(f3745,plain,(
  $false | $spl12),
  inference(resolution,[],[f3000,f1180])).
fof(f3746,plain,(
  $false | ($spl2 | $spl12)),
  inference(resolution,[],[f3000,f464])).
fof(f3747,plain,(
  $false | $spl12),
  inference(resolution,[],[f3000,f1181])).
fof(f3749,plain,(
  $false | $spl12),
  inference(resolution,[],[f3000,f579])).
fof(f3613,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))))) | ~p(s(bool,X10))) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f1666,f2991])).
fof(f3612,plain,(
  ( ! [X9] : (sP9(X9,'F_0','const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))))) | s(bool,'T_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16)))))) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f1666,f2891])).
fof(f3611,plain,(
  ( ! [X8] : (s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))))) | ~p(s(bool,X8))) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f1666,f2949])).
fof(f3348,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15))) | ~p(s(bool,X10))) ) | ($spl0 | $spl7)),
  inference(resolution,[],[f3197,f2991])).
fof(f3347,plain,(
  ( ! [X9] : (sP9(X9,'F_0','const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15))) | s(bool,'T_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15)))) ) | ($spl0 | $spl7)),
  inference(resolution,[],[f3197,f2891])).
fof(f3346,plain,(
  ( ! [X8] : (s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15))) | ~p(s(bool,X8))) ) | ($spl0 | $spl7)),
  inference(resolution,[],[f3197,f2949])).
fof(f3305,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16)))))) ) | ($spl0 | $spl4)),
  inference(resolution,[],[f3072,f2950])).
fof(f3132,plain,(
  ( ! [X24] : (p(s(bool,X24)) | s(bool,'F_0') != s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK14),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK12)))) ) | $spl0),
  inference(resolution,[],[f2950,f1241])).
fof(f3139,plain,(
  ( ! [X23] : (s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK12),s('type.num.num',X23))),s('type.realax.real',sK15)))) ) | $spl22),
  inference(cnf_transformation,[],[f3139_D])).
fof(f3139_D,plain,(
  ( ! [X23] : (s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK12),s('type.num.num',X23))),s('type.realax.real',sK15)))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f3131,plain,(
  ( ! [X23,X22] : (p(s(bool,X22)) | s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK12),s('type.num.num',X23))),s('type.realax.real',sK15)))) ) | $spl0),
  inference(resolution,[],[f2950,f1892])).
fof(f1892,plain,(
  ( ! [X182,X180,X181] : (sP6('const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK12),s('type.num.num',X180))),s('type.realax.real',sK15)),X181,X182)) )),
  inference(resolution,[],[f1181,f356])).
fof(f3138,plain,(
  ( ! [X21] : (s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK12),s('type.num.num',X21)))))) ) | $spl20),
  inference(cnf_transformation,[],[f3138_D])).
fof(f3138_D,plain,(
  ( ! [X21] : (s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK12),s('type.num.num',X21)))))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f3130,plain,(
  ( ! [X21,X20] : (p(s(bool,X20)) | s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK12),s('type.num.num',X21)))))) ) | $spl0),
  inference(resolution,[],[f2950,f1346])).
fof(f1346,plain,(
  ( ! [X182,X180,X181] : (sP6('const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK12),s('type.num.num',X180)))),X181,X182)) )),
  inference(resolution,[],[f1180,f356])).
fof(f3137,plain,(
  s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK12),s('type.num.num',sK16))),s('type.realax.real',sK15))) | $spl19),
  inference(cnf_transformation,[],[f3137_D])).
fof(f3137_D,plain,(
  s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK12),s('type.num.num',sK16))),s('type.realax.real',sK15))) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f3129,plain,(
  ( ! [X19] : (p(s(bool,X19)) | s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK12),s('type.num.num',sK16))),s('type.realax.real',sK15)))) ) | ($spl0 | $spl2)),
  inference(resolution,[],[f2950,f1564])).
fof(f1564,plain,(
  ( ! [X121,X120] : (sP6('const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK12),s('type.num.num',sK16))),s('type.realax.real',sK15)),X120,X121)) ) | $spl2),
  inference(resolution,[],[f464,f356])).
fof(f3003,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK14),s('type.realax.real',X9))),s(fun('type.num.num','type.realax.real'),sK13)))) ) | $spl16),
  inference(cnf_transformation,[],[f3003_D])).
fof(f3003_D,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK14),s('type.realax.real',X9))),s(fun('type.num.num','type.realax.real'),sK13)))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f2997,plain,(
  ( ! [X10,X9] : (s(bool,'T_0') != s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK14),s('type.realax.real',X9))),s(fun('type.num.num','type.realax.real'),sK13))) | ~p(s(bool,X10))) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f2991,f1993])).
fof(f1993,plain,(
  ( ! [X158,X159,X157] : (sP8(X157,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK14),s('type.realax.real',X158))),s(fun('type.num.num','type.realax.real'),sK13)),X159)) ) | $spl5),
  inference(resolution,[],[f1613,f347])).
fof(f1613,plain,(
  ( ! [X0] : (~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK14),s('type.realax.real',X0))),s(fun('type.num.num','type.realax.real'),sK13))))) ) | $spl5),
  inference(resolution,[],[f467,f267])).
fof(f2996,plain,(
  ( ! [X8] : (s(bool,'T_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))))) | ~p(s(bool,X8))) ) | ($spl0 | $spl5)),
  inference(resolution,[],[f2991,f1666])).
fof(f2523,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f2443,f2344])).
fof(f2443,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(resolution,[],[f2344,f451])).
fof(f2522,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f2442,f2344])).
fof(f2442,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(resolution,[],[f2344,f449])).
fof(f2521,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f2440,f2344])).
fof(f2440,plain,(
  ( ! [X230] : (~p(s(bool,X230))) ) | $spl8),
  inference(resolution,[],[f2344,f445])).
fof(f2520,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f2438,f2344])).
fof(f2438,plain,(
  ( ! [X228] : (~p(s(bool,X228))) ) | $spl8),
  inference(resolution,[],[f2344,f417])).
fof(f2441,plain,(
  $false | $spl8),
  inference(resolution,[],[f2344,f447])).
fof(f2469,plain,(
  $false | ($spl5 | $spl8)),
  inference(resolution,[],[f2344,f467])).
fof(f2473,plain,(
  $false | ($spl5 | $spl8)),
  inference(resolution,[],[f2344,f1613])).
fof(f2345,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | s(bool,'T_0') != s(bool,X2)) ) | $spl10),
  inference(cnf_transformation,[],[f2345_D])).
fof(f2345_D,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | s(bool,'T_0') != s(bool,X2)) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f2343,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) ) | $spl0),
  inference(duplicate_literal_removal,[],[f2338])).
fof(f2338,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) ) | $spl0),
  inference(resolution,[],[f2321,f360])).
fof(f2321,plain,(
  ( ! [X4,X3] : (sP9(X4,X3,'F_0') | s(bool,'T_0') != s(bool,X3) | p(s(bool,X4))) ) | $spl0),
  inference(superposition,[],[f2315,f357])).
fof(f1396,plain,(
  $false | $spl3),
  inference(resolution,[],[f962,f579])).
fof(f962,plain,(
  ( ! [X1] : (~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',sK15))),s(fun('type.num.num','type.realax.real'),sK12))))) ) | $spl3),
  inference(resolution,[],[f465,f268])).
fof(f851,plain,(
  $false | $spl1),
  inference(subsumption_resolution,[],[f579,f470])).
fof(f470,plain,(
  ( ! [X0] : (~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK14),s('type.realax.real',X0))),s(fun('type.num.num','type.realax.real'),sK12))))) ) | $spl1),
  inference(resolution,[],[f463,f267])).
fof(f463,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK12),s('type.num.num',sK16)))))) | $spl1),
  inference(cnf_transformation,[],[f463_D])).
fof(f463_D,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK12),s('type.num.num',sK16)))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f465,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK12),s('type.num.num',sK16))),s('type.realax.real',sK15)))) | $spl3),
  inference(cnf_transformation,[],[f465_D])).
fof(f465_D,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK12),s('type.num.num',sK16))),s('type.realax.real',sK15)))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f467,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16)))))) | $spl5),
  inference(cnf_transformation,[],[f467_D])).
fof(f467_D,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16)))))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f469,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15)))) | $spl7),
  inference(cnf_transformation,[],[f469_D])).
fof(f469_D,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15)))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f261,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16))),s('type.realax.real',sK15)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK13),s('type.num.num',sK16)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK12),s('type.num.num',sK16))),s('type.realax.real',sK15)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK12),s('type.num.num',sK16))))))),
  inference(cnf_transformation,[],[f161])).
% SZS output end Proof for thm.integral
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 799
% Passive clauses: 4131
% Generated clauses: 22208
% Final active clauses: 549
% Final passive clauses: 1556
% Input formulas: 48
% Initial clauses: 197
% 
% Selected by SInE selection: 45
% SInE iterations: 1
% 
% Duplicate literals: 1123
% Trivial inequalities: 193
% Fw subsumption resolutions: 3173
% Fw demodulations: 2933
% Bw demodulations: 16
% 
% Simple tautologies: 2599
% Equational tautologies: 50
% Forward subsumptions: 8090
% 
% Binary resolution: 10715
% Forward superposition: 1875
% Backward superposition: 1955
% Equality resolution: 24
% 
% Split clauses: 150
% Split components: 323
% SAT solver clauses: 199
% SAT solver unit clauses: 32
% SAT solver binary clauses: 10
% 
% Sat splits: 150
% Sat splitting refutations: 43
% 
% Memory used [KB]: 7419
% Time elapsed: 0.354 s
% ------------------------------
% ------------------------------
% Success in time 12.075 s
