Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.list
% (66810)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 128
% Passive clauses: 3246
% Generated clauses: 20657
% Final active clauses: 116
% Final passive clauses: 2893
% Input formulas: 19
% Initial clauses: 34
% 
% Duplicate literals: 877
% Trivial inequalities: 301
% Fw subsumption resolutions: 2242
% Fw demodulations: 3176
% Bw demodulations: 7
% 
% Simple tautologies: 287
% Equational tautologies: 42
% Forward subsumptions: 9130
% 
% Binary resolution: 146
% Factoring: 1
% Forward superposition: 7098
% Backward superposition: 6320
% Self superposition: 221
% Equality factoring: 101
% Equality resolution: 28
% 
% Split clauses: 346
% Split components: 706
% SAT solver clauses: 691
% SAT solver unit clauses: 9
% SAT solver binary clauses: 61
% 
% Sat splits: 346
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 684
% TWLsolver calls for satisfiability: 95
% 
% Memory used [KB]: 15863
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.list
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.list
% SZS output start Proof for thm.list
fof(f7265,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f7172,f1677,f73,f110,f79,f114,f191,f196,f82,f117,f172,f180,f69,f106,f89,f122,f3288,f66,f102,f2030,f3303,f70,f107,f84,f115])).
fof(f115,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.list.list'(X0),'const.list.DROP_2'(s('type.num.num',X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3))))) = s('type.list.list'(X0),'const.bool.COND_3'(s(bool,sK4(X3,X2,X1,X0)),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3))),s('type.list.list'(X0),'const.list.DROP_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.list.list'(X0),X3)))))) ) | $spl20),
  inference(cnf_transformation,[],[f115_D])).
fof(f115_D,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.list.list'(X0),'const.list.DROP_2'(s('type.num.num',X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3))))) = s('type.list.list'(X0),'const.bool.COND_3'(s(bool,sK4(X3,X2,X1,X0)),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3))),s('type.list.list'(X0),'const.list.DROP_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.list.list'(X0),X3)))))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f84,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.list.list'(X0),'const.list.DROP_2'(s('type.num.num',X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3))))) = s('type.list.list'(X0),'const.bool.COND_3'(s(bool,sK4(X3,X2,X1,X0)),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3))),s('type.list.list'(X0),'const.list.DROP_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.list.list'(X0),X3)))))) )),
  inference(cnf_transformation,[],[f62])).
fof(f62,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,sK4(X3,X2,X1,X0))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) & (s('type.num.num','const.num.|48|_0') != s('type.num.num',X1) | p(s(bool,sK4(X3,X2,X1,X0)))) & s('type.list.list'(X0),'const.list.DROP_2'(s('type.num.num',X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3))))) = s('type.list.list'(X0),'const.bool.COND_3'(s(bool,sK4(X3,X2,X1,X0)),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3))),s('type.list.list'(X0),'const.list.DROP_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.list.list'(X0),X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f61])).
fof(f61,plain,(
  ! [X0,X1,X2,X3] : ? [X4] : ((~p(s(bool,X4)) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) & (s('type.num.num','const.num.|48|_0') != s('type.num.num',X1) | p(s(bool,X4))) & s('type.list.list'(X0),'const.list.DROP_2'(s('type.num.num',X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3))))) = s('type.list.list'(X0),'const.bool.COND_3'(s(bool,X4),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3))),s('type.list.list'(X0),'const.list.DROP_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.list.list'(X0),X3))))))),
  inference(flattening,[],[f60])).
fof(f60,plain,(
  ! [X0,X1,X2,X3] : ? [X4] : (((~p(s(bool,X4)) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) & (s('type.num.num','const.num.|48|_0') != s('type.num.num',X1) | p(s(bool,X4)))) & s('type.list.list'(X0),'const.list.DROP_2'(s('type.num.num',X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3))))) = s('type.list.list'(X0),'const.bool.COND_3'(s(bool,X4),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3))),s('type.list.list'(X0),'const.list.DROP_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.list.list'(X0),X3))))))),
  inference(nnf_transformation,[],[f41])).
fof(f41,plain,(
  ! [X0,X1,X2,X3] : ? [X4] : ((p(s(bool,X4)) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) & s('type.list.list'(X0),'const.list.DROP_2'(s('type.num.num',X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3))))) = s('type.list.list'(X0),'const.bool.COND_3'(s(bool,X4),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3))),s('type.list.list'(X0),'const.list.DROP_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.list.list'(X0),X3))))))),
  inference(flattening,[],[f40])).
fof(f40,plain,(
  ! [X0] : ! [X1,X2,X3] : ? [X4] : ((p(s(bool,X4)) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) & s('type.list.list'(X0),'const.list.DROP_2'(s('type.num.num',X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3))))) = s('type.list.list'(X0),'const.bool.COND_3'(s(bool,X4),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3))),s('type.list.list'(X0),'const.list.DROP_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.list.list'(X0),X3))))))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X13] : ! [X7,X14,X21] : ? [X22] : ((p(s(bool,X22)) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X7)) & s('type.list.list'(X13),'const.list.DROP_2'(s('type.num.num',X7),s('type.list.list'(X13),'const.list.CONS_2'(s(X13,X14),s('type.list.list'(X13),X21))))) = s('type.list.list'(X13),'const.bool.COND_3'(s(bool,X22),s('type.list.list'(X13),'const.list.CONS_2'(s(X13,X14),s('type.list.list'(X13),X21))),s('type.list.list'(X13),'const.list.DROP_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X7),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.list.list'(X13),X21))))))),
  file('thm.list.list_455_prove_ltb|split|2','thm.list.DROP_def|split|1')).
fof(f107,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl6),
  inference(cnf_transformation,[],[f107_D])).
fof(f107_D,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f70,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(cnf_transformation,[],[f24])).
fof(f24,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X10] : s('type.num.num',X10) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  file('thm.list.list_455_prove_ltb|split|2','thm.arithmetic.SUC_SUB|49|')).
fof(f3303,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.list.list'(X1),'const.list.DROP_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X2),s('type.list.list'(X1),X3))))) = s('type.list.list'(X1),'const.bool.COND_3'(s(bool,sK4(X3,X2,'const.num.SUC_1'(s('type.num.num',X0)),X1)),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X2),s('type.list.list'(X1),X3))),s('type.list.list'(X1),'const.list.DROP_2'(s('type.num.num',X0),s('type.list.list'(X1),X3)))))) ) | $spl548),
  inference(cnf_transformation,[],[f3303_D])).
fof(f3303_D,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.list.list'(X1),'const.list.DROP_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X2),s('type.list.list'(X1),X3))))) = s('type.list.list'(X1),'const.bool.COND_3'(s(bool,sK4(X3,X2,'const.num.SUC_1'(s('type.num.num',X0)),X1)),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X2),s('type.list.list'(X1),X3))),s('type.list.list'(X1),'const.list.DROP_2'(s('type.num.num',X0),s('type.list.list'(X1),X3)))))) ) <=> ~$spl548),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl548])])).
fof(f2030,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.list.list'(X1),'const.list.DROP_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X2),s('type.list.list'(X1),X3))))) = s('type.list.list'(X1),'const.bool.COND_3'(s(bool,sK4(X3,X2,'const.num.SUC_1'(s('type.num.num',X0)),X1)),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X2),s('type.list.list'(X1),X3))),s('type.list.list'(X1),'const.list.DROP_2'(s('type.num.num',X0),s('type.list.list'(X1),X3)))))) ) | ($spl6 | $spl20)),
  inference(superposition,[],[f115,f107])).
fof(f102,plain,(
  s('type.list.list'(sK0),'const.list.DROP_2'(s('type.num.num',sK1),s('type.list.list'(sK0),sK3))) != s('type.list.list'(sK0),'const.list.DROP_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),sK3))))) | $spl1),
  inference(cnf_transformation,[],[f102_D])).
fof(f102_D,plain,(
  s('type.list.list'(sK0),'const.list.DROP_2'(s('type.num.num',sK1),s('type.list.list'(sK0),sK3))) != s('type.list.list'(sK0),'const.list.DROP_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),sK3))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f66,plain,(
  s('type.list.list'(sK0),'const.list.DROP_2'(s('type.num.num',sK1),s('type.list.list'(sK0),sK3))) != s('type.list.list'(sK0),'const.list.DROP_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),sK3)))))),
  inference(cnf_transformation,[],[f57])).
fof(f57,plain,(
  s('type.list.list'(sK0),'const.list.DROP_2'(s('type.num.num',sK1),s('type.list.list'(sK0),sK3))) != s('type.list.list'(sK0),'const.list.DROP_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),sK3)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f48])).
fof(f48,plain,(
  ? [X0,X1,X2,X3] : s('type.list.list'(X0),'const.list.DROP_2'(s('type.num.num',X1),s('type.list.list'(X0),X3))) != s('type.list.list'(X0),'const.list.DROP_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3)))))),
  inference(ennf_transformation,[],[f22])).
fof(f22,plain,(
  ~! [X0,X1,X2,X3] : s('type.list.list'(X0),'const.list.DROP_2'(s('type.num.num',X1),s('type.list.list'(X0),X3))) = s('type.list.list'(X0),'const.list.DROP_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3)))))),
  inference(flattening,[],[f21])).
fof(f21,plain,(
  ~! [X0] : ! [X1,X2,X3] : s('type.list.list'(X0),'const.list.DROP_2'(s('type.num.num',X1),s('type.list.list'(X0),X3))) = s('type.list.list'(X0),'const.list.DROP_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : s('type.list.list'(X6),'const.list.DROP_2'(s('type.num.num',X7),s('type.list.list'(X6),X9))) = s('type.list.list'(X6),'const.list.DROP_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))),s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X8),s('type.list.list'(X6),X9)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : s('type.list.list'(X6),'const.list.DROP_2'(s('type.num.num',X7),s('type.list.list'(X6),X9))) = s('type.list.list'(X6),'const.list.DROP_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))),s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X8),s('type.list.list'(X6),X9)))))),
  file('thm.list.list_455_prove_ltb|split|2','thm.list.list_455_prove_ltb|split|2')).
fof(f3288,plain,(
  $false | ($spl1 | $spl248)),
  inference(trivial_inequality_removal,[],[f3275])).
fof(f3275,plain,(
  s('type.list.list'(sK0),'const.list.DROP_2'(s('type.num.num',sK1),s('type.list.list'(sK0),sK3))) != s('type.list.list'(sK0),'const.list.DROP_2'(s('type.num.num',sK1),s('type.list.list'(sK0),sK3))) | ($spl1 | $spl248)),
  inference(superposition,[],[f102,f1677])).
fof(f122,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) = s(bool,X1)) ) | $spl34),
  inference(cnf_transformation,[],[f122_D])).
fof(f122_D,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) = s(bool,X1)) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f89,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f52])).
fof(f52,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f44])).
fof(f44,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X11,X12] : ((p(s(bool,X11)) => p(s(bool,X12))) => ((p(s(bool,X12)) => p(s(bool,X11))) => s(bool,X11) = s(bool,X12)))),
  file('thm.list.list_455_prove_ltb|split|2','thm.bool.IMP_ANTISYM_AX')).
fof(f106,plain,(
  ~p(s(bool,'F_0')) | $spl5),
  inference(cnf_transformation,[],[f106_D])).
fof(f106_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f69,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f23])).
fof(f23,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.list.list_455_prove_ltb|split|2',aHL_FALSITY)).
fof(f180,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) | $spl54),
  inference(cnf_transformation,[],[f180_D])).
fof(f180_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f172,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) | ($spl5 | $spl34)),
  inference(resolution,[],[f122,f106])).
fof(f117,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,sK4(X3,X2,X1,X0))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) ) | $spl24),
  inference(cnf_transformation,[],[f117_D])).
fof(f117_D,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,sK4(X3,X2,X1,X0))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f82,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X1) | ~p(s(bool,sK4(X3,X2,X1,X0)))) )),
  inference(cnf_transformation,[],[f62])).
fof(f196,plain,(
  ( ! [X6,X8,X7,X5] : (s(bool,'F_0') = s(bool,sK4(X5,X6,X7,X8)) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X7)) ) | $spl66),
  inference(cnf_transformation,[],[f196_D])).
fof(f196_D,plain,(
  ( ! [X6,X8,X7,X5] : (s(bool,'F_0') = s(bool,sK4(X5,X6,X7,X8)) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X7)) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f191,plain,(
  ( ! [X6,X8,X7,X5] : (s(bool,'F_0') = s(bool,sK4(X5,X6,X7,X8)) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X7)) ) | ($spl24 | $spl54)),
  inference(resolution,[],[f180,f117])).
fof(f114,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))) ) | $spl18),
  inference(cnf_transformation,[],[f114_D])).
fof(f114_D,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f79,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f37])).
fof(f37,plain,(
  ! [X0,X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f36])).
fof(f36,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X13] : ! [X11,X12] : s(X13,X12) = s(X13,'const.bool.COND_3'(s(bool,'F_0'),s(X13,X11),s(X13,X12)))),
  file('thm.list.list_455_prove_ltb|split|2','thm.bool.COND_CLAUSES|split|1')).
fof(f110,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))) ) | $spl10),
  inference(cnf_transformation,[],[f110_D])).
fof(f110_D,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f73,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f27])).
fof(f27,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),
  inference(flattening,[],[f26])).
fof(f26,plain,(
  ! [X0] : ~s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X7] : ~s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num',X7)))),
  file('thm.list.list_455_prove_ltb|split|2','thm.num.NOT_SUC')).
fof(f1677,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.list.list'(X1),'const.list.DROP_2'(s('type.num.num',X0),s('type.list.list'(X1),X3))) = s('type.list.list'(X1),'const.list.DROP_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X2),s('type.list.list'(X1),X3)))))) ) | $spl248),
  inference(cnf_transformation,[],[f1677_D])).
fof(f1677_D,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.list.list'(X1),'const.list.DROP_2'(s('type.num.num',X0),s('type.list.list'(X1),X3))) = s('type.list.list'(X1),'const.list.DROP_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X2),s('type.list.list'(X1),X3)))))) ) <=> ~$spl248),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl248])])).
fof(f7172,plain,(
  ( ! [X6,X8,X7,X5] : (s('type.list.list'(X8),'const.list.DROP_2'(s('type.num.num',X7),s('type.list.list'(X8),X5))) = s('type.list.list'(X8),'const.list.DROP_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))),s('type.list.list'(X8),'const.list.CONS_2'(s(X8,X6),s('type.list.list'(X8),X5)))))) ) | ($spl10 | $spl18 | $spl66 | $spl548)),
  inference(forward_demodulation,[],[f7171,f114])).
fof(f7171,plain,(
  ( ! [X6,X8,X7,X5] : (s('type.list.list'(X8),'const.list.DROP_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))),s('type.list.list'(X8),'const.list.CONS_2'(s(X8,X6),s('type.list.list'(X8),X5))))) = s('type.list.list'(X8),'const.bool.COND_3'(s(bool,'F_0'),s('type.list.list'(X8),'const.list.CONS_2'(s(X8,X6),s('type.list.list'(X8),X5))),s('type.list.list'(X8),'const.list.DROP_2'(s('type.num.num',X7),s('type.list.list'(X8),X5)))))) ) | ($spl10 | $spl66 | $spl548)),
  inference(subsumption_resolution,[],[f7074,f110])).
fof(f7074,plain,(
  ( ! [X6,X8,X7,X5] : (s('type.list.list'(X8),'const.list.DROP_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))),s('type.list.list'(X8),'const.list.CONS_2'(s(X8,X6),s('type.list.list'(X8),X5))))) = s('type.list.list'(X8),'const.bool.COND_3'(s(bool,'F_0'),s('type.list.list'(X8),'const.list.CONS_2'(s(X8,X6),s('type.list.list'(X8),X5))),s('type.list.list'(X8),'const.list.DROP_2'(s('type.num.num',X7),s('type.list.list'(X8),X5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num',X7)))) ) | ($spl66 | $spl548)),
  inference(superposition,[],[f3303,f196])).
% SZS output end Proof for thm.list
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 92
% Passive clauses: 506
% Generated clauses: 6974
% Final active clauses: 82
% Final passive clauses: 402
% Input formulas: 19
% Initial clauses: 34
% 
% Duplicate literals: 411
% Trivial inequalities: 81
% Fw subsumption resolutions: 432
% Fw demodulations: 468
% Bw demodulations: 14
% Condensations: 598
% Global subsumptions: 1
% 
% Simple tautologies: 136
% Equational tautologies: 33
% Forward subsumptions: 3322
% 
% Binary resolution: 114
% Factoring: 1
% Forward superposition: 2534
% Backward superposition: 1559
% Self superposition: 92
% Equality factoring: 58
% Equality resolution: 17
% 
% Split clauses: 56
% Split components: 121
% SAT solver clauses: 1010
% SAT solver unit clauses: 38
% SAT solver binary clauses: 31
% SAT solver learnt clauses: 3
% 
% Sat splits: 56
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 936
% TWLsolver calls for satisfiability: 227
% 
% Memory used [KB]: 7675
% Time elapsed: 0.397 s
% ------------------------------
% ------------------------------
% Success in time 1.015 s
