Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.semi_ring
% (37356)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 176
% Passive clauses: 2785
% Generated clauses: 24724
% Final active clauses: 89
% Final passive clauses: 1842
% Input formulas: 36
% Initial clauses: 40
% 
% Duplicate literals: 452
% Trivial inequalities: 121
% Fw subsumption resolutions: 2054
% Fw demodulations: 3764
% Bw demodulations: 52
% 
% Simple tautologies: 145
% Equational tautologies: 21
% Forward subsumptions: 13612
% 
% Binary resolution: 101
% Forward superposition: 9625
% Backward superposition: 8158
% Self superposition: 315
% Equality factoring: 60
% Equality resolution: 18
% 
% Split clauses: 857
% Split components: 1739
% SAT solver clauses: 1933
% SAT solver unit clauses: 14
% SAT solver binary clauses: 371
% 
% Sat splits: 857
% Sat splitting refutations: 33
% 
% TWLsolver clauses: 1906
% TWLsolver calls for satisfiability: 142
% 
% Memory used [KB]: 21236
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.semi_ring
% (37379)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 122
% Passive clauses: 1091
% Generated clauses: 12219
% Final active clauses: 87
% Final passive clauses: 717
% Input formulas: 36
% Initial clauses: 40
% 
% Duplicate literals: 318
% Trivial inequalities: 44
% Fw subsumption resolutions: 303
% Fw demodulations: 979
% Bw demodulations: 279
% Condensations: 578
% 
% Simple tautologies: 84
% Equational tautologies: 24
% Forward subsumptions: 7311
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 138
% Forward superposition: 4673
% Backward superposition: 3553
% Self superposition: 215
% Equality factoring: 54
% Equality resolution: 2
% 
% Split clauses: 83
% Split components: 166
% SAT solver clauses: 1468
% SAT solver unit clauses: 66
% SAT solver binary clauses: 43
% 
% Sat splits: 83
% Sat splitting refutations: 17
% 
% TWLsolver clauses: 1468
% TWLsolver calls for satisfiability: 162
% 
% Memory used [KB]: 24690
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.semi_ring
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.semi_ring
% SZS output start Proof for thm.semi_ring
fof(f5371,plain,(
  $false),
  inference(subsumption_resolution,[],[f5370,f115])).
fof(f115,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.semi_ring.semi_ring_fupdfupds_comp|split|1',aHL_TRUTH)).
fof(f5370,plain,(
  ~p(s(bool,'T_0'))),
  inference(forward_demodulation,[],[f5368,f1210])).
fof(f1210,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X1)))) )),
  inference(unit_resulting_resolution,[],[f1185,f158])).
fof(f158,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f116,f117])).
fof(f117,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.semi_ring.semi_ring_fupdfupds_comp|split|1',aHL_BOOL_CASES)).
fof(f116,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f40])).
fof(f40,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.semi_ring.semi_ring_fupdfupds_comp|split|1',aHL_FALSITY)).
fof(f1185,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X1))))) )),
  inference(backward_demodulation,[],[f1126,f583])).
fof(f583,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.combin.o_3'(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(fun(X0,X0),'_i_0_0'),s(X0,X1))))) )),
  inference(forward_demodulation,[],[f549,f138])).
fof(f138,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5)))))) )),
  inference(cnf_transformation,[],[f70])).
fof(f70,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5)))))),
  inference(flattening,[],[f69])).
fof(f69,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5)))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6,X7,X14] : ! [X10,X9,X15] : s(X6,'const.combin.o_3'(s(fun(X7,X6),X10),s(fun(X14,X7),X9),s(X14,X15))) = s(X6,happ(s(fun(X7,X6),X10),s(X7,happ(s(fun(X14,X7),X9),s(X14,X15)))))),
  file('thm.semi_ring.semi_ring_fupdfupds_comp|split|1','thm.combin.o_THM')).
fof(f549,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1))))))) )),
  inference(unit_resulting_resolution,[],[f118,f136])).
fof(f136,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) != s(X0,X2)) )),
  inference(cnf_transformation,[],[f113])).
fof(f113,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2))))))),
  inference(nnf_transformation,[],[f66])).
fof(f66,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X32,X33,X34] : (p(s(bool,happ(s(fun(X32,bool),happ(s(fun(X32,fun(X32,bool)),'=_0'),s(X32,X33))),s(X32,X34)))) <=> s(X32,X33) = s(X32,X34))),
  file('thm.semi_ring.semi_ring_fupdfupds_comp|split|1',arity_0_HL_EQ)).
fof(f118,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f42])).
fof(f42,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(flattening,[],[f41])).
fof(f41,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X23] : ! [X29] : s(X23,X29) = s(X23,happ(s(fun(X23,X23),'_i_0_0'),s(X23,X29)))),
  file('thm.semi_ring.semi_ring_fupdfupds_comp|split|1','_i0_ax')).
fof(f1126,plain,(
  ( ! [X39,X41,X42,X40] : (s(X41,happ(s(fun(X39,X41),X42),s(X39,X40))) = s(X41,'const.combin.o_3'(s(fun(X39,X41),X42),s(fun(X39,X39),'_i_0_0'),s(X39,X40)))) )),
  inference(superposition,[],[f138,f118])).
fof(f5368,plain,(
  ~p(s(bool,happ(s(fun(fun(sK0,'type.semi_ring.semi_ring'(sK1)),bool),happ(s(fun(fun(sK0,'type.semi_ring.semi_ring'(sK1)),fun(fun(sK0,'type.semi_ring.semi_ring'(sK1)),bool)),'=_0'),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK1),'type.semi_ring.semi_ring'(sK1)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(sK1,sK1),'const.combin.o_2'(s(fun(sK1,sK1),sK4),s(fun(sK1,sK1),sK3))))),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),sK2))))),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK1),'type.semi_ring.semi_ring'(sK1)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(sK1,sK1),'const.combin.o_2'(s(fun(sK1,sK1),sK4),s(fun(sK1,sK1),sK3))))),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),sK2))))))),
  inference(backward_demodulation,[],[f5324,f1482])).
fof(f1482,plain,(
  ~p(s(bool,happ(s(fun(fun(sK0,'type.semi_ring.semi_ring'(sK1)),bool),happ(s(fun(fun(sK0,'type.semi_ring.semi_ring'(sK1)),fun(fun(sK0,'type.semi_ring.semi_ring'(sK1)),bool)),'=_0'),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK1),'type.semi_ring.semi_ring'(sK1)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(sK1,sK1),sK4))),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK1),'type.semi_ring.semi_ring'(sK1)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(sK1,sK1),sK3))),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),sK2))))))),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK1),'type.semi_ring.semi_ring'(sK1)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(sK1,sK1),'const.combin.o_2'(s(fun(sK1,sK1),sK4),s(fun(sK1,sK1),sK3))))),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),sK2))))))),
  inference(unit_resulting_resolution,[],[f114,f135])).
fof(f135,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) )),
  inference(cnf_transformation,[],[f113])).
fof(f114,plain,(
  s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK1),'type.semi_ring.semi_ring'(sK1)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(sK1,sK1),'const.combin.o_2'(s(fun(sK1,sK1),sK4),s(fun(sK1,sK1),sK3))))),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),sK2))) != s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK1),'type.semi_ring.semi_ring'(sK1)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(sK1,sK1),sK4))),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK1),'type.semi_ring.semi_ring'(sK1)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(sK1,sK1),sK3))),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),sK2)))))),
  inference(cnf_transformation,[],[f103])).
fof(f103,plain,(
  s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK1),'type.semi_ring.semi_ring'(sK1)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(sK1,sK1),'const.combin.o_2'(s(fun(sK1,sK1),sK4),s(fun(sK1,sK1),sK3))))),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),sK2))) != s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK1),'type.semi_ring.semi_ring'(sK1)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(sK1,sK1),sK4))),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(sK1),'type.semi_ring.semi_ring'(sK1)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(sK1,sK1),sK3))),s(fun(sK0,'type.semi_ring.semi_ring'(sK1)),sK2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4])],[f101])).
fof(f101,plain,(
  ? [X0,X1,X2,X3,X4] : s(fun(X0,'type.semi_ring.semi_ring'(X1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X1),'type.semi_ring.semi_ring'(X1)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(X1,X1),'const.combin.o_2'(s(fun(X1,X1),X4),s(fun(X1,X1),X3))))),s(fun(X0,'type.semi_ring.semi_ring'(X1)),X2))) != s(fun(X0,'type.semi_ring.semi_ring'(X1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X1),'type.semi_ring.semi_ring'(X1)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(X1,X1),X4))),s(fun(X0,'type.semi_ring.semi_ring'(X1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X1),'type.semi_ring.semi_ring'(X1)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(X1,X1),X3))),s(fun(X0,'type.semi_ring.semi_ring'(X1)),X2)))))),
  inference(ennf_transformation,[],[f39])).
fof(f39,plain,(
  ~! [X0,X1,X2,X3,X4] : s(fun(X0,'type.semi_ring.semi_ring'(X1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X1),'type.semi_ring.semi_ring'(X1)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(X1,X1),'const.combin.o_2'(s(fun(X1,X1),X4),s(fun(X1,X1),X3))))),s(fun(X0,'type.semi_ring.semi_ring'(X1)),X2))) = s(fun(X0,'type.semi_ring.semi_ring'(X1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X1),'type.semi_ring.semi_ring'(X1)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(X1,X1),X4))),s(fun(X0,'type.semi_ring.semi_ring'(X1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X1),'type.semi_ring.semi_ring'(X1)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(X1,X1),X3))),s(fun(X0,'type.semi_ring.semi_ring'(X1)),X2)))))),
  inference(flattening,[],[f38])).
fof(f38,plain,(
  ~! [X0,X1] : ! [X2,X3,X4] : s(fun(X0,'type.semi_ring.semi_ring'(X1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X1),'type.semi_ring.semi_ring'(X1)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(X1,X1),'const.combin.o_2'(s(fun(X1,X1),X4),s(fun(X1,X1),X3))))),s(fun(X0,'type.semi_ring.semi_ring'(X1)),X2))) = s(fun(X0,'type.semi_ring.semi_ring'(X1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X1),'type.semi_ring.semi_ring'(X1)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(X1,X1),X4))),s(fun(X0,'type.semi_ring.semi_ring'(X1)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X1),'type.semi_ring.semi_ring'(X1)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(X1,X1),X3))),s(fun(X0,'type.semi_ring.semi_ring'(X1)),X2)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10] : s(fun(X6,'type.semi_ring.semi_ring'(X7)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X7),'type.semi_ring.semi_ring'(X7)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(X7,X7),'const.combin.o_2'(s(fun(X7,X7),X10),s(fun(X7,X7),X9))))),s(fun(X6,'type.semi_ring.semi_ring'(X7)),X8))) = s(fun(X6,'type.semi_ring.semi_ring'(X7)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X7),'type.semi_ring.semi_ring'(X7)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(X7,X7),X10))),s(fun(X6,'type.semi_ring.semi_ring'(X7)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X7),'type.semi_ring.semi_ring'(X7)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(X7,X7),X9))),s(fun(X6,'type.semi_ring.semi_ring'(X7)),X8)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10] : s(fun(X6,'type.semi_ring.semi_ring'(X7)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X7),'type.semi_ring.semi_ring'(X7)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(X7,X7),'const.combin.o_2'(s(fun(X7,X7),X10),s(fun(X7,X7),X9))))),s(fun(X6,'type.semi_ring.semi_ring'(X7)),X8))) = s(fun(X6,'type.semi_ring.semi_ring'(X7)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X7),'type.semi_ring.semi_ring'(X7)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(X7,X7),X10))),s(fun(X6,'type.semi_ring.semi_ring'(X7)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X7),'type.semi_ring.semi_ring'(X7)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(X7,X7),X9))),s(fun(X6,'type.semi_ring.semi_ring'(X7)),X8)))))),
  file('thm.semi_ring.semi_ring_fupdfupds_comp|split|1','thm.semi_ring.semi_ring_fupdfupds_comp|split|1')).
fof(f5324,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun(X3,'type.semi_ring.semi_ring'(X0)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(X0,X0),'const.combin.o_2'(s(fun(X0,X0),X1),s(fun(X0,X0),X2))))),s(fun(X3,'type.semi_ring.semi_ring'(X0)),X4))) = s(fun(X3,'type.semi_ring.semi_ring'(X0)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(X0,X0),X1))),s(fun(X3,'type.semi_ring.semi_ring'(X0)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(X0,X0),X2))),s(fun(X3,'type.semi_ring.semi_ring'(X0)),X4)))))) )),
  inference(superposition,[],[f140,f120])).
fof(f120,plain,(
  ( ! [X2,X0,X1] : (s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(X0,X0),'const.combin.o_2'(s(fun(X0,X0),X2),s(fun(X0,X0),X1))))) = s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(X0,X0),X2))),s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(X0,X0),X1)))))) )),
  inference(cnf_transformation,[],[f46])).
fof(f46,plain,(
  ! [X0,X1,X2] : s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(X0,X0),'const.combin.o_2'(s(fun(X0,X0),X2),s(fun(X0,X0),X1))))) = s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(X0,X0),X2))),s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(X0,X0),X1)))))),
  inference(flattening,[],[f45])).
fof(f45,plain,(
  ! [X0] : ! [X1,X2] : s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(X0,X0),'const.combin.o_2'(s(fun(X0,X0),X2),s(fun(X0,X0),X1))))) = s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(X0,X0),X2))),s(fun('type.semi_ring.semi_ring'(X0),'type.semi_ring.semi_ring'(X0)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(X0,X0),X1)))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X7] : ! [X9,X10] : s(fun('type.semi_ring.semi_ring'(X7),'type.semi_ring.semi_ring'(X7)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(X7,X7),'const.combin.o_2'(s(fun(X7,X7),X10),s(fun(X7,X7),X9))))) = s(fun('type.semi_ring.semi_ring'(X7),'type.semi_ring.semi_ring'(X7)),'const.combin.o_2'(s(fun('type.semi_ring.semi_ring'(X7),'type.semi_ring.semi_ring'(X7)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(X7,X7),X10))),s(fun('type.semi_ring.semi_ring'(X7),'type.semi_ring.semi_ring'(X7)),'const.semi_ring.semi_ring_SR|48|_fupd_1'(s(fun(X7,X7),X9)))))),
  file('thm.semi_ring.semi_ring_fupdfupds_comp|split|1','thm.semi_ring.semi_ring_fupdfupds_comp|split|0')).
fof(f140,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))) )),
  inference(cnf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))),
  inference(flattening,[],[f73])).
fof(f73,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6] : s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X6,X7,X18,X14] : ! [X10,X9,X8] : s(fun(X18,X6),'const.combin.o_2'(s(fun(X7,X6),X10),s(fun(X18,X7),'const.combin.o_2'(s(fun(X14,X7),X9),s(fun(X18,X14),X8))))) = s(fun(X18,X6),'const.combin.o_2'(s(fun(X14,X6),'const.combin.o_2'(s(fun(X7,X6),X10),s(fun(X14,X7),X9))),s(fun(X18,X14),X8)))),
  file('thm.semi_ring.semi_ring_fupdfupds_comp|split|1','thm.combin.o_ASSOC')).
% SZS output end Proof for thm.semi_ring
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 92
% Passive clauses: 645
% Generated clauses: 5465
% Final active clauses: 51
% Final passive clauses: 309
% Input formulas: 36
% Initial clauses: 40
% 
% Duplicate literals: 258
% Trivial inequalities: 31
% Fw subsumption resolutions: 142
% Fw demodulations: 813
% Bw demodulations: 247
% Condensations: 406
% Global subsumptions: 1
% 
% Simple tautologies: 38
% Equational tautologies: 42
% Forward subsumptions: 3020
% Fw demodulations to eq. taut.: 1
% Bw demodulations to eq. taut.: 13
% 
% Unit resulting resolution: 200
% Forward superposition: 1643
% Backward superposition: 1242
% Self superposition: 99
% Equality factoring: 44
% Equality resolution: 2
% 
% Split clauses: 13
% Split components: 26
% SAT solver clauses: 690
% SAT solver unit clauses: 123
% SAT solver binary clauses: 284
% 
% Sat splits: 13
% Sat splitting refutations: 16
% 
% Memory used [KB]: 6780
% Time elapsed: 0.226 s
% ------------------------------
% ------------------------------
% Success in time 1.56 s
