Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.indexedLists
% (101196)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 128
% Passive clauses: 1871
% Generated clauses: 20575
% Final active clauses: 114
% Final passive clauses: 1466
% Input formulas: 27
% Initial clauses: 50
% 
% Duplicate literals: 945
% Trivial inequalities: 402
% Fw subsumption resolutions: 3348
% Fw demodulations: 2264
% Bw demodulations: 10
% 
% Simple tautologies: 438
% Equational tautologies: 57
% Forward subsumptions: 9620
% Fw demodulations to eq. taut.: 2
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 287
% Factoring: 3
% Forward superposition: 6568
% Backward superposition: 6109
% Self superposition: 231
% Equality factoring: 159
% Equality resolution: 27
% 
% Split clauses: 1245
% Split components: 2614
% SAT solver clauses: 1847
% SAT solver unit clauses: 43
% SAT solver binary clauses: 320
% SAT solver learnt clauses: 1
% SAT solver learnt literals: 1
% 
% Sat splits: 1245
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1837
% TWLsolver calls for satisfiability: 155
% 
% Memory used [KB]: 10874
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.indexedLists
% (101211)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 77
% Passive clauses: 305
% Generated clauses: 8053
% Final active clauses: 72
% Final passive clauses: 222
% Input formulas: 27
% Initial clauses: 50
% 
% Duplicate literals: 720
% Trivial inequalities: 110
% Fw subsumption resolutions: 608
% Fw demodulations: 195
% Bw demodulations: 5
% Condensations: 1272
% Global subsumptions: 9
% 
% Simple tautologies: 188
% Equational tautologies: 64
% Forward subsumptions: 3860
% Fw demodulations to eq. taut.: 4
% 
% Binary resolution: 179
% Factoring: 2
% Forward superposition: 2309
% Backward superposition: 3246
% Self superposition: 231
% Equality factoring: 191
% Equality resolution: 11
% 
% Split clauses: 81
% Split components: 172
% SAT solver clauses: 626
% SAT solver unit clauses: 66
% SAT solver binary clauses: 43
% SAT solver learnt clauses: 19
% SAT solver learnt literals: 1
% 
% Sat splits: 81
% Sat splitting refutations: 6
% 
% TWLsolver clauses: 589
% TWLsolver calls for satisfiability: 213
% 
% Memory used [KB]: 7291
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.indexedLists
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.indexedLists
% SZS output start Proof for thm.indexedLists
fof(f1998,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f94,f148,f146,f93,f147,f145,f151,f604,f605,f612,f609,f606,f607,f615,f608,f617,f610,f614,f788,f805,f803,f812,f813,f821,f819,f814,f815,f824,f816,f823,f864,f879,f881,f880,f883,f886,f839,f826,f884,f941,f942,f943,f944,f945,f946,f1050,f1051,f1052,f1053,f1054,f1055,f1840,f1862,f1861,f1958,f1953,f1995,f1990])).
fof(f1990,plain,(
  $false | ($spl0 | $spl38)),
  inference(unit_resulting_resolution,[],[f145,f1862])).
fof(f1995,plain,(
  $false | ($spl0 | $spl38)),
  inference(trivial_inequality_removal,[],[f1994])).
fof(f1994,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') != s('type.list.list'(sK0),'const.list.NIL_0') | ($spl0 | $spl38)),
  inference(superposition,[],[f1862,f145])).
fof(f1953,plain,(
  $false | ($spl0 | $spl3 | $spl36)),
  inference(unit_resulting_resolution,[],[f145,f1921])).
fof(f1921,plain,(
  ( ! [X18] : (s('type.list.list'(sK0),'const.list.NIL_0') != s('type.list.list'(sK0),'const.indexedLists.fupdLast_2'(s(fun(sK0,sK0),X18),s('type.list.list'(sK0),sK1)))) ) | ($spl3 | $spl36)),
  inference(superposition,[],[f136,f1887])).
fof(f1887,plain,(
  ( ! [X2] : (s('type.list.list'(sK0),'const.indexedLists.fupdLast_2'(s(fun(sK0,sK0),X2),s('type.list.list'(sK0),sK1))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,happ(s(fun(sK0,sK0),X2),s(sK0,sK3(sK1,sK0)))),s('type.list.list'(sK0),'const.list.NIL_0')))) ) | ($spl3 | $spl36)),
  inference(superposition,[],[f133,f1863])).
fof(f1863,plain,(
  s('type.list.list'(sK0),sK1) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK3(sK1,sK0)),s('type.list.list'(sK0),'const.list.NIL_0'))) | ($spl3 | $spl36)),
  inference(backward_demodulation,[],[f1861,f213])).
fof(f213,plain,(
  s('type.list.list'(sK0),sK1) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK3(sK1,sK0)),s('type.list.list'(sK0),sK4(sK1,sK0)))) | $spl3),
  inference(unit_resulting_resolution,[],[f148,f130])).
fof(f130,plain,(
  ( ! [X0,X1] : (s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,sK3(X1,X0)),s('type.list.list'(X0),sK4(X1,X0)))) | s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X1)) )),
  inference(cnf_transformation,[],[f89])).
fof(f89,plain,(
  ! [X0,X1] : (s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X1) | s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,sK3(X1,X0)),s('type.list.list'(X0),sK4(X1,X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3,sK4])],[f51])).
fof(f51,plain,(
  ! [X0,X1] : (s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X1) | ? [X2,X3] : s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3))))),
  inference(flattening,[],[f50])).
fof(f50,plain,(
  ! [X0] : ! [X1] : (s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X1) | ? [X2,X3] : s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3))))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X6] : ! [X16] : (s('type.list.list'(X6),'const.list.NIL_0') = s('type.list.list'(X6),X16) | ? [X13,X11] : s('type.list.list'(X6),X16) = s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X13),s('type.list.list'(X6),X11))))),
  file('thm.indexedLists.fupdLast_EQ_NIL|split|1','thm.list.list_CASES')).
fof(f133,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'(X0),'const.list.CONS_2'(s(X0,happ(s(fun(X0,X0),X2),s(X0,X1))),s('type.list.list'(X0),'const.list.NIL_0'))) = s('type.list.list'(X0),'const.indexedLists.fupdLast_2'(s(fun(X0,X0),X2),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),'const.list.NIL_0')))))) )),
  inference(cnf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0,X1,X2] : s('type.list.list'(X0),'const.list.CONS_2'(s(X0,happ(s(fun(X0,X0),X2),s(X0,X1))),s('type.list.list'(X0),'const.list.NIL_0'))) = s('type.list.list'(X0),'const.indexedLists.fupdLast_2'(s(fun(X0,X0),X2),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),'const.list.NIL_0')))))),
  inference(flattening,[],[f54])).
fof(f54,plain,(
  ! [X0] : ! [X1,X2] : s('type.list.list'(X0),'const.list.CONS_2'(s(X0,happ(s(fun(X0,X0),X2),s(X0,X1))),s('type.list.list'(X0),'const.list.NIL_0'))) = s('type.list.list'(X0),'const.indexedLists.fupdLast_2'(s(fun(X0,X0),X2),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),'const.list.NIL_0')))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X6] : ! [X13,X8] : s('type.list.list'(X6),'const.list.CONS_2'(s(X6,happ(s(fun(X6,X6),X8),s(X6,X13))),s('type.list.list'(X6),'const.list.NIL_0'))) = s('type.list.list'(X6),'const.indexedLists.fupdLast_2'(s(fun(X6,X6),X8),s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X13),s('type.list.list'(X6),'const.list.NIL_0')))))),
  file('thm.indexedLists.fupdLast_EQ_NIL|split|1','thm.indexedLists.fupdLast_def|split|1')).
fof(f136,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'(X0),'const.list.NIL_0') != s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X1)))) )),
  inference(cnf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0,X1,X2] : s('type.list.list'(X0),'const.list.NIL_0') != s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X1)))),
  inference(flattening,[],[f58])).
fof(f58,plain,(
  ! [X0] : ! [X1,X2] : ~s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X1)))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X6] : ! [X17,X18] : ~s('type.list.list'(X6),'const.list.NIL_0') = s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X18),s('type.list.list'(X6),X17)))),
  file('thm.indexedLists.fupdLast_EQ_NIL|split|1','thm.list.list_distinct')).
fof(f1958,plain,(
  $false | ($spl0 | $spl3 | $spl36)),
  inference(trivial_inequality_removal,[],[f1957])).
fof(f1957,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') != s('type.list.list'(sK0),'const.list.NIL_0') | ($spl0 | $spl3 | $spl36)),
  inference(superposition,[],[f1921,f145])).
fof(f1861,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),sK4(sK1,sK0)) | $spl36),
  inference(cnf_transformation,[],[f1861_D])).
fof(f1861_D,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),sK4(sK1,sK0)) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f1862,plain,(
  ( ! [X31] : (s('type.list.list'(sK0),'const.list.NIL_0') != s('type.list.list'(sK0),'const.indexedLists.fupdLast_2'(s(fun(sK0,sK0),X31),s('type.list.list'(sK0),sK1)))) ) | $spl38),
  inference(cnf_transformation,[],[f1862_D])).
fof(f1862_D,plain,(
  ( ! [X31] : (s('type.list.list'(sK0),'const.list.NIL_0') != s('type.list.list'(sK0),'const.indexedLists.fupdLast_2'(s(fun(sK0,sK0),X31),s('type.list.list'(sK0),sK1)))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f1840,plain,(
  ( ! [X31] : (s('type.list.list'(sK0),'const.list.NIL_0') != s('type.list.list'(sK0),'const.indexedLists.fupdLast_2'(s(fun(sK0,sK0),X31),s('type.list.list'(sK0),sK1))) | s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),sK4(sK1,sK0))) ) | $spl3),
  inference(superposition,[],[f718,f213])).
fof(f718,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.list.list'(X0),'const.list.NIL_0') != s('type.list.list'(X0),'const.indexedLists.fupdLast_2'(s(fun(X0,X0),X2),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X1))))) | s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X1)) )),
  inference(superposition,[],[f656,f130])).
fof(f656,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.list.list'(X0),'const.list.NIL_0') != s('type.list.list'(X0),'const.indexedLists.fupdLast_2'(s(fun(X0,X0),X2),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X4)))))))) )),
  inference(superposition,[],[f136,f137])).
fof(f137,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.list.list'(X0),'const.indexedLists.fupdLast_2'(s(fun(X0,X0),X4),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X1))))))) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),'const.indexedLists.fupdLast_2'(s(fun(X0,X0),X4),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X1)))))))) )),
  inference(cnf_transformation,[],[f61])).
fof(f61,plain,(
  ! [X0,X1,X2,X3,X4] : s('type.list.list'(X0),'const.indexedLists.fupdLast_2'(s(fun(X0,X0),X4),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X1))))))) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),'const.indexedLists.fupdLast_2'(s(fun(X0,X0),X4),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X1)))))))),
  inference(flattening,[],[f60])).
fof(f60,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : s('type.list.list'(X0),'const.indexedLists.fupdLast_2'(s(fun(X0,X0),X4),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X1))))))) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),'const.indexedLists.fupdLast_2'(s(fun(X0,X0),X4),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X1)))))))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X6] : ! [X14,X15,X13,X8] : s('type.list.list'(X6),'const.indexedLists.fupdLast_2'(s(fun(X6,X6),X8),s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X13),s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X15),s('type.list.list'(X6),X14))))))) = s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X13),s('type.list.list'(X6),'const.indexedLists.fupdLast_2'(s(fun(X6,X6),X8),s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X15),s('type.list.list'(X6),X14)))))))),
  file('thm.indexedLists.fupdLast_EQ_NIL|split|1','thm.indexedLists.fupdLast_def|split|2')).
fof(f1055,plain,(
  $false | ($spl19 | $spl24 | $spl29)),
  inference(subsumption_resolution,[],[f997,f826])).
fof(f997,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | ($spl19 | $spl24)),
  inference(unit_resulting_resolution,[],[f887,f102])).
fof(f102,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.indexedLists.fupdLast_EQ_NIL|split|1',aHL_BOOL_CASES)).
fof(f887,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl19 | $spl24)),
  inference(forward_demodulation,[],[f805,f823])).
fof(f1054,plain,(
  $false | ($spl19 | $spl24 | $spl29)),
  inference(subsumption_resolution,[],[f999,f826])).
fof(f999,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | ($spl19 | $spl24)),
  inference(unit_resulting_resolution,[],[f152,f887,f161])).
fof(f161,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f102,f102])).
fof(f152,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f149,f104])).
fof(f104,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X11] : (s(bool,'T_0') = s(bool,X11) <=> p(s(bool,X11)))),
  file('thm.indexedLists.fupdLast_EQ_NIL|split|1','thm.bool.EQ_CLAUSES|split|1')).
fof(f149,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f98,f95])).
fof(f95,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.indexedLists.fupdLast_EQ_NIL|split|1',aHL_TRUTH)).
fof(f98,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f31])).
fof(f31,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f16])).
fof(f16,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.indexedLists.fupdLast_EQ_NIL|split|1','thm.bool.NOT_CLAUSES|split|1')).
fof(f1053,plain,(
  $false | ($spl19 | $spl24 | $spl29)),
  inference(subsumption_resolution,[],[f1002,f826])).
fof(f1002,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | ($spl19 | $spl24)),
  inference(unit_resulting_resolution,[],[f152,f887,f161])).
fof(f1052,plain,(
  $false | ($spl19 | $spl24 | $spl29)),
  inference(subsumption_resolution,[],[f1006,f826])).
fof(f1006,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | ($spl19 | $spl24)),
  inference(unit_resulting_resolution,[],[f149,f887,f190])).
fof(f190,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f187,f149])).
fof(f187,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f138,f102])).
fof(f138,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f68])).
fof(f68,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f67])).
fof(f67,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f62])).
fof(f62,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X9,X10] : ((p(s(bool,X9)) => p(s(bool,X10))) => ((p(s(bool,X10)) => p(s(bool,X9))) => s(bool,X9) = s(bool,X10)))),
  file('thm.indexedLists.fupdLast_EQ_NIL|split|1','thm.bool.IMP_ANTISYM_AX')).
fof(f1051,plain,(
  $false | ($spl19 | $spl24 | $spl29)),
  inference(subsumption_resolution,[],[f1008,f826])).
fof(f1008,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | ($spl19 | $spl24)),
  inference(unit_resulting_resolution,[],[f95,f152,f887,f469])).
fof(f469,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,X24)) | s(bool,X23) = s(bool,X24) | s(bool,X23) = s(bool,X25) | s(bool,'T_0') = s(bool,X25)) )),
  inference(subsumption_resolution,[],[f318,f95])).
fof(f318,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,'T_0')) | ~p(s(bool,X24)) | s(bool,X23) = s(bool,X24) | s(bool,X23) = s(bool,X25) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f141,f161])).
fof(f141,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f68])).
fof(f1050,plain,(
  $false | ($spl19 | $spl24 | $spl29)),
  inference(subsumption_resolution,[],[f1030,f826])).
fof(f1030,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | ($spl19 | $spl24)),
  inference(unit_resulting_resolution,[],[f95,f152,f887,f469])).
fof(f946,plain,(
  $false | ($spl23 | $spl29)),
  inference(subsumption_resolution,[],[f888,f826])).
fof(f888,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | $spl23),
  inference(unit_resulting_resolution,[],[f821,f102])).
fof(f945,plain,(
  $false | ($spl23 | $spl29)),
  inference(subsumption_resolution,[],[f890,f826])).
fof(f890,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | $spl23),
  inference(unit_resulting_resolution,[],[f152,f821,f161])).
fof(f944,plain,(
  $false | ($spl23 | $spl29)),
  inference(subsumption_resolution,[],[f893,f826])).
fof(f893,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | $spl23),
  inference(unit_resulting_resolution,[],[f152,f821,f161])).
fof(f943,plain,(
  $false | ($spl23 | $spl29)),
  inference(subsumption_resolution,[],[f897,f826])).
fof(f897,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | $spl23),
  inference(unit_resulting_resolution,[],[f149,f821,f190])).
fof(f942,plain,(
  $false | ($spl23 | $spl29)),
  inference(subsumption_resolution,[],[f899,f826])).
fof(f899,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | $spl23),
  inference(unit_resulting_resolution,[],[f95,f152,f821,f469])).
fof(f941,plain,(
  $false | ($spl23 | $spl29)),
  inference(subsumption_resolution,[],[f921,f826])).
fof(f921,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | $spl23),
  inference(unit_resulting_resolution,[],[f95,f152,f821,f469])).
fof(f884,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,g(s(bool,sK5(X0,X1,bool,bool)))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1)) ) | $spl34),
  inference(cnf_transformation,[],[f884_D])).
fof(f884_D,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,g(s(bool,sK5(X0,X1,bool,bool)))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1)) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f826,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | $spl29),
  inference(cnf_transformation,[],[f826_D])).
fof(f826_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f839,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X0,X1,bool,bool))))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f827])).
fof(f827,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X0,X1,bool,bool)))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1)) ) | $spl6),
  inference(superposition,[],[f485,f610])).
fof(f485,plain,(
  ( ! [X17,X18,X16] : (s(bool,'F_0') != s(bool,f(s(X16,sK5(X17,X18,bool,X16)))) | s(fun(X16,bool),X17) = s(fun(X16,bool),X18) | s(bool,'T_0') = s(bool,g(s(X16,sK5(X17,X18,bool,X16))))) )),
  inference(superposition,[],[f142,f102])).
fof(f142,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK5(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK5(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f92])).
fof(f92,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK5(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK5(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5])],[f69])).
fof(f69,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f64])).
fof(f64,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f63])).
fof(f63,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.indexedLists.fupdLast_EQ_NIL|split|1',aHL_EXT)).
fof(f886,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,g(s(bool,sK5(X0,X1,bool,bool)))) | s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1)) ) | $spl6),
  inference(condensation,[],[f885])).
fof(f885,plain,(
  ( ! [X4,X2,X3] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X2,X3,bool,bool)))) | s(bool,'T_0') = s(bool,X4)) ) | $spl6),
  inference(subsumption_resolution,[],[f828,f798])).
fof(f798,plain,(
  ( ! [X30,X31,X29] : (s(bool,X31) != s(bool,sK5(X29,X30,bool,bool)) | s(bool,'T_0') = s(bool,X31) | s(fun(bool,bool),X29) = s(fun(bool,bool),X30)) ) | $spl6),
  inference(subsumption_resolution,[],[f775,f95])).
fof(f775,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,'T_0')) | s(bool,X31) != s(bool,sK5(X29,X30,bool,bool)) | s(bool,'T_0') = s(bool,X31) | s(fun(bool,bool),X29) = s(fun(bool,bool),X30)) ) | $spl6),
  inference(superposition,[],[f163,f610])).
fof(f163,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f108,f102])).
fof(f108,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f37])).
fof(f37,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X11] : (s(bool,'F_0') = s(bool,X11) <=> ~p(s(bool,X11)))),
  file('thm.indexedLists.fupdLast_EQ_NIL|split|1','thm.bool.EQ_CLAUSES|split|2')).
fof(f828,plain,(
  ( ! [X4,X2,X3] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X2,X3,bool,bool)))) | s(bool,X4) = s(bool,sK5(X2,X3,bool,bool)) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f485,f161])).
fof(f883,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,g(s(bool,sK5(X0,X1,bool,bool)))) | s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1)) ) | $spl6),
  inference(condensation,[],[f882])).
fof(f882,plain,(
  ( ! [X6,X7,X5] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X5) = s(fun(bool,bool),X6) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X5,X6,bool,bool)))) | s(bool,'T_0') = s(bool,X7)) ) | $spl6),
  inference(subsumption_resolution,[],[f829,f798])).
fof(f829,plain,(
  ( ! [X6,X7,X5] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X5) = s(fun(bool,bool),X6) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X5,X6,bool,bool)))) | s(bool,X7) = s(bool,sK5(X5,X6,bool,bool)) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f485,f161])).
fof(f880,plain,(
  ( ! [X59,X60] : (s(fun(fun(bool,bool),bool),X59) = s(fun(fun(bool,bool),bool),X60) | s(bool,'T_0') = s(bool,g(s(fun(bool,bool),sK5(X59,X60,bool,fun(bool,bool)))))) ) | $spl30),
  inference(cnf_transformation,[],[f880_D])).
fof(f880_D,plain,(
  ( ! [X59,X60] : (s(fun(fun(bool,bool),bool),X59) = s(fun(fun(bool,bool),bool),X60) | s(bool,'T_0') = s(bool,g(s(fun(bool,bool),sK5(X59,X60,bool,fun(bool,bool)))))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f881,plain,(
  ( ! [X61] : (s(bool,'F_0') != s(bool,f(s(fun(bool,bool),X61)))) ) | $spl32),
  inference(cnf_transformation,[],[f881_D])).
fof(f881_D,plain,(
  ( ! [X61] : (s(bool,'F_0') != s(bool,f(s(fun(bool,bool),X61)))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f879,plain,(
  ( ! [X61,X59,X60] : (s(bool,'F_0') != s(bool,f(s(fun(bool,bool),X61))) | s(fun(fun(bool,bool),bool),X59) = s(fun(fun(bool,bool),bool),X60) | s(bool,'T_0') = s(bool,g(s(fun(bool,bool),sK5(X59,X60,bool,fun(bool,bool)))))) ) | $spl16),
  inference(superposition,[],[f485,f803])).
fof(f864,plain,(
  ( ! [X61,X59,X60] : (s(bool,'F_0') != s(bool,f(s(fun(bool,bool),X61))) | s(fun(fun(bool,bool),bool),X59) = s(fun(fun(bool,bool),bool),X60) | s(bool,'T_0') = s(bool,g(s(fun(bool,bool),sK5(X59,X60,bool,fun(bool,bool)))))) ) | $spl16),
  inference(superposition,[],[f485,f803])).
fof(f823,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | $spl24),
  inference(cnf_transformation,[],[f823_D])).
fof(f823_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f816,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | $spl19),
  inference(superposition,[],[f805,f102])).
fof(f824,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl26),
  inference(cnf_transformation,[],[f824_D])).
fof(f824_D,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f815,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,g(s(bool,'T_0')))) ) | $spl19),
  inference(superposition,[],[f805,f161])).
fof(f814,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl19),
  inference(superposition,[],[f805,f161])).
fof(f819,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl20),
  inference(cnf_transformation,[],[f819_D])).
fof(f819_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f821,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | $spl23),
  inference(cnf_transformation,[],[f821_D])).
fof(f821_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f813,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | s(bool,X1) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X1)) ) | $spl19),
  inference(superposition,[],[f805,f161])).
fof(f812,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | s(bool,X0) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl19),
  inference(superposition,[],[f805,f161])).
fof(f803,plain,(
  ( ! [X50,X51] : (s(fun(bool,bool),X50) = s(fun(bool,bool),X51)) ) | $spl16),
  inference(cnf_transformation,[],[f803_D])).
fof(f803_D,plain,(
  ( ! [X50,X51] : (s(fun(bool,bool),X50) = s(fun(bool,bool),X51)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f805,plain,(
  s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | $spl19),
  inference(cnf_transformation,[],[f805_D])).
fof(f805_D,plain,(
  s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f788,plain,(
  ( ! [X50,X51] : (s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X50) = s(fun(bool,bool),X51)) ) | $spl6),
  inference(duplicate_literal_removal,[],[f781])).
fof(f781,plain,(
  ( ! [X50,X51] : (s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X50) = s(fun(bool,bool),X51) | s(fun(bool,bool),X50) = s(fun(bool,bool),X51)) ) | $spl6),
  inference(superposition,[],[f142,f610])).
fof(f614,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl10),
  inference(cnf_transformation,[],[f614_D])).
fof(f614_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f610,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK5(X0,X1,bool,bool)) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1)) ) | $spl6),
  inference(cnf_transformation,[],[f610_D])).
fof(f610_D,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK5(X0,X1,bool,bool)) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1)) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f617,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl15),
  inference(cnf_transformation,[],[f617_D])).
fof(f617_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f608,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X12) = s(fun(bool,bool),X13) | s(bool,'T_0') = s(bool,sK5(X12,X13,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f480,f102])).
fof(f480,plain,(
  ( ! [X17,X18,X16] : (s(X18,f(s(bool,'F_0'))) != s(X18,g(s(bool,'F_0'))) | s(fun(bool,X18),X16) = s(fun(bool,X18),X17) | s(bool,'T_0') = s(bool,sK5(X16,X17,X18,bool))) )),
  inference(superposition,[],[f142,f102])).
fof(f615,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl12),
  inference(cnf_transformation,[],[f615_D])).
fof(f615_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f607,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X10) = s(fun(bool,bool),X11) | s(bool,'T_0') = s(bool,sK5(X10,X11,bool,bool)) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f480,f161])).
fof(f606,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X7) = s(fun(bool,bool),X8) | s(bool,'T_0') = s(bool,sK5(X7,X8,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f480,f161])).
fof(f609,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl4),
  inference(cnf_transformation,[],[f609_D])).
fof(f609_D,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f612,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | $spl9),
  inference(cnf_transformation,[],[f612_D])).
fof(f612_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f605,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,sK5(X3,X4,bool,bool)) | s(bool,X5) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f480,f161])).
fof(f604,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK5(X0,X1,bool,bool)) | s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f480,f161])).
fof(f151,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f150,f129])).
fof(f129,plain,(
  ( ! [X0,X1] : (s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),'const.indexedLists.fupdLast_2'(s(fun(X0,X0),X1),s('type.list.list'(X0),'const.list.NIL_0')))) )),
  inference(cnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0,X1] : s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),'const.indexedLists.fupdLast_2'(s(fun(X0,X0),X1),s('type.list.list'(X0),'const.list.NIL_0')))),
  inference(flattening,[],[f48])).
fof(f48,plain,(
  ! [X0] : ! [X1] : s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),'const.indexedLists.fupdLast_2'(s(fun(X0,X0),X1),s('type.list.list'(X0),'const.list.NIL_0')))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X6] : ! [X8] : s('type.list.list'(X6),'const.list.NIL_0') = s('type.list.list'(X6),'const.indexedLists.fupdLast_2'(s(fun(X6,X6),X8),s('type.list.list'(X6),'const.list.NIL_0')))),
  file('thm.indexedLists.fupdLast_EQ_NIL|split|1','thm.indexedLists.fupdLast_def|split|0')).
fof(f150,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') != s('type.list.list'(sK0),'const.indexedLists.fupdLast_2'(s(fun(sK0,sK0),sK2),s('type.list.list'(sK0),'const.list.NIL_0'))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f146,f147])).
fof(f145,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),'const.indexedLists.fupdLast_2'(s(fun(sK0,sK0),sK2),s('type.list.list'(sK0),sK1))) | $spl0),
  inference(cnf_transformation,[],[f145_D])).
fof(f145_D,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),'const.indexedLists.fupdLast_2'(s(fun(sK0,sK0),sK2),s('type.list.list'(sK0),sK1))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f147,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),sK1) | $spl2),
  inference(cnf_transformation,[],[f147_D])).
fof(f147_D,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),sK1) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f93,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),sK1) | s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),'const.indexedLists.fupdLast_2'(s(fun(sK0,sK0),sK2),s('type.list.list'(sK0),sK1)))),
  inference(cnf_transformation,[],[f71])).
fof(f71,plain,(
  (s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),'const.indexedLists.fupdLast_2'(s(fun(sK0,sK0),sK2),s('type.list.list'(sK0),sK1))) | s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),sK1)) & (s('type.list.list'(sK0),'const.list.NIL_0') != s('type.list.list'(sK0),'const.indexedLists.fupdLast_2'(s(fun(sK0,sK0),sK2),s('type.list.list'(sK0),sK1))) | s('type.list.list'(sK0),'const.list.NIL_0') != s('type.list.list'(sK0),sK1))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f70])).
fof(f70,plain,(
  ? [X0,X1,X2] : ((s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),'const.indexedLists.fupdLast_2'(s(fun(X0,X0),X2),s('type.list.list'(X0),X1))) | s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X1)) & (s('type.list.list'(X0),'const.list.NIL_0') != s('type.list.list'(X0),'const.indexedLists.fupdLast_2'(s(fun(X0,X0),X2),s('type.list.list'(X0),X1))) | s('type.list.list'(X0),'const.list.NIL_0') != s('type.list.list'(X0),X1)))),
  inference(nnf_transformation,[],[f65])).
fof(f65,plain,(
  ? [X0,X1,X2] : (s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),'const.indexedLists.fupdLast_2'(s(fun(X0,X0),X2),s('type.list.list'(X0),X1))) <~> s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X1))),
  inference(ennf_transformation,[],[f30])).
fof(f30,plain,(
  ~! [X0,X1,X2] : (s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),'const.indexedLists.fupdLast_2'(s(fun(X0,X0),X2),s('type.list.list'(X0),X1))) <=> s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X1))),
  inference(flattening,[],[f29])).
fof(f29,plain,(
  ~! [X0] : ! [X1,X2] : (s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),'const.indexedLists.fupdLast_2'(s(fun(X0,X0),X2),s('type.list.list'(X0),X1))) <=> s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X1))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : (s('type.list.list'(X6),'const.list.NIL_0') = s('type.list.list'(X6),'const.indexedLists.fupdLast_2'(s(fun(X6,X6),X8),s('type.list.list'(X6),X7))) <=> s('type.list.list'(X6),'const.list.NIL_0') = s('type.list.list'(X6),X7))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : (s('type.list.list'(X6),'const.list.NIL_0') = s('type.list.list'(X6),'const.indexedLists.fupdLast_2'(s(fun(X6,X6),X8),s('type.list.list'(X6),X7))) <=> s('type.list.list'(X6),'const.list.NIL_0') = s('type.list.list'(X6),X7))),
  file('thm.indexedLists.fupdLast_EQ_NIL|split|1','thm.indexedLists.fupdLast_EQ_NIL|split|1')).
fof(f146,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') != s('type.list.list'(sK0),'const.indexedLists.fupdLast_2'(s(fun(sK0,sK0),sK2),s('type.list.list'(sK0),sK1))) | $spl1),
  inference(cnf_transformation,[],[f146_D])).
fof(f146_D,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') != s('type.list.list'(sK0),'const.indexedLists.fupdLast_2'(s(fun(sK0,sK0),sK2),s('type.list.list'(sK0),sK1))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f148,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') != s('type.list.list'(sK0),sK1) | $spl3),
  inference(cnf_transformation,[],[f148_D])).
fof(f148_D,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') != s('type.list.list'(sK0),sK1) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f94,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') != s('type.list.list'(sK0),sK1) | s('type.list.list'(sK0),'const.list.NIL_0') != s('type.list.list'(sK0),'const.indexedLists.fupdLast_2'(s(fun(sK0,sK0),sK2),s('type.list.list'(sK0),sK1)))),
  inference(cnf_transformation,[],[f71])).
% SZS output end Proof for thm.indexedLists
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 57
% Passive clauses: 129
% Generated clauses: 1860
% Final active clauses: 47
% Final passive clauses: 50
% Input formulas: 27
% Initial clauses: 50
% 
% Duplicate literals: 97
% Trivial inequalities: 22
% Fw subsumption resolutions: 78
% Fw demodulations: 85
% Bw demodulations: 6
% Condensations: 173
% 
% Simple tautologies: 56
% Equational tautologies: 76
% Forward subsumptions: 1096
% Fw demodulations to eq. taut.: 12
% 
% Unit resulting resolution: 206
% Factoring: 1
% Forward superposition: 538
% Backward superposition: 502
% Self superposition: 44
% Equality factoring: 28
% Equality resolution: 4
% 
% Split clauses: 19
% Split components: 43
% SAT solver clauses: 186
% SAT solver unit clauses: 17
% SAT solver binary clauses: 59
% 
% Sat splits: 19
% Sat splitting refutations: 17
% 
% Memory used [KB]: 3837
% Time elapsed: 0.069 s
% ------------------------------
% ------------------------------
% Success in time 1.399 s
