Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.rich_list
% (98080)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 76
% Passive clauses: 2450
% Generated clauses: 16289
% Final active clauses: 68
% Final passive clauses: 2257
% Input formulas: 10
% Initial clauses: 16
% 
% Duplicate literals: 1119
% Trivial inequalities: 219
% Fw subsumption resolutions: 1880
% Fw demodulations: 1270
% Bw demodulations: 5
% 
% Simple tautologies: 295
% Equational tautologies: 103
% Forward subsumptions: 7026
% 
% Binary resolution: 54
% Forward superposition: 6329
% Backward superposition: 4567
% Self superposition: 380
% Equality factoring: 291
% Equality resolution: 11
% 
% Split clauses: 833
% Split components: 1678
% SAT solver clauses: 2000
% SAT solver unit clauses: 248
% SAT solver binary clauses: 649
% 
% Sat splits: 833
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 1992
% TWLsolver calls for satisfiability: 71
% 
% Memory used [KB]: 7803
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.rich_list
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.rich_list
% SZS output start Proof for thm.rich_list
fof(f427,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f426,f337,f338,f41,f62,f51,f69,f81,f83,f38,f60,f43,f64,f44,f65,f348,f351,f341,f343,f40,f56,f413,f422,f394,f395,f47,f67,f84,f86,f39,f58])).
fof(f58,plain,(
  p(s(bool,'const.operator.ASSOC_1'(s(fun(sK0,fun(sK0,sK0)),sK1)))) | $spl2),
  inference(cnf_transformation,[],[f58_D])).
fof(f58_D,plain,(
  p(s(bool,'const.operator.ASSOC_1'(s(fun(sK0,fun(sK0,sK0)),sK1)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f39,plain,(
  p(s(bool,'const.operator.ASSOC_1'(s(fun(sK0,fun(sK0,sK0)),sK1))))),
  inference(cnf_transformation,[],[f28])).
fof(f28,plain,(
  p(s(bool,'const.operator.COMM_1'(s(fun(sK0,fun(sK0,sK0)),sK1)))) & p(s(bool,'const.operator.ASSOC_1'(s(fun(sK0,fun(sK0,sK0)),sK1)))) & s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,sK2))),s(sK0,sK3))))),s(sK0,sK4))) != s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,sK2))),s(sK0,sK4))))),s(sK0,sK3)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4])],[f26])).
fof(f26,plain,(
  ? [X0,X1] : (p(s(bool,'const.operator.COMM_1'(s(fun(X0,fun(X0,X0)),X1)))) & p(s(bool,'const.operator.ASSOC_1'(s(fun(X0,fun(X0,X0)),X1)))) & ? [X2,X3,X4] : s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X2))),s(X0,X3))))),s(X0,X4))) != s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X2))),s(X0,X4))))),s(X0,X3))))),
  inference(flattening,[],[f25])).
fof(f25,plain,(
  ? [X0,X1] : (p(s(bool,'const.operator.COMM_1'(s(fun(X0,fun(X0,X0)),X1)))) & (p(s(bool,'const.operator.ASSOC_1'(s(fun(X0,fun(X0,X0)),X1)))) & ? [X2,X3,X4] : s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X2))),s(X0,X3))))),s(X0,X4))) != s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X2))),s(X0,X4))))),s(X0,X3)))))),
  inference(ennf_transformation,[],[f13])).
fof(f13,plain,(
  ~! [X0,X1] : (p(s(bool,'const.operator.COMM_1'(s(fun(X0,fun(X0,X0)),X1)))) => (p(s(bool,'const.operator.ASSOC_1'(s(fun(X0,fun(X0,X0)),X1)))) => ! [X2,X3,X4] : s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X2))),s(X0,X3))))),s(X0,X4))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X2))),s(X0,X4))))),s(X0,X3)))))),
  inference(flattening,[],[f12])).
fof(f12,plain,(
  ~! [X0] : ! [X1] : (p(s(bool,'const.operator.COMM_1'(s(fun(X0,fun(X0,X0)),X1)))) => (p(s(bool,'const.operator.ASSOC_1'(s(fun(X0,fun(X0,X0)),X1)))) => ! [X2,X3,X4] : s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X2))),s(X0,X3))))),s(X0,X4))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X2))),s(X0,X4))))),s(X0,X3)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7] : (p(s(bool,'const.operator.COMM_1'(s(fun(X6,fun(X6,X6)),X7)))) => (p(s(bool,'const.operator.ASSOC_1'(s(fun(X6,fun(X6,X6)),X7)))) => ! [X8,X9,X10] : s(X6,happ(s(fun(X6,X6),happ(s(fun(X6,fun(X6,X6)),X7),s(X6,happ(s(fun(X6,X6),happ(s(fun(X6,fun(X6,X6)),X7),s(X6,X8))),s(X6,X9))))),s(X6,X10))) = s(X6,happ(s(fun(X6,X6),happ(s(fun(X6,fun(X6,X6)),X7),s(X6,happ(s(fun(X6,X6),happ(s(fun(X6,fun(X6,X6)),X7),s(X6,X8))),s(X6,X10))))),s(X6,X9)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7] : (p(s(bool,'const.operator.COMM_1'(s(fun(X6,fun(X6,X6)),X7)))) => (p(s(bool,'const.operator.ASSOC_1'(s(fun(X6,fun(X6,X6)),X7)))) => ! [X8,X9,X10] : s(X6,happ(s(fun(X6,X6),happ(s(fun(X6,fun(X6,X6)),X7),s(X6,happ(s(fun(X6,X6),happ(s(fun(X6,fun(X6,X6)),X7),s(X6,X8))),s(X6,X9))))),s(X6,X10))) = s(X6,happ(s(fun(X6,X6),happ(s(fun(X6,fun(X6,X6)),X7),s(X6,happ(s(fun(X6,X6),happ(s(fun(X6,fun(X6,X6)),X7),s(X6,X8))),s(X6,X10))))),s(X6,X9)))))),
  file('thm.rich_list.rich_list_2|48|2_prove_ltb','thm.rich_list.rich_list_2|48|2_prove_ltb')).
fof(f86,plain,(
  s(bool,'T_0') = s(bool,'const.operator.ASSOC_1'(s(fun(sK0,fun(sK0,sK0)),sK1))) | $spl28),
  inference(cnf_transformation,[],[f86_D])).
fof(f86_D,plain,(
  s(bool,'T_0') = s(bool,'const.operator.ASSOC_1'(s(fun(sK0,fun(sK0,sK0)),sK1))) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f84,plain,(
  s(bool,'T_0') = s(bool,'const.operator.ASSOC_1'(s(fun(sK0,fun(sK0,sK0)),sK1))) | ($spl2 | $spl9 | $spl10)),
  inference(subsumption_resolution,[],[f76,f64])).
fof(f76,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.operator.ASSOC_1'(s(fun(sK0,fun(sK0,sK0)),sK1))) | ($spl2 | $spl10)),
  inference(superposition,[],[f58,f65])).
fof(f67,plain,(
  ( ! [X6,X0,X7,X5,X1] : (~p(s(bool,'const.operator.ASSOC_1'(s(fun(X0,fun(X0,X0)),X1)))) | s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X5))),s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X6))),s(X0,X7))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X5))),s(X0,X6))))),s(X0,X7)))) ) | $spl14),
  inference(cnf_transformation,[],[f67_D])).
fof(f67_D,plain,(
  ( ! [X6,X0,X7,X5,X1] : (~p(s(bool,'const.operator.ASSOC_1'(s(fun(X0,fun(X0,X0)),X1)))) | s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X5))),s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X6))),s(X0,X7))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X5))),s(X0,X6))))),s(X0,X7)))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f47,plain,(
  ( ! [X6,X0,X7,X5,X1] : (s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X5))),s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X6))),s(X0,X7))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X5))),s(X0,X6))))),s(X0,X7))) | ~p(s(bool,'const.operator.ASSOC_1'(s(fun(X0,fun(X0,X0)),X1))))) )),
  inference(cnf_transformation,[],[f32])).
fof(f32,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.operator.ASSOC_1'(s(fun(X0,fun(X0,X0)),X1)))) | ! [X5,X6,X7] : s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X5))),s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X6))),s(X0,X7))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X5))),s(X0,X6))))),s(X0,X7)))) & (s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,sK5(X1,X0)))),s(X0,sK6(X1,X0)))))),s(X0,sK7(X1,X0)))) != s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,sK5(X1,X0)))),s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,sK6(X1,X0)))),s(X0,sK7(X1,X0)))))) | p(s(bool,'const.operator.ASSOC_1'(s(fun(X0,fun(X0,X0)),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5,sK6,sK7])],[f31])).
fof(f31,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.operator.ASSOC_1'(s(fun(X0,fun(X0,X0)),X1)))) | ! [X5,X6,X7] : s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X5))),s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X6))),s(X0,X7))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X5))),s(X0,X6))))),s(X0,X7)))) & (? [X2,X3,X4] : s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X2))),s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X3))),s(X0,X4))))) != s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X2))),s(X0,X3))))),s(X0,X4))) | p(s(bool,'const.operator.ASSOC_1'(s(fun(X0,fun(X0,X0)),X1))))))),
  inference(rectify,[],[f30])).
fof(f30,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.operator.ASSOC_1'(s(fun(X0,fun(X0,X0)),X1)))) | ! [X2,X3,X4] : s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X2))),s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X3))),s(X0,X4))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X2))),s(X0,X3))))),s(X0,X4)))) & (? [X2,X3,X4] : s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X2))),s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X3))),s(X0,X4))))) != s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X2))),s(X0,X3))))),s(X0,X4))) | p(s(bool,'const.operator.ASSOC_1'(s(fun(X0,fun(X0,X0)),X1))))))),
  inference(nnf_transformation,[],[f18])).
fof(f18,plain,(
  ! [X0,X1] : (p(s(bool,'const.operator.ASSOC_1'(s(fun(X0,fun(X0,X0)),X1)))) <=> ! [X2,X3,X4] : s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X2))),s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X3))),s(X0,X4))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X2))),s(X0,X3))))),s(X0,X4))))),
  inference(flattening,[],[f17])).
fof(f17,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.operator.ASSOC_1'(s(fun(X0,fun(X0,X0)),X1)))) <=> ! [X2,X3,X4] : s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X2))),s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X3))),s(X0,X4))))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),X1),s(X0,X2))),s(X0,X3))))),s(X0,X4))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X6] : ! [X7] : (p(s(bool,'const.operator.ASSOC_1'(s(fun(X6,fun(X6,X6)),X7)))) <=> ! [X11,X12,X13] : s(X6,happ(s(fun(X6,X6),happ(s(fun(X6,fun(X6,X6)),X7),s(X6,X11))),s(X6,happ(s(fun(X6,X6),happ(s(fun(X6,fun(X6,X6)),X7),s(X6,X12))),s(X6,X13))))) = s(X6,happ(s(fun(X6,X6),happ(s(fun(X6,fun(X6,X6)),X7),s(X6,happ(s(fun(X6,X6),happ(s(fun(X6,fun(X6,X6)),X7),s(X6,X11))),s(X6,X12))))),s(X6,X13))))),
  file('thm.rich_list.rich_list_2|48|2_prove_ltb','thm.operator.ASSOC_DEF')).
fof(f395,plain,(
  ( ! [X2,X0,X1] : (s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,X0))),s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,X1))),s(sK0,X2))))) = s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,X0))),s(sK0,X1))))),s(sK0,X2)))) ) | $spl68),
  inference(cnf_transformation,[],[f395_D])).
fof(f395_D,plain,(
  ( ! [X2,X0,X1] : (s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,X0))),s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,X1))),s(sK0,X2))))) = s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,X0))),s(sK0,X1))))),s(sK0,X2)))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f394,plain,(
  ( ! [X2,X0,X1] : (s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,X0))),s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,X1))),s(sK0,X2))))) = s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,X0))),s(sK0,X1))))),s(sK0,X2)))) ) | ($spl6 | $spl14 | $spl28)),
  inference(subsumption_resolution,[],[f387,f62])).
fof(f387,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,X0))),s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,X1))),s(sK0,X2))))) = s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,X0))),s(sK0,X1))))),s(sK0,X2)))) ) | ($spl14 | $spl28)),
  inference(superposition,[],[f67,f86])).
fof(f422,plain,(
  ( ! [X4,X5,X3] : (s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,X3))),s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,X4))),s(sK0,X5))))) = s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,X5))),s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,X3))),s(sK0,X4)))))) ) | $spl74),
  inference(cnf_transformation,[],[f422_D])).
fof(f422_D,plain,(
  ( ! [X4,X5,X3] : (s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,X3))),s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,X4))),s(sK0,X5))))) = s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,X5))),s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,X3))),s(sK0,X4)))))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f413,plain,(
  ( ! [X4,X5,X3] : (s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,X3))),s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,X4))),s(sK0,X5))))) = s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,X5))),s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,X3))),s(sK0,X4)))))) ) | ($spl54 | $spl68)),
  inference(superposition,[],[f395,f338])).
fof(f56,plain,(
  s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,sK2))),s(sK0,sK3))))),s(sK0,sK4))) != s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,sK2))),s(sK0,sK4))))),s(sK0,sK3))) | $spl1),
  inference(cnf_transformation,[],[f56_D])).
fof(f56_D,plain,(
  s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,sK2))),s(sK0,sK3))))),s(sK0,sK4))) != s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,sK2))),s(sK0,sK4))))),s(sK0,sK3))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f40,plain,(
  s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,sK2))),s(sK0,sK3))))),s(sK0,sK4))) != s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,sK2))),s(sK0,sK4))))),s(sK0,sK3)))),
  inference(cnf_transformation,[],[f28])).
fof(f343,plain,(
  s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,sK2))),s(sK0,sK3))))),s(sK0,sK4))) != s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,sK3))),s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,sK2))),s(sK0,sK4))))) | $spl57),
  inference(cnf_transformation,[],[f343_D])).
fof(f343_D,plain,(
  s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,sK2))),s(sK0,sK3))))),s(sK0,sK4))) != s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,sK3))),s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,sK2))),s(sK0,sK4))))) <=> ~$spl57),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl57])])).
fof(f341,plain,(
  s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,sK2))),s(sK0,sK3))))),s(sK0,sK4))) != s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,sK3))),s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,sK2))),s(sK0,sK4))))) | ($spl1 | $spl54)),
  inference(backward_demodulation,[],[f338,f56])).
fof(f351,plain,(
  s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,sK3))),s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,sK2))),s(sK0,sK4))))) != s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,sK4))),s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,sK2))),s(sK0,sK3))))) | $spl59),
  inference(cnf_transformation,[],[f351_D])).
fof(f351_D,plain,(
  s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,sK3))),s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,sK2))),s(sK0,sK4))))) != s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,sK4))),s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,sK2))),s(sK0,sK3))))) <=> ~$spl59),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl59])])).
fof(f348,plain,(
  s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,sK3))),s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,sK2))),s(sK0,sK4))))) != s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,sK4))),s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,sK2))),s(sK0,sK3))))) | ($spl54 | $spl57)),
  inference(superposition,[],[f343,f338])).
fof(f65,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl10),
  inference(cnf_transformation,[],[f65_D])).
fof(f65_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f44,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.rich_list.rich_list_2|48|2_prove_ltb',aHL_BOOL_CASES)).
fof(f64,plain,(
  ~p(s(bool,'F_0')) | $spl9),
  inference(cnf_transformation,[],[f64_D])).
fof(f64_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f43,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f14])).
fof(f14,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.rich_list.rich_list_2|48|2_prove_ltb',aHL_FALSITY)).
fof(f60,plain,(
  p(s(bool,'const.operator.COMM_1'(s(fun(sK0,fun(sK0,sK0)),sK1)))) | $spl4),
  inference(cnf_transformation,[],[f60_D])).
fof(f60_D,plain,(
  p(s(bool,'const.operator.COMM_1'(s(fun(sK0,fun(sK0,sK0)),sK1)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f38,plain,(
  p(s(bool,'const.operator.COMM_1'(s(fun(sK0,fun(sK0,sK0)),sK1))))),
  inference(cnf_transformation,[],[f28])).
fof(f83,plain,(
  s(bool,'T_0') = s(bool,'const.operator.COMM_1'(s(fun(sK0,fun(sK0,sK0)),sK1))) | $spl26),
  inference(cnf_transformation,[],[f83_D])).
fof(f83_D,plain,(
  s(bool,'T_0') = s(bool,'const.operator.COMM_1'(s(fun(sK0,fun(sK0,sK0)),sK1))) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f81,plain,(
  s(bool,'T_0') = s(bool,'const.operator.COMM_1'(s(fun(sK0,fun(sK0,sK0)),sK1))) | ($spl4 | $spl9 | $spl10)),
  inference(subsumption_resolution,[],[f75,f64])).
fof(f75,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.operator.COMM_1'(s(fun(sK0,fun(sK0,sK0)),sK1))) | ($spl4 | $spl10)),
  inference(superposition,[],[f60,f65])).
fof(f69,plain,(
  ( ! [X6,X2,X0,X5,X1] : (~p(s(bool,'const.operator.COMM_1'(s(fun(X1,fun(X1,X0)),X2)))) | s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,X5))),s(X1,X6))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,X6))),s(X1,X5)))) ) | $spl18),
  inference(cnf_transformation,[],[f69_D])).
fof(f69_D,plain,(
  ( ! [X6,X2,X0,X5,X1] : (~p(s(bool,'const.operator.COMM_1'(s(fun(X1,fun(X1,X0)),X2)))) | s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,X5))),s(X1,X6))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,X6))),s(X1,X5)))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f51,plain,(
  ( ! [X6,X2,X0,X5,X1] : (s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,X5))),s(X1,X6))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,X6))),s(X1,X5))) | ~p(s(bool,'const.operator.COMM_1'(s(fun(X1,fun(X1,X0)),X2))))) )),
  inference(cnf_transformation,[],[f36])).
fof(f36,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.operator.COMM_1'(s(fun(X1,fun(X1,X0)),X2)))) | ! [X5,X6] : s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,X5))),s(X1,X6))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,X6))),s(X1,X5)))) & (s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,sK8(X2,X1,X0)))),s(X1,sK9(X2,X1,X0)))) != s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,sK9(X2,X1,X0)))),s(X1,sK8(X2,X1,X0)))) | p(s(bool,'const.operator.COMM_1'(s(fun(X1,fun(X1,X0)),X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK8,sK9])],[f35])).
fof(f35,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.operator.COMM_1'(s(fun(X1,fun(X1,X0)),X2)))) | ! [X5,X6] : s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,X5))),s(X1,X6))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,X6))),s(X1,X5)))) & (? [X3,X4] : s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,X3))),s(X1,X4))) != s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,X4))),s(X1,X3))) | p(s(bool,'const.operator.COMM_1'(s(fun(X1,fun(X1,X0)),X2))))))),
  inference(rectify,[],[f34])).
fof(f34,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.operator.COMM_1'(s(fun(X1,fun(X1,X0)),X2)))) | ! [X3,X4] : s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,X3))),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,X4))),s(X1,X3)))) & (? [X3,X4] : s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,X3))),s(X1,X4))) != s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,X4))),s(X1,X3))) | p(s(bool,'const.operator.COMM_1'(s(fun(X1,fun(X1,X0)),X2))))))),
  inference(nnf_transformation,[],[f22])).
fof(f22,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.operator.COMM_1'(s(fun(X1,fun(X1,X0)),X2)))) <=> ! [X3,X4] : s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,X3))),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,X4))),s(X1,X3))))),
  inference(flattening,[],[f21])).
fof(f21,plain,(
  ! [X0,X1] : ! [X2] : (p(s(bool,'const.operator.COMM_1'(s(fun(X1,fun(X1,X0)),X2)))) <=> ! [X3,X4] : s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,X3))),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,X4))),s(X1,X3))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X14,X6] : ! [X7] : (p(s(bool,'const.operator.COMM_1'(s(fun(X6,fun(X6,X14)),X7)))) <=> ! [X11,X12] : s(X14,happ(s(fun(X6,X14),happ(s(fun(X6,fun(X6,X14)),X7),s(X6,X11))),s(X6,X12))) = s(X14,happ(s(fun(X6,X14),happ(s(fun(X6,fun(X6,X14)),X7),s(X6,X12))),s(X6,X11))))),
  file('thm.rich_list.rich_list_2|48|2_prove_ltb','thm.operator.COMM_DEF')).
fof(f62,plain,(
  p(s(bool,'T_0')) | $spl6),
  inference(cnf_transformation,[],[f62_D])).
fof(f62_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f41,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.rich_list.rich_list_2|48|2_prove_ltb',aHL_TRUTH)).
fof(f338,plain,(
  ( ! [X0,X1] : (s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,X0))),s(sK0,X1))) = s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,X1))),s(sK0,X0)))) ) | $spl54),
  inference(cnf_transformation,[],[f338_D])).
fof(f338_D,plain,(
  ( ! [X0,X1] : (s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,X0))),s(sK0,X1))) = s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,X1))),s(sK0,X0)))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f337,plain,(
  ( ! [X0,X1] : (s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,X0))),s(sK0,X1))) = s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,X1))),s(sK0,X0)))) ) | ($spl6 | $spl18 | $spl26)),
  inference(subsumption_resolution,[],[f330,f62])).
fof(f330,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,X0))),s(sK0,X1))) = s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,X1))),s(sK0,X0)))) ) | ($spl18 | $spl26)),
  inference(superposition,[],[f69,f83])).
fof(f426,plain,(
  $false | ($spl54 | $spl59 | $spl74)),
  inference(trivial_inequality_removal,[],[f425])).
fof(f425,plain,(
  s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,sK2))),s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,sK3))),s(sK0,sK4))))) != s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,sK2))),s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,sK3))),s(sK0,sK4))))) | ($spl54 | $spl59 | $spl74)),
  inference(forward_demodulation,[],[f424,f338])).
fof(f424,plain,(
  s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,sK2))),s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,sK3))),s(sK0,sK4))))) != s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,sK2))),s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,sK4))),s(sK0,sK3))))) | ($spl59 | $spl74)),
  inference(forward_demodulation,[],[f423,f422])).
fof(f423,plain,(
  s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,sK2))),s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,sK4))),s(sK0,sK3))))) != s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,sK4))),s(sK0,happ(s(fun(sK0,sK0),happ(s(fun(sK0,fun(sK0,sK0)),sK1),s(sK0,sK2))),s(sK0,sK3))))) | ($spl59 | $spl74)),
  inference(backward_demodulation,[],[f422,f351])).
% SZS output end Proof for thm.rich_list
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 22
% Passive clauses: 42
% Generated clauses: 375
% Final active clauses: 17
% Final passive clauses: 13
% Input formulas: 10
% Initial clauses: 16
% 
% Duplicate literals: 52
% Trivial inequalities: 4
% Fw subsumption resolutions: 15
% Fw demodulations: 4
% Bw demodulations: 11
% Condensations: 46
% 
% Simple tautologies: 18
% Equational tautologies: 15
% Forward subsumptions: 110
% Fw demodulations to eq. taut.: 1
% 
% Binary resolution: 7
% Forward superposition: 65
% Backward superposition: 52
% Self superposition: 30
% Equality factoring: 26
% 
% Split clauses: 1
% Split components: 2
% SAT solver clauses: 76
% SAT solver unit clauses: 20
% SAT solver binary clauses: 11
% SAT solver learnt clauses: 1
% 
% Sat splits: 1
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 75
% TWLsolver calls for satisfiability: 50
% 
% Memory used [KB]: 1151
% Time elapsed: 0.016 s
% ------------------------------
% ------------------------------
% Success in time 0.641 s
