Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.quantHeuristics
% (84537)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 739
% Passive clauses: 8422
% Generated clauses: 19673
% Final active clauses: 733
% Final passive clauses: 7247
% Input formulas: 59
% Initial clauses: 229
% 
% Duplicate literals: 418
% Trivial inequalities: 69
% Fw subsumption resolutions: 1748
% Fw demodulations: 682
% Bw demodulations: 6
% 
% Simple tautologies: 266
% Equational tautologies: 67
% Forward subsumptions: 7703
% 
% Binary resolution: 9827
% Factoring: 1625
% Forward superposition: 2035
% Backward superposition: 2786
% Self superposition: 35
% Equality factoring: 28
% Equality resolution: 22
% 
% Split clauses: 378
% Split components: 758
% SAT solver clauses: 701
% SAT solver unit clauses: 6
% SAT solver binary clauses: 200
% 
% Sat splits: 378
% 
% TWLsolver clauses: 701
% TWLsolver calls for satisfiability: 436
% 
% Memory used [KB]: 9210
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.quantHeuristics
% (84558)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 437
% Passive clauses: 3117
% Generated clauses: 12528
% Final active clauses: 423
% Final passive clauses: 2635
% Input formulas: 59
% Initial clauses: 229
% 
% Duplicate literals: 328
% Trivial inequalities: 18
% Fw subsumption resolutions: 774
% Fw demodulations: 407
% Bw demodulations: 55
% Condensations: 287
% Global subsumptions: 22
% 
% Simple tautologies: 263
% Equational tautologies: 16
% Forward subsumptions: 3578
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 3245
% Factoring: 62
% Forward superposition: 1836
% Backward superposition: 2051
% Self superposition: 32
% Equality factoring: 26
% Equality resolution: 6
% 
% Split clauses: 219
% Split components: 438
% SAT solver clauses: 3941
% SAT solver unit clauses: 232
% SAT solver binary clauses: 98
% SAT solver learnt clauses: 44
% SAT solver learnt literals: 10
% 
% Sat splits: 219
% Sat splitting refutations: 10
% 
% TWLsolver clauses: 3853
% TWLsolver calls for satisfiability: 1975
% 
% Memory used [KB]: 14583
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.quantHeuristics
% (84576)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1163
% Passive clauses: 10461
% Generated clauses: 55498
% Final active clauses: 681
% Final passive clauses: 1497
% Input formulas: 59
% Initial clauses: 229
% Discarded non-redundant clauses: 4518
% 
% Duplicate literals: 1102
% Trivial inequalities: 81
% Fw subsumption resolutions: 2414
% Fw demodulations: 6492
% Bw demodulations: 4867
% Condensations: 1860
% Global subsumptions: 965
% 
% Simple tautologies: 135
% Equational tautologies: 156
% Forward subsumptions: 29073
% Bw demodulations to eq. taut.: 286
% 
% Unit resulting resolution: 9047
% Factoring: 3
% Forward superposition: 17652
% Backward superposition: 9823
% Self superposition: 94
% Equality factoring: 28
% Equality resolution: 1
% 
% Split clauses: 77
% Split components: 159
% SAT solver clauses: 12850
% SAT solver unit clauses: 2762
% SAT solver binary clauses: 5000
% 
% Sat splits: 77
% Sat splitting refutations: 1267
% 
% Memory used [KB]: 205583
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.quantHeuristics
% (84858)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 2364
% Passive clauses: 39076
% Generated clauses: 235469
% Final active clauses: 2015
% Final passive clauses: 28840
% Input formulas: 59
% Initial clauses: 213
% 
% Selected by SInE selection: 51
% SInE iterations: 1
% 
% Duplicate literals: 7229
% Trivial inequalities: 213
% Fw subsumption resolutions: 27049
% Fw demodulations: 4279
% Bw demodulations: 20
% 
% Simple tautologies: 3403
% Equational tautologies: 310
% Forward subsumptions: 146610
% 
% Binary resolution: 185124
% Factoring: 963
% Forward superposition: 6708
% Backward superposition: 2296
% Equality resolution: 20
% 
% Split clauses: 10115
% Split components: 21916
% SAT solver clauses: 10466
% SAT solver unit clauses: 24
% SAT solver binary clauses: 197
% 
% Sat splits: 10115
% Sat splitting refutations: 90
% 
% Memory used [KB]: 67418
% Time elapsed: 4.0000 s
% ------------------------------
% ------------------------------
% remaining time: 142 next slice time: 6
lrs+1002_2:3_br=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:sas=minisat:stl=30:sac=on:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none:sp=reverse_arity:urr=on_4 on thm.quantHeuristics
% (84935)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 336
% Passive clauses: 3691
% Generated clauses: 16736
% Final active clauses: 325
% Final passive clauses: 3246
% Input formulas: 59
% Initial clauses: 229
% 
% Duplicate literals: 530
% Trivial inequalities: 70
% Fw subsumption resolutions: 1383
% Fw demodulations: 235
% Bw demodulations: 61
% Global subsumptions: 7
% 
% Simple tautologies: 408
% Equational tautologies: 64
% Forward subsumptions: 10259
% 
% Unit resulting resolution: 2379
% Factoring: 8
% Forward superposition: 4917
% Backward superposition: 6721
% Self superposition: 114
% Equality factoring: 6
% Equality resolution: 10
% 
% Split clauses: 18
% Split components: 37
% SAT solver clauses: 3782
% SAT solver unit clauses: 501
% SAT solver binary clauses: 311
% 
% Sat splits: 18
% Sat splitting refutations: 64
% 
% Memory used [KB]: 13816
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 135 next slice time: 15
lrs+1011_12_bs=on:bsr=unit_only:cond=on:gs=on:gsaa=from_current:gsssp=full:nwc=1.1:sas=minisat:stl=60:sos=all:sac=on:sdd=off:sser=off:sfr=on:ssfp=100000:ssfq=1.2:ssnc=none:sp=reverse_arity:updr=off_14 on thm.quantHeuristics
% (84951)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 919
% Passive clauses: 7010
% Generated clauses: 46835
% Final active clauses: 890
% Final passive clauses: 5937
% Input formulas: 59
% Initial clauses: 229
% 
% Duplicate literals: 1406
% Trivial inequalities: 145
% Fw subsumption resolutions: 2179
% Bw subsumption resolutions: 15
% Fw demodulations: 5476
% Bw demodulations: 23
% Condensations: 3370
% 
% Simple tautologies: 3177
% Equational tautologies: 11
% Forward subsumptions: 24155
% Backward subsumptions: 299
% 
% Binary resolution: 11184
% Factoring: 2
% Forward superposition: 18231
% Backward superposition: 4837
% Self superposition: 4
% Equality resolution: 7
% 
% SAT solver clauses: 7012
% SAT solver unit clauses: 184
% SAT solver binary clauses: 1559
% 
% Memory used [KB]: 22643
% Time elapsed: 1.600 s
% ------------------------------
% ------------------------------
% remaining time: 119 next slice time: 9
lrs+10_8:1_bd=preordered:bs=on:ccuc=first:fsr=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:nicw=on:sas=minisat:stl=120:sos=on:sscc=on:sser=off:sfr=on:ssfp=4000:ssfq=1.0:ssnc=none:sp=reverse_arity:urr=on_7 on thm.quantHeuristics
% (84999)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1373
% Passive clauses: 4148
% Generated clauses: 36129
% Final active clauses: 748
% Final passive clauses: 2247
% Input formulas: 59
% Initial clauses: 229
% 
% Duplicate literals: 345
% Trivial inequalities: 82
% Fw demodulations: 5804
% Bw demodulations: 365
% Global subsumptions: 32
% 
% Simple tautologies: 768
% Equational tautologies: 38
% Forward subsumptions: 22873
% Backward subsumptions: 86
% 
% Binary resolution: 5853
% Unit resulting resolution: 14227
% Forward superposition: 2599
% Backward superposition: 4905
% Self superposition: 2
% Equality resolution: 4
% 
% Split clauses: 165
% Split components: 347
% SAT solver clauses: 8289
% SAT solver unit clauses: 6
% SAT solver binary clauses: 4341
% 
% Sat splits: 165
% Sat splitting refutations: 26
% 
% Memory used [KB]: 17270
% Time elapsed: 1.0000 s
% ------------------------------
% ------------------------------
% remaining time: 109 next slice time: 85
lrs+1_5:4_cond=on:fsr=off:fde=none:gs=on:gsem=on:lwlo=on:nm=64:nwc=1:stl=60:sos=all:spl=off_80 on thm.quantHeuristics
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.quantHeuristics
% SZS output start Proof for thm.quantHeuristics
fof(f12634,plain,(
  $false),
  inference(global_subsumption,[],[f12633,f888])).
fof(f888,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(forward_demodulation,[],[f817,f815])).
fof(f815,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(sK1,sK2),sK4),s(fun(sK2,bool),sK5)))),
  inference(resolution,[],[f271,f288])).
fof(f288,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f185])).
fof(f185,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X15] : (s(bool,'T_0') = s(bool,X15) <=> p(s(bool,X15)))),
  file('thm.quantHeuristics.GUESS_RULES_DISJ|split|6','thm.bool.EQ_CLAUSES|split|0')).
fof(f271,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(sK1,sK2),sK4),s(fun(sK2,bool),sK5))))),
  inference(cnf_transformation,[],[f182])).
fof(f182,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(sK1,sK2),sK4),s(fun(sK2,bool),sK5)))) & ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(sK1,sK2),sK4),s(fun(sK2,bool),'_dst_x0x1_2'(s(fun(bool,bool),'|bslash||slash|_1'(s(bool,sK3))),s(fun(sK2,bool),sK5))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK1,sK2,sK3,sK4,sK5])],[f146])).
fof(f146,plain,(
  ? [X0,X1,X2,X3,X4] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X3),s(fun(X1,bool),X4)))) & ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X3),s(fun(X1,bool),'_dst_x0x1_2'(s(fun(bool,bool),'|bslash||slash|_1'(s(bool,X2))),s(fun(X1,bool),X4)))))))),
  inference(ennf_transformation,[],[f62])).
fof(f62,plain,(
  ~! [X0,X1,X2,X3,X4] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X3),s(fun(X1,bool),X4)))) => p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X3),s(fun(X1,bool),'_dst_x0x1_2'(s(fun(bool,bool),'|bslash||slash|_1'(s(bool,X2))),s(fun(X1,bool),X4)))))))),
  inference(flattening,[],[f61])).
fof(f61,plain,(
  ~! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X3),s(fun(X1,bool),X4)))) => p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X3),s(fun(X1,bool),'_dst_x0x1_2'(s(fun(bool,bool),'|bslash||slash|_1'(s(bool,X2))),s(fun(X1,bool),X4)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X6,X7),X9),s(fun(X7,bool),X10)))) => p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X6,X7),X9),s(fun(X7,bool),'_dst_x0x1_2'(s(fun(bool,bool),'|bslash||slash|_1'(s(bool,X8))),s(fun(X7,bool),X10)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X6,X7),X9),s(fun(X7,bool),X10)))) => p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X6,X7),X9),s(fun(X7,bool),'_dst_x0x1_2'(s(fun(bool,bool),'|bslash||slash|_1'(s(bool,X8))),s(fun(X7,bool),X10)))))))),
  file('thm.quantHeuristics.GUESS_RULES_DISJ|split|6','thm.quantHeuristics.GUESS_RULES_DISJ|split|6')).
fof(f817,plain,(
  s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(sK1,sK2),sK4),s(fun(sK2,bool),sK5)))),
  inference(resolution,[],[f271,f291])).
fof(f291,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f187])).
fof(f187,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f73])).
fof(f73,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X15] : (s(bool,'F_0') = s(bool,X15) <=> ~p(s(bool,X15)))),
  file('thm.quantHeuristics.GUESS_RULES_DISJ|split|6','thm.bool.EQ_CLAUSES|split|3')).
fof(f12633,plain,(
  s(bool,'T_0') = s(bool,'F_0')),
  inference(forward_demodulation,[],[f12604,f1545])).
fof(f1545,plain,(
  ( ! [X8] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'|bslash||slash|_1'(s(bool,X8))),s(bool,'T_0')))) )),
  inference(resolution,[],[f893,f288])).
fof(f893,plain,(
  ( ! [X4] : (p(s(bool,happ(s(fun(bool,bool),'|bslash||slash|_1'(s(bool,X4))),s(bool,'T_0'))))) )),
  inference(forward_demodulation,[],[f828,f815])).
fof(f828,plain,(
  ( ! [X4] : (p(s(bool,happ(s(fun(bool,bool),'|bslash||slash|_1'(s(bool,X4))),s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(sK1,sK2),sK4),s(fun(sK2,bool),sK5))))))) )),
  inference(resolution,[],[f271,f360])).
fof(f360,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),'|bslash||slash|_1'(s(bool,X0))),s(bool,X1)))) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f227])).
fof(f227,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(bool,bool),'|bslash||slash|_1'(s(bool,X0))),s(bool,X1)))) | p(s(bool,X0)) | p(s(bool,X1))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,happ(s(fun(bool,bool),'|bslash||slash|_1'(s(bool,X0))),s(bool,X1))))))),
  inference(flattening,[],[f226])).
fof(f226,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(bool,bool),'|bslash||slash|_1'(s(bool,X0))),s(bool,X1)))) | (p(s(bool,X0)) | p(s(bool,X1)))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,happ(s(fun(bool,bool),'|bslash||slash|_1'(s(bool,X0))),s(bool,X1))))))),
  inference(nnf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),'|bslash||slash|_1'(s(bool,X0))),s(bool,X1)))) <=> (p(s(bool,X0)) | p(s(bool,X1))))),
  inference(rectify,[],[f59])).
fof(f59,axiom,(
  ! [X26,X27] : (p(s(bool,happ(s(fun(bool,bool),'|bslash||slash|_1'(s(bool,X26))),s(bool,X27)))) <=> (p(s(bool,X26)) | p(s(bool,X27))))),
  file('thm.quantHeuristics.GUESS_RULES_DISJ|split|6','|bslash||slash|_HL_1')).
fof(f12604,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'|bslash||slash|_1'(s(bool,'F_0'))),s(bool,'T_0')))),
  inference(backward_demodulation,[],[f12500,f6774])).
fof(f6774,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),'|bslash||slash|_1'(s(bool,'F_0'))),s(bool,happ(s(fun(sK2,bool),sK5),s(sK2,sK15('_dst_x0x1_2'(s(fun(bool,bool),'|bslash||slash|_1'(s(bool,'F_0'))),s(fun(sK2,bool),sK5)),sK4,sK1,sK2))))))),
  inference(resolution,[],[f3752,f292])).
fof(f292,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f187])).
fof(f3752,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),'|bslash||slash|_1'(s(bool,'F_0'))),s(bool,happ(s(fun(sK2,bool),sK5),s(sK2,sK15('_dst_x0x1_2'(s(fun(bool,bool),'|bslash||slash|_1'(s(bool,'F_0'))),s(fun(sK2,bool),sK5)),sK4,sK1,sK2)))))))),
  inference(backward_demodulation,[],[f3730,f1064])).
fof(f1064,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),'|bslash||slash|_1'(s(bool,sK3))),s(bool,happ(s(fun(sK2,bool),sK5),s(sK2,sK15('_dst_x0x1_2'(s(fun(bool,bool),'|bslash||slash|_1'(s(bool,sK3))),s(fun(sK2,bool),sK5)),sK4,sK1,sK2)))))))),
  inference(forward_demodulation,[],[f943,f620])).
fof(f620,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))) )),
  inference(cnf_transformation,[],[f144])).
fof(f144,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(flattening,[],[f143])).
fof(f143,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(rectify,[],[f60])).
fof(f60,axiom,(
  ! [X28,X29,X30] : ! [X31,X32,X33] : s(X28,happ(s(fun(X29,X28),X31),s(X29,happ(s(fun(X30,X29),X32),s(X30,X33))))) = s(X28,happ(s(fun(X30,X28),'_dst_x0x1_2'(s(fun(X29,X28),X31),s(fun(X30,X29),X32))),s(X30,X33)))),
  file('thm.quantHeuristics.GUESS_RULES_DISJ|split|6','_dst_x0x1_ax')).
fof(f943,plain,(
  ~p(s(bool,happ(s(fun(sK2,bool),'_dst_x0x1_2'(s(fun(bool,bool),'|bslash||slash|_1'(s(bool,sK3))),s(fun(sK2,bool),sK5))),s(sK2,sK15('_dst_x0x1_2'(s(fun(bool,bool),'|bslash||slash|_1'(s(bool,sK3))),s(fun(sK2,bool),sK5)),sK4,sK1,sK2)))))),
  inference(resolution,[],[f272,f409])).
fof(f409,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,sK15(X3,X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f253])).
fof(f253,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X6] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK14(X6,X3,X2,X1,X0))))))))) & ((~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,sK15(X3,X2,X1,X0))))) & ! [X5] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK14,sK15])],[f252])).
fof(f252,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X6] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | ? [X7] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X7)))))))) & (? [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) & ! [X5] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f251])).
fof(f251,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) | ? [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))))) & (? [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) & ! [X5] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(nnf_transformation,[],[f174])).
fof(f174,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) | ? [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))))),
  inference(ennf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) => ? [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))))),
  inference(flattening,[],[f130])).
fof(f130,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) => ? [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X7,X6] : ! [X9,X11] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X6,X7),X9),s(fun(X7,bool),X11)))) <=> ! [X20] : (~p(s(bool,happ(s(fun(X7,bool),X11),s(X7,X20)))) => ? [X21] : ~p(s(bool,happ(s(fun(X7,bool),X11),s(X7,happ(s(fun(X6,X7),X9),s(X6,X21))))))))),
  file('thm.quantHeuristics.GUESS_RULES_DISJ|split|6','thm.quantHeuristics.GUESS_EXISTS_FORALL_REWRITES|split|1')).
fof(f272,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(sK1,sK2),sK4),s(fun(sK2,bool),'_dst_x0x1_2'(s(fun(bool,bool),'|bslash||slash|_1'(s(bool,sK3))),s(fun(sK2,bool),sK5))))))),
  inference(cnf_transformation,[],[f182])).
fof(f3730,plain,(
  s(bool,'F_0') = s(bool,sK3)),
  inference(trivial_inequality_removal,[],[f3706])).
fof(f3706,plain,(
  s(bool,'T_0') != s(bool,'T_0') | s(bool,'F_0') = s(bool,sK3)),
  inference(superposition,[],[f3419,f280])).
fof(f280,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.quantHeuristics.GUESS_RULES_DISJ|split|6',aHL_BOOL_CASES)).
fof(f3419,plain,(
  s(bool,'T_0') != s(bool,sK3)),
  inference(resolution,[],[f3197,f273])).
fof(f273,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f9])).
fof(f9,axiom,(
  p(s(bool,'T_0'))),
  file('thm.quantHeuristics.GUESS_RULES_DISJ|split|6','thm.bool.TRUTH')).
fof(f3197,plain,(
  ~p(s(bool,'T_0')) | s(bool,'T_0') != s(bool,sK3)),
  inference(resolution,[],[f3193,f287])).
fof(f287,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f185])).
fof(f3193,plain,(
  ~p(s(bool,sK3)) | ~p(s(bool,'T_0'))),
  inference(forward_demodulation,[],[f3192,f1221])).
fof(f1221,plain,(
  ( ! [X20] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'|bslash||slash|_1'(s(bool,'T_0'))),s(bool,X20)))) )),
  inference(resolution,[],[f892,f288])).
fof(f892,plain,(
  ( ! [X3] : (p(s(bool,happ(s(fun(bool,bool),'|bslash||slash|_1'(s(bool,'T_0'))),s(bool,X3))))) )),
  inference(forward_demodulation,[],[f827,f815])).
fof(f827,plain,(
  ( ! [X3] : (p(s(bool,happ(s(fun(bool,bool),'|bslash||slash|_1'(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(sK1,sK2),sK4),s(fun(sK2,bool),sK5))))),s(bool,X3))))) )),
  inference(resolution,[],[f271,f359])).
fof(f359,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),'|bslash||slash|_1'(s(bool,X0))),s(bool,X1)))) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f227])).
fof(f3192,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),'|bslash||slash|_1'(s(bool,'T_0'))),s(bool,happ(s(fun(sK2,bool),sK5),s(sK2,sK15('_dst_x0x1_2'(s(fun(bool,bool),'|bslash||slash|_1'(s(bool,'T_0'))),s(fun(sK2,bool),sK5)),sK4,sK1,sK2))))))) | ~p(s(bool,sK3))),
  inference(forward_demodulation,[],[f3071,f620])).
fof(f3071,plain,(
  ~p(s(bool,sK3)) | ~p(s(bool,happ(s(fun(sK2,bool),'_dst_x0x1_2'(s(fun(bool,bool),'|bslash||slash|_1'(s(bool,'T_0'))),s(fun(sK2,bool),sK5))),s(sK2,sK15('_dst_x0x1_2'(s(fun(bool,bool),'|bslash||slash|_1'(s(bool,'T_0'))),s(fun(sK2,bool),sK5)),sK4,sK1,sK2)))))),
  inference(resolution,[],[f1006,f409])).
fof(f1006,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(sK1,sK2),sK4),s(fun(sK2,bool),'_dst_x0x1_2'(s(fun(bool,bool),'|bslash||slash|_1'(s(bool,'T_0'))),s(fun(sK2,bool),sK5)))))) | ~p(s(bool,sK3))),
  inference(superposition,[],[f272,f288])).
fof(f12500,plain,(
  ( ! [X19] : (s(bool,'T_0') = s(bool,happ(s(fun(sK2,bool),sK5),s(sK2,X19)))) )),
  inference(resolution,[],[f12355,f288])).
fof(f12355,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK2,bool),sK5),s(sK2,X0))))) )),
  inference(resolution,[],[f3880,f814])).
fof(f814,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK2,bool),sK5),s(sK2,happ(s(fun(sK1,sK2),sK4),s(sK1,sK14(X0,sK5,sK4,sK1,sK2))))))) | p(s(bool,happ(s(fun(sK2,bool),sK5),s(sK2,X0))))) )),
  inference(resolution,[],[f271,f408])).
fof(f408,plain,(
  ( ! [X6,X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK14(X6,X3,X2,X1,X0))))))) | p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f253])).
fof(f3880,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK2,bool),sK5),s(sK2,happ(s(fun(sK1,sK2),sK4),s(sK1,X0))))))) )),
  inference(condensation,[],[f3879])).
fof(f3879,plain,(
  ( ! [X6,X7] : (p(s(bool,X7)) | p(s(bool,happ(s(fun(sK2,bool),sK5),s(sK2,happ(s(fun(sK1,sK2),sK4),s(sK1,X6))))))) )),
  inference(global_subsumption,[],[f3865,f1069])).
fof(f1069,plain,(
  ( ! [X3] : (~p(s(bool,happ(s(fun(bool,bool),'|bslash||slash|_1'(s(bool,'F_0'))),s(bool,X3)))) | p(s(bool,X3))) )),
  inference(forward_demodulation,[],[f957,f948])).
fof(f948,plain,(
  s(bool,'F_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(sK1,sK2),sK4),s(fun(sK2,bool),'_dst_x0x1_2'(s(fun(bool,bool),'|bslash||slash|_1'(s(bool,sK3))),s(fun(sK2,bool),sK5)))))),
  inference(resolution,[],[f272,f292])).
fof(f957,plain,(
  ( ! [X3] : (p(s(bool,X3)) | ~p(s(bool,happ(s(fun(bool,bool),'|bslash||slash|_1'(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(sK1,sK2),sK4),s(fun(sK2,bool),'_dst_x0x1_2'(s(fun(bool,bool),'|bslash||slash|_1'(s(bool,sK3))),s(fun(sK2,bool),sK5))))))),s(bool,X3))))) )),
  inference(resolution,[],[f272,f358])).
fof(f358,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,happ(s(fun(bool,bool),'|bslash||slash|_1'(s(bool,X0))),s(bool,X1))))) )),
  inference(cnf_transformation,[],[f227])).
fof(f3865,plain,(
  ( ! [X6,X7] : (p(s(bool,happ(s(fun(bool,bool),'|bslash||slash|_1'(s(bool,'F_0'))),s(bool,X7)))) | p(s(bool,X7)) | p(s(bool,happ(s(fun(sK2,bool),sK5),s(sK2,happ(s(fun(sK1,sK2),sK4),s(sK1,X6))))))) )),
  inference(backward_demodulation,[],[f3730,f3610])).
fof(f3610,plain,(
  ( ! [X6,X7] : (p(s(bool,happ(s(fun(bool,bool),'|bslash||slash|_1'(s(bool,sK3))),s(bool,X7)))) | p(s(bool,X7)) | p(s(bool,happ(s(fun(sK2,bool),sK5),s(sK2,happ(s(fun(sK1,sK2),sK4),s(sK1,X6))))))) )),
  inference(superposition,[],[f1065,f352])).
fof(f352,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f159])).
fof(f159,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f158])).
fof(f158,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f105])).
fof(f105,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X13,X14] : ((p(s(bool,X13)) => p(s(bool,X14))) => ((p(s(bool,X14)) => p(s(bool,X13))) => s(bool,X13) = s(bool,X14)))),
  file('thm.quantHeuristics.GUESS_RULES_DISJ|split|6','thm.bool.IMP_ANTISYM_AX')).
fof(f1065,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),'|bslash||slash|_1'(s(bool,sK3))),s(bool,happ(s(fun(sK2,bool),sK5),s(sK2,happ(s(fun(sK1,sK2),sK4),s(sK1,X0))))))))) )),
  inference(forward_demodulation,[],[f944,f620])).
fof(f944,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK2,bool),'_dst_x0x1_2'(s(fun(bool,bool),'|bslash||slash|_1'(s(bool,sK3))),s(fun(sK2,bool),sK5))),s(sK2,happ(s(fun(sK1,sK2),sK4),s(sK1,X0))))))) )),
  inference(resolution,[],[f272,f410])).
fof(f410,plain,(
  ( ! [X2,X0,X5,X3,X1] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) )),
  inference(cnf_transformation,[],[f253])).
% SZS output end Proof for thm.quantHeuristics
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 185
% Passive clauses: 1163
% Generated clauses: 11794
% Final active clauses: 157
% Final passive clauses: 857
% Input formulas: 59
% Initial clauses: 358
% 
% Duplicate literals: 655
% Trivial inequalities: 84
% Fw demodulations: 1622
% Bw demodulations: 230
% Condensations: 741
% Global subsumptions: 3
% 
% Simple tautologies: 797
% Forward subsumptions: 7155
% Bw demodulations to eq. taut.: 6
% 
% Binary resolution: 3011
% Forward superposition: 3966
% Backward superposition: 1715
% Self superposition: 3
% Equality resolution: 4
% 
% SAT solver clauses: 1175
% SAT solver unit clauses: 120
% SAT solver binary clauses: 335
% SAT solver learnt clauses: 5
% SAT solver learnt literals: 2
% 
% TWLsolver clauses: 1166
% TWLsolver calls for satisfiability: 16031
% 
% Memory used [KB]: 6908
% Time elapsed: 0.260 s
% ------------------------------
% ------------------------------
% Success in time 19.325 s
