Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.integer
% (29082)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 738
% Passive clauses: 5108
% Generated clauses: 16705
% Final active clauses: 725
% Final passive clauses: 4118
% Input formulas: 261
% Initial clauses: 507
% 
% Duplicate literals: 148
% Trivial inequalities: 19
% Fw subsumption resolutions: 1236
% Fw demodulations: 2121
% Bw demodulations: 12
% 
% Simple tautologies: 230
% Equational tautologies: 25
% Forward subsumptions: 7760
% 
% Binary resolution: 5085
% Factoring: 484
% Forward superposition: 3937
% Backward superposition: 3043
% Self superposition: 24
% Equality factoring: 2
% Equality resolution: 9
% 
% Split clauses: 193
% Split components: 388
% SAT solver clauses: 314
% SAT solver unit clauses: 17
% SAT solver binary clauses: 100
% 
% Sat splits: 193
% 
% TWLsolver clauses: 314
% TWLsolver calls for satisfiability: 247
% 
% Memory used [KB]: 7675
% Time elapsed: 0.621 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.integer
% (29101)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 414
% Passive clauses: 1703
% Generated clauses: 9618
% Final active clauses: 385
% Final passive clauses: 1201
% Input formulas: 261
% Initial clauses: 507
% 
% Duplicate literals: 78
% Trivial inequalities: 25
% Fw subsumption resolutions: 420
% Fw demodulations: 1315
% Bw demodulations: 115
% Condensations: 54
% Global subsumptions: 25
% 
% Simple tautologies: 208
% Equational tautologies: 6
% Forward subsumptions: 3726
% 
% Binary resolution: 1642
% Factoring: 86
% Forward superposition: 2244
% Backward superposition: 1303
% Self superposition: 31
% Equality factoring: 2
% Equality resolution: 6
% 
% Split clauses: 27
% Split components: 54
% SAT solver clauses: 2559
% SAT solver unit clauses: 619
% SAT solver binary clauses: 184
% SAT solver learnt clauses: 50
% SAT solver learnt literals: 14
% 
% Sat splits: 27
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 2459
% TWLsolver calls for satisfiability: 3457
% 
% Memory used [KB]: 9850
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.integer
% (29122)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 2344
% Passive clauses: 18020
% Generated clauses: 99877
% Final active clauses: 2161
% Final passive clauses: 4825
% Input formulas: 261
% Initial clauses: 507
% Discarded non-redundant clauses: 33781
% 
% Duplicate literals: 8
% Trivial inequalities: 44
% Fw subsumption resolutions: 883
% Fw demodulations: 16878
% Bw demodulations: 720
% Condensations: 40
% 
% Simple tautologies: 88
% Equational tautologies: 108
% Forward subsumptions: 46213
% Fw demodulations to eq. taut.: 191
% 
% Unit resulting resolution: 54268
% Factoring: 3
% Forward superposition: 15055
% Backward superposition: 10914
% Self superposition: 211
% Equality factoring: 2
% Equality resolution: 11
% 
% Split clauses: 43
% Split components: 92
% SAT solver clauses: 18877
% SAT solver unit clauses: 4996
% SAT solver binary clauses: 10104
% 
% Sat splits: 43
% Sat splitting refutations: 769
% 
% Memory used [KB]: 96331
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.integer
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.integer
% SZS output start Proof for thm.integer
fof(f119341,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f668,f670,f671,f672,f675,f741,f786,f787,f791,f793,f794,f795,f798,f864,f998,f1066,f1067,f1068,f1069,f1070,f1071,f1072,f1073,f1074,f1075,f1401,f1405,f1403,f1459,f1461,f1462,f1463,f1466,f1532,f1576,f1577,f1581,f1582,f1583,f1585,f1586,f1588,f1590,f1659,f2016,f2020,f2018,f2084,f2086,f2087,f2088,f2091,f2157,f2208,f2210,f2211,f2213,f2215,f2284,f2338,f2340,f2341,f2343,f2345,f2414,f2459,f2461,f2462,f2463,f2466,f2532,f2582,f2584,f2585,f2587,f2589,f2658,f2915,f2917,f2918,f2920,f2922,f2991,f3256,f3258,f3259,f3260,f3263,f3329,f2914,f2913,f2909,f2908,f3371,f2080,f2079,f1455,f1454,f3558,f3747,f3749,f3750,f3751,f3754,f3820,f3871,f3873,f3874,f3876,f3878,f3947,f4000,f4002,f4003,f4005,f4007,f4076,f4340,f4342,f4343,f4344,f4347,f4413,f4458,f4460,f4461,f4462,f4465,f4531,f4789,f4791,f4792,f4794,f4796,f4865,f4454,f4453,f3870,f3869,f3865,f3864,f2207,f2206,f1703,f2202,f2201,f1701,f1699,f5101,f5102,f5103,f5277,f5278,f5282,f5284,f5285,f5286,f5289,f5355,f5424,f5426,f5427,f5429,f5431,f5500,f5553,f5555,f5556,f5558,f5560,f5629,f5893,f5895,f5896,f5897,f5900,f5966,f6011,f6013,f6014,f6015,f6018,f6084,f6341,f6343,f6344,f6346,f6348,f6417,f6470,f6472,f6473,f6474,f6477,f6543,f6598,f6600,f6601,f6602,f6605,f6671,f6852,f6851,f6845,f6844,f6843,f6840,f6913,f6914,f6915,f6916,f6917,f6918,f6919,f6920,f6921,f6922,f6933,f6934,f6938,f6939,f6940,f6942,f6943,f6945,f6947,f7016,f7086,f7088,f7089,f7090,f7093,f7159,f7213,f7215,f7216,f7217,f7220,f7286,f7476,f7475,f7474,f7473,f7461,f7460,f7459,f7458,f7457,f7542,f7543,f7544,f7545,f7546,f7549,f7550,f7551,f7552,f7553,f7554,f7555,f8539,f8541,f8542,f8544,f8546,f8615,f9698,f9700,f9701,f9703,f9705,f9774,f10249,f10251,f10252,f10253,f10256,f10322,f10398,f10400,f10401,f10402,f10405,f10471,f10526,f10528,f10529,f10530,f10533,f10599,f10651,f10653,f10654,f10656,f10658,f10727,f10771,f10773,f10774,f10775,f10778,f10844,f10893,f10895,f10896,f10898,f10900,f10969,f11225,f11227,f11228,f11230,f11232,f11301,f11565,f11567,f11568,f11569,f11572,f11638,f11689,f11691,f11692,f11694,f11696,f11765,f12814,f12816,f12817,f12819,f12821,f12890,f13201,f13203,f13204,f13205,f13208,f13274,f13330,f13332,f13333,f13334,f13337,f13403,f13454,f13456,f13457,f13459,f13461,f13530,f13574,f13576,f13577,f13578,f13581,f13647,f13698,f13700,f13701,f13703,f13705,f13774,f13822,f13824,f13825,f13826,f13829,f13895,f14152,f14154,f14155,f14157,f14159,f14228,f13818,f13817,f13697,f13696,f13692,f13691,f14713,f14715,f14716,f14717,f14720,f14786,f14843,f14845,f14846,f14847,f14850,f14916,f15001,f15003,f15004,f15005,f15008,f15074,f15119,f15121,f15122,f15123,f15126,f15192,f15244,f15246,f15247,f15249,f15251,f15320,f15364,f15366,f15367,f15368,f15371,f15437,f15486,f15488,f15489,f15491,f15493,f15562,f15818,f15820,f15821,f15823,f15825,f15894,f16158,f16160,f16161,f16162,f16165,f16231,f16285,f16287,f16288,f16289,f16292,f16358,f17404,f17406,f17407,f17409,f17411,f17480,f17797,f17799,f17800,f17802,f17804,f17873,f17917,f17919,f17920,f17921,f17924,f17990,f18039,f18041,f18042,f18044,f18046,f18115,f18590,f18592,f18593,f18594,f18597,f18663,f18720,f18722,f18723,f18724,f18727,f18793,f18844,f18846,f18847,f18849,f18851,f18920,f19004,f19006,f19007,f19008,f19011,f19077,f19122,f19124,f19125,f19126,f19129,f19195,f19247,f19249,f19250,f19252,f19254,f19323,f19367,f19369,f19370,f19371,f19374,f19440,f19489,f19491,f19492,f19494,f19496,f19565,f19821,f19823,f19824,f19826,f19828,f19897,f20161,f20163,f20164,f20165,f20168,f20234,f20285,f20287,f20288,f20290,f20292,f20361,f21440,f21442,f21443,f21445,f21447,f21516,f21797,f21799,f21800,f21801,f21804,f21870,f22127,f22129,f22130,f22132,f22134,f22203,f22258,f22260,f22261,f22262,f22265,f22331,f21793,f21792,f20284,f20283,f7060,f22641,f22820,f22822,f22823,f22824,f22827,f22893,f22950,f22952,f22953,f22954,f22957,f23023,f23074,f23076,f23077,f23079,f23081,f23150,f23194,f23196,f23197,f23198,f23201,f23267,f20279,f20278,f7058,f7056,f23351,f23353,f23354,f23355,f23358,f23424,f23469,f23471,f23472,f23473,f23476,f23542,f23594,f23596,f23597,f23599,f23601,f23670,f23714,f23716,f23717,f23718,f23721,f23787,f23836,f23838,f23839,f23841,f23843,f23912,f24168,f24170,f24171,f24173,f24175,f24244,f24508,f24510,f24511,f24512,f24515,f24581,f24632,f24634,f24635,f24637,f24639,f24708,f25787,f25789,f25790,f25792,f25794,f25863,f26144,f26146,f26147,f26148,f26151,f26217,f26474,f26476,f26477,f26479,f26481,f26550,f26605,f26607,f26608,f26609,f26612,f26678,f26729,f26731,f26732,f26734,f26736,f26805,f26849,f26851,f26852,f26853,f26856,f26922,f26971,f26973,f26974,f26976,f26978,f27047,f27091,f27092,f27214,f27096,f27097,f27218,f27216,f27098,f27100,f27101,f27103,f27105,f27174,f27244,f27246,f27247,f27248,f27251,f27317,f27368,f27370,f27371,f27373,f27375,f27444,f27488,f27490,f27491,f27492,f27495,f27561,f27610,f27612,f27613,f27615,f27617,f27686,f33931,f33933,f33934,f33936,f33938,f34007,f34956,f34958,f34959,f34961,f34963,f35032,f36421,f36423,f36424,f36425,f36428,f36494,f37520,f37523,f37524,f37525,f37522,f45038,f45040,f45041,f45042,f45045,f45111,f45315,f45317,f45318,f45320,f45322,f45391,f45596,f45598,f45599,f45600,f45603,f45669,f45876,f45878,f45879,f45880,f45883,f45949,f46154,f46156,f46157,f46159,f46161,f46230,f46435,f46437,f46438,f46440,f46442,f46511,f46729,f46731,f46732,f46733,f46736,f46802,f47009,f47011,f47012,f47014,f47016,f47085,f47221,f47224,f47225,f47226,f47375,f47500,f47499,f47498,f47497,f47496,f47495,f47483,f47482,f47481,f47480,f47479,f47566,f47567,f47568,f47569,f47570,f47573,f47574,f47575,f47576,f47577,f47578,f47579,f47589,f47598,f47605,f47617,f47658,f47671,f47739,f47740,f47743,f47744,f47745,f47746,f47747,f47748,f47749,f47750,f47751,f47752,f47753,f47754,f47755,f47756,f47757,f47758,f47775,f47776,f47777,f47778,f47285,f47284,f47283,f47282,f47781,f47782,f47783,f47784,f47785,f47786,f47787,f47788,f47789,f47790,f47791,f47792,f47793,f47794,f47795,f47796,f48432,f48434,f48435,f48436,f48439,f48505,f48567,f48569,f48570,f48572,f48574,f48643,f48732,f48734,f48735,f48736,f48739,f48805,f48857,f48859,f48860,f48862,f48864,f48933,f48988,f48990,f48991,f48993,f48995,f49064,f49109,f49111,f49112,f49113,f49116,f49182,f49233,f49235,f49236,f49238,f49240,f49309,f49569,f49571,f49572,f49573,f49576,f49642,f49908,f49910,f49911,f49913,f49915,f49984,f50038,f50040,f50041,f50042,f50045,f50111,f50197,f50199,f50200,f50201,f50204,f50270,f50323,f50325,f50326,f50328,f50330,f50399,f51967,f51969,f51970,f51972,f51974,f52043,f52088,f52090,f52091,f52092,f52095,f52161,f52212,f52214,f52215,f52217,f52219,f52288,f53025,f53027,f53028,f53029,f53032,f53098,f53156,f53158,f53159,f53160,f53163,f53229,f53282,f53284,f53285,f53287,f53289,f53358,f53403,f53405,f53406,f53407,f53410,f53476,f53561,f53563,f53564,f53566,f53568,f53637,f53963,f53975,f53977,f53978,f53980,f53982,f54051,f27219,f54100,f54102,f54103,f54104,f54107,f54173,f53141,f54547,f54549,f54550,f54552,f54554,f54623,f54663,f54678,f54680,f54681,f54682,f54685,f54751,f54804,f54806,f54807,f54809,f54811,f54880,f18959,f54925,f54927,f54928,f54929,f54932,f54998,f23306,f55082,f55084,f55085,f55087,f55089,f55158,f56862,f56864,f56865,f56867,f56869,f56938,f57137,f57139,f57140,f57141,f57144,f57210,f68285,f68287,f68288,f68289,f68292,f68358,f68565,f68567,f68568,f68570,f68572,f68641,f68850,f68852,f68853,f68854,f68857,f68923,f69131,f69133,f69134,f69136,f69138,f69207,f69428,f69430,f69431,f69432,f69435,f69501,f69710,f69712,f69713,f69715,f69717,f69786,f70013,f70027,f70036,f70043,f70055,f70064,f70076,f70106,f70107,f70108,f70109,f70110,f70111,f70112,f70113,f70114,f70115,f70116,f70117,f70118,f70119,f70120,f70121,f70122,f70123,f70154,f70156,f70157,f70158,f70161,f70227,f70887,f70889,f70890,f70891,f70894,f70960,f71012,f71014,f71015,f71017,f71019,f71088,f71176,f71178,f71179,f71180,f71183,f71249,f781,f71296,f71298,f71299,f71300,f71303,f71369,f71423,f71425,f71426,f71428,f71430,f71499,f2453,f71545,f71547,f71548,f71549,f71552,f71618,f2571,f71669,f71671,f71672,f71674,f71676,f71745,f2697,f72011,f72013,f72014,f72016,f72018,f72087,f72361,f72363,f72364,f72365,f72368,f72434,f72487,f72489,f72490,f72492,f72494,f72563,f73740,f73742,f73743,f73745,f73747,f73816,f74235,f74237,f74238,f74239,f74242,f74308,f74580,f74582,f74583,f74585,f74587,f74656,f74713,f74715,f74716,f74717,f74720,f74786,f75612,f75614,f75615,f75617,f75619,f75688,f77331,f77333,f77334,f77336,f77338,f77407,f77467,f77469,f77470,f77471,f77474,f77540,f14956,f77582,f77583,f77587,f77589,f77590,f77591,f77594,f77660,f10363,f77712,f77714,f77715,f77717,f77719,f77788,f75611,f75610,f75606,f75605,f77450,f74110,f74109,f47223,f74229,f78192,f78193,f78194,f78197,f78198,f78225,f78226,f78227,f78432,f78433,f78434,f78435,f27213,f69924,f69923,f69922,f69921,f47780,f78451,f78452,f78454,f78459,f78460,f904,f906,f78464,f78466,f78467,f78468,f78471,f78537,f907,f78616,f78627,f78629,f78630,f78632,f78634,f78703,f1704,f78766,f78768,f78769,f78771,f78773,f78842,f3034,f79106,f79108,f79109,f79110,f79113,f79179,f3369,f79224,f79226,f79227,f79228,f79231,f79297,f4572,f79553,f79555,f79556,f79558,f79560,f79629,f4904,f79682,f79684,f79685,f79686,f79689,f79755,f6583,f79809,f79811,f79812,f79813,f79816,f79882,f79936,f79938,f79939,f79940,f79943,f80009,f7326,f80318,f80320,f80321,f80323,f80325,f80394,f8654,f81567,f81569,f81570,f81572,f81574,f81643,f81687,f81689,f81690,f81691,f81694,f81760,f81809,f81811,f81812,f81814,f81816,f81885,f82384,f82386,f82387,f82389,f82391,f82460,f27086,f83204,f83206,f83207,f83209,f83211,f83280,f85308,f85310,f85311,f85313,f85315,f85384,f77448,f88569,f88571,f88572,f88574,f88576,f88645,f34078,f89239,f89241,f89242,f89244,f89246,f89315,f35103,f90030,f90032,f90033,f90034,f90037,f90103,f36568,f90490,f90492,f90494,f90496,f90497,f104098,f104100,f104101,f104102,f104105,f104171,f45185,f104379,f104381,f104382,f104384,f104386,f104455,f45464,f104665,f104667,f104668,f104669,f104672,f104738,f45744,f104950,f104952,f104953,f104954,f104957,f105023,f46024,f105232,f105234,f105235,f105237,f105239,f105308,f46305,f105517,f105519,f105520,f105522,f105524,f105593,f46586,f105816,f105818,f105819,f105820,f105823,f105889,f46879,f106100,f106102,f106103,f106105,f106107,f106176,f47161,f106347,f106361,f106370,f106377,f106389,f106398,f106410,f106441,f106442,f106466,f106445,f106446,f47762,f106447,f106448,f47764,f106449,f106450,f70125,f106451,f106452,f47766,f106453,f106454,f106468,f106455,f106456,f106457,f106458,f47772,f106459,f106460,f106461,f106462,f70127,f106463,f106464,f47774,f106502,f106516,f106525,f106532,f106544,f106553,f106565,f106597,f106598,f106622,f106599,f106600,f47381,f106601,f106602,f106624,f106603,f106604,f106626,f106605,f106606,f106628,f106607,f106608,f106630,f106609,f106610,f106632,f106611,f106612,f106613,f106614,f106636,f106617,f106618,f106638,f106619,f106620,f106640,f106652,f106654,f106655,f106657,f106659,f106728,f7061,f106787,f106789,f106790,f106792,f106794,f106863,f77829,f106912,f106914,f106915,f106916,f106919,f106985,f13315,f107281,f107283,f107284,f107286,f107288,f107357,f14267,f107417,f107419,f107420,f107421,f107424,f107490,f22371,f107546,f107548,f107549,f107551,f107553,f107622,f48687,f107737,f107739,f107740,f107741,f107744,f107810,f107856,f107862,f107864,f107865,f107866,f107869,f107935,f107980,f108929,f112837,f112838,f112937,f47768,f112938,f112880,f114037,f114038,f114140,f47760,f114141,f114082,f115283,f115284,f115391,f115330,f115392,f115332,f116114,f116115,f116222,f116161,f116223,f116163,f116907,f116909,f116911,f116913,f116914,f116971,f116973,f116975,f116977,f116978,f117233,f117235,f117236,f117237,f117240,f117306,f117386,f117439,f117505,f117516,f117528,f117536,f117546,f117556,f117565,f117574,f117510,f117609,f117610,f117623,f117624,f117632,f117633,f117639,f117640,f117651,f117652,f117660,f117661,f117672,f117673,f117704,f117705,f117729,f117706,f117707,f47770,f117708,f117709,f117731,f117710,f117711,f117733,f117712,f117713,f117735,f117714,f117715,f117737,f117716,f117717,f117739,f117720,f117721,f117741,f117722,f117723,f106634,f117724,f117725,f117743,f117726,f117727,f117745,f117746,f117747,f117748,f117749,f1402,f117777,f117779,f117780,f117781,f117784,f117850,f6711,f118032,f118034,f118035,f118037,f118039,f118108,f9813,f118213,f118215,f118216,f118217,f118220,f118286,f13686,f118341,f118343,f118344,f118346,f118348,f118417,f18154,f118836,f118838,f118839,f782,f118841,f784,f118843,f783,f118912,f785,f71131,f119337])).
fof(f119337,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f119187,f459])).
fof(f459,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))))),
  inference(cnf_transformation,[],[f375])).
fof(f375,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) & ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK11,sK12,sK13])],[f336])).
fof(f336,plain,(
  ? [X0,X1,X2] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X2)))) & ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X2)))))),
  inference(flattening,[],[f335])).
fof(f335,plain,(
  ? [X0,X1,X2] : ((p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X2))))) & ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X2)))))),
  inference(ennf_transformation,[],[f263])).
fof(f263,plain,(
  ~! [X0,X1,X2] : ((p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X2))))) => p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X2)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : ((p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X6),s('type.integer.int',X7)))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X7),s('type.integer.int',X8))))) => p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X6),s('type.integer.int',X8)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : ((p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X6),s('type.integer.int',X7)))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X7),s('type.integer.int',X8))))) => p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X6),s('type.integer.int',X8)))))),
  file('thm.integer.INT_LTE_TRANS','thm.integer.INT_LTE_TRANS')).
fof(f119187,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13)))) | $spl12),
  inference(resolution,[],[f661,f906])).
fof(f661,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',X0))))) )),
  inference(resolution,[],[f457,f622])).
fof(f622,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X2)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X2)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f351])).
fof(f351,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X2)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X2)))))),
  inference(flattening,[],[f350])).
fof(f350,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X2))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X2)))))),
  inference(ennf_transformation,[],[f309])).
fof(f309,plain,(
  ! [X0,X1,X2] : ((p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X2))))) => p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X2)))))),
  inference(rectify,[],[f55])).
fof(f55,axiom,(
  ! [X6,X7,X8] : ((p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X6),s('type.integer.int',X7)))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X7),s('type.integer.int',X8))))) => p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X6),s('type.integer.int',X8)))))),
  file('thm.integer.INT_LTE_TRANS','thm.integer.INT_LT_TRANS')).
fof(f457,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))))),
  inference(cnf_transformation,[],[f375])).
fof(f71131,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) | ~p(s(bool,X2))) ) | $spl150),
  inference(cnf_transformation,[],[f71131_D])).
fof(f71131_D,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) | ~p(s(bool,X2))) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f785,plain,(
  ( ! [X15] : (p(s(bool,X15))) ) | $spl8),
  inference(cnf_transformation,[],[f785_D])).
fof(f785_D,plain,(
  ( ! [X15] : (p(s(bool,X15))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f118912,plain,(
  ( ! [X147,X148] : (p(s(bool,X147)) | ~p(s(bool,X148)) | s(bool,X148) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl66),
  inference(resolution,[],[f74111,f648])).
fof(f648,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f481])).
fof(f481,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f384])).
fof(f384,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f383])).
fof(f383,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f339])).
fof(f339,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f338])).
fof(f338,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f278])).
fof(f278,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f198])).
fof(f198,axiom,(
  ! [X6,X62,X7,X63] : ((s(bool,X6) = s(bool,X62) & (p(s(bool,X62)) => s(bool,X7) = s(bool,X63))) => ((p(s(bool,X6)) => p(s(bool,X7))) <=> (p(s(bool,X62)) => p(s(bool,X63)))))),
  file('thm.integer.INT_LTE_TRANS','thm.bool.IMP_CONG')).
fof(f74111,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) | $spl66),
  inference(resolution,[],[f37512,f464])).
fof(f464,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f378])).
fof(f378,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) & (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))),
  inference(nnf_transformation,[],[f269])).
fof(f269,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) <=> ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(flattening,[],[f268])).
fof(f268,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) <=> ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6,X7] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X6),s('type.integer.int',X7)))) <=> ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X7),s('type.integer.int',X6)))))),
  file('thm.integer.INT_LTE_TRANS','thm.integer.int_le')).
fof(f37512,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) | $spl66),
  inference(equality_resolution,[],[f10363])).
fof(f783,plain,(
  ( ! [X10,X11] : (s(bool,X10) != s(bool,X11) | ~p(s(bool,X10))) ) | $spl4),
  inference(cnf_transformation,[],[f783_D])).
fof(f783_D,plain,(
  ( ! [X10,X11] : (s(bool,X10) != s(bool,X11) | ~p(s(bool,X10))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f118843,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X22) != s(bool,X24)) ) | $spl66),
  inference(resolution,[],[f74111,f484])).
fof(f484,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f384])).
fof(f784,plain,(
  ( ! [X16] : (~p(s(bool,X16))) ) | $spl6),
  inference(cnf_transformation,[],[f784_D])).
fof(f784_D,plain,(
  ( ! [X16] : (~p(s(bool,X16))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f118841,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl66),
  inference(resolution,[],[f74111,f483])).
fof(f483,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f384])).
fof(f782,plain,(
  ( ! [X5,X3] : (s(bool,X3) != s(bool,X5) | p(s(bool,X3))) ) | $spl2),
  inference(cnf_transformation,[],[f782_D])).
fof(f782_D,plain,(
  ( ! [X5,X3] : (s(bool,X3) != s(bool,X5) | p(s(bool,X3))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f118839,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl66),
  inference(resolution,[],[f74111,f482])).
fof(f482,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f384])).
fof(f118838,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X7) != s(bool,X9)) ) | $spl66),
  inference(resolution,[],[f74111,f480])).
fof(f480,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f384])).
fof(f118836,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl66),
  inference(resolution,[],[f74111,f478])).
fof(f478,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f384])).
fof(f18154,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl76),
  inference(cnf_transformation,[],[f18154_D])).
fof(f18154_D,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f118417,plain,(
  ( ! [X147,X148] : (p(s(bool,X147)) | ~p(s(bool,X148)) | s(bool,X148) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl16),
  inference(resolution,[],[f6594,f648])).
fof(f6594,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) | $spl16),
  inference(resolution,[],[f1403,f626])).
fof(f626,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f353])).
fof(f353,plain,(
  ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(ennf_transformation,[],[f313])).
fof(f313,plain,(
  ! [X0,X1] : ~(p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X6,X7] : ~(p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X6),s('type.integer.int',X7)))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X7),s('type.integer.int',X6)))))),
  file('thm.integer.INT_LTE_TRANS','thm.integer.INT_LT_ANTISYM')).
fof(f118348,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | s(bool,X22) != s(bool,X24)) ) | $spl16),
  inference(resolution,[],[f6594,f484])).
fof(f118346,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl16),
  inference(resolution,[],[f6594,f483])).
fof(f118344,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl16),
  inference(resolution,[],[f6594,f482])).
fof(f118343,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | s(bool,X7) != s(bool,X9)) ) | $spl16),
  inference(resolution,[],[f6594,f480])).
fof(f118341,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl16),
  inference(resolution,[],[f6594,f478])).
fof(f13686,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl70),
  inference(cnf_transformation,[],[f13686_D])).
fof(f13686_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f118286,plain,(
  ( ! [X140,X141] : (p(s(bool,X140)) | p(s(bool,X141)) | s(bool,X141) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl16),
  inference(resolution,[],[f6587,f648])).
fof(f6587,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) | $spl16),
  inference(resolution,[],[f1403,f467])).
fof(f467,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f380])).
fof(f380,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s('type.integer.int',X0) = s('type.integer.int',X1)) & ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & s('type.integer.int',X0) != s('type.integer.int',X1)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))),
  inference(flattening,[],[f379])).
fof(f379,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s('type.integer.int',X0) = s('type.integer.int',X1))) & ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & s('type.integer.int',X0) != s('type.integer.int',X1)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))),
  inference(nnf_transformation,[],[f270])).
fof(f270,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) <=> (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s('type.integer.int',X0) = s('type.integer.int',X1)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X6,X7] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X6),s('type.integer.int',X7)))) <=> (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X6),s('type.integer.int',X7)))) | s('type.integer.int',X6) = s('type.integer.int',X7)))),
  file('thm.integer.INT_LTE_TRANS','thm.integer.INT_LE_LT')).
fof(f118220,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X22) != s(bool,X23)) ) | $spl16),
  inference(resolution,[],[f6587,f484])).
fof(f118217,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl16),
  inference(resolution,[],[f6587,f483])).
fof(f118216,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl16),
  inference(resolution,[],[f6587,f482])).
fof(f118215,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X7) != s(bool,X8)) ) | $spl16),
  inference(resolution,[],[f6587,f480])).
fof(f118213,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl16),
  inference(resolution,[],[f6587,f478])).
fof(f9813,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl64),
  inference(cnf_transformation,[],[f9813_D])).
fof(f9813_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f118108,plain,(
  ( ! [X146,X147] : (p(s(bool,X146)) | ~p(s(bool,X147)) | s(bool,X147) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl16),
  inference(resolution,[],[f6584,f648])).
fof(f6584,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) | $spl16),
  inference(resolution,[],[f1403,f461])).
fof(f461,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) )),
  inference(cnf_transformation,[],[f376])).
fof(f376,plain,(
  ! [X0,X1] : ((p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) & (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))),
  inference(nnf_transformation,[],[f265])).
fof(f265,plain,(
  ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) <=> p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(flattening,[],[f264])).
fof(f264,plain,(
  ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) <=> p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X6,X7] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X6),s('type.integer.int',X7)))) <=> p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X7),s('type.integer.int',X6)))))),
  file('thm.integer.INT_LTE_TRANS','thm.integer.INT_NOT_LT')).
fof(f118039,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | s(bool,X21) != s(bool,X23)) ) | $spl16),
  inference(resolution,[],[f6584,f484])).
fof(f118037,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl16),
  inference(resolution,[],[f6584,f483])).
fof(f118035,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl16),
  inference(resolution,[],[f6584,f482])).
fof(f118034,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | s(bool,X6) != s(bool,X8)) ) | $spl16),
  inference(resolution,[],[f6584,f480])).
fof(f118032,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl16),
  inference(resolution,[],[f6584,f478])).
fof(f6711,plain,(
  ( ! [X4] : (s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X4))) ) | $spl50),
  inference(cnf_transformation,[],[f6711_D])).
fof(f6711_D,plain,(
  ( ! [X4] : (s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X4))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f117850,plain,(
  ( ! [X144,X143] : (p(s(bool,X143)) | p(s(bool,X144)) | s(bool,X144) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl16),
  inference(resolution,[],[f1403,f648])).
fof(f117784,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X25) != s(bool,X26)) ) | $spl16),
  inference(resolution,[],[f1403,f484])).
fof(f117781,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl16),
  inference(resolution,[],[f1403,f483])).
fof(f117780,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl16),
  inference(resolution,[],[f1403,f482])).
fof(f117779,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X10) != s(bool,X11)) ) | $spl16),
  inference(resolution,[],[f1403,f480])).
fof(f117777,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl16),
  inference(resolution,[],[f1403,f478])).
fof(f1402,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) | $spl17),
  inference(cnf_transformation,[],[f1402_D])).
fof(f1402_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f117749,plain,(
  ( ! [X12,X13] : (~p(s(bool,X13)) | ~p(s(bool,X12)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))))) ) | $spl58),
  inference(subsumption_resolution,[],[f117018,f590])).
fof(f590,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f433])).
fof(f433,plain,(
  ! [X0,X1] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X0)) & p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X1)))) & (((~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X1)))) | s(bool,X0) = s(bool,X1)))),
  inference(flattening,[],[f432])).
fof(f432,plain,(
  ! [X0,X1] : ((s(bool,X0) != s(bool,X1) | ((p(s(bool,X0)) & p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X1))))) & (((~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X1)))) | s(bool,X0) = s(bool,X1)))),
  inference(nnf_transformation,[],[f296])).
fof(f296,plain,(
  ! [X0,X1] : (s(bool,X0) = s(bool,X1) <=> ((p(s(bool,X0)) & p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X1)))))),
  inference(flattening,[],[f295])).
fof(f295,plain,(
  ! [X0,X1] : (s(bool,X0) = s(bool,X1) <=> ((p(s(bool,X0)) & p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X1)))))),
  inference(rectify,[],[f244])).
fof(f244,axiom,(
  ! [X14,X15] : (s(bool,X14) = s(bool,X15) <=> ((p(s(bool,X14)) & p(s(bool,X15))) | (~p(s(bool,X14)) & ~p(s(bool,X15)))))),
  file('thm.integer.INT_LTE_TRANS','thm.bool.EQ_EXPAND')).
fof(f117018,plain,(
  ( ! [X12,X13] : (s(bool,X12) != s(bool,X13) | ~p(s(bool,X13)) | ~p(s(bool,X12)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))))) ) | $spl58),
  inference(superposition,[],[f7061,f590])).
fof(f117748,plain,(
  ( ! [X14,X15] : (~p(s(bool,X15)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) | ~p(s(bool,X14))) ) | $spl58),
  inference(subsumption_resolution,[],[f117019,f590])).
fof(f117019,plain,(
  ( ! [X14,X15] : (s(bool,X14) != s(bool,X15) | ~p(s(bool,X15)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) | ~p(s(bool,X14))) ) | $spl58),
  inference(superposition,[],[f7061,f590])).
fof(f117747,plain,(
  ( ! [X24,X25] : (~p(s(bool,X25)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) | ~p(s(bool,X24))) ) | $spl58),
  inference(subsumption_resolution,[],[f117024,f590])).
fof(f117024,plain,(
  ( ! [X24,X25] : (s(bool,X24) != s(bool,X25) | ~p(s(bool,X25)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) | ~p(s(bool,X24))) ) | $spl58),
  inference(superposition,[],[f7061,f595])).
fof(f595,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f346])).
fof(f346,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f345])).
fof(f345,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f297])).
fof(f297,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X14,X15] : ((p(s(bool,X14)) => p(s(bool,X15))) => ((p(s(bool,X15)) => p(s(bool,X14))) => s(bool,X14) = s(bool,X15)))),
  file('thm.integer.INT_LTE_TRANS','thm.bool.IMP_ANTISYM_AX')).
fof(f117746,plain,(
  ( ! [X26,X27] : (~p(s(bool,X27)) | ~p(s(bool,X26)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))))) ) | $spl58),
  inference(subsumption_resolution,[],[f117025,f590])).
fof(f117025,plain,(
  ( ! [X26,X27] : (s(bool,X26) != s(bool,X27) | ~p(s(bool,X27)) | ~p(s(bool,X26)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))))) ) | $spl58),
  inference(superposition,[],[f7061,f595])).
fof(f117745,plain,(
  s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) | $spl220),
  inference(cnf_transformation,[],[f117745_D])).
fof(f117745_D,plain,(
  s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) <=> ~$spl220),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl220])])).
fof(f117727,plain,(
  ( ! [X55] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) | ~p(s(bool,X55))) ) | $spl101),
  inference(resolution,[],[f45376,f107913])).
fof(f107913,plain,(
  ( ! [X111,X110] : (sP8('const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')),X110,X111) | ~p(s(bool,X111))) ) | $spl101),
  inference(resolution,[],[f77710,f558])).
fof(f558,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f419])).
fof(f419,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | (p(s(bool,X2)) & p(s(bool,X0))) | (p(s(bool,X1)) & p(s(bool,X0)))) & (((~p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f418])).
fof(f418,plain,(
  ! [X0,X2,X1] : ((~sP8(X0,X2,X1) | (p(s(bool,X1)) & p(s(bool,X0))) | (p(s(bool,X2)) & p(s(bool,X0)))) & (((~p(s(bool,X1)) | ~p(s(bool,X0))) & (~p(s(bool,X2)) | ~p(s(bool,X0)))) | sP8(X0,X2,X1)))),
  inference(flattening,[],[f417])).
fof(f417,plain,(
  ! [X0,X2,X1] : ((~sP8(X0,X2,X1) | ((p(s(bool,X1)) & p(s(bool,X0))) | (p(s(bool,X2)) & p(s(bool,X0))))) & (((~p(s(bool,X1)) | ~p(s(bool,X0))) & (~p(s(bool,X2)) | ~p(s(bool,X0)))) | sP8(X0,X2,X1)))),
  inference(nnf_transformation,[],[f369])).
fof(f369,plain,(
  ! [X0,X2,X1] : (sP8(X0,X2,X1) <=> ((p(s(bool,X1)) & p(s(bool,X0))) | (p(s(bool,X2)) & p(s(bool,X0)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f77710,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) | $spl101),
  inference(resolution,[],[f37521,f475])).
fof(f475,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f276])).
fof(f276,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X6,X7] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X6),s('type.integer.int',X7)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X7),s('type.integer.int',X6)))))),
  file('thm.integer.INT_LTE_TRANS','thm.integer.INT_LE_TOTAL')).
fof(f37521,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) | $spl101),
  inference(cnf_transformation,[],[f37521_D])).
fof(f37521_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) <=> ~$spl101),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl101])])).
fof(f45376,plain,(
  ( ! [X191,X192,X193] : (~sP8(X191,X192,X193) | s(bool,X191) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f2528,f563])).
fof(f563,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~sP8(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f421])).
fof(f421,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ~p(s(bool,X2))) | ~p(s(bool,X0)) | sP8(X0,X2,X1)) & (~sP8(X0,X2,X1) | ((p(s(bool,X1)) | p(s(bool,X2))) & p(s(bool,X0)))))),
  inference(flattening,[],[f420])).
fof(f420,plain,(
  ! [X0,X1,X2] : ((((~p(s(bool,X1)) & ~p(s(bool,X2))) | ~p(s(bool,X0))) | sP8(X0,X2,X1)) & (~sP8(X0,X2,X1) | ((p(s(bool,X1)) | p(s(bool,X2))) & p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f370])).
fof(f370,plain,(
  ! [X0,X1,X2] : (((p(s(bool,X1)) | p(s(bool,X2))) & p(s(bool,X0))) <=> sP8(X0,X2,X1))),
  inference(definition_folding,[],[f287,f369])).
fof(f287,plain,(
  ! [X0,X1,X2] : (((p(s(bool,X1)) | p(s(bool,X2))) & p(s(bool,X0))) <=> ((p(s(bool,X1)) & p(s(bool,X0))) | (p(s(bool,X2)) & p(s(bool,X0)))))),
  inference(rectify,[],[f124])).
fof(f124,axiom,(
  ! [X9,X18,X17] : (((p(s(bool,X18)) | p(s(bool,X17))) & p(s(bool,X9))) <=> ((p(s(bool,X18)) & p(s(bool,X9))) | (p(s(bool,X17)) & p(s(bool,X9)))))),
  file('thm.integer.INT_LTE_TRANS','thm.bool.RIGHT_AND_OVER_OR')).
fof(f2528,plain,(
  ( ! [X136] : (~p(s(bool,X136)) | s(bool,X136) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f657,f590])).
fof(f657,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))))),
  inference(resolution,[],[f457,f467])).
fof(f117726,plain,(
  ( ! [X54] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) | ~p(s(bool,X54))) ) | $spl101),
  inference(resolution,[],[f45376,f107915])).
fof(f107915,plain,(
  ( ! [X114,X115] : (sP8('const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')),X114,X115) | ~p(s(bool,X114))) ) | $spl101),
  inference(resolution,[],[f77710,f559])).
fof(f559,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f419])).
fof(f117743,plain,(
  s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | $spl218),
  inference(cnf_transformation,[],[f117743_D])).
fof(f117743_D,plain,(
  s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) <=> ~$spl218),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl218])])).
fof(f117725,plain,(
  ( ! [X53] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | ~p(s(bool,X53))) )),
  inference(resolution,[],[f45376,f4509])).
fof(f4509,plain,(
  ( ! [X111,X110] : (sP8('const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11)),X110,X111) | ~p(s(bool,X111))) )),
  inference(resolution,[],[f1573,f558])).
fof(f1573,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))))),
  inference(resolution,[],[f459,f460])).
fof(f460,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f376])).
fof(f117724,plain,(
  ( ! [X52] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | ~p(s(bool,X52))) )),
  inference(resolution,[],[f45376,f4511])).
fof(f4511,plain,(
  ( ! [X114,X115] : (sP8('const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11)),X114,X115) | ~p(s(bool,X114))) )),
  inference(resolution,[],[f1573,f559])).
fof(f106634,plain,(
  s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | $spl176),
  inference(cnf_transformation,[],[f106634_D])).
fof(f106634_D,plain,(
  s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f117723,plain,(
  ( ! [X51] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | ~p(s(bool,X51))) )),
  inference(resolution,[],[f45376,f842])).
fof(f842,plain,(
  ( ! [X111,X110] : (sP8('const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)),X110,X111) | ~p(s(bool,X111))) )),
  inference(resolution,[],[f458,f558])).
fof(f458,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))),
  inference(cnf_transformation,[],[f375])).
fof(f117722,plain,(
  ( ! [X50] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | ~p(s(bool,X50))) )),
  inference(resolution,[],[f45376,f844])).
fof(f844,plain,(
  ( ! [X114,X115] : (sP8('const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)),X114,X115) | ~p(s(bool,X114))) )),
  inference(resolution,[],[f458,f559])).
fof(f117741,plain,(
  s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | $spl216),
  inference(cnf_transformation,[],[f117741_D])).
fof(f117741_D,plain,(
  s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) <=> ~$spl216),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl216])])).
fof(f117721,plain,(
  ( ! [X49] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | ~p(s(bool,X49))) ) | $spl28),
  inference(resolution,[],[f45376,f82313])).
fof(f82313,plain,(
  ( ! [X111,X110] : (sP8('const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')),X110,X111) | ~p(s(bool,X111))) ) | $spl28),
  inference(resolution,[],[f7202,f558])).
fof(f7202,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) | $spl28),
  inference(resolution,[],[f2018,f467])).
fof(f117720,plain,(
  ( ! [X48] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | ~p(s(bool,X48))) ) | $spl28),
  inference(resolution,[],[f45376,f82315])).
fof(f82315,plain,(
  ( ! [X114,X115] : (sP8('const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')),X114,X115) | ~p(s(bool,X114))) ) | $spl28),
  inference(resolution,[],[f7202,f559])).
fof(f117739,plain,(
  s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | $spl214),
  inference(cnf_transformation,[],[f117739_D])).
fof(f117739_D,plain,(
  s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) <=> ~$spl214),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl214])])).
fof(f117717,plain,(
  ( ! [X45] : (s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | ~p(s(bool,X45))) ) | $spl17),
  inference(resolution,[],[f45376,f106963])).
fof(f106963,plain,(
  ( ! [X111,X110] : (sP8('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)),X110,X111) | ~p(s(bool,X111))) ) | $spl17),
  inference(resolution,[],[f6930,f558])).
fof(f6930,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) | $spl17),
  inference(resolution,[],[f1402,f460])).
fof(f117716,plain,(
  ( ! [X44] : (s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | ~p(s(bool,X44))) ) | $spl17),
  inference(resolution,[],[f45376,f106965])).
fof(f106965,plain,(
  ( ! [X114,X115] : (sP8('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)),X114,X115) | ~p(s(bool,X114))) ) | $spl17),
  inference(resolution,[],[f6930,f559])).
fof(f117737,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | $spl212),
  inference(cnf_transformation,[],[f117737_D])).
fof(f117737_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) <=> ~$spl212),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl212])])).
fof(f117715,plain,(
  ( ! [X43] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | ~p(s(bool,X43))) ) | $spl101),
  inference(resolution,[],[f45376,f107788])).
fof(f107788,plain,(
  ( ! [X114,X113] : (sP8('const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')),X113,X114) | ~p(s(bool,X114))) ) | $spl101),
  inference(resolution,[],[f77701,f558])).
fof(f77701,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) | $spl101),
  inference(resolution,[],[f37521,f460])).
fof(f117714,plain,(
  ( ! [X42] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | ~p(s(bool,X42))) ) | $spl101),
  inference(resolution,[],[f45376,f107790])).
fof(f107790,plain,(
  ( ! [X118,X117] : (sP8('const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')),X117,X118) | ~p(s(bool,X117))) ) | $spl101),
  inference(resolution,[],[f77701,f559])).
fof(f117735,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | $spl210),
  inference(cnf_transformation,[],[f117735_D])).
fof(f117735_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f117713,plain,(
  ( ! [X41] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | ~p(s(bool,X41))) ) | $spl24),
  inference(resolution,[],[f45376,f6521])).
fof(f6521,plain,(
  ( ! [X114,X113] : (sP8('const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11)),X113,X114) | ~p(s(bool,X114))) ) | $spl24),
  inference(resolution,[],[f1703,f558])).
fof(f117712,plain,(
  ( ! [X40] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | ~p(s(bool,X40))) ) | $spl24),
  inference(resolution,[],[f45376,f6523])).
fof(f6523,plain,(
  ( ! [X118,X117] : (sP8('const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11)),X117,X118) | ~p(s(bool,X117))) ) | $spl24),
  inference(resolution,[],[f1703,f559])).
fof(f117733,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | $spl208),
  inference(cnf_transformation,[],[f117733_D])).
fof(f117733_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f117711,plain,(
  ( ! [X39] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | ~p(s(bool,X39))) ) | $spl12),
  inference(resolution,[],[f45376,f4391])).
fof(f4391,plain,(
  ( ! [X114,X113] : (sP8('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)),X113,X114) | ~p(s(bool,X114))) ) | $spl12),
  inference(resolution,[],[f906,f558])).
fof(f117710,plain,(
  ( ! [X38] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | ~p(s(bool,X38))) ) | $spl12),
  inference(resolution,[],[f45376,f4393])).
fof(f4393,plain,(
  ( ! [X118,X117] : (sP8('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)),X117,X118) | ~p(s(bool,X117))) ) | $spl12),
  inference(resolution,[],[f906,f559])).
fof(f117731,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | $spl206),
  inference(cnf_transformation,[],[f117731_D])).
fof(f117731_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f117709,plain,(
  ( ! [X37] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | ~p(s(bool,X37))) ) | $spl28),
  inference(resolution,[],[f45376,f79987])).
fof(f79987,plain,(
  ( ! [X114,X113] : (sP8('const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')),X113,X114) | ~p(s(bool,X114))) ) | $spl28),
  inference(resolution,[],[f2018,f558])).
fof(f117708,plain,(
  ( ! [X36] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | ~p(s(bool,X36))) ) | $spl28),
  inference(resolution,[],[f45376,f79989])).
fof(f79989,plain,(
  ( ! [X118,X117] : (sP8('const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')),X117,X118) | ~p(s(bool,X117))) ) | $spl28),
  inference(resolution,[],[f2018,f559])).
fof(f47770,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | $spl132),
  inference(cnf_transformation,[],[f47770_D])).
fof(f47770_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f117707,plain,(
  ( ! [X35] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | ~p(s(bool,X35))) )),
  inference(resolution,[],[f45376,f719])).
fof(f719,plain,(
  ( ! [X114,X113] : (sP8('const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)),X113,X114) | ~p(s(bool,X114))) )),
  inference(resolution,[],[f457,f558])).
fof(f117706,plain,(
  ( ! [X34] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | ~p(s(bool,X34))) )),
  inference(resolution,[],[f45376,f721])).
fof(f721,plain,(
  ( ! [X118,X117] : (sP8('const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)),X117,X118) | ~p(s(bool,X117))) )),
  inference(resolution,[],[f457,f559])).
fof(f117729,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | $spl204),
  inference(cnf_transformation,[],[f117729_D])).
fof(f117729_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f117705,plain,(
  ( ! [X33] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | ~p(s(bool,X33))) ) | $spl56),
  inference(resolution,[],[f45376,f107468])).
fof(f107468,plain,(
  ( ! [X116,X115] : (sP8('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)),X115,X116) | ~p(s(bool,X116))) ) | $spl56),
  inference(resolution,[],[f7060,f558])).
fof(f117704,plain,(
  ( ! [X32] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | ~p(s(bool,X32))) ) | $spl56),
  inference(resolution,[],[f45376,f107470])).
fof(f107470,plain,(
  ( ! [X120,X119] : (sP8('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)),X119,X120) | ~p(s(bool,X119))) ) | $spl56),
  inference(resolution,[],[f7060,f559])).
fof(f117673,plain,(
  ( ! [X21,X20] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | p(s(bool,X20)) | ~p(s(bool,X21))) )),
  inference(resolution,[],[f45362,f834])).
fof(f834,plain,(
  ( ! [X94,X95] : (~sP5('const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)),X94,X95) | p(s(bool,X95)) | ~p(s(bool,X94))) )),
  inference(resolution,[],[f458,f530])).
fof(f530,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f409])).
fof(f409,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f408])).
fof(f408,plain,(
  ! [X1,X0,X2] : ((~sP5(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP5(X1,X0,X2)))),
  inference(flattening,[],[f407])).
fof(f407,plain,(
  ! [X1,X0,X2] : ((~sP5(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP5(X1,X0,X2)))),
  inference(nnf_transformation,[],[f365])).
fof(f365,plain,(
  ! [X1,X0,X2] : (sP5(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f45362,plain,(
  ( ! [X151,X149,X150] : (sP5(X149,X150,X151) | s(bool,X149) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f2528,f533])).
fof(f533,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f409])).
fof(f117672,plain,(
  ( ! [X19,X18] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | p(s(bool,X18)) | ~p(s(bool,X19))) )),
  inference(resolution,[],[f45362,f711])).
fof(f711,plain,(
  ( ! [X97,X98] : (~sP5('const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)),X97,X98) | p(s(bool,X98)) | ~p(s(bool,X97))) )),
  inference(resolution,[],[f457,f530])).
fof(f117661,plain,(
  ( ! [X17,X16] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | p(s(bool,X16)) | ~p(s(bool,X17))) )),
  inference(resolution,[],[f45361,f835])).
fof(f835,plain,(
  ( ! [X97,X96] : (~sP5(X96,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)),X97) | p(s(bool,X97)) | ~p(s(bool,X96))) )),
  inference(resolution,[],[f458,f530])).
fof(f45361,plain,(
  ( ! [X146,X147,X148] : (sP5(X147,X146,X148) | s(bool,X146) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f2528,f532])).
fof(f532,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f409])).
fof(f117660,plain,(
  ( ! [X14,X15] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | p(s(bool,X14)) | ~p(s(bool,X15))) )),
  inference(resolution,[],[f45361,f712])).
fof(f712,plain,(
  ( ! [X99,X100] : (~sP5(X99,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)),X100) | p(s(bool,X100)) | ~p(s(bool,X99))) )),
  inference(resolution,[],[f457,f530])).
fof(f117652,plain,(
  ( ! [X21,X20] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | p(s(bool,X20)) | ~p(s(bool,X21))) )),
  inference(resolution,[],[f45359,f831])).
fof(f831,plain,(
  ( ! [X88,X89] : (~sP6('const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)),X88,X89) | p(s(bool,X89)) | ~p(s(bool,X88))) )),
  inference(resolution,[],[f458,f526])).
fof(f526,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f406])).
fof(f406,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(rectify,[],[f405])).
fof(f405,plain,(
  ! [X2,X0,X1] : ((~sP6(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP6(X2,X0,X1)))),
  inference(flattening,[],[f404])).
fof(f404,plain,(
  ! [X2,X0,X1] : ((~sP6(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP6(X2,X0,X1)))),
  inference(nnf_transformation,[],[f366])).
fof(f366,plain,(
  ! [X2,X0,X1] : (sP6(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f45359,plain,(
  ( ! [X142,X140,X141] : (sP6(X140,X141,X142) | s(bool,X140) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f2528,f529])).
fof(f529,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f406])).
fof(f117651,plain,(
  ( ! [X19,X18] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | p(s(bool,X18)) | ~p(s(bool,X19))) )),
  inference(resolution,[],[f45359,f708])).
fof(f708,plain,(
  ( ! [X92,X91] : (~sP6('const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)),X91,X92) | p(s(bool,X92)) | ~p(s(bool,X91))) )),
  inference(resolution,[],[f457,f526])).
fof(f117640,plain,(
  ( ! [X17,X16] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | p(s(bool,X16)) | ~p(s(bool,X17))) )),
  inference(resolution,[],[f45358,f832])).
fof(f832,plain,(
  ( ! [X90,X91] : (~sP6(X90,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)),X91) | p(s(bool,X91)) | ~p(s(bool,X90))) )),
  inference(resolution,[],[f458,f526])).
fof(f45358,plain,(
  ( ! [X138,X139,X137] : (sP6(X138,X137,X139) | s(bool,X137) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f2528,f528])).
fof(f528,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f406])).
fof(f117639,plain,(
  ( ! [X14,X15] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | p(s(bool,X14)) | ~p(s(bool,X15))) )),
  inference(resolution,[],[f45358,f709])).
fof(f709,plain,(
  ( ! [X94,X93] : (~sP6(X93,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)),X94) | p(s(bool,X94)) | ~p(s(bool,X93))) )),
  inference(resolution,[],[f457,f526])).
fof(f117633,plain,(
  ( ! [X21,X22] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | p(s(bool,X21)) | p(s(bool,X22))) )),
  inference(resolution,[],[f45346,f818])).
fof(f818,plain,(
  ( ! [X62,X63] : (~sP2('const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)),X62,X63) | p(s(bool,X63)) | p(s(bool,X62))) )),
  inference(resolution,[],[f458,f509])).
fof(f509,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f398])).
fof(f398,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(rectify,[],[f397])).
fof(f397,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(flattening,[],[f396])).
fof(f396,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(nnf_transformation,[],[f361])).
fof(f361,plain,(
  ! [X2,X0,X1] : (sP2(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f45346,plain,(
  ( ! [X103,X101,X102] : (sP2(X101,X102,X103) | s(bool,X101) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f2528,f512])).
fof(f512,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f398])).
fof(f117632,plain,(
  ( ! [X19,X20] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | p(s(bool,X19)) | p(s(bool,X20))) )),
  inference(resolution,[],[f45346,f695])).
fof(f695,plain,(
  ( ! [X66,X65] : (~sP2('const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)),X65,X66) | p(s(bool,X66)) | p(s(bool,X65))) )),
  inference(resolution,[],[f457,f509])).
fof(f117624,plain,(
  ( ! [X24,X23] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | ~p(s(bool,X23)) | p(s(bool,X24))) )),
  inference(resolution,[],[f45330,f804])).
fof(f804,plain,(
  ( ! [X35,X34] : (~sP0('const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)),X34,X35) | ~p(s(bool,X35)) | p(s(bool,X34))) )),
  inference(resolution,[],[f458,f492])).
fof(f492,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP0(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f390])).
fof(f390,plain,(
  ! [X0,X1,X2] : ((~sP0(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP0(X0,X1,X2)))),
  inference(rectify,[],[f389])).
fof(f389,plain,(
  ! [X2,X0,X1] : ((~sP0(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP0(X2,X0,X1)))),
  inference(flattening,[],[f388])).
fof(f388,plain,(
  ! [X2,X0,X1] : ((~sP0(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP0(X2,X0,X1)))),
  inference(nnf_transformation,[],[f358])).
fof(f358,plain,(
  ! [X2,X0,X1] : (sP0(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f45330,plain,(
  ( ! [X54,X55,X53] : (sP0(X53,X54,X55) | s(bool,X53) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f2528,f495])).
fof(f495,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f390])).
fof(f117623,plain,(
  ( ! [X21,X22] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | ~p(s(bool,X21)) | p(s(bool,X22))) )),
  inference(resolution,[],[f45330,f681])).
fof(f681,plain,(
  ( ! [X37,X38] : (~sP0('const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)),X37,X38) | ~p(s(bool,X38)) | p(s(bool,X37))) )),
  inference(resolution,[],[f457,f492])).
fof(f117610,plain,(
  ( ! [X17,X18] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | p(s(bool,X17)) | ~p(s(bool,X18))) )),
  inference(resolution,[],[f45329,f805])).
fof(f805,plain,(
  ( ! [X37,X36] : (~sP0(X36,X37,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | p(s(bool,X37)) | ~p(s(bool,X36))) )),
  inference(resolution,[],[f458,f492])).
fof(f45329,plain,(
  ( ! [X52,X50,X51] : (sP0(X51,X52,X50) | s(bool,X50) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f2528,f493])).
fof(f493,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f390])).
fof(f117609,plain,(
  ( ! [X15,X16] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | p(s(bool,X15)) | ~p(s(bool,X16))) )),
  inference(resolution,[],[f45329,f682])).
fof(f682,plain,(
  ( ! [X39,X40] : (~sP0(X39,X40,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | p(s(bool,X40)) | ~p(s(bool,X39))) )),
  inference(resolution,[],[f457,f492])).
fof(f117510,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | $spl202),
  inference(cnf_transformation,[],[f117510_D])).
fof(f117510_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f117574,plain,(
  ( ! [X17,X16] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | p(s(bool,X16)) | p(s(bool,X17))) )),
  inference(resolution,[],[f45086,f1632])).
fof(f1632,plain,(
  ( ! [X107,X108] : (~sP4('const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13)),X107,X108) | p(s(bool,X108)) | p(s(bool,X107))) )),
  inference(resolution,[],[f459,f534])).
fof(f534,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f411])).
fof(f411,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(flattening,[],[f410])).
fof(f410,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(nnf_transformation,[],[f364])).
fof(f364,plain,(
  ! [X0,X1,X2] : (sP4(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f45086,plain,(
  ( ! [X154,X152,X153] : (sP4(X152,X153,X154) | s(bool,X152) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f2410,f537])).
fof(f537,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f411])).
fof(f2410,plain,(
  ( ! [X142] : (p(s(bool,X142)) | s(bool,X142) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f654,f591])).
fof(f591,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f433])).
fof(f654,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))))),
  inference(resolution,[],[f457,f461])).
fof(f117565,plain,(
  ( ! [X14,X15] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | p(s(bool,X14)) | p(s(bool,X15))) )),
  inference(resolution,[],[f45085,f1631])).
fof(f1631,plain,(
  ( ! [X105,X106] : (~sP4(X105,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13)),X106) | p(s(bool,X106)) | p(s(bool,X105))) )),
  inference(resolution,[],[f459,f534])).
fof(f45085,plain,(
  ( ! [X151,X149,X150] : (sP4(X150,X149,X151) | s(bool,X149) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f2410,f536])).
fof(f536,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f411])).
fof(f117556,plain,(
  ( ! [X12,X13] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | p(s(bool,X12)) | p(s(bool,X13))) )),
  inference(resolution,[],[f45084,f1633])).
fof(f1633,plain,(
  ( ! [X109,X110] : (~sP4(X109,X110,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | p(s(bool,X110)) | p(s(bool,X109))) )),
  inference(resolution,[],[f459,f534])).
fof(f45084,plain,(
  ( ! [X146,X147,X148] : (sP4(X147,X148,X146) | s(bool,X146) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f2410,f535])).
fof(f535,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f411])).
fof(f117546,plain,(
  ( ! [X12,X13] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | ~p(s(bool,X12)) | ~p(s(bool,X13))) )),
  inference(resolution,[],[f45083,f1628])).
fof(f1628,plain,(
  ( ! [X99,X100] : (~sP5(X99,X100,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | ~p(s(bool,X100)) | ~p(s(bool,X99))) )),
  inference(resolution,[],[f459,f530])).
fof(f45083,plain,(
  ( ! [X144,X145,X143] : (sP5(X144,X145,X143) | s(bool,X143) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f2410,f531])).
fof(f531,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f409])).
fof(f117536,plain,(
  ( ! [X12,X13] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | ~p(s(bool,X12)) | ~p(s(bool,X13))) )),
  inference(resolution,[],[f45080,f1625])).
fof(f1625,plain,(
  ( ! [X94,X93] : (~sP6(X93,X94,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | ~p(s(bool,X94)) | ~p(s(bool,X93))) )),
  inference(resolution,[],[f459,f526])).
fof(f45080,plain,(
  ( ! [X134,X136,X135] : (sP6(X135,X136,X134) | s(bool,X134) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f2410,f527])).
fof(f527,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f406])).
fof(f117528,plain,(
  ( ! [X17,X18] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | p(s(bool,X17)) | ~p(s(bool,X18))) )),
  inference(resolution,[],[f45067,f1612])).
fof(f1612,plain,(
  ( ! [X68,X67] : (~sP2(X67,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13)),X68) | p(s(bool,X68)) | ~p(s(bool,X67))) )),
  inference(resolution,[],[f459,f509])).
fof(f45067,plain,(
  ( ! [X97,X95,X96] : (sP2(X96,X95,X97) | s(bool,X95) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f2410,f511])).
fof(f511,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f398])).
fof(f117516,plain,(
  ( ! [X15,X16] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | p(s(bool,X15)) | ~p(s(bool,X16))) )),
  inference(resolution,[],[f45066,f1613])).
fof(f1613,plain,(
  ( ! [X70,X69] : (~sP2(X69,X70,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | p(s(bool,X70)) | ~p(s(bool,X69))) )),
  inference(resolution,[],[f459,f509])).
fof(f45066,plain,(
  ( ! [X94,X92,X93] : (sP2(X93,X94,X92) | s(bool,X92) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f2410,f510])).
fof(f510,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f398])).
fof(f117505,plain,(
  ( ! [X19,X20] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | ~p(s(bool,X19)) | ~p(s(bool,X20))) )),
  inference(resolution,[],[f45053,f1596])).
fof(f1596,plain,(
  ( ! [X35,X36] : (~sP0(X35,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13)),X36) | ~p(s(bool,X36)) | ~p(s(bool,X35))) )),
  inference(resolution,[],[f459,f492])).
fof(f45053,plain,(
  ( ! [X54,X55,X53] : (sP0(X54,X53,X55) | s(bool,X53) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f2410,f494])).
fof(f494,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f390])).
fof(f117439,plain,(
  $false | ($spl60 | $spl92)),
  inference(equality_resolution,[],[f117391])).
fof(f117391,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) ) | ($spl60 | $spl92)),
  inference(subsumption_resolution,[],[f27219,f7326])).
fof(f117386,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) | $spl200),
  inference(cnf_transformation,[],[f117386_D])).
fof(f117386_D,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f117306,plain,(
  ( ! [X215,X213,X214] : (s(bool,X213) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | p(s(bool,X214)) | p(s(bool,X215)) | s(bool,X213) != s(bool,X215)) ) | $spl12),
  inference(resolution,[],[f8611,f648])).
fof(f8611,plain,(
  ( ! [X142] : (p(s(bool,X142)) | s(bool,X142) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) | $spl12),
  inference(resolution,[],[f3242,f591])).
fof(f3242,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | $spl12),
  inference(resolution,[],[f906,f461])).
fof(f117240,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) ) | $spl12),
  inference(resolution,[],[f8611,f484])).
fof(f117237,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,X16) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) ) | $spl12),
  inference(resolution,[],[f8611,f483])).
fof(f117236,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl12),
  inference(resolution,[],[f8611,f482])).
fof(f117235,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) ) | $spl12),
  inference(resolution,[],[f8611,f480])).
fof(f117233,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl12),
  inference(resolution,[],[f8611,f478])).
fof(f116978,plain,(
  $false | ($spl17 | $spl50)),
  inference(subsumption_resolution,[],[f116962,f1402])).
fof(f116962,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) | $spl50),
  inference(equality_resolution,[],[f6711])).
fof(f116977,plain,(
  ( ! [X23,X22] : (p(s(bool,X23)) | p(s(bool,X22))) ) | ($spl17 | $spl50)),
  inference(subsumption_resolution,[],[f116976,f1402])).
fof(f116976,plain,(
  ( ! [X23,X22] : (p(s(bool,X23)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,X22))) ) | $spl50),
  inference(subsumption_resolution,[],[f116959,f591])).
fof(f116959,plain,(
  ( ! [X23,X22] : (s(bool,X22) != s(bool,X23) | p(s(bool,X23)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,X22))) ) | $spl50),
  inference(superposition,[],[f6711,f592])).
fof(f592,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f346])).
fof(f116975,plain,(
  ( ! [X21,X20] : (p(s(bool,X21)) | p(s(bool,X20))) ) | ($spl17 | $spl50)),
  inference(subsumption_resolution,[],[f116974,f1402])).
fof(f116974,plain,(
  ( ! [X21,X20] : (p(s(bool,X21)) | p(s(bool,X20)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))))) ) | $spl50),
  inference(subsumption_resolution,[],[f116958,f591])).
fof(f116958,plain,(
  ( ! [X21,X20] : (s(bool,X20) != s(bool,X21) | p(s(bool,X21)) | p(s(bool,X20)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))))) ) | $spl50),
  inference(superposition,[],[f6711,f592])).
fof(f116973,plain,(
  ( ! [X19,X18] : (p(s(bool,X19)) | p(s(bool,X18))) ) | ($spl17 | $spl50)),
  inference(subsumption_resolution,[],[f116972,f1402])).
fof(f116972,plain,(
  ( ! [X19,X18] : (p(s(bool,X19)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,X18))) ) | $spl50),
  inference(subsumption_resolution,[],[f116957,f591])).
fof(f116957,plain,(
  ( ! [X19,X18] : (s(bool,X18) != s(bool,X19) | p(s(bool,X19)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,X18))) ) | $spl50),
  inference(superposition,[],[f6711,f591])).
fof(f116971,plain,(
  ( ! [X17,X16] : (p(s(bool,X17)) | p(s(bool,X16))) ) | ($spl17 | $spl50)),
  inference(subsumption_resolution,[],[f116970,f1402])).
fof(f116970,plain,(
  ( ! [X17,X16] : (p(s(bool,X17)) | p(s(bool,X16)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))))) ) | $spl50),
  inference(subsumption_resolution,[],[f116956,f591])).
fof(f116956,plain,(
  ( ! [X17,X16] : (s(bool,X16) != s(bool,X17) | p(s(bool,X17)) | p(s(bool,X16)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))))) ) | $spl50),
  inference(superposition,[],[f6711,f591])).
fof(f116914,plain,(
  $false | ($spl101 | $spl150)),
  inference(subsumption_resolution,[],[f116898,f77701])).
fof(f116898,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) | $spl150),
  inference(equality_resolution,[],[f71131])).
fof(f116913,plain,(
  ( ! [X26,X27] : (~p(s(bool,X27)) | ~p(s(bool,X26))) ) | ($spl101 | $spl150)),
  inference(subsumption_resolution,[],[f116912,f77701])).
fof(f116912,plain,(
  ( ! [X26,X27] : (~p(s(bool,X27)) | ~p(s(bool,X26)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))))) ) | $spl150),
  inference(subsumption_resolution,[],[f116897,f590])).
fof(f116897,plain,(
  ( ! [X26,X27] : (s(bool,X26) != s(bool,X27) | ~p(s(bool,X27)) | ~p(s(bool,X26)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))))) ) | $spl150),
  inference(superposition,[],[f71131,f595])).
fof(f116911,plain,(
  ( ! [X24,X25] : (~p(s(bool,X25)) | ~p(s(bool,X24))) ) | ($spl101 | $spl150)),
  inference(subsumption_resolution,[],[f116910,f77701])).
fof(f116910,plain,(
  ( ! [X24,X25] : (~p(s(bool,X25)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) | ~p(s(bool,X24))) ) | $spl150),
  inference(subsumption_resolution,[],[f116896,f590])).
fof(f116896,plain,(
  ( ! [X24,X25] : (s(bool,X24) != s(bool,X25) | ~p(s(bool,X25)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) | ~p(s(bool,X24))) ) | $spl150),
  inference(superposition,[],[f71131,f595])).
fof(f116909,plain,(
  ( ! [X14,X15] : (~p(s(bool,X15)) | ~p(s(bool,X14))) ) | ($spl101 | $spl150)),
  inference(subsumption_resolution,[],[f116908,f77701])).
fof(f116908,plain,(
  ( ! [X14,X15] : (~p(s(bool,X15)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) | ~p(s(bool,X14))) ) | $spl150),
  inference(subsumption_resolution,[],[f116891,f590])).
fof(f116891,plain,(
  ( ! [X14,X15] : (s(bool,X14) != s(bool,X15) | ~p(s(bool,X15)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) | ~p(s(bool,X14))) ) | $spl150),
  inference(superposition,[],[f71131,f590])).
fof(f116907,plain,(
  ( ! [X12,X13] : (~p(s(bool,X13)) | ~p(s(bool,X12))) ) | ($spl101 | $spl150)),
  inference(subsumption_resolution,[],[f116906,f77701])).
fof(f116906,plain,(
  ( ! [X12,X13] : (~p(s(bool,X13)) | ~p(s(bool,X12)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))))) ) | $spl150),
  inference(subsumption_resolution,[],[f116890,f590])).
fof(f116890,plain,(
  ( ! [X12,X13] : (s(bool,X12) != s(bool,X13) | ~p(s(bool,X13)) | ~p(s(bool,X12)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))))) ) | $spl150),
  inference(superposition,[],[f71131,f590])).
fof(f116163,plain,(
  s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) | $spl198),
  inference(cnf_transformation,[],[f116163_D])).
fof(f116163_D,plain,(
  s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f116223,plain,(
  ( ! [X10] : (~p(s(bool,X10)) | s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl101),
  inference(resolution,[],[f107915,f35017])).
fof(f35017,plain,(
  ( ! [X191,X192,X193] : (~sP8(X191,X192,X193) | s(bool,X191) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f860,f563])).
fof(f860,plain,(
  ( ! [X136] : (~p(s(bool,X136)) | s(bool,X136) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f458,f590])).
fof(f116161,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) | $spl196),
  inference(cnf_transformation,[],[f116161_D])).
fof(f116161_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f116222,plain,(
  ( ! [X9] : (~p(s(bool,X9)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl101),
  inference(resolution,[],[f107915,f33992])).
fof(f33992,plain,(
  ( ! [X191,X192,X193] : (~sP8(X191,X192,X193) | s(bool,X191) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f737,f563])).
fof(f737,plain,(
  ( ! [X139] : (~p(s(bool,X139)) | s(bool,X139) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f457,f590])).
fof(f116115,plain,(
  ( ! [X10] : (~p(s(bool,X10)) | s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl101),
  inference(resolution,[],[f107913,f35017])).
fof(f116114,plain,(
  ( ! [X9] : (~p(s(bool,X9)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl101),
  inference(resolution,[],[f107913,f33992])).
fof(f115332,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | $spl194),
  inference(cnf_transformation,[],[f115332_D])).
fof(f115332_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f115392,plain,(
  ( ! [X10] : (~p(s(bool,X10)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) | $spl101),
  inference(resolution,[],[f107790,f35017])).
fof(f115330,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) | $spl192),
  inference(cnf_transformation,[],[f115330_D])).
fof(f115330_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f115391,plain,(
  ( ! [X9] : (~p(s(bool,X9)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl101),
  inference(resolution,[],[f107790,f33992])).
fof(f115284,plain,(
  ( ! [X10] : (~p(s(bool,X10)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) | $spl101),
  inference(resolution,[],[f107788,f35017])).
fof(f115283,plain,(
  ( ! [X9] : (~p(s(bool,X9)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl101),
  inference(resolution,[],[f107788,f33992])).
fof(f114082,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | $spl190),
  inference(cnf_transformation,[],[f114082_D])).
fof(f114082_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f114141,plain,(
  ( ! [X10] : (~p(s(bool,X10)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) | $spl56),
  inference(resolution,[],[f107470,f35017])).
fof(f47760,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | $spl122),
  inference(cnf_transformation,[],[f47760_D])).
fof(f47760_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f114140,plain,(
  ( ! [X9] : (~p(s(bool,X9)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) ) | $spl56),
  inference(resolution,[],[f107470,f33992])).
fof(f114038,plain,(
  ( ! [X10] : (~p(s(bool,X10)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) | $spl56),
  inference(resolution,[],[f107468,f35017])).
fof(f114037,plain,(
  ( ! [X9] : (~p(s(bool,X9)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) ) | $spl56),
  inference(resolution,[],[f107468,f33992])).
fof(f112880,plain,(
  s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | $spl188),
  inference(cnf_transformation,[],[f112880_D])).
fof(f112880_D,plain,(
  s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f112938,plain,(
  ( ! [X10] : (~p(s(bool,X10)) | s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) | $spl17),
  inference(resolution,[],[f106965,f35017])).
fof(f47768,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | $spl130),
  inference(cnf_transformation,[],[f47768_D])).
fof(f47768_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f112937,plain,(
  ( ! [X9] : (~p(s(bool,X9)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl17),
  inference(resolution,[],[f106965,f33992])).
fof(f112838,plain,(
  ( ! [X10] : (~p(s(bool,X10)) | s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) | $spl17),
  inference(resolution,[],[f106963,f35017])).
fof(f112837,plain,(
  ( ! [X9] : (~p(s(bool,X9)) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl17),
  inference(resolution,[],[f106963,f33992])).
fof(f108929,plain,(
  $false | ($spl101 | $spl157)),
  inference(subsumption_resolution,[],[f77450,f77710])).
fof(f107980,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl186),
  inference(cnf_transformation,[],[f107980_D])).
fof(f107980_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) ) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f107935,plain,(
  ( ! [X140,X141] : (p(s(bool,X140)) | p(s(bool,X141)) | s(bool,X141) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl101),
  inference(resolution,[],[f77710,f648])).
fof(f107869,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X22) != s(bool,X23)) ) | $spl101),
  inference(resolution,[],[f77710,f484])).
fof(f107866,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl101),
  inference(resolution,[],[f77710,f483])).
fof(f107865,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl101),
  inference(resolution,[],[f77710,f482])).
fof(f107864,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X7) != s(bool,X8)) ) | $spl101),
  inference(resolution,[],[f77710,f480])).
fof(f107862,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl101),
  inference(resolution,[],[f77710,f478])).
fof(f107856,plain,(
  ( ! [X4] : (p(s(bool,X4)) | s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl184),
  inference(cnf_transformation,[],[f107856_D])).
fof(f107856_D,plain,(
  ( ! [X4] : (p(s(bool,X4)) | s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) ) <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f107810,plain,(
  ( ! [X144,X143] : (p(s(bool,X143)) | p(s(bool,X144)) | s(bool,X144) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl101),
  inference(resolution,[],[f77701,f648])).
fof(f107744,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X25) != s(bool,X26)) ) | $spl101),
  inference(resolution,[],[f77701,f484])).
fof(f107741,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl101),
  inference(resolution,[],[f77701,f483])).
fof(f107740,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl101),
  inference(resolution,[],[f77701,f482])).
fof(f107739,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X10) != s(bool,X11)) ) | $spl101),
  inference(resolution,[],[f77701,f480])).
fof(f107737,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl101),
  inference(resolution,[],[f77701,f478])).
fof(f48687,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | ~p(s(bool,X2))) ) | $spl140),
  inference(cnf_transformation,[],[f48687_D])).
fof(f48687_D,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | ~p(s(bool,X2))) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f107622,plain,(
  ( ! [X147,X148] : (p(s(bool,X147)) | ~p(s(bool,X148)) | s(bool,X148) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | $spl119),
  inference(resolution,[],[f47222,f648])).
fof(f47222,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) | $spl119),
  inference(cnf_transformation,[],[f47222_D])).
fof(f47222_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) <=> ~$spl119),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl119])])).
fof(f107553,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | s(bool,X22) != s(bool,X24)) ) | $spl119),
  inference(resolution,[],[f47222,f484])).
fof(f107551,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | $spl119),
  inference(resolution,[],[f47222,f483])).
fof(f107549,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | $spl119),
  inference(resolution,[],[f47222,f482])).
fof(f107548,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | s(bool,X7) != s(bool,X9)) ) | $spl119),
  inference(resolution,[],[f47222,f480])).
fof(f107546,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | $spl119),
  inference(resolution,[],[f47222,f478])).
fof(f22371,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl80),
  inference(cnf_transformation,[],[f22371_D])).
fof(f22371_D,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f107490,plain,(
  ( ! [X146,X145] : (p(s(bool,X145)) | p(s(bool,X146)) | s(bool,X146) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl56),
  inference(resolution,[],[f7060,f648])).
fof(f107424,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | s(bool,X27) != s(bool,X28)) ) | $spl56),
  inference(resolution,[],[f7060,f484])).
fof(f107421,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X19) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl56),
  inference(resolution,[],[f7060,f483])).
fof(f107420,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl56),
  inference(resolution,[],[f7060,f482])).
fof(f107419,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | s(bool,X12) != s(bool,X13)) ) | $spl56),
  inference(resolution,[],[f7060,f480])).
fof(f107417,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl56),
  inference(resolution,[],[f7060,f478])).
fof(f14267,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl72),
  inference(cnf_transformation,[],[f14267_D])).
fof(f14267_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f107357,plain,(
  ( ! [X146,X147] : (p(s(bool,X146)) | ~p(s(bool,X147)) | s(bool,X147) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl53),
  inference(resolution,[],[f7056,f648])).
fof(f107288,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X21) != s(bool,X23)) ) | $spl53),
  inference(resolution,[],[f7056,f484])).
fof(f107286,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl53),
  inference(resolution,[],[f7056,f483])).
fof(f107284,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl53),
  inference(resolution,[],[f7056,f482])).
fof(f107283,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X6) != s(bool,X8)) ) | $spl53),
  inference(resolution,[],[f7056,f480])).
fof(f107281,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl53),
  inference(resolution,[],[f7056,f478])).
fof(f13315,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl68),
  inference(cnf_transformation,[],[f13315_D])).
fof(f13315_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f106985,plain,(
  ( ! [X140,X141] : (p(s(bool,X140)) | p(s(bool,X141)) | s(bool,X141) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl17),
  inference(resolution,[],[f6930,f648])).
fof(f106919,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | s(bool,X22) != s(bool,X23)) ) | $spl17),
  inference(resolution,[],[f6930,f484])).
fof(f106916,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl17),
  inference(resolution,[],[f6930,f483])).
fof(f106915,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl17),
  inference(resolution,[],[f6930,f482])).
fof(f106914,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | s(bool,X7) != s(bool,X8)) ) | $spl17),
  inference(resolution,[],[f6930,f480])).
fof(f106912,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl17),
  inference(resolution,[],[f6930,f478])).
fof(f77829,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | $spl158),
  inference(cnf_transformation,[],[f77829_D])).
fof(f77829_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f106863,plain,(
  ( ! [X146,X147] : (p(s(bool,X146)) | ~p(s(bool,X147)) | s(bool,X147) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | $spl101),
  inference(resolution,[],[f37521,f648])).
fof(f106794,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | s(bool,X21) != s(bool,X23)) ) | $spl101),
  inference(resolution,[],[f37521,f484])).
fof(f106792,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | $spl101),
  inference(resolution,[],[f37521,f483])).
fof(f106790,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | $spl101),
  inference(resolution,[],[f37521,f482])).
fof(f106789,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | s(bool,X6) != s(bool,X8)) ) | $spl101),
  inference(resolution,[],[f37521,f480])).
fof(f106787,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | $spl101),
  inference(resolution,[],[f37521,f478])).
fof(f7061,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | ~p(s(bool,X2))) ) | $spl58),
  inference(cnf_transformation,[],[f7061_D])).
fof(f7061_D,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | ~p(s(bool,X2))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f106728,plain,(
  ( ! [X147,X148] : (p(s(bool,X147)) | ~p(s(bool,X148)) | s(bool,X148) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl17),
  inference(resolution,[],[f1402,f648])).
fof(f106659,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X22) != s(bool,X24)) ) | $spl17),
  inference(resolution,[],[f1402,f484])).
fof(f106657,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl17),
  inference(resolution,[],[f1402,f483])).
fof(f106655,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl17),
  inference(resolution,[],[f1402,f482])).
fof(f106654,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X7) != s(bool,X9)) ) | $spl17),
  inference(resolution,[],[f1402,f480])).
fof(f106652,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl17),
  inference(resolution,[],[f1402,f478])).
fof(f106640,plain,(
  s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | $spl182),
  inference(cnf_transformation,[],[f106640_D])).
fof(f106640_D,plain,(
  s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f106620,plain,(
  ( ! [X55] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | ~p(s(bool,X55))) )),
  inference(resolution,[],[f35017,f4509])).
fof(f106619,plain,(
  ( ! [X54] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | ~p(s(bool,X54))) )),
  inference(resolution,[],[f35017,f4511])).
fof(f106638,plain,(
  s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | $spl180),
  inference(cnf_transformation,[],[f106638_D])).
fof(f106638_D,plain,(
  s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f106618,plain,(
  ( ! [X53] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | ~p(s(bool,X53))) ) | $spl16),
  inference(resolution,[],[f35017,f81738])).
fof(f81738,plain,(
  ( ! [X111,X110] : (sP8('const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')),X110,X111) | ~p(s(bool,X111))) ) | $spl16),
  inference(resolution,[],[f6587,f558])).
fof(f106617,plain,(
  ( ! [X52] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | ~p(s(bool,X52))) ) | $spl16),
  inference(resolution,[],[f35017,f81740])).
fof(f81740,plain,(
  ( ! [X114,X115] : (sP8('const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')),X114,X115) | ~p(s(bool,X114))) ) | $spl16),
  inference(resolution,[],[f6587,f559])).
fof(f106636,plain,(
  s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | $spl178),
  inference(cnf_transformation,[],[f106636_D])).
fof(f106636_D,plain,(
  s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f106614,plain,(
  ( ! [X49] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | ~p(s(bool,X49))) ) | $spl28),
  inference(resolution,[],[f35017,f82313])).
fof(f106613,plain,(
  ( ! [X48] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | ~p(s(bool,X48))) ) | $spl28),
  inference(resolution,[],[f35017,f82315])).
fof(f106612,plain,(
  ( ! [X47] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | ~p(s(bool,X47))) )),
  inference(resolution,[],[f35017,f2510])).
fof(f2510,plain,(
  ( ! [X111,X110] : (sP8('const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)),X110,X111) | ~p(s(bool,X111))) )),
  inference(resolution,[],[f657,f558])).
fof(f106611,plain,(
  ( ! [X46] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | ~p(s(bool,X46))) )),
  inference(resolution,[],[f35017,f2512])).
fof(f2512,plain,(
  ( ! [X114,X115] : (sP8('const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)),X114,X115) | ~p(s(bool,X114))) )),
  inference(resolution,[],[f657,f559])).
fof(f106632,plain,(
  s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | $spl174),
  inference(cnf_transformation,[],[f106632_D])).
fof(f106632_D,plain,(
  s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f106610,plain,(
  ( ! [X45] : (s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | ~p(s(bool,X45))) ) | $spl66),
  inference(resolution,[],[f35017,f74165])).
fof(f74165,plain,(
  ( ! [X111,X110] : (sP8('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)),X110,X111) | ~p(s(bool,X111))) ) | $spl66),
  inference(resolution,[],[f37512,f558])).
fof(f106609,plain,(
  ( ! [X44] : (s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | ~p(s(bool,X44))) ) | $spl66),
  inference(resolution,[],[f35017,f74167])).
fof(f74167,plain,(
  ( ! [X114,X115] : (sP8('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)),X114,X115) | ~p(s(bool,X114))) ) | $spl66),
  inference(resolution,[],[f37512,f559])).
fof(f106630,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | $spl172),
  inference(cnf_transformation,[],[f106630_D])).
fof(f106630_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f106608,plain,(
  ( ! [X43] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | ~p(s(bool,X43))) ) | $spl24),
  inference(resolution,[],[f35017,f6521])).
fof(f106607,plain,(
  ( ! [X42] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | ~p(s(bool,X42))) ) | $spl24),
  inference(resolution,[],[f35017,f6523])).
fof(f106628,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | $spl170),
  inference(cnf_transformation,[],[f106628_D])).
fof(f106628_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f106606,plain,(
  ( ! [X41] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | ~p(s(bool,X41))) ) | $spl16),
  inference(resolution,[],[f35017,f79860])).
fof(f79860,plain,(
  ( ! [X114,X113] : (sP8('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')),X113,X114) | ~p(s(bool,X114))) ) | $spl16),
  inference(resolution,[],[f1403,f558])).
fof(f106605,plain,(
  ( ! [X40] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | ~p(s(bool,X40))) ) | $spl16),
  inference(resolution,[],[f35017,f79862])).
fof(f79862,plain,(
  ( ! [X118,X117] : (sP8('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')),X117,X118) | ~p(s(bool,X117))) ) | $spl16),
  inference(resolution,[],[f1403,f559])).
fof(f106626,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | $spl168),
  inference(cnf_transformation,[],[f106626_D])).
fof(f106626_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f106604,plain,(
  ( ! [X39] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | ~p(s(bool,X39))) ) | $spl12),
  inference(resolution,[],[f35017,f4391])).
fof(f106603,plain,(
  ( ! [X38] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | ~p(s(bool,X38))) ) | $spl12),
  inference(resolution,[],[f35017,f4393])).
fof(f106624,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | $spl166),
  inference(cnf_transformation,[],[f106624_D])).
fof(f106624_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f106602,plain,(
  ( ! [X37] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | ~p(s(bool,X37))) ) | $spl28),
  inference(resolution,[],[f35017,f79987])).
fof(f106601,plain,(
  ( ! [X36] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | ~p(s(bool,X36))) ) | $spl28),
  inference(resolution,[],[f35017,f79989])).
fof(f47381,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | $spl120),
  inference(cnf_transformation,[],[f47381_D])).
fof(f47381_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f106600,plain,(
  ( ! [X35] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | ~p(s(bool,X35))) )),
  inference(resolution,[],[f35017,f719])).
fof(f106599,plain,(
  ( ! [X34] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | ~p(s(bool,X34))) )),
  inference(resolution,[],[f35017,f721])).
fof(f106622,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | $spl164),
  inference(cnf_transformation,[],[f106622_D])).
fof(f106622_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f106598,plain,(
  ( ! [X33] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | ~p(s(bool,X33))) ) | $spl74),
  inference(resolution,[],[f35017,f75275])).
fof(f75275,plain,(
  ( ! [X116,X115] : (sP8('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)),X115,X116) | ~p(s(bool,X116))) ) | $spl74),
  inference(resolution,[],[f47213,f558])).
fof(f47213,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) | $spl74),
  inference(equality_resolution,[],[f14956])).
fof(f106597,plain,(
  ( ! [X32] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | ~p(s(bool,X32))) ) | $spl74),
  inference(resolution,[],[f35017,f75277])).
fof(f75277,plain,(
  ( ! [X120,X119] : (sP8('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)),X119,X120) | ~p(s(bool,X119))) ) | $spl74),
  inference(resolution,[],[f47213,f559])).
fof(f106565,plain,(
  ( ! [X19,X18] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | p(s(bool,X18)) | ~p(s(bool,X19))) )),
  inference(resolution,[],[f35003,f711])).
fof(f35003,plain,(
  ( ! [X151,X149,X150] : (sP5(X149,X150,X151) | s(bool,X149) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f860,f533])).
fof(f106553,plain,(
  ( ! [X14,X15] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | p(s(bool,X14)) | ~p(s(bool,X15))) )),
  inference(resolution,[],[f35002,f712])).
fof(f35002,plain,(
  ( ! [X146,X147,X148] : (sP5(X147,X146,X148) | s(bool,X146) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f860,f532])).
fof(f106544,plain,(
  ( ! [X19,X18] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | p(s(bool,X18)) | ~p(s(bool,X19))) )),
  inference(resolution,[],[f35000,f708])).
fof(f35000,plain,(
  ( ! [X142,X140,X141] : (sP6(X140,X141,X142) | s(bool,X140) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f860,f529])).
fof(f106532,plain,(
  ( ! [X14,X15] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | p(s(bool,X14)) | ~p(s(bool,X15))) )),
  inference(resolution,[],[f34999,f709])).
fof(f34999,plain,(
  ( ! [X138,X139,X137] : (sP6(X138,X137,X139) | s(bool,X137) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f860,f528])).
fof(f106525,plain,(
  ( ! [X19,X20] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | p(s(bool,X19)) | p(s(bool,X20))) )),
  inference(resolution,[],[f34987,f695])).
fof(f34987,plain,(
  ( ! [X103,X101,X102] : (sP2(X101,X102,X103) | s(bool,X101) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f860,f512])).
fof(f106516,plain,(
  ( ! [X21,X22] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | ~p(s(bool,X21)) | p(s(bool,X22))) )),
  inference(resolution,[],[f34971,f681])).
fof(f34971,plain,(
  ( ! [X54,X55,X53] : (sP0(X53,X54,X55) | s(bool,X53) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f860,f495])).
fof(f106502,plain,(
  ( ! [X15,X16] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | p(s(bool,X15)) | ~p(s(bool,X16))) )),
  inference(resolution,[],[f34970,f682])).
fof(f34970,plain,(
  ( ! [X52,X50,X51] : (sP0(X51,X52,X50) | s(bool,X50) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f860,f493])).
fof(f47774,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | $spl136),
  inference(cnf_transformation,[],[f47774_D])).
fof(f47774_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f106464,plain,(
  ( ! [X55] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | ~p(s(bool,X55))) )),
  inference(resolution,[],[f33992,f4509])).
fof(f106463,plain,(
  ( ! [X54] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | ~p(s(bool,X54))) )),
  inference(resolution,[],[f33992,f4511])).
fof(f70127,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | $spl148),
  inference(cnf_transformation,[],[f70127_D])).
fof(f70127_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f106462,plain,(
  ( ! [X53] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | ~p(s(bool,X53))) ) | $spl16),
  inference(resolution,[],[f33992,f81738])).
fof(f106461,plain,(
  ( ! [X52] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | ~p(s(bool,X52))) ) | $spl16),
  inference(resolution,[],[f33992,f81740])).
fof(f106460,plain,(
  ( ! [X51] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | ~p(s(bool,X51))) )),
  inference(resolution,[],[f33992,f842])).
fof(f106459,plain,(
  ( ! [X50] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | ~p(s(bool,X50))) )),
  inference(resolution,[],[f33992,f844])).
fof(f47772,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))) | $spl134),
  inference(cnf_transformation,[],[f47772_D])).
fof(f47772_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f106458,plain,(
  ( ! [X49] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))) | ~p(s(bool,X49))) ) | $spl28),
  inference(resolution,[],[f33992,f82313])).
fof(f106457,plain,(
  ( ! [X48] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))) | ~p(s(bool,X48))) ) | $spl28),
  inference(resolution,[],[f33992,f82315])).
fof(f106456,plain,(
  ( ! [X47] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | ~p(s(bool,X47))) )),
  inference(resolution,[],[f33992,f2510])).
fof(f106455,plain,(
  ( ! [X46] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | ~p(s(bool,X46))) )),
  inference(resolution,[],[f33992,f2512])).
fof(f106468,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | $spl162),
  inference(cnf_transformation,[],[f106468_D])).
fof(f106468_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f106454,plain,(
  ( ! [X45] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | ~p(s(bool,X45))) ) | $spl66),
  inference(resolution,[],[f33992,f74165])).
fof(f106453,plain,(
  ( ! [X44] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | ~p(s(bool,X44))) ) | $spl66),
  inference(resolution,[],[f33992,f74167])).
fof(f47766,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | $spl128),
  inference(cnf_transformation,[],[f47766_D])).
fof(f47766_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f106452,plain,(
  ( ! [X43] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | ~p(s(bool,X43))) ) | $spl24),
  inference(resolution,[],[f33992,f6521])).
fof(f106451,plain,(
  ( ! [X42] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | ~p(s(bool,X42))) ) | $spl24),
  inference(resolution,[],[f33992,f6523])).
fof(f70125,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | $spl146),
  inference(cnf_transformation,[],[f70125_D])).
fof(f70125_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f106450,plain,(
  ( ! [X41] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | ~p(s(bool,X41))) ) | $spl16),
  inference(resolution,[],[f33992,f79860])).
fof(f106449,plain,(
  ( ! [X40] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | ~p(s(bool,X40))) ) | $spl16),
  inference(resolution,[],[f33992,f79862])).
fof(f47764,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | $spl126),
  inference(cnf_transformation,[],[f47764_D])).
fof(f47764_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f106448,plain,(
  ( ! [X39] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | ~p(s(bool,X39))) ) | $spl12),
  inference(resolution,[],[f33992,f4391])).
fof(f106447,plain,(
  ( ! [X38] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | ~p(s(bool,X38))) ) | $spl12),
  inference(resolution,[],[f33992,f4393])).
fof(f47762,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | $spl124),
  inference(cnf_transformation,[],[f47762_D])).
fof(f47762_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f106446,plain,(
  ( ! [X37] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | ~p(s(bool,X37))) ) | $spl28),
  inference(resolution,[],[f33992,f79987])).
fof(f106445,plain,(
  ( ! [X36] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | ~p(s(bool,X36))) ) | $spl28),
  inference(resolution,[],[f33992,f79989])).
fof(f106466,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | $spl160),
  inference(cnf_transformation,[],[f106466_D])).
fof(f106466_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f106442,plain,(
  ( ! [X33] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | ~p(s(bool,X33))) ) | $spl74),
  inference(resolution,[],[f33992,f75275])).
fof(f106441,plain,(
  ( ! [X32] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | ~p(s(bool,X32))) ) | $spl74),
  inference(resolution,[],[f33992,f75277])).
fof(f106410,plain,(
  ( ! [X21,X20] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | p(s(bool,X20)) | ~p(s(bool,X21))) )),
  inference(resolution,[],[f33978,f834])).
fof(f33978,plain,(
  ( ! [X151,X149,X150] : (sP5(X149,X150,X151) | s(bool,X149) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f737,f533])).
fof(f106398,plain,(
  ( ! [X17,X16] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | p(s(bool,X16)) | ~p(s(bool,X17))) )),
  inference(resolution,[],[f33977,f835])).
fof(f33977,plain,(
  ( ! [X146,X147,X148] : (sP5(X147,X146,X148) | s(bool,X146) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f737,f532])).
fof(f106389,plain,(
  ( ! [X21,X20] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | p(s(bool,X20)) | ~p(s(bool,X21))) )),
  inference(resolution,[],[f33975,f831])).
fof(f33975,plain,(
  ( ! [X142,X140,X141] : (sP6(X140,X141,X142) | s(bool,X140) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f737,f529])).
fof(f106377,plain,(
  ( ! [X17,X16] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | p(s(bool,X16)) | ~p(s(bool,X17))) )),
  inference(resolution,[],[f33974,f832])).
fof(f33974,plain,(
  ( ! [X138,X139,X137] : (sP6(X138,X137,X139) | s(bool,X137) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f737,f528])).
fof(f106370,plain,(
  ( ! [X21,X22] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | p(s(bool,X21)) | p(s(bool,X22))) )),
  inference(resolution,[],[f33962,f818])).
fof(f33962,plain,(
  ( ! [X103,X101,X102] : (sP2(X101,X102,X103) | s(bool,X101) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f737,f512])).
fof(f106361,plain,(
  ( ! [X24,X23] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | ~p(s(bool,X23)) | p(s(bool,X24))) )),
  inference(resolution,[],[f33946,f804])).
fof(f33946,plain,(
  ( ! [X54,X55,X53] : (sP0(X53,X54,X55) | s(bool,X53) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f737,f495])).
fof(f106347,plain,(
  ( ! [X17,X18] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | p(s(bool,X17)) | ~p(s(bool,X18))) )),
  inference(resolution,[],[f33945,f805])).
fof(f33945,plain,(
  ( ! [X52,X50,X51] : (sP0(X51,X52,X50) | s(bool,X50) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f737,f493])).
fof(f47161,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) | $spl116),
  inference(cnf_transformation,[],[f47161_D])).
fof(f47161_D,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f106176,plain,(
  ( ! [X222,X223,X224] : (s(bool,X222) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | p(s(bool,X223)) | ~p(s(bool,X224)) | s(bool,X222) != s(bool,X224)) ) | $spl24),
  inference(resolution,[],[f6539,f648])).
fof(f6539,plain,(
  ( ! [X139] : (~p(s(bool,X139)) | s(bool,X139) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11)))) ) | $spl24),
  inference(resolution,[],[f1703,f590])).
fof(f106107,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) ) | $spl24),
  inference(resolution,[],[f6539,f484])).
fof(f106105,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,X20) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X20) != s(bool,X22)) ) | $spl24),
  inference(resolution,[],[f6539,f483])).
fof(f106103,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl24),
  inference(resolution,[],[f6539,f482])).
fof(f106102,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) ) | $spl24),
  inference(resolution,[],[f6539,f480])).
fof(f106100,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl24),
  inference(resolution,[],[f6539,f478])).
fof(f46879,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) | $spl114),
  inference(cnf_transformation,[],[f46879_D])).
fof(f46879_D,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f105889,plain,(
  ( ! [X215,X213,X214] : (s(bool,X213) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | p(s(bool,X214)) | p(s(bool,X215)) | s(bool,X213) != s(bool,X215)) ) | $spl21),
  inference(resolution,[],[f6413,f648])).
fof(f6413,plain,(
  ( ! [X142] : (p(s(bool,X142)) | s(bool,X142) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK13)))) ) | $spl21),
  inference(resolution,[],[f1699,f591])).
fof(f105823,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) ) | $spl21),
  inference(resolution,[],[f6413,f484])).
fof(f105820,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,X16) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) ) | $spl21),
  inference(resolution,[],[f6413,f483])).
fof(f105819,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl21),
  inference(resolution,[],[f6413,f482])).
fof(f105818,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) ) | $spl21),
  inference(resolution,[],[f6413,f480])).
fof(f105816,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl21),
  inference(resolution,[],[f6413,f478])).
fof(f46586,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) | $spl112),
  inference(cnf_transformation,[],[f46586_D])).
fof(f46586_D,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f105593,plain,(
  ( ! [X222,X223,X224] : (s(bool,X222) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | p(s(bool,X223)) | ~p(s(bool,X224)) | s(bool,X222) != s(bool,X224)) )),
  inference(resolution,[],[f4527,f648])).
fof(f4527,plain,(
  ( ! [X136] : (~p(s(bool,X136)) | s(bool,X136) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f1573,f590])).
fof(f105524,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) )),
  inference(resolution,[],[f4527,f484])).
fof(f105522,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,X20) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X20) != s(bool,X22)) )),
  inference(resolution,[],[f4527,f483])).
fof(f105520,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) )),
  inference(resolution,[],[f4527,f482])).
fof(f105519,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) )),
  inference(resolution,[],[f4527,f480])).
fof(f105517,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) )),
  inference(resolution,[],[f4527,f478])).
fof(f46305,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) | $spl110),
  inference(cnf_transformation,[],[f46305_D])).
fof(f46305_D,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f105308,plain,(
  ( ! [X222,X223,X224] : (s(bool,X222) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | p(s(bool,X223)) | ~p(s(bool,X224)) | s(bool,X222) != s(bool,X224)) ) | $spl12),
  inference(resolution,[],[f4409,f648])).
fof(f4409,plain,(
  ( ! [X139] : (~p(s(bool,X139)) | s(bool,X139) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) | $spl12),
  inference(resolution,[],[f906,f590])).
fof(f105239,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) ) | $spl12),
  inference(resolution,[],[f4409,f484])).
fof(f105237,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,X20) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X20) != s(bool,X22)) ) | $spl12),
  inference(resolution,[],[f4409,f483])).
fof(f105235,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl12),
  inference(resolution,[],[f4409,f482])).
fof(f105234,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) ) | $spl12),
  inference(resolution,[],[f4409,f480])).
fof(f105232,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl12),
  inference(resolution,[],[f4409,f478])).
fof(f46024,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) | $spl108),
  inference(cnf_transformation,[],[f46024_D])).
fof(f46024_D,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f105023,plain,(
  ( ! [X215,X213,X214] : (s(bool,X213) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | p(s(bool,X214)) | p(s(bool,X215)) | s(bool,X213) != s(bool,X215)) )),
  inference(resolution,[],[f2987,f648])).
fof(f2987,plain,(
  ( ! [X143] : (p(s(bool,X143)) | s(bool,X143) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f788,f591])).
fof(f788,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))))),
  inference(resolution,[],[f458,f464])).
fof(f104957,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) )),
  inference(resolution,[],[f2987,f484])).
fof(f104954,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,X16) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) )),
  inference(resolution,[],[f2987,f483])).
fof(f104953,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) )),
  inference(resolution,[],[f2987,f482])).
fof(f104952,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) )),
  inference(resolution,[],[f2987,f480])).
fof(f104950,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) )),
  inference(resolution,[],[f2987,f478])).
fof(f45744,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) | $spl106),
  inference(cnf_transformation,[],[f45744_D])).
fof(f45744_D,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f104738,plain,(
  ( ! [X215,X213,X214] : (s(bool,X213) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | p(s(bool,X214)) | p(s(bool,X215)) | s(bool,X213) != s(bool,X215)) )),
  inference(resolution,[],[f2654,f648])).
fof(f2654,plain,(
  ( ! [X143] : (p(s(bool,X143)) | s(bool,X143) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f664,f591])).
fof(f664,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))))),
  inference(resolution,[],[f457,f626])).
fof(f104672,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) )),
  inference(resolution,[],[f2654,f484])).
fof(f104669,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,X16) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) )),
  inference(resolution,[],[f2654,f483])).
fof(f104668,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) )),
  inference(resolution,[],[f2654,f482])).
fof(f104667,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) )),
  inference(resolution,[],[f2654,f480])).
fof(f104665,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) )),
  inference(resolution,[],[f2654,f478])).
fof(f45464,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) | $spl104),
  inference(cnf_transformation,[],[f45464_D])).
fof(f45464_D,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f104455,plain,(
  ( ! [X222,X223,X224] : (s(bool,X222) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | p(s(bool,X223)) | ~p(s(bool,X224)) | s(bool,X222) != s(bool,X224)) )),
  inference(resolution,[],[f2528,f648])).
fof(f104386,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) )),
  inference(resolution,[],[f2528,f484])).
fof(f104384,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,X20) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X20) != s(bool,X22)) )),
  inference(resolution,[],[f2528,f483])).
fof(f104382,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) )),
  inference(resolution,[],[f2528,f482])).
fof(f104381,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) )),
  inference(resolution,[],[f2528,f480])).
fof(f104379,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) )),
  inference(resolution,[],[f2528,f478])).
fof(f45185,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) | $spl102),
  inference(cnf_transformation,[],[f45185_D])).
fof(f45185_D,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f104171,plain,(
  ( ! [X215,X213,X214] : (s(bool,X213) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | p(s(bool,X214)) | p(s(bool,X215)) | s(bool,X213) != s(bool,X215)) )),
  inference(resolution,[],[f2410,f648])).
fof(f104105,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) )),
  inference(resolution,[],[f2410,f484])).
fof(f104102,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,X16) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) )),
  inference(resolution,[],[f2410,f483])).
fof(f104101,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) )),
  inference(resolution,[],[f2410,f482])).
fof(f104100,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) )),
  inference(resolution,[],[f2410,f480])).
fof(f104098,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) )),
  inference(resolution,[],[f2410,f478])).
fof(f90497,plain,(
  $false | ($spl74 | $spl140)),
  inference(subsumption_resolution,[],[f90481,f47213])).
fof(f90481,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) | $spl140),
  inference(equality_resolution,[],[f48687])).
fof(f90496,plain,(
  ( ! [X26,X27] : (~p(s(bool,X27)) | ~p(s(bool,X26))) ) | ($spl74 | $spl140)),
  inference(subsumption_resolution,[],[f90495,f47213])).
fof(f90495,plain,(
  ( ! [X26,X27] : (~p(s(bool,X27)) | ~p(s(bool,X26)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))))) ) | $spl140),
  inference(subsumption_resolution,[],[f90480,f590])).
fof(f90480,plain,(
  ( ! [X26,X27] : (s(bool,X26) != s(bool,X27) | ~p(s(bool,X27)) | ~p(s(bool,X26)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))))) ) | $spl140),
  inference(superposition,[],[f48687,f595])).
fof(f90494,plain,(
  ( ! [X24,X25] : (~p(s(bool,X25)) | ~p(s(bool,X24))) ) | ($spl74 | $spl140)),
  inference(subsumption_resolution,[],[f90493,f47213])).
fof(f90493,plain,(
  ( ! [X24,X25] : (~p(s(bool,X25)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) | ~p(s(bool,X24))) ) | $spl140),
  inference(subsumption_resolution,[],[f90479,f590])).
fof(f90479,plain,(
  ( ! [X24,X25] : (s(bool,X24) != s(bool,X25) | ~p(s(bool,X25)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) | ~p(s(bool,X24))) ) | $spl140),
  inference(superposition,[],[f48687,f595])).
fof(f90492,plain,(
  ( ! [X14,X15] : (~p(s(bool,X15)) | ~p(s(bool,X14))) ) | ($spl74 | $spl140)),
  inference(subsumption_resolution,[],[f90491,f47213])).
fof(f90491,plain,(
  ( ! [X14,X15] : (~p(s(bool,X15)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) | ~p(s(bool,X14))) ) | $spl140),
  inference(subsumption_resolution,[],[f90474,f590])).
fof(f90474,plain,(
  ( ! [X14,X15] : (s(bool,X14) != s(bool,X15) | ~p(s(bool,X15)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) | ~p(s(bool,X14))) ) | $spl140),
  inference(superposition,[],[f48687,f590])).
fof(f90490,plain,(
  ( ! [X12,X13] : (~p(s(bool,X13)) | ~p(s(bool,X12))) ) | ($spl74 | $spl140)),
  inference(subsumption_resolution,[],[f90489,f47213])).
fof(f90489,plain,(
  ( ! [X12,X13] : (~p(s(bool,X13)) | ~p(s(bool,X12)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))))) ) | $spl140),
  inference(subsumption_resolution,[],[f90473,f590])).
fof(f90473,plain,(
  ( ! [X12,X13] : (s(bool,X12) != s(bool,X13) | ~p(s(bool,X13)) | ~p(s(bool,X12)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))))) ) | $spl140),
  inference(superposition,[],[f48687,f590])).
fof(f36568,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) | $spl98),
  inference(cnf_transformation,[],[f36568_D])).
fof(f36568_D,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f90103,plain,(
  ( ! [X215,X213,X214] : (s(bool,X213) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | p(s(bool,X214)) | p(s(bool,X215)) | s(bool,X213) != s(bool,X215)) )),
  inference(resolution,[],[f1655,f648])).
fof(f1655,plain,(
  ( ! [X143] : (p(s(bool,X143)) | s(bool,X143) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f459,f591])).
fof(f90037,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) )),
  inference(resolution,[],[f1655,f484])).
fof(f90034,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,X16) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) )),
  inference(resolution,[],[f1655,f483])).
fof(f90033,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) )),
  inference(resolution,[],[f1655,f482])).
fof(f90032,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) )),
  inference(resolution,[],[f1655,f480])).
fof(f90030,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) )),
  inference(resolution,[],[f1655,f478])).
fof(f35103,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) | $spl96),
  inference(cnf_transformation,[],[f35103_D])).
fof(f35103_D,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f89315,plain,(
  ( ! [X222,X223,X224] : (s(bool,X222) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | p(s(bool,X223)) | ~p(s(bool,X224)) | s(bool,X222) != s(bool,X224)) )),
  inference(resolution,[],[f860,f648])).
fof(f89246,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) )),
  inference(resolution,[],[f860,f484])).
fof(f89244,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,X20) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X20) != s(bool,X22)) )),
  inference(resolution,[],[f860,f483])).
fof(f89242,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) )),
  inference(resolution,[],[f860,f482])).
fof(f89241,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) )),
  inference(resolution,[],[f860,f480])).
fof(f89239,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) )),
  inference(resolution,[],[f860,f478])).
fof(f34078,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) | $spl94),
  inference(cnf_transformation,[],[f34078_D])).
fof(f34078_D,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f88645,plain,(
  ( ! [X222,X223,X224] : (s(bool,X222) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | p(s(bool,X223)) | ~p(s(bool,X224)) | s(bool,X222) != s(bool,X224)) )),
  inference(resolution,[],[f737,f648])).
fof(f88576,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) )),
  inference(resolution,[],[f737,f484])).
fof(f88574,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,X20) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X20) != s(bool,X22)) )),
  inference(resolution,[],[f737,f483])).
fof(f88572,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) )),
  inference(resolution,[],[f737,f482])).
fof(f88571,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) )),
  inference(resolution,[],[f737,f480])).
fof(f88569,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) )),
  inference(resolution,[],[f737,f478])).
fof(f77448,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl154),
  inference(cnf_transformation,[],[f77448_D])).
fof(f77448_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f85384,plain,(
  ( ! [X146,X147] : (p(s(bool,X146)) | ~p(s(bool,X147)) | s(bool,X147) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl74),
  inference(resolution,[],[f75210,f648])).
fof(f75210,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) | $spl74),
  inference(resolution,[],[f47213,f461])).
fof(f85315,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X21) != s(bool,X23)) ) | $spl74),
  inference(resolution,[],[f75210,f484])).
fof(f85313,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl74),
  inference(resolution,[],[f75210,f483])).
fof(f85311,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl74),
  inference(resolution,[],[f75210,f482])).
fof(f85310,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X6) != s(bool,X8)) ) | $spl74),
  inference(resolution,[],[f75210,f480])).
fof(f85308,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl74),
  inference(resolution,[],[f75210,f478])).
fof(f83280,plain,(
  ( ! [X147,X148] : (p(s(bool,X147)) | ~p(s(bool,X148)) | s(bool,X148) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl66),
  inference(resolution,[],[f74111,f648])).
fof(f83211,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X22) != s(bool,X24)) ) | $spl66),
  inference(resolution,[],[f74111,f484])).
fof(f83209,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl66),
  inference(resolution,[],[f74111,f483])).
fof(f83207,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl66),
  inference(resolution,[],[f74111,f482])).
fof(f83206,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X7) != s(bool,X9)) ) | $spl66),
  inference(resolution,[],[f74111,f480])).
fof(f83204,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl66),
  inference(resolution,[],[f74111,f478])).
fof(f27086,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) ) | $spl84),
  inference(cnf_transformation,[],[f27086_D])).
fof(f27086_D,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f82460,plain,(
  ( ! [X147,X148] : (p(s(bool,X147)) | ~p(s(bool,X148)) | s(bool,X148) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) ) | $spl28),
  inference(resolution,[],[f7209,f648])).
fof(f7209,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) | $spl28),
  inference(resolution,[],[f2018,f626])).
fof(f82391,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11))) | s(bool,X22) != s(bool,X24)) ) | $spl28),
  inference(resolution,[],[f7209,f484])).
fof(f82389,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) ) | $spl28),
  inference(resolution,[],[f7209,f483])).
fof(f82387,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) ) | $spl28),
  inference(resolution,[],[f7209,f482])).
fof(f82386,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11))) | s(bool,X7) != s(bool,X9)) ) | $spl28),
  inference(resolution,[],[f7209,f480])).
fof(f82384,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) ) | $spl28),
  inference(resolution,[],[f7209,f478])).
fof(f81885,plain,(
  ( ! [X147,X148] : (p(s(bool,X147)) | ~p(s(bool,X148)) | s(bool,X148) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl16),
  inference(resolution,[],[f6594,f648])).
fof(f81816,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | s(bool,X22) != s(bool,X24)) ) | $spl16),
  inference(resolution,[],[f6594,f484])).
fof(f81814,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl16),
  inference(resolution,[],[f6594,f483])).
fof(f81812,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl16),
  inference(resolution,[],[f6594,f482])).
fof(f81811,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | s(bool,X7) != s(bool,X9)) ) | $spl16),
  inference(resolution,[],[f6594,f480])).
fof(f81809,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl16),
  inference(resolution,[],[f6594,f478])).
fof(f81760,plain,(
  ( ! [X140,X141] : (p(s(bool,X140)) | p(s(bool,X141)) | s(bool,X141) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl16),
  inference(resolution,[],[f6587,f648])).
fof(f81694,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X22) != s(bool,X23)) ) | $spl16),
  inference(resolution,[],[f6587,f484])).
fof(f81691,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl16),
  inference(resolution,[],[f6587,f483])).
fof(f81690,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl16),
  inference(resolution,[],[f6587,f482])).
fof(f81689,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X7) != s(bool,X8)) ) | $spl16),
  inference(resolution,[],[f6587,f480])).
fof(f81687,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl16),
  inference(resolution,[],[f6587,f478])).
fof(f81643,plain,(
  ( ! [X146,X147] : (p(s(bool,X146)) | ~p(s(bool,X147)) | s(bool,X147) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl16),
  inference(resolution,[],[f6584,f648])).
fof(f81574,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | s(bool,X21) != s(bool,X23)) ) | $spl16),
  inference(resolution,[],[f6584,f484])).
fof(f81572,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl16),
  inference(resolution,[],[f6584,f483])).
fof(f81570,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl16),
  inference(resolution,[],[f6584,f482])).
fof(f81569,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | s(bool,X6) != s(bool,X8)) ) | $spl16),
  inference(resolution,[],[f6584,f480])).
fof(f81567,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl16),
  inference(resolution,[],[f6584,f478])).
fof(f8654,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) | $spl62),
  inference(cnf_transformation,[],[f8654_D])).
fof(f8654_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f80394,plain,(
  ( ! [X146,X147] : (p(s(bool,X146)) | ~p(s(bool,X147)) | s(bool,X147) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) | $spl12),
  inference(resolution,[],[f3242,f648])).
fof(f80325,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | s(bool,X21) != s(bool,X23)) ) | $spl12),
  inference(resolution,[],[f3242,f484])).
fof(f80323,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) | $spl12),
  inference(resolution,[],[f3242,f483])).
fof(f80321,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) | $spl12),
  inference(resolution,[],[f3242,f482])).
fof(f80320,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | s(bool,X6) != s(bool,X8)) ) | $spl12),
  inference(resolution,[],[f3242,f480])).
fof(f80318,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) | $spl12),
  inference(resolution,[],[f3242,f478])).
fof(f7326,plain,(
  ( ! [X4] : (s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X4))) ) | $spl60),
  inference(cnf_transformation,[],[f7326_D])).
fof(f7326_D,plain,(
  ( ! [X4] : (s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X4))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f80009,plain,(
  ( ! [X144,X143] : (p(s(bool,X143)) | p(s(bool,X144)) | s(bool,X144) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl28),
  inference(resolution,[],[f2018,f648])).
fof(f79943,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X25) != s(bool,X26)) ) | $spl28),
  inference(resolution,[],[f2018,f484])).
fof(f79940,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl28),
  inference(resolution,[],[f2018,f483])).
fof(f79939,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl28),
  inference(resolution,[],[f2018,f482])).
fof(f79938,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X10) != s(bool,X11)) ) | $spl28),
  inference(resolution,[],[f2018,f480])).
fof(f79936,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl28),
  inference(resolution,[],[f2018,f478])).
fof(f79882,plain,(
  ( ! [X144,X143] : (p(s(bool,X143)) | p(s(bool,X144)) | s(bool,X144) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl16),
  inference(resolution,[],[f1403,f648])).
fof(f79816,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X25) != s(bool,X26)) ) | $spl16),
  inference(resolution,[],[f1403,f484])).
fof(f79813,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl16),
  inference(resolution,[],[f1403,f483])).
fof(f79812,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl16),
  inference(resolution,[],[f1403,f482])).
fof(f79811,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X10) != s(bool,X11)) ) | $spl16),
  inference(resolution,[],[f1403,f480])).
fof(f79809,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl16),
  inference(resolution,[],[f1403,f478])).
fof(f6583,plain,(
  ( ! [X4] : (p(s(bool,X4)) | s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11)))) ) | $spl48),
  inference(cnf_transformation,[],[f6583_D])).
fof(f6583_D,plain,(
  ( ! [X4] : (p(s(bool,X4)) | s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11)))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f79755,plain,(
  ( ! [X144,X143] : (p(s(bool,X143)) | p(s(bool,X144)) | s(bool,X144) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11)))) ) | $spl24),
  inference(resolution,[],[f1703,f648])).
fof(f79689,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | s(bool,X25) != s(bool,X26)) ) | $spl24),
  inference(resolution,[],[f1703,f484])).
fof(f79686,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11)))) ) | $spl24),
  inference(resolution,[],[f1703,f483])).
fof(f79685,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11)))) ) | $spl24),
  inference(resolution,[],[f1703,f482])).
fof(f79684,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | s(bool,X10) != s(bool,X11)) ) | $spl24),
  inference(resolution,[],[f1703,f480])).
fof(f79682,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11)))) ) | $spl24),
  inference(resolution,[],[f1703,f478])).
fof(f4904,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | ~p(s(bool,X1))) ) | $spl46),
  inference(cnf_transformation,[],[f4904_D])).
fof(f4904_D,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | ~p(s(bool,X1))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f79629,plain,(
  ( ! [X146,X147] : (p(s(bool,X146)) | ~p(s(bool,X147)) | s(bool,X147) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK13)))) ) | $spl21),
  inference(resolution,[],[f1699,f648])).
fof(f79560,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | s(bool,X21) != s(bool,X23)) ) | $spl21),
  inference(resolution,[],[f1699,f484])).
fof(f79558,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK13)))) ) | $spl21),
  inference(resolution,[],[f1699,f483])).
fof(f79556,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK13)))) ) | $spl21),
  inference(resolution,[],[f1699,f482])).
fof(f79555,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | s(bool,X6) != s(bool,X8)) ) | $spl21),
  inference(resolution,[],[f1699,f480])).
fof(f79553,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK13)))) ) | $spl21),
  inference(resolution,[],[f1699,f478])).
fof(f4572,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11)))) ) | $spl44),
  inference(cnf_transformation,[],[f4572_D])).
fof(f4572_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11)))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f79297,plain,(
  ( ! [X140,X141] : (p(s(bool,X140)) | p(s(bool,X141)) | s(bool,X141) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f1573,f648])).
fof(f79231,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | s(bool,X22) != s(bool,X23)) )),
  inference(resolution,[],[f1573,f484])).
fof(f79228,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f1573,f483])).
fof(f79227,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f1573,f482])).
fof(f79226,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | s(bool,X7) != s(bool,X8)) )),
  inference(resolution,[],[f1573,f480])).
fof(f79224,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f1573,f478])).
fof(f3369,plain,(
  ( ! [X4] : (s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | p(s(bool,X4))) ) | $spl40),
  inference(cnf_transformation,[],[f3369_D])).
fof(f3369_D,plain,(
  ( ! [X4] : (s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | p(s(bool,X4))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f79179,plain,(
  ( ! [X144,X143] : (p(s(bool,X143)) | p(s(bool,X144)) | s(bool,X144) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) | $spl12),
  inference(resolution,[],[f906,f648])).
fof(f79113,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | s(bool,X25) != s(bool,X26)) ) | $spl12),
  inference(resolution,[],[f906,f484])).
fof(f79110,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) | $spl12),
  inference(resolution,[],[f906,f483])).
fof(f79109,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) | $spl12),
  inference(resolution,[],[f906,f482])).
fof(f79108,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | s(bool,X10) != s(bool,X11)) ) | $spl12),
  inference(resolution,[],[f906,f480])).
fof(f79106,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) | $spl12),
  inference(resolution,[],[f906,f478])).
fof(f3034,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) | $spl38),
  inference(cnf_transformation,[],[f3034_D])).
fof(f3034_D,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f78842,plain,(
  ( ! [X147,X148] : (p(s(bool,X147)) | ~p(s(bool,X148)) | s(bool,X148) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f788,f648])).
fof(f78773,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | s(bool,X22) != s(bool,X24)) )),
  inference(resolution,[],[f788,f484])).
fof(f78771,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f788,f483])).
fof(f78769,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f788,f482])).
fof(f78768,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | s(bool,X7) != s(bool,X9)) )),
  inference(resolution,[],[f788,f480])).
fof(f78766,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f788,f478])).
fof(f1704,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13)))) ) | $spl26),
  inference(cnf_transformation,[],[f1704_D])).
fof(f1704_D,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13)))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f78703,plain,(
  ( ! [X147,X148] : (p(s(bool,X147)) | ~p(s(bool,X148)) | s(bool,X148) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f459,f648])).
fof(f78634,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | s(bool,X22) != s(bool,X24)) )),
  inference(resolution,[],[f459,f484])).
fof(f78632,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f459,f483])).
fof(f78630,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f459,f482])).
fof(f78629,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | s(bool,X7) != s(bool,X9)) )),
  inference(resolution,[],[f459,f480])).
fof(f78627,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f459,f478])).
fof(f78616,plain,(
  $false | ($spl28 | $spl88)),
  inference(subsumption_resolution,[],[f27216,f7204])).
fof(f7204,plain,(
  s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int',sK11) | $spl28),
  inference(resolution,[],[f2018,f470])).
fof(f470,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X0) != s('type.integer.int',X1) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f382])).
fof(f382,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & s('type.integer.int',X0) != s('type.integer.int',X1))) & (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s('type.integer.int',X0) = s('type.integer.int',X1) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))),
  inference(flattening,[],[f381])).
fof(f381,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & s('type.integer.int',X0) != s('type.integer.int',X1))) & ((~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s('type.integer.int',X0) = s('type.integer.int',X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))),
  inference(nnf_transformation,[],[f272])).
fof(f272,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) <=> (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & s('type.integer.int',X0) != s('type.integer.int',X1)))),
  inference(flattening,[],[f271])).
fof(f271,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) <=> (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & ~s('type.integer.int',X0) = s('type.integer.int',X1)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X6,X7] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X6),s('type.integer.int',X7)))) <=> (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X6),s('type.integer.int',X7)))) & ~s('type.integer.int',X6) = s('type.integer.int',X7)))),
  file('thm.integer.INT_LTE_TRANS','thm.integer.INT_LT_LE')).
fof(f907,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) | $spl14),
  inference(cnf_transformation,[],[f907_D])).
fof(f907_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f78537,plain,(
  ( ! [X140,X141] : (p(s(bool,X140)) | p(s(bool,X141)) | s(bool,X141) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f458,f648])).
fof(f78471,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | s(bool,X22) != s(bool,X23)) )),
  inference(resolution,[],[f458,f484])).
fof(f78468,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f458,f483])).
fof(f78467,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f458,f482])).
fof(f78466,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | s(bool,X7) != s(bool,X8)) )),
  inference(resolution,[],[f458,f480])).
fof(f78464,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f458,f478])).
fof(f906,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | $spl12),
  inference(cnf_transformation,[],[f906_D])).
fof(f906_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f904,plain,(
  s('type.integer.int',sK12) = s('type.integer.int',sK13) | $spl10),
  inference(cnf_transformation,[],[f904_D])).
fof(f904_D,plain,(
  s('type.integer.int',sK12) = s('type.integer.int',sK13) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f78460,plain,(
  s('type.integer.int',sK12) = s('type.integer.int',sK13) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))),
  inference(resolution,[],[f458,f466])).
fof(f466,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X0) = s('type.integer.int',X1) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f380])).
fof(f78459,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | s('type.integer.int',sK12) = s('type.integer.int',sK13)),
  inference(resolution,[],[f458,f471])).
fof(f471,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s('type.integer.int',X0) = s('type.integer.int',X1) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f382])).
fof(f78454,plain,(
  $false | ($spl21 | $spl88 | $spl152)),
  inference(subsumption_resolution,[],[f78453,f476])).
fof(f476,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X0))))) )),
  inference(cnf_transformation,[],[f277])).
fof(f277,plain,(
  ! [X0] : p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X0))))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X6] : p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X6),s('type.integer.int',X6))))),
  file('thm.integer.INT_LTE_TRANS','thm.integer.INT_LE_REFL')).
fof(f78453,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))) | ($spl21 | $spl88 | $spl152)),
  inference(forward_demodulation,[],[f47894,f74229])).
fof(f47894,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) | ($spl21 | $spl88)),
  inference(backward_demodulation,[],[f27216,f1699])).
fof(f78452,plain,(
  $false | ($spl23 | $spl88 | $spl152)),
  inference(subsumption_resolution,[],[f47895,f74229])).
fof(f47895,plain,(
  s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int',sK13) | ($spl23 | $spl88)),
  inference(backward_demodulation,[],[f27216,f1700])).
fof(f1700,plain,(
  s('type.integer.int',sK11) != s('type.integer.int',sK13) | $spl23),
  inference(cnf_transformation,[],[f1700_D])).
fof(f1700_D,plain,(
  s('type.integer.int',sK11) != s('type.integer.int',sK13) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f78451,plain,(
  $false | ($spl24 | $spl88 | $spl152)),
  inference(subsumption_resolution,[],[f78450,f631])).
fof(f631,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X0))))) )),
  inference(cnf_transformation,[],[f320])).
fof(f320,plain,(
  ! [X0] : ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X0))))),
  inference(flattening,[],[f319])).
fof(f319,plain,(
  ! [X0] : ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X0))))),
  inference(rectify,[],[f54])).
fof(f54,axiom,(
  ! [X6] : ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X6),s('type.integer.int',X6))))),
  file('thm.integer.INT_LTE_TRANS','thm.integer.INT_LT_REFL')).
fof(f78450,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))) | ($spl24 | $spl88 | $spl152)),
  inference(forward_demodulation,[],[f47896,f74229])).
fof(f47896,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) | ($spl24 | $spl88)),
  inference(backward_demodulation,[],[f27216,f1703])).
fof(f47780,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) | $spl139),
  inference(cnf_transformation,[],[f47780_D])).
fof(f47780_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) <=> ~$spl139),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl139])])).
fof(f69921,plain,(
  ( ! [X12,X13] : (~p(s(bool,X13)) | ~p(s(bool,X12)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11))))) ) | $spl78),
  inference(subsumption_resolution,[],[f69905,f590])).
fof(f69905,plain,(
  ( ! [X12,X13] : (s(bool,X12) != s(bool,X13) | ~p(s(bool,X13)) | ~p(s(bool,X12)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11))))) ) | $spl78),
  inference(superposition,[],[f18959,f590])).
fof(f69922,plain,(
  ( ! [X14,X15] : (~p(s(bool,X15)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) | ~p(s(bool,X14))) ) | $spl78),
  inference(subsumption_resolution,[],[f69906,f590])).
fof(f69906,plain,(
  ( ! [X14,X15] : (s(bool,X14) != s(bool,X15) | ~p(s(bool,X15)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) | ~p(s(bool,X14))) ) | $spl78),
  inference(superposition,[],[f18959,f590])).
fof(f69923,plain,(
  ( ! [X24,X25] : (~p(s(bool,X25)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) | ~p(s(bool,X24))) ) | $spl78),
  inference(subsumption_resolution,[],[f69911,f590])).
fof(f69911,plain,(
  ( ! [X24,X25] : (s(bool,X24) != s(bool,X25) | ~p(s(bool,X25)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) | ~p(s(bool,X24))) ) | $spl78),
  inference(superposition,[],[f18959,f595])).
fof(f69924,plain,(
  ( ! [X26,X27] : (~p(s(bool,X27)) | ~p(s(bool,X26)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11))))) ) | $spl78),
  inference(subsumption_resolution,[],[f69912,f590])).
fof(f69912,plain,(
  ( ! [X26,X27] : (s(bool,X26) != s(bool,X27) | ~p(s(bool,X27)) | ~p(s(bool,X26)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11))))) ) | $spl78),
  inference(superposition,[],[f18959,f595])).
fof(f27213,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) | $spl86),
  inference(cnf_transformation,[],[f27213_D])).
fof(f27213_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f78435,plain,(
  ( ! [X17,X16] : (p(s(bool,X17)) | p(s(bool,X16)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))))) ) | $spl82),
  inference(subsumption_resolution,[],[f69966,f591])).
fof(f69966,plain,(
  ( ! [X17,X16] : (s(bool,X16) != s(bool,X17) | p(s(bool,X17)) | p(s(bool,X16)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))))) ) | $spl82),
  inference(superposition,[],[f23306,f591])).
fof(f78434,plain,(
  ( ! [X19,X18] : (p(s(bool,X19)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,X18))) ) | $spl82),
  inference(subsumption_resolution,[],[f69967,f591])).
fof(f69967,plain,(
  ( ! [X19,X18] : (s(bool,X18) != s(bool,X19) | p(s(bool,X19)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,X18))) ) | $spl82),
  inference(superposition,[],[f23306,f591])).
fof(f78433,plain,(
  ( ! [X21,X20] : (p(s(bool,X21)) | p(s(bool,X20)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))))) ) | $spl82),
  inference(subsumption_resolution,[],[f69968,f591])).
fof(f69968,plain,(
  ( ! [X21,X20] : (s(bool,X20) != s(bool,X21) | p(s(bool,X21)) | p(s(bool,X20)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))))) ) | $spl82),
  inference(superposition,[],[f23306,f592])).
fof(f78432,plain,(
  ( ! [X23,X22] : (p(s(bool,X23)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,X22))) ) | $spl82),
  inference(subsumption_resolution,[],[f69969,f591])).
fof(f69969,plain,(
  ( ! [X23,X22] : (s(bool,X22) != s(bool,X23) | p(s(bool,X23)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,X22))) ) | $spl82),
  inference(superposition,[],[f23306,f592])).
fof(f78227,plain,(
  $false | ($spl12 | $spl17 | $spl152)),
  inference(subsumption_resolution,[],[f77917,f6930])).
fof(f77917,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) | ($spl12 | $spl152)),
  inference(backward_demodulation,[],[f74229,f3242])).
fof(f78226,plain,(
  $false | ($spl24 | $spl28 | $spl152)),
  inference(subsumption_resolution,[],[f77908,f7209])).
fof(f77908,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) | ($spl24 | $spl152)),
  inference(backward_demodulation,[],[f74229,f1703])).
fof(f78225,plain,(
  $false | ($spl21 | $spl28 | $spl152)),
  inference(subsumption_resolution,[],[f77906,f7202])).
fof(f77906,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) | ($spl21 | $spl152)),
  inference(backward_demodulation,[],[f74229,f1699])).
fof(f78198,plain,(
  $false | ($spl28 | $spl152)),
  inference(subsumption_resolution,[],[f77861,f7199])).
fof(f7199,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) | $spl28),
  inference(resolution,[],[f2018,f461])).
fof(f77861,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) | $spl152),
  inference(backward_demodulation,[],[f74229,f1573])).
fof(f78197,plain,(
  $false | ($spl12 | $spl17 | $spl152)),
  inference(subsumption_resolution,[],[f77860,f1402])).
fof(f77860,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) | ($spl12 | $spl152)),
  inference(backward_demodulation,[],[f74229,f906])).
fof(f78194,plain,(
  $false | ($spl56 | $spl152)),
  inference(subsumption_resolution,[],[f77832,f7060])).
fof(f77832,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) | $spl152),
  inference(backward_demodulation,[],[f74229,f788])).
fof(f78193,plain,(
  $false | ($spl28 | $spl152)),
  inference(subsumption_resolution,[],[f77831,f2018])).
fof(f77831,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) | $spl152),
  inference(backward_demodulation,[],[f74229,f459])).
fof(f78192,plain,(
  $false | ($spl53 | $spl152)),
  inference(subsumption_resolution,[],[f77830,f7056])).
fof(f77830,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) | $spl152),
  inference(backward_demodulation,[],[f74229,f458])).
fof(f74229,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13) | $spl152),
  inference(cnf_transformation,[],[f74229_D])).
fof(f74229_D,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f47223,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) | $spl118),
  inference(cnf_transformation,[],[f47223_D])).
fof(f47223_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f74109,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13) | $spl66),
  inference(resolution,[],[f37512,f471])).
fof(f74110,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) | $spl66),
  inference(resolution,[],[f37512,f466])).
fof(f77450,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) | $spl157),
  inference(cnf_transformation,[],[f77450_D])).
fof(f77450_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) <=> ~$spl157),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl157])])).
fof(f75605,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) | $spl66),
  inference(resolution,[],[f74111,f466])).
fof(f75606,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) | $spl66),
  inference(resolution,[],[f74111,f471])).
fof(f75610,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13) | $spl66),
  inference(resolution,[],[f74111,f629])).
fof(f629,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s('type.integer.int',X0) = s('type.integer.int',X1)) )),
  inference(cnf_transformation,[],[f317])).
fof(f317,plain,(
  ! [X0,X1] : (s('type.integer.int',X0) = s('type.integer.int',X1) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X6,X7] : (s('type.integer.int',X6) = s('type.integer.int',X7) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X6),s('type.integer.int',X7)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X7),s('type.integer.int',X6)))))),
  file('thm.integer.INT_LTE_TRANS','thm.integer.INT_LT_TOTAL')).
fof(f75611,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13) | $spl66),
  inference(resolution,[],[f74111,f629])).
fof(f77788,plain,(
  ( ! [X146,X147] : (p(s(bool,X146)) | ~p(s(bool,X147)) | s(bool,X147) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | $spl101),
  inference(resolution,[],[f37521,f648])).
fof(f77719,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | s(bool,X21) != s(bool,X23)) ) | $spl101),
  inference(resolution,[],[f37521,f484])).
fof(f77717,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | $spl101),
  inference(resolution,[],[f37521,f483])).
fof(f77715,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | $spl101),
  inference(resolution,[],[f37521,f482])).
fof(f77714,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | s(bool,X6) != s(bool,X8)) ) | $spl101),
  inference(resolution,[],[f37521,f480])).
fof(f77712,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | $spl101),
  inference(resolution,[],[f37521,f478])).
fof(f10363,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | p(s(bool,X1))) ) | $spl66),
  inference(cnf_transformation,[],[f10363_D])).
fof(f10363_D,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | p(s(bool,X1))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f77660,plain,(
  ( ! [X140,X141] : (p(s(bool,X140)) | p(s(bool,X141)) | s(bool,X141) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | $spl100),
  inference(resolution,[],[f37522,f648])).
fof(f77594,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | s(bool,X22) != s(bool,X23)) ) | $spl100),
  inference(resolution,[],[f37522,f484])).
fof(f77591,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | $spl100),
  inference(resolution,[],[f37522,f483])).
fof(f77590,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | $spl100),
  inference(resolution,[],[f37522,f482])).
fof(f77589,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | s(bool,X7) != s(bool,X8)) ) | $spl100),
  inference(resolution,[],[f37522,f480])).
fof(f77587,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | $spl100),
  inference(resolution,[],[f37522,f478])).
fof(f77583,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) | $spl100),
  inference(resolution,[],[f37522,f466])).
fof(f77582,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13) | $spl100),
  inference(resolution,[],[f37522,f471])).
fof(f14956,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | p(s(bool,X6))) ) | $spl74),
  inference(cnf_transformation,[],[f14956_D])).
fof(f14956_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | p(s(bool,X6))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f77540,plain,(
  ( ! [X146,X145] : (p(s(bool,X145)) | p(s(bool,X146)) | s(bool,X146) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | $spl118),
  inference(resolution,[],[f47223,f648])).
fof(f77474,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | s(bool,X27) != s(bool,X28)) ) | $spl118),
  inference(resolution,[],[f47223,f484])).
fof(f77471,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X19) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | $spl118),
  inference(resolution,[],[f47223,f483])).
fof(f77470,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | $spl118),
  inference(resolution,[],[f47223,f482])).
fof(f77469,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | s(bool,X12) != s(bool,X13)) ) | $spl118),
  inference(resolution,[],[f47223,f480])).
fof(f77467,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | $spl118),
  inference(resolution,[],[f47223,f478])).
fof(f77407,plain,(
  ( ! [X146,X147] : (p(s(bool,X146)) | ~p(s(bool,X147)) | s(bool,X147) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl74),
  inference(resolution,[],[f75210,f648])).
fof(f77338,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X21) != s(bool,X23)) ) | $spl74),
  inference(resolution,[],[f75210,f484])).
fof(f77336,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl74),
  inference(resolution,[],[f75210,f483])).
fof(f77334,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl74),
  inference(resolution,[],[f75210,f482])).
fof(f77333,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X6) != s(bool,X8)) ) | $spl74),
  inference(resolution,[],[f75210,f480])).
fof(f77331,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl74),
  inference(resolution,[],[f75210,f478])).
fof(f75688,plain,(
  ( ! [X147,X148] : (p(s(bool,X147)) | ~p(s(bool,X148)) | s(bool,X148) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl66),
  inference(resolution,[],[f74111,f648])).
fof(f75619,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X22) != s(bool,X24)) ) | $spl66),
  inference(resolution,[],[f74111,f484])).
fof(f75617,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl66),
  inference(resolution,[],[f74111,f483])).
fof(f75615,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl66),
  inference(resolution,[],[f74111,f482])).
fof(f75614,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X7) != s(bool,X9)) ) | $spl66),
  inference(resolution,[],[f74111,f480])).
fof(f75612,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl66),
  inference(resolution,[],[f74111,f478])).
fof(f74786,plain,(
  ( ! [X146,X145] : (p(s(bool,X145)) | p(s(bool,X146)) | s(bool,X146) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl56),
  inference(resolution,[],[f7060,f648])).
fof(f74720,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | s(bool,X27) != s(bool,X28)) ) | $spl56),
  inference(resolution,[],[f7060,f484])).
fof(f74717,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X19) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl56),
  inference(resolution,[],[f7060,f483])).
fof(f74716,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl56),
  inference(resolution,[],[f7060,f482])).
fof(f74715,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | s(bool,X12) != s(bool,X13)) ) | $spl56),
  inference(resolution,[],[f7060,f480])).
fof(f74713,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl56),
  inference(resolution,[],[f7060,f478])).
fof(f74656,plain,(
  ( ! [X146,X147] : (p(s(bool,X146)) | ~p(s(bool,X147)) | s(bool,X147) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl53),
  inference(resolution,[],[f7056,f648])).
fof(f74587,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X21) != s(bool,X23)) ) | $spl53),
  inference(resolution,[],[f7056,f484])).
fof(f74585,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl53),
  inference(resolution,[],[f7056,f483])).
fof(f74583,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl53),
  inference(resolution,[],[f7056,f482])).
fof(f74582,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X6) != s(bool,X8)) ) | $spl53),
  inference(resolution,[],[f7056,f480])).
fof(f74580,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl53),
  inference(resolution,[],[f7056,f478])).
fof(f74308,plain,(
  ( ! [X140,X141] : (p(s(bool,X140)) | p(s(bool,X141)) | s(bool,X141) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl17),
  inference(resolution,[],[f6930,f648])).
fof(f74242,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | s(bool,X22) != s(bool,X23)) ) | $spl17),
  inference(resolution,[],[f6930,f484])).
fof(f74239,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl17),
  inference(resolution,[],[f6930,f483])).
fof(f74238,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl17),
  inference(resolution,[],[f6930,f482])).
fof(f74237,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | s(bool,X7) != s(bool,X8)) ) | $spl17),
  inference(resolution,[],[f6930,f480])).
fof(f74235,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl17),
  inference(resolution,[],[f6930,f478])).
fof(f73816,plain,(
  ( ! [X146,X147] : (p(s(bool,X146)) | ~p(s(bool,X147)) | s(bool,X147) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) | $spl12),
  inference(resolution,[],[f3242,f648])).
fof(f73747,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | s(bool,X21) != s(bool,X23)) ) | $spl12),
  inference(resolution,[],[f3242,f484])).
fof(f73745,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) | $spl12),
  inference(resolution,[],[f3242,f483])).
fof(f73743,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) | $spl12),
  inference(resolution,[],[f3242,f482])).
fof(f73742,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | s(bool,X6) != s(bool,X8)) ) | $spl12),
  inference(resolution,[],[f3242,f480])).
fof(f73740,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) | $spl12),
  inference(resolution,[],[f3242,f478])).
fof(f72563,plain,(
  ( ! [X147,X148] : (p(s(bool,X147)) | ~p(s(bool,X148)) | s(bool,X148) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl17),
  inference(resolution,[],[f1402,f648])).
fof(f72494,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X22) != s(bool,X24)) ) | $spl17),
  inference(resolution,[],[f1402,f484])).
fof(f72492,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl17),
  inference(resolution,[],[f1402,f483])).
fof(f72490,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl17),
  inference(resolution,[],[f1402,f482])).
fof(f72489,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X7) != s(bool,X9)) ) | $spl17),
  inference(resolution,[],[f1402,f480])).
fof(f72487,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl17),
  inference(resolution,[],[f1402,f478])).
fof(f72434,plain,(
  ( ! [X144,X143] : (p(s(bool,X143)) | p(s(bool,X144)) | s(bool,X144) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) | $spl12),
  inference(resolution,[],[f906,f648])).
fof(f72368,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | s(bool,X25) != s(bool,X26)) ) | $spl12),
  inference(resolution,[],[f906,f484])).
fof(f72365,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) | $spl12),
  inference(resolution,[],[f906,f483])).
fof(f72364,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) | $spl12),
  inference(resolution,[],[f906,f482])).
fof(f72363,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | s(bool,X10) != s(bool,X11)) ) | $spl12),
  inference(resolution,[],[f906,f480])).
fof(f72361,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) | $spl12),
  inference(resolution,[],[f906,f478])).
fof(f72087,plain,(
  ( ! [X147,X148] : (p(s(bool,X147)) | ~p(s(bool,X148)) | s(bool,X148) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f788,f648])).
fof(f72018,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | s(bool,X22) != s(bool,X24)) )),
  inference(resolution,[],[f788,f484])).
fof(f72016,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f788,f483])).
fof(f72014,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f788,f482])).
fof(f72013,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | s(bool,X7) != s(bool,X9)) )),
  inference(resolution,[],[f788,f480])).
fof(f72011,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f788,f478])).
fof(f2697,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) ) | $spl36),
  inference(cnf_transformation,[],[f2697_D])).
fof(f2697_D,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f71745,plain,(
  ( ! [X147,X148] : (p(s(bool,X147)) | ~p(s(bool,X148)) | s(bool,X148) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f664,f648])).
fof(f71676,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | s(bool,X22) != s(bool,X24)) )),
  inference(resolution,[],[f664,f484])).
fof(f71674,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f664,f483])).
fof(f71672,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f664,f482])).
fof(f71671,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | s(bool,X7) != s(bool,X9)) )),
  inference(resolution,[],[f664,f480])).
fof(f71669,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f664,f478])).
fof(f2571,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) ) | $spl34),
  inference(cnf_transformation,[],[f2571_D])).
fof(f2571_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f71618,plain,(
  ( ! [X140,X141] : (p(s(bool,X140)) | p(s(bool,X141)) | s(bool,X141) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f657,f648])).
fof(f71552,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | s(bool,X22) != s(bool,X23)) )),
  inference(resolution,[],[f657,f484])).
fof(f71549,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f657,f483])).
fof(f71548,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f657,f482])).
fof(f71547,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | s(bool,X7) != s(bool,X8)) )),
  inference(resolution,[],[f657,f480])).
fof(f71545,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f657,f478])).
fof(f2453,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) ) | $spl32),
  inference(cnf_transformation,[],[f2453_D])).
fof(f2453_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f71499,plain,(
  ( ! [X146,X147] : (p(s(bool,X146)) | ~p(s(bool,X147)) | s(bool,X147) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f654,f648])).
fof(f71430,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | s(bool,X21) != s(bool,X23)) )),
  inference(resolution,[],[f654,f484])).
fof(f71428,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f654,f483])).
fof(f71426,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f654,f482])).
fof(f71425,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | s(bool,X6) != s(bool,X8)) )),
  inference(resolution,[],[f654,f480])).
fof(f71423,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f654,f478])).
fof(f71369,plain,(
  ( ! [X140,X141] : (p(s(bool,X140)) | p(s(bool,X141)) | s(bool,X141) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f458,f648])).
fof(f71303,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | s(bool,X22) != s(bool,X23)) )),
  inference(resolution,[],[f458,f484])).
fof(f71300,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f458,f483])).
fof(f71299,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f458,f482])).
fof(f71298,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | s(bool,X7) != s(bool,X8)) )),
  inference(resolution,[],[f458,f480])).
fof(f71296,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f458,f478])).
fof(f781,plain,(
  ( ! [X4] : (p(s(bool,X4)) | s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) ) | $spl0),
  inference(cnf_transformation,[],[f781_D])).
fof(f781_D,plain,(
  ( ! [X4] : (p(s(bool,X4)) | s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f71249,plain,(
  ( ! [X144,X143] : (p(s(bool,X143)) | p(s(bool,X144)) | s(bool,X144) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f457,f648])).
fof(f71183,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | s(bool,X25) != s(bool,X26)) )),
  inference(resolution,[],[f457,f484])).
fof(f71180,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f457,f483])).
fof(f71179,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f457,f482])).
fof(f71178,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | s(bool,X10) != s(bool,X11)) )),
  inference(resolution,[],[f457,f480])).
fof(f71176,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f457,f478])).
fof(f71088,plain,(
  ( ! [X147,X148] : (p(s(bool,X147)) | ~p(s(bool,X148)) | s(bool,X148) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl54),
  inference(resolution,[],[f70325,f648])).
fof(f70325,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) | $spl54),
  inference(backward_demodulation,[],[f7058,f788])).
fof(f71019,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X22) != s(bool,X24)) ) | $spl54),
  inference(resolution,[],[f70325,f484])).
fof(f71017,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl54),
  inference(resolution,[],[f70325,f483])).
fof(f71015,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl54),
  inference(resolution,[],[f70325,f482])).
fof(f71014,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X7) != s(bool,X9)) ) | $spl54),
  inference(resolution,[],[f70325,f480])).
fof(f71012,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl54),
  inference(resolution,[],[f70325,f478])).
fof(f70960,plain,(
  ( ! [X140,X141] : (p(s(bool,X140)) | p(s(bool,X141)) | s(bool,X141) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | $spl54),
  inference(resolution,[],[f70274,f648])).
fof(f70274,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) | $spl54),
  inference(backward_demodulation,[],[f7058,f458])).
fof(f70894,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | s(bool,X22) != s(bool,X23)) ) | $spl54),
  inference(resolution,[],[f70274,f484])).
fof(f70891,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | $spl54),
  inference(resolution,[],[f70274,f483])).
fof(f70890,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | $spl54),
  inference(resolution,[],[f70274,f482])).
fof(f70889,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | s(bool,X7) != s(bool,X8)) ) | $spl54),
  inference(resolution,[],[f70274,f480])).
fof(f70887,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | $spl54),
  inference(resolution,[],[f70274,f478])).
fof(f70227,plain,(
  ( ! [X146,X145] : (p(s(bool,X145)) | p(s(bool,X146)) | s(bool,X146) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | ($spl12 | $spl54)),
  inference(resolution,[],[f9846,f648])).
fof(f9846,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) | ($spl12 | $spl54)),
  inference(backward_demodulation,[],[f7058,f906])).
fof(f70161,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | s(bool,X27) != s(bool,X28)) ) | ($spl12 | $spl54)),
  inference(resolution,[],[f9846,f484])).
fof(f70158,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X19) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | ($spl12 | $spl54)),
  inference(resolution,[],[f9846,f483])).
fof(f70157,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | ($spl12 | $spl54)),
  inference(resolution,[],[f9846,f482])).
fof(f70156,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | s(bool,X12) != s(bool,X13)) ) | ($spl12 | $spl54)),
  inference(resolution,[],[f9846,f480])).
fof(f70154,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | ($spl12 | $spl54)),
  inference(resolution,[],[f9846,f478])).
fof(f70123,plain,(
  ( ! [X47] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | ~p(s(bool,X47))) )),
  inference(resolution,[],[f33992,f4509])).
fof(f70122,plain,(
  ( ! [X46] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | ~p(s(bool,X46))) )),
  inference(resolution,[],[f33992,f4511])).
fof(f70121,plain,(
  ( ! [X45] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | ~p(s(bool,X45))) ) | $spl16),
  inference(resolution,[],[f33992,f52139])).
fof(f52139,plain,(
  ( ! [X111,X110] : (sP8('const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')),X110,X111) | ~p(s(bool,X111))) ) | $spl16),
  inference(resolution,[],[f6587,f558])).
fof(f70120,plain,(
  ( ! [X44] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | ~p(s(bool,X44))) ) | $spl16),
  inference(resolution,[],[f33992,f52141])).
fof(f52141,plain,(
  ( ! [X114,X115] : (sP8('const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')),X114,X115) | ~p(s(bool,X114))) ) | $spl16),
  inference(resolution,[],[f6587,f559])).
fof(f70119,plain,(
  ( ! [X43] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | ~p(s(bool,X43))) )),
  inference(resolution,[],[f33992,f842])).
fof(f70118,plain,(
  ( ! [X42] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | ~p(s(bool,X42))) )),
  inference(resolution,[],[f33992,f844])).
fof(f70117,plain,(
  ( ! [X41] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))) | ~p(s(bool,X41))) ) | $spl28),
  inference(resolution,[],[f33992,f54976])).
fof(f54976,plain,(
  ( ! [X111,X110] : (sP8('const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')),X110,X111) | ~p(s(bool,X111))) ) | $spl28),
  inference(resolution,[],[f7202,f558])).
fof(f70116,plain,(
  ( ! [X40] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))) | ~p(s(bool,X40))) ) | $spl28),
  inference(resolution,[],[f33992,f54978])).
fof(f54978,plain,(
  ( ! [X114,X115] : (sP8('const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')),X114,X115) | ~p(s(bool,X114))) ) | $spl28),
  inference(resolution,[],[f7202,f559])).
fof(f70115,plain,(
  ( ! [X39] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | ~p(s(bool,X39))) )),
  inference(resolution,[],[f33992,f2510])).
fof(f70114,plain,(
  ( ! [X38] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | ~p(s(bool,X38))) )),
  inference(resolution,[],[f33992,f2512])).
fof(f70113,plain,(
  ( ! [X37] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | ~p(s(bool,X37))) ) | $spl24),
  inference(resolution,[],[f33992,f6521])).
fof(f70112,plain,(
  ( ! [X36] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | ~p(s(bool,X36))) ) | $spl24),
  inference(resolution,[],[f33992,f6523])).
fof(f70111,plain,(
  ( ! [X35] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | ~p(s(bool,X35))) ) | $spl16),
  inference(resolution,[],[f33992,f50248])).
fof(f50248,plain,(
  ( ! [X114,X113] : (sP8('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')),X113,X114) | ~p(s(bool,X114))) ) | $spl16),
  inference(resolution,[],[f1403,f558])).
fof(f70110,plain,(
  ( ! [X34] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | ~p(s(bool,X34))) ) | $spl16),
  inference(resolution,[],[f33992,f50250])).
fof(f50250,plain,(
  ( ! [X118,X117] : (sP8('const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')),X117,X118) | ~p(s(bool,X117))) ) | $spl16),
  inference(resolution,[],[f1403,f559])).
fof(f70109,plain,(
  ( ! [X33] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | ~p(s(bool,X33))) ) | $spl12),
  inference(resolution,[],[f33992,f4391])).
fof(f70108,plain,(
  ( ! [X32] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | ~p(s(bool,X32))) ) | $spl12),
  inference(resolution,[],[f33992,f4393])).
fof(f70107,plain,(
  ( ! [X31] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | ~p(s(bool,X31))) ) | $spl28),
  inference(resolution,[],[f33992,f54729])).
fof(f54729,plain,(
  ( ! [X114,X113] : (sP8('const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')),X113,X114) | ~p(s(bool,X114))) ) | $spl28),
  inference(resolution,[],[f2018,f558])).
fof(f70106,plain,(
  ( ! [X30] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | ~p(s(bool,X30))) ) | $spl28),
  inference(resolution,[],[f33992,f54731])).
fof(f54731,plain,(
  ( ! [X118,X117] : (sP8('const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')),X117,X118) | ~p(s(bool,X117))) ) | $spl28),
  inference(resolution,[],[f2018,f559])).
fof(f70076,plain,(
  ( ! [X19,X18] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | p(s(bool,X18)) | ~p(s(bool,X19))) )),
  inference(resolution,[],[f33978,f834])).
fof(f70064,plain,(
  ( ! [X14,X15] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | p(s(bool,X14)) | ~p(s(bool,X15))) )),
  inference(resolution,[],[f33977,f835])).
fof(f70055,plain,(
  ( ! [X19,X18] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | p(s(bool,X18)) | ~p(s(bool,X19))) )),
  inference(resolution,[],[f33975,f831])).
fof(f70043,plain,(
  ( ! [X14,X15] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | p(s(bool,X14)) | ~p(s(bool,X15))) )),
  inference(resolution,[],[f33974,f832])).
fof(f70036,plain,(
  ( ! [X19,X20] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | p(s(bool,X19)) | p(s(bool,X20))) )),
  inference(resolution,[],[f33962,f818])).
fof(f70027,plain,(
  ( ! [X21,X22] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | ~p(s(bool,X21)) | p(s(bool,X22))) )),
  inference(resolution,[],[f33946,f804])).
fof(f70013,plain,(
  ( ! [X15,X16] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | p(s(bool,X15)) | ~p(s(bool,X16))) )),
  inference(resolution,[],[f33945,f805])).
fof(f69786,plain,(
  ( ! [X222,X223,X224] : (s(bool,X222) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | p(s(bool,X223)) | ~p(s(bool,X224)) | s(bool,X222) != s(bool,X224)) ) | $spl24),
  inference(resolution,[],[f6539,f648])).
fof(f69717,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) ) | $spl24),
  inference(resolution,[],[f6539,f484])).
fof(f69715,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,X20) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X20) != s(bool,X22)) ) | $spl24),
  inference(resolution,[],[f6539,f483])).
fof(f69713,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl24),
  inference(resolution,[],[f6539,f482])).
fof(f69712,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) ) | $spl24),
  inference(resolution,[],[f6539,f480])).
fof(f69710,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl24),
  inference(resolution,[],[f6539,f478])).
fof(f69501,plain,(
  ( ! [X215,X213,X214] : (s(bool,X213) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | p(s(bool,X214)) | p(s(bool,X215)) | s(bool,X213) != s(bool,X215)) ) | $spl21),
  inference(resolution,[],[f6413,f648])).
fof(f69435,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) ) | $spl21),
  inference(resolution,[],[f6413,f484])).
fof(f69432,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,X16) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) ) | $spl21),
  inference(resolution,[],[f6413,f483])).
fof(f69431,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl21),
  inference(resolution,[],[f6413,f482])).
fof(f69430,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) ) | $spl21),
  inference(resolution,[],[f6413,f480])).
fof(f69428,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl21),
  inference(resolution,[],[f6413,f478])).
fof(f69207,plain,(
  ( ! [X222,X223,X224] : (s(bool,X222) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | p(s(bool,X223)) | ~p(s(bool,X224)) | s(bool,X222) != s(bool,X224)) )),
  inference(resolution,[],[f4527,f648])).
fof(f69138,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) )),
  inference(resolution,[],[f4527,f484])).
fof(f69136,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,X20) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X20) != s(bool,X22)) )),
  inference(resolution,[],[f4527,f483])).
fof(f69134,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) )),
  inference(resolution,[],[f4527,f482])).
fof(f69133,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) )),
  inference(resolution,[],[f4527,f480])).
fof(f69131,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) )),
  inference(resolution,[],[f4527,f478])).
fof(f68923,plain,(
  ( ! [X215,X213,X214] : (s(bool,X213) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | p(s(bool,X214)) | p(s(bool,X215)) | s(bool,X213) != s(bool,X215)) )),
  inference(resolution,[],[f2654,f648])).
fof(f68857,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) )),
  inference(resolution,[],[f2654,f484])).
fof(f68854,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,X16) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) )),
  inference(resolution,[],[f2654,f483])).
fof(f68853,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) )),
  inference(resolution,[],[f2654,f482])).
fof(f68852,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) )),
  inference(resolution,[],[f2654,f480])).
fof(f68850,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) )),
  inference(resolution,[],[f2654,f478])).
fof(f68641,plain,(
  ( ! [X222,X223,X224] : (s(bool,X222) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | p(s(bool,X223)) | ~p(s(bool,X224)) | s(bool,X222) != s(bool,X224)) )),
  inference(resolution,[],[f2528,f648])).
fof(f68572,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) )),
  inference(resolution,[],[f2528,f484])).
fof(f68570,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,X20) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X20) != s(bool,X22)) )),
  inference(resolution,[],[f2528,f483])).
fof(f68568,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) )),
  inference(resolution,[],[f2528,f482])).
fof(f68567,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) )),
  inference(resolution,[],[f2528,f480])).
fof(f68565,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) )),
  inference(resolution,[],[f2528,f478])).
fof(f68358,plain,(
  ( ! [X215,X213,X214] : (s(bool,X213) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | p(s(bool,X214)) | p(s(bool,X215)) | s(bool,X213) != s(bool,X215)) )),
  inference(resolution,[],[f2410,f648])).
fof(f68292,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) )),
  inference(resolution,[],[f2410,f484])).
fof(f68289,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,X16) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) )),
  inference(resolution,[],[f2410,f483])).
fof(f68288,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) )),
  inference(resolution,[],[f2410,f482])).
fof(f68287,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) )),
  inference(resolution,[],[f2410,f480])).
fof(f68285,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) )),
  inference(resolution,[],[f2410,f478])).
fof(f57210,plain,(
  ( ! [X215,X213,X214] : (s(bool,X213) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | p(s(bool,X214)) | p(s(bool,X215)) | s(bool,X213) != s(bool,X215)) )),
  inference(resolution,[],[f1655,f648])).
fof(f57144,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) )),
  inference(resolution,[],[f1655,f484])).
fof(f57141,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,X16) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) )),
  inference(resolution,[],[f1655,f483])).
fof(f57140,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) )),
  inference(resolution,[],[f1655,f482])).
fof(f57139,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) )),
  inference(resolution,[],[f1655,f480])).
fof(f57137,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) )),
  inference(resolution,[],[f1655,f478])).
fof(f56938,plain,(
  ( ! [X222,X223,X224] : (s(bool,X222) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | p(s(bool,X223)) | ~p(s(bool,X224)) | s(bool,X222) != s(bool,X224)) )),
  inference(resolution,[],[f737,f648])).
fof(f56869,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) )),
  inference(resolution,[],[f737,f484])).
fof(f56867,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,X20) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X20) != s(bool,X22)) )),
  inference(resolution,[],[f737,f483])).
fof(f56865,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) )),
  inference(resolution,[],[f737,f482])).
fof(f56864,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) )),
  inference(resolution,[],[f737,f480])).
fof(f56862,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) )),
  inference(resolution,[],[f737,f478])).
fof(f55158,plain,(
  ( ! [X147,X148] : (p(s(bool,X147)) | ~p(s(bool,X148)) | s(bool,X148) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) ) | $spl28),
  inference(resolution,[],[f7209,f648])).
fof(f55089,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11))) | s(bool,X22) != s(bool,X24)) ) | $spl28),
  inference(resolution,[],[f7209,f484])).
fof(f55087,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) ) | $spl28),
  inference(resolution,[],[f7209,f483])).
fof(f55085,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) ) | $spl28),
  inference(resolution,[],[f7209,f482])).
fof(f55084,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11))) | s(bool,X7) != s(bool,X9)) ) | $spl28),
  inference(resolution,[],[f7209,f480])).
fof(f55082,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) ) | $spl28),
  inference(resolution,[],[f7209,f478])).
fof(f23306,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X1))) ) | $spl82),
  inference(cnf_transformation,[],[f23306_D])).
fof(f23306_D,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X1))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f54998,plain,(
  ( ! [X140,X141] : (p(s(bool,X140)) | p(s(bool,X141)) | s(bool,X141) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl28),
  inference(resolution,[],[f7202,f648])).
fof(f54932,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X22) != s(bool,X23)) ) | $spl28),
  inference(resolution,[],[f7202,f484])).
fof(f54929,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl28),
  inference(resolution,[],[f7202,f483])).
fof(f54928,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl28),
  inference(resolution,[],[f7202,f482])).
fof(f54927,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X7) != s(bool,X8)) ) | $spl28),
  inference(resolution,[],[f7202,f480])).
fof(f54925,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl28),
  inference(resolution,[],[f7202,f478])).
fof(f18959,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11))) | ~p(s(bool,X1))) ) | $spl78),
  inference(cnf_transformation,[],[f18959_D])).
fof(f18959_D,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11))) | ~p(s(bool,X1))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f54880,plain,(
  ( ! [X146,X147] : (p(s(bool,X146)) | ~p(s(bool,X147)) | s(bool,X147) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) ) | $spl28),
  inference(resolution,[],[f7199,f648])).
fof(f54811,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11))) | s(bool,X21) != s(bool,X23)) ) | $spl28),
  inference(resolution,[],[f7199,f484])).
fof(f54809,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) ) | $spl28),
  inference(resolution,[],[f7199,f483])).
fof(f54807,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) ) | $spl28),
  inference(resolution,[],[f7199,f482])).
fof(f54806,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11))) | s(bool,X6) != s(bool,X8)) ) | $spl28),
  inference(resolution,[],[f7199,f480])).
fof(f54804,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) ) | $spl28),
  inference(resolution,[],[f7199,f478])).
fof(f54751,plain,(
  ( ! [X144,X143] : (p(s(bool,X143)) | p(s(bool,X144)) | s(bool,X144) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl28),
  inference(resolution,[],[f2018,f648])).
fof(f54685,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X25) != s(bool,X26)) ) | $spl28),
  inference(resolution,[],[f2018,f484])).
fof(f54682,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl28),
  inference(resolution,[],[f2018,f483])).
fof(f54681,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl28),
  inference(resolution,[],[f2018,f482])).
fof(f54680,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X10) != s(bool,X11)) ) | $spl28),
  inference(resolution,[],[f2018,f480])).
fof(f54678,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl28),
  inference(resolution,[],[f2018,f478])).
fof(f54663,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl144),
  inference(cnf_transformation,[],[f54663_D])).
fof(f54663_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f54623,plain,(
  ( ! [X146,X147] : (p(s(bool,X146)) | ~p(s(bool,X147)) | s(bool,X147) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl87),
  inference(resolution,[],[f27214,f648])).
fof(f54554,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X21) != s(bool,X23)) ) | $spl87),
  inference(resolution,[],[f27214,f484])).
fof(f54552,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl87),
  inference(resolution,[],[f27214,f483])).
fof(f54550,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl87),
  inference(resolution,[],[f27214,f482])).
fof(f54549,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X6) != s(bool,X8)) ) | $spl87),
  inference(resolution,[],[f27214,f480])).
fof(f54547,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl87),
  inference(resolution,[],[f27214,f478])).
fof(f53141,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) ) | $spl142),
  inference(cnf_transformation,[],[f53141_D])).
fof(f53141_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f54173,plain,(
  ( ! [X140,X141] : (p(s(bool,X140)) | p(s(bool,X141)) | s(bool,X141) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) ) | $spl29),
  inference(resolution,[],[f27088,f648])).
fof(f27088,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) | $spl29),
  inference(resolution,[],[f2017,f460])).
fof(f2017,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) | $spl29),
  inference(cnf_transformation,[],[f2017_D])).
fof(f2017_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f54107,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11))) | s(bool,X22) != s(bool,X23)) ) | $spl29),
  inference(resolution,[],[f27088,f484])).
fof(f54104,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) ) | $spl29),
  inference(resolution,[],[f27088,f483])).
fof(f54103,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) ) | $spl29),
  inference(resolution,[],[f27088,f482])).
fof(f54102,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11))) | s(bool,X7) != s(bool,X8)) ) | $spl29),
  inference(resolution,[],[f27088,f480])).
fof(f54100,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) ) | $spl29),
  inference(resolution,[],[f27088,f478])).
fof(f27219,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl92),
  inference(cnf_transformation,[],[f27219_D])).
fof(f27219_D,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f54051,plain,(
  ( ! [X147,X148] : (p(s(bool,X147)) | ~p(s(bool,X148)) | s(bool,X148) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl29),
  inference(resolution,[],[f2017,f648])).
fof(f53982,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X22) != s(bool,X24)) ) | $spl29),
  inference(resolution,[],[f2017,f484])).
fof(f53980,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl29),
  inference(resolution,[],[f2017,f483])).
fof(f53978,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl29),
  inference(resolution,[],[f2017,f482])).
fof(f53977,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X7) != s(bool,X9)) ) | $spl29),
  inference(resolution,[],[f2017,f480])).
fof(f53975,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl29),
  inference(resolution,[],[f2017,f478])).
fof(f53963,plain,(
  $false | ($spl28 | $spl90)),
  inference(subsumption_resolution,[],[f27218,f7209])).
fof(f53637,plain,(
  ( ! [X147,X148] : (p(s(bool,X147)) | ~p(s(bool,X148)) | s(bool,X148) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) ) | $spl28),
  inference(resolution,[],[f7209,f648])).
fof(f53568,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11))) | s(bool,X22) != s(bool,X24)) ) | $spl28),
  inference(resolution,[],[f7209,f484])).
fof(f53566,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) ) | $spl28),
  inference(resolution,[],[f7209,f483])).
fof(f53564,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) ) | $spl28),
  inference(resolution,[],[f7209,f482])).
fof(f53563,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11))) | s(bool,X7) != s(bool,X9)) ) | $spl28),
  inference(resolution,[],[f7209,f480])).
fof(f53561,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) ) | $spl28),
  inference(resolution,[],[f7209,f478])).
fof(f53476,plain,(
  ( ! [X140,X141] : (p(s(bool,X140)) | p(s(bool,X141)) | s(bool,X141) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl28),
  inference(resolution,[],[f7202,f648])).
fof(f53410,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X22) != s(bool,X23)) ) | $spl28),
  inference(resolution,[],[f7202,f484])).
fof(f53407,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl28),
  inference(resolution,[],[f7202,f483])).
fof(f53406,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl28),
  inference(resolution,[],[f7202,f482])).
fof(f53405,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X7) != s(bool,X8)) ) | $spl28),
  inference(resolution,[],[f7202,f480])).
fof(f53403,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl28),
  inference(resolution,[],[f7202,f478])).
fof(f53358,plain,(
  ( ! [X146,X147] : (p(s(bool,X146)) | ~p(s(bool,X147)) | s(bool,X147) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) ) | $spl28),
  inference(resolution,[],[f7199,f648])).
fof(f53289,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11))) | s(bool,X21) != s(bool,X23)) ) | $spl28),
  inference(resolution,[],[f7199,f484])).
fof(f53287,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) ) | $spl28),
  inference(resolution,[],[f7199,f483])).
fof(f53285,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) ) | $spl28),
  inference(resolution,[],[f7199,f482])).
fof(f53284,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11))) | s(bool,X6) != s(bool,X8)) ) | $spl28),
  inference(resolution,[],[f7199,f480])).
fof(f53282,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) ) | $spl28),
  inference(resolution,[],[f7199,f478])).
fof(f53229,plain,(
  ( ! [X144,X143] : (p(s(bool,X143)) | p(s(bool,X144)) | s(bool,X144) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl28),
  inference(resolution,[],[f2018,f648])).
fof(f53163,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X25) != s(bool,X26)) ) | $spl28),
  inference(resolution,[],[f2018,f484])).
fof(f53160,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl28),
  inference(resolution,[],[f2018,f483])).
fof(f53159,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl28),
  inference(resolution,[],[f2018,f482])).
fof(f53158,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X10) != s(bool,X11)) ) | $spl28),
  inference(resolution,[],[f2018,f480])).
fof(f53156,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl28),
  inference(resolution,[],[f2018,f478])).
fof(f53098,plain,(
  ( ! [X140,X141] : (p(s(bool,X140)) | p(s(bool,X141)) | s(bool,X141) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) ) | $spl29),
  inference(resolution,[],[f27088,f648])).
fof(f53032,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11))) | s(bool,X22) != s(bool,X23)) ) | $spl29),
  inference(resolution,[],[f27088,f484])).
fof(f53029,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) ) | $spl29),
  inference(resolution,[],[f27088,f483])).
fof(f53028,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) ) | $spl29),
  inference(resolution,[],[f27088,f482])).
fof(f53027,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11))) | s(bool,X7) != s(bool,X8)) ) | $spl29),
  inference(resolution,[],[f27088,f480])).
fof(f53025,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) ) | $spl29),
  inference(resolution,[],[f27088,f478])).
fof(f52288,plain,(
  ( ! [X147,X148] : (p(s(bool,X147)) | ~p(s(bool,X148)) | s(bool,X148) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl16),
  inference(resolution,[],[f6594,f648])).
fof(f52219,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | s(bool,X22) != s(bool,X24)) ) | $spl16),
  inference(resolution,[],[f6594,f484])).
fof(f52217,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl16),
  inference(resolution,[],[f6594,f483])).
fof(f52215,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl16),
  inference(resolution,[],[f6594,f482])).
fof(f52214,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | s(bool,X7) != s(bool,X9)) ) | $spl16),
  inference(resolution,[],[f6594,f480])).
fof(f52212,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl16),
  inference(resolution,[],[f6594,f478])).
fof(f52161,plain,(
  ( ! [X140,X141] : (p(s(bool,X140)) | p(s(bool,X141)) | s(bool,X141) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl16),
  inference(resolution,[],[f6587,f648])).
fof(f52095,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X22) != s(bool,X23)) ) | $spl16),
  inference(resolution,[],[f6587,f484])).
fof(f52092,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl16),
  inference(resolution,[],[f6587,f483])).
fof(f52091,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl16),
  inference(resolution,[],[f6587,f482])).
fof(f52090,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X7) != s(bool,X8)) ) | $spl16),
  inference(resolution,[],[f6587,f480])).
fof(f52088,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl16),
  inference(resolution,[],[f6587,f478])).
fof(f52043,plain,(
  ( ! [X146,X147] : (p(s(bool,X146)) | ~p(s(bool,X147)) | s(bool,X147) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl16),
  inference(resolution,[],[f6584,f648])).
fof(f51974,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | s(bool,X21) != s(bool,X23)) ) | $spl16),
  inference(resolution,[],[f6584,f484])).
fof(f51972,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl16),
  inference(resolution,[],[f6584,f483])).
fof(f51970,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl16),
  inference(resolution,[],[f6584,f482])).
fof(f51969,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | s(bool,X6) != s(bool,X8)) ) | $spl16),
  inference(resolution,[],[f6584,f480])).
fof(f51967,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl16),
  inference(resolution,[],[f6584,f478])).
fof(f50399,plain,(
  ( ! [X147,X148] : (p(s(bool,X147)) | ~p(s(bool,X148)) | s(bool,X148) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl29),
  inference(resolution,[],[f2017,f648])).
fof(f50330,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X22) != s(bool,X24)) ) | $spl29),
  inference(resolution,[],[f2017,f484])).
fof(f50328,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl29),
  inference(resolution,[],[f2017,f483])).
fof(f50326,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl29),
  inference(resolution,[],[f2017,f482])).
fof(f50325,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X7) != s(bool,X9)) ) | $spl29),
  inference(resolution,[],[f2017,f480])).
fof(f50323,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl29),
  inference(resolution,[],[f2017,f478])).
fof(f50270,plain,(
  ( ! [X144,X143] : (p(s(bool,X143)) | p(s(bool,X144)) | s(bool,X144) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl16),
  inference(resolution,[],[f1403,f648])).
fof(f50204,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X25) != s(bool,X26)) ) | $spl16),
  inference(resolution,[],[f1403,f484])).
fof(f50201,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl16),
  inference(resolution,[],[f1403,f483])).
fof(f50200,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl16),
  inference(resolution,[],[f1403,f482])).
fof(f50199,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X10) != s(bool,X11)) ) | $spl16),
  inference(resolution,[],[f1403,f480])).
fof(f50197,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl16),
  inference(resolution,[],[f1403,f478])).
fof(f50111,plain,(
  ( ! [X144,X143] : (p(s(bool,X143)) | p(s(bool,X144)) | s(bool,X144) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11)))) ) | $spl24),
  inference(resolution,[],[f1703,f648])).
fof(f50045,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | s(bool,X25) != s(bool,X26)) ) | $spl24),
  inference(resolution,[],[f1703,f484])).
fof(f50042,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11)))) ) | $spl24),
  inference(resolution,[],[f1703,f483])).
fof(f50041,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11)))) ) | $spl24),
  inference(resolution,[],[f1703,f482])).
fof(f50040,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | s(bool,X10) != s(bool,X11)) ) | $spl24),
  inference(resolution,[],[f1703,f480])).
fof(f50038,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11)))) ) | $spl24),
  inference(resolution,[],[f1703,f478])).
fof(f49984,plain,(
  ( ! [X146,X147] : (p(s(bool,X146)) | ~p(s(bool,X147)) | s(bool,X147) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK13)))) ) | $spl21),
  inference(resolution,[],[f1699,f648])).
fof(f49915,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | s(bool,X21) != s(bool,X23)) ) | $spl21),
  inference(resolution,[],[f1699,f484])).
fof(f49913,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK13)))) ) | $spl21),
  inference(resolution,[],[f1699,f483])).
fof(f49911,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK13)))) ) | $spl21),
  inference(resolution,[],[f1699,f482])).
fof(f49910,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | s(bool,X6) != s(bool,X8)) ) | $spl21),
  inference(resolution,[],[f1699,f480])).
fof(f49908,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK13)))) ) | $spl21),
  inference(resolution,[],[f1699,f478])).
fof(f49642,plain,(
  ( ! [X140,X141] : (p(s(bool,X140)) | p(s(bool,X141)) | s(bool,X141) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f1573,f648])).
fof(f49576,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | s(bool,X22) != s(bool,X23)) )),
  inference(resolution,[],[f1573,f484])).
fof(f49573,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f1573,f483])).
fof(f49572,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f1573,f482])).
fof(f49571,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | s(bool,X7) != s(bool,X8)) )),
  inference(resolution,[],[f1573,f480])).
fof(f49569,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f1573,f478])).
fof(f49309,plain,(
  ( ! [X147,X148] : (p(s(bool,X147)) | ~p(s(bool,X148)) | s(bool,X148) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f664,f648])).
fof(f49240,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | s(bool,X22) != s(bool,X24)) )),
  inference(resolution,[],[f664,f484])).
fof(f49238,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f664,f483])).
fof(f49236,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f664,f482])).
fof(f49235,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | s(bool,X7) != s(bool,X9)) )),
  inference(resolution,[],[f664,f480])).
fof(f49233,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f664,f478])).
fof(f49182,plain,(
  ( ! [X140,X141] : (p(s(bool,X140)) | p(s(bool,X141)) | s(bool,X141) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f657,f648])).
fof(f49116,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | s(bool,X22) != s(bool,X23)) )),
  inference(resolution,[],[f657,f484])).
fof(f49113,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f657,f483])).
fof(f49112,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f657,f482])).
fof(f49111,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | s(bool,X7) != s(bool,X8)) )),
  inference(resolution,[],[f657,f480])).
fof(f49109,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f657,f478])).
fof(f49064,plain,(
  ( ! [X146,X147] : (p(s(bool,X146)) | ~p(s(bool,X147)) | s(bool,X147) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f654,f648])).
fof(f48995,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | s(bool,X21) != s(bool,X23)) )),
  inference(resolution,[],[f654,f484])).
fof(f48993,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f654,f483])).
fof(f48991,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f654,f482])).
fof(f48990,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | s(bool,X6) != s(bool,X8)) )),
  inference(resolution,[],[f654,f480])).
fof(f48988,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f654,f478])).
fof(f48933,plain,(
  ( ! [X147,X148] : (p(s(bool,X147)) | ~p(s(bool,X148)) | s(bool,X148) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f459,f648])).
fof(f48864,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | s(bool,X22) != s(bool,X24)) )),
  inference(resolution,[],[f459,f484])).
fof(f48862,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f459,f483])).
fof(f48860,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f459,f482])).
fof(f48859,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | s(bool,X7) != s(bool,X9)) )),
  inference(resolution,[],[f459,f480])).
fof(f48857,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f459,f478])).
fof(f48805,plain,(
  ( ! [X144,X143] : (p(s(bool,X143)) | p(s(bool,X144)) | s(bool,X144) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f457,f648])).
fof(f48739,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | s(bool,X25) != s(bool,X26)) )),
  inference(resolution,[],[f457,f484])).
fof(f48736,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f457,f483])).
fof(f48735,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f457,f482])).
fof(f48734,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | s(bool,X10) != s(bool,X11)) )),
  inference(resolution,[],[f457,f480])).
fof(f48732,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f457,f478])).
fof(f48643,plain,(
  ( ! [X147,X148] : (p(s(bool,X147)) | ~p(s(bool,X148)) | s(bool,X148) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | $spl88),
  inference(resolution,[],[f47798,f648])).
fof(f47798,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) | $spl88),
  inference(backward_demodulation,[],[f27216,f459])).
fof(f48574,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | s(bool,X22) != s(bool,X24)) ) | $spl88),
  inference(resolution,[],[f47798,f484])).
fof(f48572,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | $spl88),
  inference(resolution,[],[f47798,f483])).
fof(f48570,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | $spl88),
  inference(resolution,[],[f47798,f482])).
fof(f48569,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | s(bool,X7) != s(bool,X9)) ) | $spl88),
  inference(resolution,[],[f47798,f480])).
fof(f48567,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | $spl88),
  inference(resolution,[],[f47798,f478])).
fof(f48505,plain,(
  ( ! [X146,X145] : (p(s(bool,X145)) | p(s(bool,X146)) | s(bool,X146) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl88),
  inference(resolution,[],[f47797,f648])).
fof(f47797,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) | $spl88),
  inference(backward_demodulation,[],[f27216,f457])).
fof(f48439,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | s(bool,X27) != s(bool,X28)) ) | $spl88),
  inference(resolution,[],[f47797,f484])).
fof(f48436,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X19) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl88),
  inference(resolution,[],[f47797,f483])).
fof(f48435,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl88),
  inference(resolution,[],[f47797,f482])).
fof(f48434,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | s(bool,X12) != s(bool,X13)) ) | $spl88),
  inference(resolution,[],[f47797,f480])).
fof(f48432,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl88),
  inference(resolution,[],[f47797,f478])).
fof(f47796,plain,(
  $false | ($spl16 | $spl53)),
  inference(subsumption_resolution,[],[f6587,f7056])).
fof(f47795,plain,(
  $false | ($spl16 | $spl53)),
  inference(subsumption_resolution,[],[f7075,f7056])).
fof(f7075,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) | $spl16),
  inference(resolution,[],[f1403,f467])).
fof(f47794,plain,(
  $false | ($spl16 | $spl53)),
  inference(subsumption_resolution,[],[f7076,f7056])).
fof(f7076,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) | $spl16),
  inference(resolution,[],[f1403,f469])).
fof(f469,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f382])).
fof(f47793,plain,(
  $false | ($spl16 | $spl53)),
  inference(subsumption_resolution,[],[f7078,f7056])).
fof(f7078,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) | $spl16),
  inference(resolution,[],[f1403,f472])).
fof(f472,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f337])).
fof(f337,plain,(
  ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))),
  inference(ennf_transformation,[],[f273])).
fof(f273,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) => p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X6,X7] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X6),s('type.integer.int',X7)))) => p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X6),s('type.integer.int',X7)))))),
  file('thm.integer.INT_LTE_TRANS','thm.integer.INT_LT_IMP_LE')).
fof(f47792,plain,(
  $false | ($spl16 | $spl53)),
  inference(subsumption_resolution,[],[f9696,f7056])).
fof(f9696,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) | $spl16),
  inference(resolution,[],[f6584,f475])).
fof(f47791,plain,(
  $false | ($spl16 | $spl53)),
  inference(subsumption_resolution,[],[f9697,f7056])).
fof(f9697,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) | $spl16),
  inference(resolution,[],[f6584,f475])).
fof(f47790,plain,(
  $false | ($spl16 | $spl53)),
  inference(subsumption_resolution,[],[f13319,f7056])).
fof(f13319,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) | $spl16),
  inference(resolution,[],[f1403,f467])).
fof(f47789,plain,(
  $false | ($spl16 | $spl53)),
  inference(subsumption_resolution,[],[f13320,f7056])).
fof(f13320,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) | $spl16),
  inference(resolution,[],[f1403,f469])).
fof(f47788,plain,(
  $false | ($spl16 | $spl53)),
  inference(subsumption_resolution,[],[f13322,f7056])).
fof(f13322,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) | $spl16),
  inference(resolution,[],[f1403,f472])).
fof(f47787,plain,(
  $false | ($spl16 | $spl53)),
  inference(subsumption_resolution,[],[f13452,f7056])).
fof(f13452,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) | $spl16),
  inference(resolution,[],[f6584,f475])).
fof(f47786,plain,(
  $false | ($spl16 | $spl53)),
  inference(subsumption_resolution,[],[f13453,f7056])).
fof(f13453,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) | $spl16),
  inference(resolution,[],[f6584,f475])).
fof(f47785,plain,(
  $false | ($spl16 | $spl53)),
  inference(subsumption_resolution,[],[f16274,f7056])).
fof(f16274,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) | $spl16),
  inference(resolution,[],[f1403,f467])).
fof(f47784,plain,(
  $false | ($spl16 | $spl53)),
  inference(subsumption_resolution,[],[f16275,f7056])).
fof(f16275,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) | $spl16),
  inference(resolution,[],[f1403,f469])).
fof(f47783,plain,(
  $false | ($spl16 | $spl53)),
  inference(subsumption_resolution,[],[f16277,f7056])).
fof(f16277,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) | $spl16),
  inference(resolution,[],[f1403,f472])).
fof(f47782,plain,(
  $false | ($spl16 | $spl53)),
  inference(subsumption_resolution,[],[f17795,f7056])).
fof(f17795,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) | $spl16),
  inference(resolution,[],[f6584,f475])).
fof(f47781,plain,(
  $false | ($spl16 | $spl53)),
  inference(subsumption_resolution,[],[f17796,f7056])).
fof(f17796,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) | $spl16),
  inference(resolution,[],[f6584,f475])).
fof(f47282,plain,(
  ( ! [X12,X13] : (~p(s(bool,X13)) | ~p(s(bool,X12)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11))))) ) | $spl78),
  inference(subsumption_resolution,[],[f47266,f590])).
fof(f47266,plain,(
  ( ! [X12,X13] : (s(bool,X12) != s(bool,X13) | ~p(s(bool,X13)) | ~p(s(bool,X12)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11))))) ) | $spl78),
  inference(superposition,[],[f18959,f590])).
fof(f47283,plain,(
  ( ! [X14,X15] : (~p(s(bool,X15)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) | ~p(s(bool,X14))) ) | $spl78),
  inference(subsumption_resolution,[],[f47267,f590])).
fof(f47267,plain,(
  ( ! [X14,X15] : (s(bool,X14) != s(bool,X15) | ~p(s(bool,X15)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) | ~p(s(bool,X14))) ) | $spl78),
  inference(superposition,[],[f18959,f590])).
fof(f47284,plain,(
  ( ! [X24,X25] : (~p(s(bool,X25)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) | ~p(s(bool,X24))) ) | $spl78),
  inference(subsumption_resolution,[],[f47272,f590])).
fof(f47272,plain,(
  ( ! [X24,X25] : (s(bool,X24) != s(bool,X25) | ~p(s(bool,X25)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) | ~p(s(bool,X24))) ) | $spl78),
  inference(superposition,[],[f18959,f595])).
fof(f47285,plain,(
  ( ! [X26,X27] : (~p(s(bool,X27)) | ~p(s(bool,X26)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11))))) ) | $spl78),
  inference(subsumption_resolution,[],[f47273,f590])).
fof(f47273,plain,(
  ( ! [X26,X27] : (s(bool,X26) != s(bool,X27) | ~p(s(bool,X27)) | ~p(s(bool,X26)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11))))) ) | $spl78),
  inference(superposition,[],[f18959,f595])).
fof(f47778,plain,(
  ( ! [X17,X16] : (p(s(bool,X17)) | p(s(bool,X16)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))))) ) | $spl82),
  inference(subsumption_resolution,[],[f47327,f591])).
fof(f47327,plain,(
  ( ! [X17,X16] : (s(bool,X16) != s(bool,X17) | p(s(bool,X17)) | p(s(bool,X16)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))))) ) | $spl82),
  inference(superposition,[],[f23306,f591])).
fof(f47777,plain,(
  ( ! [X19,X18] : (p(s(bool,X19)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,X18))) ) | $spl82),
  inference(subsumption_resolution,[],[f47328,f591])).
fof(f47328,plain,(
  ( ! [X19,X18] : (s(bool,X18) != s(bool,X19) | p(s(bool,X19)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,X18))) ) | $spl82),
  inference(superposition,[],[f23306,f591])).
fof(f47776,plain,(
  ( ! [X21,X20] : (p(s(bool,X21)) | p(s(bool,X20)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))))) ) | $spl82),
  inference(subsumption_resolution,[],[f47329,f591])).
fof(f47329,plain,(
  ( ! [X21,X20] : (s(bool,X20) != s(bool,X21) | p(s(bool,X21)) | p(s(bool,X20)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))))) ) | $spl82),
  inference(superposition,[],[f23306,f592])).
fof(f47775,plain,(
  ( ! [X23,X22] : (p(s(bool,X23)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,X22))) ) | $spl82),
  inference(subsumption_resolution,[],[f47330,f591])).
fof(f47330,plain,(
  ( ! [X23,X22] : (s(bool,X22) != s(bool,X23) | p(s(bool,X23)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,X22))) ) | $spl82),
  inference(superposition,[],[f23306,f592])).
fof(f47758,plain,(
  ( ! [X47] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | ~p(s(bool,X47))) )),
  inference(resolution,[],[f33992,f4509])).
fof(f47757,plain,(
  ( ! [X46] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | ~p(s(bool,X46))) )),
  inference(resolution,[],[f33992,f4511])).
fof(f47756,plain,(
  ( ! [X45] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | ~p(s(bool,X45))) )),
  inference(resolution,[],[f33992,f842])).
fof(f47755,plain,(
  ( ! [X44] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | ~p(s(bool,X44))) )),
  inference(resolution,[],[f33992,f844])).
fof(f47754,plain,(
  ( ! [X43] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))) | ~p(s(bool,X43))) ) | $spl28),
  inference(resolution,[],[f33992,f27539])).
fof(f27539,plain,(
  ( ! [X111,X110] : (sP8('const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')),X110,X111) | ~p(s(bool,X111))) ) | $spl28),
  inference(resolution,[],[f7202,f558])).
fof(f47753,plain,(
  ( ! [X42] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))) | ~p(s(bool,X42))) ) | $spl28),
  inference(resolution,[],[f33992,f27541])).
fof(f27541,plain,(
  ( ! [X114,X115] : (sP8('const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')),X114,X115) | ~p(s(bool,X114))) ) | $spl28),
  inference(resolution,[],[f7202,f559])).
fof(f47752,plain,(
  ( ! [X41] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | ~p(s(bool,X41))) )),
  inference(resolution,[],[f33992,f2510])).
fof(f47751,plain,(
  ( ! [X40] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | ~p(s(bool,X40))) )),
  inference(resolution,[],[f33992,f2512])).
fof(f47750,plain,(
  ( ! [X39] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | ~p(s(bool,X39))) ) | $spl17),
  inference(resolution,[],[f33992,f21848])).
fof(f21848,plain,(
  ( ! [X111,X110] : (sP8('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)),X110,X111) | ~p(s(bool,X111))) ) | $spl17),
  inference(resolution,[],[f6930,f558])).
fof(f47749,plain,(
  ( ! [X38] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | ~p(s(bool,X38))) ) | $spl17),
  inference(resolution,[],[f33992,f21850])).
fof(f21850,plain,(
  ( ! [X114,X115] : (sP8('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)),X114,X115) | ~p(s(bool,X114))) ) | $spl17),
  inference(resolution,[],[f6930,f559])).
fof(f47748,plain,(
  ( ! [X37] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | ~p(s(bool,X37))) ) | $spl24),
  inference(resolution,[],[f33992,f6521])).
fof(f47747,plain,(
  ( ! [X36] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | ~p(s(bool,X36))) ) | $spl24),
  inference(resolution,[],[f33992,f6523])).
fof(f47746,plain,(
  ( ! [X35] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | ~p(s(bool,X35))) ) | $spl12),
  inference(resolution,[],[f33992,f4391])).
fof(f47745,plain,(
  ( ! [X34] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | ~p(s(bool,X34))) ) | $spl12),
  inference(resolution,[],[f33992,f4393])).
fof(f47744,plain,(
  ( ! [X33] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | ~p(s(bool,X33))) ) | $spl28),
  inference(resolution,[],[f33992,f27295])).
fof(f27295,plain,(
  ( ! [X114,X113] : (sP8('const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')),X113,X114) | ~p(s(bool,X114))) ) | $spl28),
  inference(resolution,[],[f2018,f558])).
fof(f47743,plain,(
  ( ! [X32] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | ~p(s(bool,X32))) ) | $spl28),
  inference(resolution,[],[f33992,f27297])).
fof(f27297,plain,(
  ( ! [X118,X117] : (sP8('const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')),X117,X118) | ~p(s(bool,X117))) ) | $spl28),
  inference(resolution,[],[f2018,f559])).
fof(f47740,plain,(
  ( ! [X29] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | ~p(s(bool,X29))) ) | $spl56),
  inference(resolution,[],[f33992,f26656])).
fof(f26656,plain,(
  ( ! [X116,X115] : (sP8('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)),X115,X116) | ~p(s(bool,X116))) ) | $spl56),
  inference(resolution,[],[f7060,f558])).
fof(f47739,plain,(
  ( ! [X28] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | ~p(s(bool,X28))) ) | $spl56),
  inference(resolution,[],[f33992,f26658])).
fof(f26658,plain,(
  ( ! [X120,X119] : (sP8('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)),X119,X120) | ~p(s(bool,X119))) ) | $spl56),
  inference(resolution,[],[f7060,f559])).
fof(f47671,plain,(
  ( ! [X21,X20] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | p(s(bool,X20)) | ~p(s(bool,X21))) )),
  inference(resolution,[],[f33978,f834])).
fof(f47658,plain,(
  ( ! [X17,X16] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | p(s(bool,X16)) | ~p(s(bool,X17))) )),
  inference(resolution,[],[f33977,f835])).
fof(f47617,plain,(
  ( ! [X21,X20] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | p(s(bool,X20)) | ~p(s(bool,X21))) )),
  inference(resolution,[],[f33975,f831])).
fof(f47605,plain,(
  ( ! [X17,X16] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | p(s(bool,X16)) | ~p(s(bool,X17))) )),
  inference(resolution,[],[f33974,f832])).
fof(f47598,plain,(
  ( ! [X21,X22] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | p(s(bool,X21)) | p(s(bool,X22))) )),
  inference(resolution,[],[f33962,f818])).
fof(f47589,plain,(
  ( ! [X24,X23] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | ~p(s(bool,X23)) | p(s(bool,X24))) )),
  inference(resolution,[],[f33946,f804])).
fof(f47579,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f47494,f784])).
fof(f47494,plain,(
  ( ! [X280,X281] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X280),s('type.integer.int',X281))))) ) | $spl6),
  inference(resolution,[],[f784,f475])).
fof(f47578,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f47493,f784])).
fof(f47493,plain,(
  ( ! [X279,X278] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X278),s('type.integer.int',X279))))) ) | $spl6),
  inference(resolution,[],[f784,f475])).
fof(f47577,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f47492,f784])).
fof(f47492,plain,(
  ( ! [X277,X276] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X276),s('type.integer.int',X277))))) ) | $spl6),
  inference(resolution,[],[f784,f474])).
fof(f474,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f275])).
fof(f275,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X6,X7] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X6),s('type.integer.int',X7)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X7),s('type.integer.int',X6)))))),
  file('thm.integer.INT_LTE_TRANS','thm.integer.INT_LET_TOTAL')).
fof(f47576,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f47491,f784])).
fof(f47491,plain,(
  ( ! [X275,X274] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X274),s('type.integer.int',X275))))) ) | $spl6),
  inference(resolution,[],[f784,f473])).
fof(f473,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f274])).
fof(f274,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X6,X7] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X6),s('type.integer.int',X7)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X7),s('type.integer.int',X6)))))),
  file('thm.integer.INT_LTE_TRANS','thm.integer.INT_LTE_TOTAL')).
fof(f47575,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f47486,f784])).
fof(f47486,plain,(
  ( ! [X265,X264] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X264),s('type.integer.int',X265))))) ) | $spl6),
  inference(resolution,[],[f784,f465])).
fof(f465,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) )),
  inference(cnf_transformation,[],[f378])).
fof(f47574,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f47485,f784])).
fof(f47485,plain,(
  ( ! [X263,X262] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X262),s('type.integer.int',X263))))) ) | $spl6),
  inference(resolution,[],[f784,f462])).
fof(f462,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f377])).
fof(f377,plain,(
  ! [X0,X1] : ((p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) & (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))),
  inference(nnf_transformation,[],[f267])).
fof(f267,plain,(
  ! [X0,X1] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) <=> p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(flattening,[],[f266])).
fof(f266,plain,(
  ! [X0,X1] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) <=> p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X6,X7] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X6),s('type.integer.int',X7)))) <=> p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X7),s('type.integer.int',X6)))))),
  file('thm.integer.INT_LTE_TRANS','thm.integer.INT_NOT_LE')).
fof(f47573,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f47484,f784])).
fof(f47484,plain,(
  ( ! [X261,X260] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X260),s('type.integer.int',X261))))) ) | $spl6),
  inference(resolution,[],[f784,f460])).
fof(f47570,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f47472,f784])).
fof(f47472,plain,(
  ( ! [X246,X247] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X246),s('type.integer.int',X247))))) ) | $spl6),
  inference(resolution,[],[f784,f474])).
fof(f47569,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f47471,f784])).
fof(f47471,plain,(
  ( ! [X244,X245] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X244),s('type.integer.int',X245))))) ) | $spl6),
  inference(resolution,[],[f784,f473])).
fof(f47568,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f47468,f784])).
fof(f47468,plain,(
  ( ! [X239,X238] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X238),s('type.integer.int',X239))))) ) | $spl6),
  inference(resolution,[],[f784,f465])).
fof(f47567,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f47467,f784])).
fof(f47467,plain,(
  ( ! [X237,X236] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X236),s('type.integer.int',X237))))) ) | $spl6),
  inference(resolution,[],[f784,f462])).
fof(f47566,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f47466,f784])).
fof(f47466,plain,(
  ( ! [X235,X234] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X234),s('type.integer.int',X235))))) ) | $spl6),
  inference(resolution,[],[f784,f460])).
fof(f47479,plain,(
  $false | ($spl6 | $spl56)),
  inference(resolution,[],[f784,f7060])).
fof(f47480,plain,(
  $false | $spl6),
  inference(resolution,[],[f784,f457])).
fof(f47481,plain,(
  $false | ($spl6 | $spl28)),
  inference(resolution,[],[f784,f2018])).
fof(f47482,plain,(
  $false | ($spl6 | $spl12)),
  inference(resolution,[],[f784,f906])).
fof(f47483,plain,(
  $false | ($spl6 | $spl24)),
  inference(resolution,[],[f784,f1703])).
fof(f47495,plain,(
  $false | $spl6),
  inference(resolution,[],[f784,f476])).
fof(f47496,plain,(
  $false | ($spl6 | $spl17)),
  inference(resolution,[],[f784,f6930])).
fof(f47497,plain,(
  $false | $spl6),
  inference(resolution,[],[f784,f657])).
fof(f47498,plain,(
  $false | ($spl6 | $spl28)),
  inference(resolution,[],[f784,f7202])).
fof(f47499,plain,(
  $false | $spl6),
  inference(resolution,[],[f784,f458])).
fof(f47500,plain,(
  $false | $spl6),
  inference(resolution,[],[f784,f1573])).
fof(f47375,plain,(
  ( ! [X17,X18] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | p(s(bool,X17)) | ~p(s(bool,X18))) )),
  inference(resolution,[],[f33945,f805])).
fof(f47226,plain,(
  ( ! [X23,X22] : (p(s(bool,X23)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) | p(s(bool,X22))) ) | $spl74),
  inference(subsumption_resolution,[],[f47210,f591])).
fof(f47210,plain,(
  ( ! [X23,X22] : (s(bool,X22) != s(bool,X23) | p(s(bool,X23)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) | p(s(bool,X22))) ) | $spl74),
  inference(superposition,[],[f14956,f592])).
fof(f47225,plain,(
  ( ! [X21,X20] : (p(s(bool,X21)) | p(s(bool,X20)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))))) ) | $spl74),
  inference(subsumption_resolution,[],[f47209,f591])).
fof(f47209,plain,(
  ( ! [X21,X20] : (s(bool,X20) != s(bool,X21) | p(s(bool,X21)) | p(s(bool,X20)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))))) ) | $spl74),
  inference(superposition,[],[f14956,f592])).
fof(f47224,plain,(
  ( ! [X19,X18] : (p(s(bool,X19)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) | p(s(bool,X18))) ) | $spl74),
  inference(subsumption_resolution,[],[f47208,f591])).
fof(f47208,plain,(
  ( ! [X19,X18] : (s(bool,X18) != s(bool,X19) | p(s(bool,X19)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) | p(s(bool,X18))) ) | $spl74),
  inference(superposition,[],[f14956,f591])).
fof(f47221,plain,(
  ( ! [X17,X16] : (p(s(bool,X17)) | p(s(bool,X16)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))))) ) | $spl74),
  inference(subsumption_resolution,[],[f47207,f591])).
fof(f47207,plain,(
  ( ! [X17,X16] : (s(bool,X16) != s(bool,X17) | p(s(bool,X17)) | p(s(bool,X16)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))))) ) | $spl74),
  inference(superposition,[],[f14956,f591])).
fof(f47085,plain,(
  ( ! [X222,X223,X224] : (s(bool,X222) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | p(s(bool,X223)) | ~p(s(bool,X224)) | s(bool,X222) != s(bool,X224)) ) | $spl24),
  inference(resolution,[],[f6539,f648])).
fof(f47016,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) ) | $spl24),
  inference(resolution,[],[f6539,f484])).
fof(f47014,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,X20) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X20) != s(bool,X22)) ) | $spl24),
  inference(resolution,[],[f6539,f483])).
fof(f47012,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl24),
  inference(resolution,[],[f6539,f482])).
fof(f47011,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) ) | $spl24),
  inference(resolution,[],[f6539,f480])).
fof(f47009,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl24),
  inference(resolution,[],[f6539,f478])).
fof(f46802,plain,(
  ( ! [X215,X213,X214] : (s(bool,X213) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | p(s(bool,X214)) | p(s(bool,X215)) | s(bool,X213) != s(bool,X215)) ) | $spl21),
  inference(resolution,[],[f6413,f648])).
fof(f46736,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) ) | $spl21),
  inference(resolution,[],[f6413,f484])).
fof(f46733,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,X16) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) ) | $spl21),
  inference(resolution,[],[f6413,f483])).
fof(f46732,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl21),
  inference(resolution,[],[f6413,f482])).
fof(f46731,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) ) | $spl21),
  inference(resolution,[],[f6413,f480])).
fof(f46729,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl21),
  inference(resolution,[],[f6413,f478])).
fof(f46511,plain,(
  ( ! [X222,X223,X224] : (s(bool,X222) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | p(s(bool,X223)) | ~p(s(bool,X224)) | s(bool,X222) != s(bool,X224)) )),
  inference(resolution,[],[f4527,f648])).
fof(f46442,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) )),
  inference(resolution,[],[f4527,f484])).
fof(f46440,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,X20) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X20) != s(bool,X22)) )),
  inference(resolution,[],[f4527,f483])).
fof(f46438,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) )),
  inference(resolution,[],[f4527,f482])).
fof(f46437,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) )),
  inference(resolution,[],[f4527,f480])).
fof(f46435,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) )),
  inference(resolution,[],[f4527,f478])).
fof(f46230,plain,(
  ( ! [X222,X223,X224] : (s(bool,X222) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | p(s(bool,X223)) | ~p(s(bool,X224)) | s(bool,X222) != s(bool,X224)) ) | $spl12),
  inference(resolution,[],[f4409,f648])).
fof(f46161,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) ) | $spl12),
  inference(resolution,[],[f4409,f484])).
fof(f46159,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,X20) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X20) != s(bool,X22)) ) | $spl12),
  inference(resolution,[],[f4409,f483])).
fof(f46157,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl12),
  inference(resolution,[],[f4409,f482])).
fof(f46156,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) ) | $spl12),
  inference(resolution,[],[f4409,f480])).
fof(f46154,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl12),
  inference(resolution,[],[f4409,f478])).
fof(f45949,plain,(
  ( ! [X215,X213,X214] : (s(bool,X213) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | p(s(bool,X214)) | p(s(bool,X215)) | s(bool,X213) != s(bool,X215)) )),
  inference(resolution,[],[f2987,f648])).
fof(f45883,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) )),
  inference(resolution,[],[f2987,f484])).
fof(f45880,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,X16) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) )),
  inference(resolution,[],[f2987,f483])).
fof(f45879,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) )),
  inference(resolution,[],[f2987,f482])).
fof(f45878,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) )),
  inference(resolution,[],[f2987,f480])).
fof(f45876,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) )),
  inference(resolution,[],[f2987,f478])).
fof(f45669,plain,(
  ( ! [X215,X213,X214] : (s(bool,X213) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | p(s(bool,X214)) | p(s(bool,X215)) | s(bool,X213) != s(bool,X215)) )),
  inference(resolution,[],[f2654,f648])).
fof(f45603,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) )),
  inference(resolution,[],[f2654,f484])).
fof(f45600,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,X16) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) )),
  inference(resolution,[],[f2654,f483])).
fof(f45599,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) )),
  inference(resolution,[],[f2654,f482])).
fof(f45598,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) )),
  inference(resolution,[],[f2654,f480])).
fof(f45596,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) )),
  inference(resolution,[],[f2654,f478])).
fof(f45391,plain,(
  ( ! [X222,X223,X224] : (s(bool,X222) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | p(s(bool,X223)) | ~p(s(bool,X224)) | s(bool,X222) != s(bool,X224)) )),
  inference(resolution,[],[f2528,f648])).
fof(f45322,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) )),
  inference(resolution,[],[f2528,f484])).
fof(f45320,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,X20) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X20) != s(bool,X22)) )),
  inference(resolution,[],[f2528,f483])).
fof(f45318,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) )),
  inference(resolution,[],[f2528,f482])).
fof(f45317,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) )),
  inference(resolution,[],[f2528,f480])).
fof(f45315,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) )),
  inference(resolution,[],[f2528,f478])).
fof(f45111,plain,(
  ( ! [X215,X213,X214] : (s(bool,X213) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | p(s(bool,X214)) | p(s(bool,X215)) | s(bool,X213) != s(bool,X215)) )),
  inference(resolution,[],[f2410,f648])).
fof(f45045,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) )),
  inference(resolution,[],[f2410,f484])).
fof(f45042,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,X16) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) )),
  inference(resolution,[],[f2410,f483])).
fof(f45041,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) )),
  inference(resolution,[],[f2410,f482])).
fof(f45040,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) )),
  inference(resolution,[],[f2410,f480])).
fof(f45038,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) )),
  inference(resolution,[],[f2410,f478])).
fof(f37522,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) | $spl100),
  inference(cnf_transformation,[],[f37522_D])).
fof(f37522_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f37525,plain,(
  ( ! [X23,X22] : (p(s(bool,X23)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) | p(s(bool,X22))) ) | $spl66),
  inference(subsumption_resolution,[],[f37509,f591])).
fof(f37509,plain,(
  ( ! [X23,X22] : (s(bool,X22) != s(bool,X23) | p(s(bool,X23)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) | p(s(bool,X22))) ) | $spl66),
  inference(superposition,[],[f10363,f592])).
fof(f37524,plain,(
  ( ! [X21,X20] : (p(s(bool,X21)) | p(s(bool,X20)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))))) ) | $spl66),
  inference(subsumption_resolution,[],[f37508,f591])).
fof(f37508,plain,(
  ( ! [X21,X20] : (s(bool,X20) != s(bool,X21) | p(s(bool,X21)) | p(s(bool,X20)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))))) ) | $spl66),
  inference(superposition,[],[f10363,f592])).
fof(f37523,plain,(
  ( ! [X19,X18] : (p(s(bool,X19)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) | p(s(bool,X18))) ) | $spl66),
  inference(subsumption_resolution,[],[f37507,f591])).
fof(f37507,plain,(
  ( ! [X19,X18] : (s(bool,X18) != s(bool,X19) | p(s(bool,X19)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) | p(s(bool,X18))) ) | $spl66),
  inference(superposition,[],[f10363,f591])).
fof(f37520,plain,(
  ( ! [X17,X16] : (p(s(bool,X17)) | p(s(bool,X16)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))))) ) | $spl66),
  inference(subsumption_resolution,[],[f37506,f591])).
fof(f37506,plain,(
  ( ! [X17,X16] : (s(bool,X16) != s(bool,X17) | p(s(bool,X17)) | p(s(bool,X16)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))))) ) | $spl66),
  inference(superposition,[],[f10363,f591])).
fof(f36494,plain,(
  ( ! [X215,X213,X214] : (s(bool,X213) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | p(s(bool,X214)) | p(s(bool,X215)) | s(bool,X213) != s(bool,X215)) )),
  inference(resolution,[],[f1655,f648])).
fof(f36428,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) )),
  inference(resolution,[],[f1655,f484])).
fof(f36425,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,X16) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) )),
  inference(resolution,[],[f1655,f483])).
fof(f36424,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) )),
  inference(resolution,[],[f1655,f482])).
fof(f36423,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) )),
  inference(resolution,[],[f1655,f480])).
fof(f36421,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) )),
  inference(resolution,[],[f1655,f478])).
fof(f35032,plain,(
  ( ! [X222,X223,X224] : (s(bool,X222) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | p(s(bool,X223)) | ~p(s(bool,X224)) | s(bool,X222) != s(bool,X224)) )),
  inference(resolution,[],[f860,f648])).
fof(f34963,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) )),
  inference(resolution,[],[f860,f484])).
fof(f34961,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,X20) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X20) != s(bool,X22)) )),
  inference(resolution,[],[f860,f483])).
fof(f34959,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) )),
  inference(resolution,[],[f860,f482])).
fof(f34958,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) )),
  inference(resolution,[],[f860,f480])).
fof(f34956,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) )),
  inference(resolution,[],[f860,f478])).
fof(f34007,plain,(
  ( ! [X222,X223,X224] : (s(bool,X222) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | p(s(bool,X223)) | ~p(s(bool,X224)) | s(bool,X222) != s(bool,X224)) )),
  inference(resolution,[],[f737,f648])).
fof(f33938,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) )),
  inference(resolution,[],[f737,f484])).
fof(f33936,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,X20) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X20) != s(bool,X22)) )),
  inference(resolution,[],[f737,f483])).
fof(f33934,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) )),
  inference(resolution,[],[f737,f482])).
fof(f33933,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) )),
  inference(resolution,[],[f737,f480])).
fof(f33931,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) )),
  inference(resolution,[],[f737,f478])).
fof(f27686,plain,(
  ( ! [X147,X148] : (p(s(bool,X147)) | ~p(s(bool,X148)) | s(bool,X148) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) ) | $spl28),
  inference(resolution,[],[f7209,f648])).
fof(f27617,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11))) | s(bool,X22) != s(bool,X24)) ) | $spl28),
  inference(resolution,[],[f7209,f484])).
fof(f27615,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) ) | $spl28),
  inference(resolution,[],[f7209,f483])).
fof(f27613,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) ) | $spl28),
  inference(resolution,[],[f7209,f482])).
fof(f27612,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11))) | s(bool,X7) != s(bool,X9)) ) | $spl28),
  inference(resolution,[],[f7209,f480])).
fof(f27610,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) ) | $spl28),
  inference(resolution,[],[f7209,f478])).
fof(f27561,plain,(
  ( ! [X140,X141] : (p(s(bool,X140)) | p(s(bool,X141)) | s(bool,X141) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl28),
  inference(resolution,[],[f7202,f648])).
fof(f27495,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X22) != s(bool,X23)) ) | $spl28),
  inference(resolution,[],[f7202,f484])).
fof(f27492,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl28),
  inference(resolution,[],[f7202,f483])).
fof(f27491,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl28),
  inference(resolution,[],[f7202,f482])).
fof(f27490,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X7) != s(bool,X8)) ) | $spl28),
  inference(resolution,[],[f7202,f480])).
fof(f27488,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl28),
  inference(resolution,[],[f7202,f478])).
fof(f27444,plain,(
  ( ! [X146,X147] : (p(s(bool,X146)) | ~p(s(bool,X147)) | s(bool,X147) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) ) | $spl28),
  inference(resolution,[],[f7199,f648])).
fof(f27375,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11))) | s(bool,X21) != s(bool,X23)) ) | $spl28),
  inference(resolution,[],[f7199,f484])).
fof(f27373,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) ) | $spl28),
  inference(resolution,[],[f7199,f483])).
fof(f27371,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) ) | $spl28),
  inference(resolution,[],[f7199,f482])).
fof(f27370,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11))) | s(bool,X6) != s(bool,X8)) ) | $spl28),
  inference(resolution,[],[f7199,f480])).
fof(f27368,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) ) | $spl28),
  inference(resolution,[],[f7199,f478])).
fof(f27317,plain,(
  ( ! [X144,X143] : (p(s(bool,X143)) | p(s(bool,X144)) | s(bool,X144) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl28),
  inference(resolution,[],[f2018,f648])).
fof(f27251,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X25) != s(bool,X26)) ) | $spl28),
  inference(resolution,[],[f2018,f484])).
fof(f27248,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl28),
  inference(resolution,[],[f2018,f483])).
fof(f27247,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl28),
  inference(resolution,[],[f2018,f482])).
fof(f27246,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X10) != s(bool,X11)) ) | $spl28),
  inference(resolution,[],[f2018,f480])).
fof(f27244,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl28),
  inference(resolution,[],[f2018,f478])).
fof(f27174,plain,(
  ( ! [X147,X148] : (p(s(bool,X147)) | ~p(s(bool,X148)) | s(bool,X148) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl29),
  inference(resolution,[],[f2017,f648])).
fof(f27105,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X22) != s(bool,X24)) ) | $spl29),
  inference(resolution,[],[f2017,f484])).
fof(f27103,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl29),
  inference(resolution,[],[f2017,f483])).
fof(f27101,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl29),
  inference(resolution,[],[f2017,f482])).
fof(f27100,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X7) != s(bool,X9)) ) | $spl29),
  inference(resolution,[],[f2017,f480])).
fof(f27098,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl29),
  inference(resolution,[],[f2017,f478])).
fof(f27216,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK11) | $spl88),
  inference(cnf_transformation,[],[f27216_D])).
fof(f27216_D,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK11) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f27218,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) | $spl90),
  inference(cnf_transformation,[],[f27218_D])).
fof(f27218_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f27097,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK11) | $spl29),
  inference(resolution,[],[f2017,f629])).
fof(f27096,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK11) | $spl29),
  inference(resolution,[],[f2017,f629])).
fof(f27214,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) | $spl87),
  inference(cnf_transformation,[],[f27214_D])).
fof(f27214_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) <=> ~$spl87),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl87])])).
fof(f27092,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK11) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) | $spl29),
  inference(resolution,[],[f2017,f471])).
fof(f27091,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK11) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) | $spl29),
  inference(resolution,[],[f2017,f466])).
fof(f27047,plain,(
  ( ! [X147,X148] : (p(s(bool,X147)) | ~p(s(bool,X148)) | s(bool,X148) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) ) | $spl28),
  inference(resolution,[],[f7209,f648])).
fof(f26978,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11))) | s(bool,X22) != s(bool,X24)) ) | $spl28),
  inference(resolution,[],[f7209,f484])).
fof(f26976,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) ) | $spl28),
  inference(resolution,[],[f7209,f483])).
fof(f26974,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) ) | $spl28),
  inference(resolution,[],[f7209,f482])).
fof(f26973,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11))) | s(bool,X7) != s(bool,X9)) ) | $spl28),
  inference(resolution,[],[f7209,f480])).
fof(f26971,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) ) | $spl28),
  inference(resolution,[],[f7209,f478])).
fof(f26922,plain,(
  ( ! [X140,X141] : (p(s(bool,X140)) | p(s(bool,X141)) | s(bool,X141) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl28),
  inference(resolution,[],[f7202,f648])).
fof(f26856,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X22) != s(bool,X23)) ) | $spl28),
  inference(resolution,[],[f7202,f484])).
fof(f26853,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl28),
  inference(resolution,[],[f7202,f483])).
fof(f26852,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl28),
  inference(resolution,[],[f7202,f482])).
fof(f26851,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X7) != s(bool,X8)) ) | $spl28),
  inference(resolution,[],[f7202,f480])).
fof(f26849,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl28),
  inference(resolution,[],[f7202,f478])).
fof(f26805,plain,(
  ( ! [X146,X147] : (p(s(bool,X146)) | ~p(s(bool,X147)) | s(bool,X147) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) ) | $spl28),
  inference(resolution,[],[f7199,f648])).
fof(f26736,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11))) | s(bool,X21) != s(bool,X23)) ) | $spl28),
  inference(resolution,[],[f7199,f484])).
fof(f26734,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) ) | $spl28),
  inference(resolution,[],[f7199,f483])).
fof(f26732,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) ) | $spl28),
  inference(resolution,[],[f7199,f482])).
fof(f26731,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11))) | s(bool,X6) != s(bool,X8)) ) | $spl28),
  inference(resolution,[],[f7199,f480])).
fof(f26729,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) ) | $spl28),
  inference(resolution,[],[f7199,f478])).
fof(f26678,plain,(
  ( ! [X146,X145] : (p(s(bool,X145)) | p(s(bool,X146)) | s(bool,X146) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl56),
  inference(resolution,[],[f7060,f648])).
fof(f26612,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | s(bool,X27) != s(bool,X28)) ) | $spl56),
  inference(resolution,[],[f7060,f484])).
fof(f26609,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X19) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl56),
  inference(resolution,[],[f7060,f483])).
fof(f26608,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl56),
  inference(resolution,[],[f7060,f482])).
fof(f26607,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | s(bool,X12) != s(bool,X13)) ) | $spl56),
  inference(resolution,[],[f7060,f480])).
fof(f26605,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl56),
  inference(resolution,[],[f7060,f478])).
fof(f26550,plain,(
  ( ! [X146,X147] : (p(s(bool,X146)) | ~p(s(bool,X147)) | s(bool,X147) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl53),
  inference(resolution,[],[f7056,f648])).
fof(f26481,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X21) != s(bool,X23)) ) | $spl53),
  inference(resolution,[],[f7056,f484])).
fof(f26479,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl53),
  inference(resolution,[],[f7056,f483])).
fof(f26477,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl53),
  inference(resolution,[],[f7056,f482])).
fof(f26476,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X6) != s(bool,X8)) ) | $spl53),
  inference(resolution,[],[f7056,f480])).
fof(f26474,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl53),
  inference(resolution,[],[f7056,f478])).
fof(f26217,plain,(
  ( ! [X140,X141] : (p(s(bool,X140)) | p(s(bool,X141)) | s(bool,X141) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl17),
  inference(resolution,[],[f6930,f648])).
fof(f26151,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | s(bool,X22) != s(bool,X23)) ) | $spl17),
  inference(resolution,[],[f6930,f484])).
fof(f26148,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl17),
  inference(resolution,[],[f6930,f483])).
fof(f26147,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl17),
  inference(resolution,[],[f6930,f482])).
fof(f26146,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | s(bool,X7) != s(bool,X8)) ) | $spl17),
  inference(resolution,[],[f6930,f480])).
fof(f26144,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl17),
  inference(resolution,[],[f6930,f478])).
fof(f25863,plain,(
  ( ! [X146,X147] : (p(s(bool,X146)) | ~p(s(bool,X147)) | s(bool,X147) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) | $spl12),
  inference(resolution,[],[f3242,f648])).
fof(f25794,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | s(bool,X21) != s(bool,X23)) ) | $spl12),
  inference(resolution,[],[f3242,f484])).
fof(f25792,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) | $spl12),
  inference(resolution,[],[f3242,f483])).
fof(f25790,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) | $spl12),
  inference(resolution,[],[f3242,f482])).
fof(f25789,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | s(bool,X6) != s(bool,X8)) ) | $spl12),
  inference(resolution,[],[f3242,f480])).
fof(f25787,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) | $spl12),
  inference(resolution,[],[f3242,f478])).
fof(f24708,plain,(
  ( ! [X147,X148] : (p(s(bool,X147)) | ~p(s(bool,X148)) | s(bool,X148) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl17),
  inference(resolution,[],[f1402,f648])).
fof(f24639,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X22) != s(bool,X24)) ) | $spl17),
  inference(resolution,[],[f1402,f484])).
fof(f24637,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl17),
  inference(resolution,[],[f1402,f483])).
fof(f24635,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl17),
  inference(resolution,[],[f1402,f482])).
fof(f24634,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X7) != s(bool,X9)) ) | $spl17),
  inference(resolution,[],[f1402,f480])).
fof(f24632,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl17),
  inference(resolution,[],[f1402,f478])).
fof(f24581,plain,(
  ( ! [X144,X143] : (p(s(bool,X143)) | p(s(bool,X144)) | s(bool,X144) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) | $spl12),
  inference(resolution,[],[f906,f648])).
fof(f24515,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | s(bool,X25) != s(bool,X26)) ) | $spl12),
  inference(resolution,[],[f906,f484])).
fof(f24512,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) | $spl12),
  inference(resolution,[],[f906,f483])).
fof(f24511,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) | $spl12),
  inference(resolution,[],[f906,f482])).
fof(f24510,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | s(bool,X10) != s(bool,X11)) ) | $spl12),
  inference(resolution,[],[f906,f480])).
fof(f24508,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) | $spl12),
  inference(resolution,[],[f906,f478])).
fof(f24244,plain,(
  ( ! [X147,X148] : (p(s(bool,X147)) | ~p(s(bool,X148)) | s(bool,X148) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f788,f648])).
fof(f24175,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | s(bool,X22) != s(bool,X24)) )),
  inference(resolution,[],[f788,f484])).
fof(f24173,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f788,f483])).
fof(f24171,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f788,f482])).
fof(f24170,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | s(bool,X7) != s(bool,X9)) )),
  inference(resolution,[],[f788,f480])).
fof(f24168,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f788,f478])).
fof(f23912,plain,(
  ( ! [X147,X148] : (p(s(bool,X147)) | ~p(s(bool,X148)) | s(bool,X148) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f664,f648])).
fof(f23843,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | s(bool,X22) != s(bool,X24)) )),
  inference(resolution,[],[f664,f484])).
fof(f23841,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f664,f483])).
fof(f23839,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f664,f482])).
fof(f23838,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | s(bool,X7) != s(bool,X9)) )),
  inference(resolution,[],[f664,f480])).
fof(f23836,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f664,f478])).
fof(f23787,plain,(
  ( ! [X140,X141] : (p(s(bool,X140)) | p(s(bool,X141)) | s(bool,X141) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f657,f648])).
fof(f23721,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | s(bool,X22) != s(bool,X23)) )),
  inference(resolution,[],[f657,f484])).
fof(f23718,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f657,f483])).
fof(f23717,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f657,f482])).
fof(f23716,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | s(bool,X7) != s(bool,X8)) )),
  inference(resolution,[],[f657,f480])).
fof(f23714,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f657,f478])).
fof(f23670,plain,(
  ( ! [X146,X147] : (p(s(bool,X146)) | ~p(s(bool,X147)) | s(bool,X147) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f654,f648])).
fof(f23601,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | s(bool,X21) != s(bool,X23)) )),
  inference(resolution,[],[f654,f484])).
fof(f23599,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f654,f483])).
fof(f23597,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f654,f482])).
fof(f23596,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | s(bool,X6) != s(bool,X8)) )),
  inference(resolution,[],[f654,f480])).
fof(f23594,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f654,f478])).
fof(f23542,plain,(
  ( ! [X140,X141] : (p(s(bool,X140)) | p(s(bool,X141)) | s(bool,X141) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f458,f648])).
fof(f23476,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | s(bool,X22) != s(bool,X23)) )),
  inference(resolution,[],[f458,f484])).
fof(f23473,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f458,f483])).
fof(f23472,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f458,f482])).
fof(f23471,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | s(bool,X7) != s(bool,X8)) )),
  inference(resolution,[],[f458,f480])).
fof(f23469,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f458,f478])).
fof(f23424,plain,(
  ( ! [X144,X143] : (p(s(bool,X143)) | p(s(bool,X144)) | s(bool,X144) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f457,f648])).
fof(f23358,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | s(bool,X25) != s(bool,X26)) )),
  inference(resolution,[],[f457,f484])).
fof(f23355,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f457,f483])).
fof(f23354,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f457,f482])).
fof(f23353,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | s(bool,X10) != s(bool,X11)) )),
  inference(resolution,[],[f457,f480])).
fof(f23351,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f457,f478])).
fof(f7056,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) | $spl53),
  inference(cnf_transformation,[],[f7056_D])).
fof(f7056_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) <=> ~$spl53),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl53])])).
fof(f7058,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | $spl54),
  inference(cnf_transformation,[],[f7058_D])).
fof(f7058_D,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f20278,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) | $spl17),
  inference(resolution,[],[f1402,f466])).
fof(f20279,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) | $spl17),
  inference(resolution,[],[f1402,f471])).
fof(f23267,plain,(
  ( ! [X140,X141] : (p(s(bool,X140)) | p(s(bool,X141)) | s(bool,X141) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl54),
  inference(resolution,[],[f22385,f648])).
fof(f22385,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) | $spl54),
  inference(backward_demodulation,[],[f7058,f657])).
fof(f23201,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X22) != s(bool,X23)) ) | $spl54),
  inference(resolution,[],[f22385,f484])).
fof(f23198,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl54),
  inference(resolution,[],[f22385,f483])).
fof(f23197,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl54),
  inference(resolution,[],[f22385,f482])).
fof(f23196,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X7) != s(bool,X8)) ) | $spl54),
  inference(resolution,[],[f22385,f480])).
fof(f23194,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl54),
  inference(resolution,[],[f22385,f478])).
fof(f23150,plain,(
  ( ! [X146,X147] : (p(s(bool,X146)) | ~p(s(bool,X147)) | s(bool,X147) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) ) | $spl54),
  inference(resolution,[],[f22384,f648])).
fof(f22384,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) | $spl54),
  inference(backward_demodulation,[],[f7058,f654])).
fof(f23081,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11))) | s(bool,X21) != s(bool,X23)) ) | $spl54),
  inference(resolution,[],[f22384,f484])).
fof(f23079,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) ) | $spl54),
  inference(resolution,[],[f22384,f483])).
fof(f23077,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) ) | $spl54),
  inference(resolution,[],[f22384,f482])).
fof(f23076,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11))) | s(bool,X6) != s(bool,X8)) ) | $spl54),
  inference(resolution,[],[f22384,f480])).
fof(f23074,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) ) | $spl54),
  inference(resolution,[],[f22384,f478])).
fof(f23023,plain,(
  ( ! [X146,X145] : (p(s(bool,X145)) | p(s(bool,X146)) | s(bool,X146) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | ($spl12 | $spl54)),
  inference(resolution,[],[f9846,f648])).
fof(f22957,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | s(bool,X27) != s(bool,X28)) ) | ($spl12 | $spl54)),
  inference(resolution,[],[f9846,f484])).
fof(f22954,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X19) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | ($spl12 | $spl54)),
  inference(resolution,[],[f9846,f483])).
fof(f22953,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | ($spl12 | $spl54)),
  inference(resolution,[],[f9846,f482])).
fof(f22952,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | s(bool,X12) != s(bool,X13)) ) | ($spl12 | $spl54)),
  inference(resolution,[],[f9846,f480])).
fof(f22950,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | ($spl12 | $spl54)),
  inference(resolution,[],[f9846,f478])).
fof(f22893,plain,(
  ( ! [X140,X141] : (p(s(bool,X140)) | p(s(bool,X141)) | s(bool,X141) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | $spl54),
  inference(resolution,[],[f22383,f648])).
fof(f22383,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) | $spl54),
  inference(backward_demodulation,[],[f7058,f458])).
fof(f22827,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | s(bool,X22) != s(bool,X23)) ) | $spl54),
  inference(resolution,[],[f22383,f484])).
fof(f22824,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | $spl54),
  inference(resolution,[],[f22383,f483])).
fof(f22823,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | $spl54),
  inference(resolution,[],[f22383,f482])).
fof(f22822,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | s(bool,X7) != s(bool,X8)) ) | $spl54),
  inference(resolution,[],[f22383,f480])).
fof(f22820,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | $spl54),
  inference(resolution,[],[f22383,f478])).
fof(f22641,plain,(
  $false | ($spl53 | $spl54)),
  inference(subsumption_resolution,[],[f22466,f476])).
fof(f22466,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))) | ($spl53 | $spl54)),
  inference(backward_demodulation,[],[f7058,f7056])).
fof(f7060,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) | $spl56),
  inference(cnf_transformation,[],[f7060_D])).
fof(f7060_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f20283,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | $spl17),
  inference(resolution,[],[f1402,f629])).
fof(f20284,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | $spl17),
  inference(resolution,[],[f1402,f629])).
fof(f21792,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | $spl17),
  inference(resolution,[],[f6930,f471])).
fof(f21793,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) | $spl17),
  inference(resolution,[],[f6930,f466])).
fof(f22331,plain,(
  ( ! [X146,X145] : (p(s(bool,X145)) | p(s(bool,X146)) | s(bool,X146) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl56),
  inference(resolution,[],[f7060,f648])).
fof(f22265,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | s(bool,X27) != s(bool,X28)) ) | $spl56),
  inference(resolution,[],[f7060,f484])).
fof(f22262,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X19) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl56),
  inference(resolution,[],[f7060,f483])).
fof(f22261,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl56),
  inference(resolution,[],[f7060,f482])).
fof(f22260,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | s(bool,X12) != s(bool,X13)) ) | $spl56),
  inference(resolution,[],[f7060,f480])).
fof(f22258,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl56),
  inference(resolution,[],[f7060,f478])).
fof(f22203,plain,(
  ( ! [X146,X147] : (p(s(bool,X146)) | ~p(s(bool,X147)) | s(bool,X147) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl53),
  inference(resolution,[],[f7056,f648])).
fof(f22134,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X21) != s(bool,X23)) ) | $spl53),
  inference(resolution,[],[f7056,f484])).
fof(f22132,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl53),
  inference(resolution,[],[f7056,f483])).
fof(f22130,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl53),
  inference(resolution,[],[f7056,f482])).
fof(f22129,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X6) != s(bool,X8)) ) | $spl53),
  inference(resolution,[],[f7056,f480])).
fof(f22127,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl53),
  inference(resolution,[],[f7056,f478])).
fof(f21870,plain,(
  ( ! [X140,X141] : (p(s(bool,X140)) | p(s(bool,X141)) | s(bool,X141) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl17),
  inference(resolution,[],[f6930,f648])).
fof(f21804,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | s(bool,X22) != s(bool,X23)) ) | $spl17),
  inference(resolution,[],[f6930,f484])).
fof(f21801,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl17),
  inference(resolution,[],[f6930,f483])).
fof(f21800,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl17),
  inference(resolution,[],[f6930,f482])).
fof(f21799,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | s(bool,X7) != s(bool,X8)) ) | $spl17),
  inference(resolution,[],[f6930,f480])).
fof(f21797,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl17),
  inference(resolution,[],[f6930,f478])).
fof(f21516,plain,(
  ( ! [X146,X147] : (p(s(bool,X146)) | ~p(s(bool,X147)) | s(bool,X147) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) | $spl12),
  inference(resolution,[],[f3242,f648])).
fof(f21447,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | s(bool,X21) != s(bool,X23)) ) | $spl12),
  inference(resolution,[],[f3242,f484])).
fof(f21445,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) | $spl12),
  inference(resolution,[],[f3242,f483])).
fof(f21443,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) | $spl12),
  inference(resolution,[],[f3242,f482])).
fof(f21442,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | s(bool,X6) != s(bool,X8)) ) | $spl12),
  inference(resolution,[],[f3242,f480])).
fof(f21440,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) | $spl12),
  inference(resolution,[],[f3242,f478])).
fof(f20361,plain,(
  ( ! [X147,X148] : (p(s(bool,X147)) | ~p(s(bool,X148)) | s(bool,X148) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl17),
  inference(resolution,[],[f1402,f648])).
fof(f20292,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X22) != s(bool,X24)) ) | $spl17),
  inference(resolution,[],[f1402,f484])).
fof(f20290,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl17),
  inference(resolution,[],[f1402,f483])).
fof(f20288,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl17),
  inference(resolution,[],[f1402,f482])).
fof(f20287,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X7) != s(bool,X9)) ) | $spl17),
  inference(resolution,[],[f1402,f480])).
fof(f20285,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl17),
  inference(resolution,[],[f1402,f478])).
fof(f20234,plain,(
  ( ! [X144,X143] : (p(s(bool,X143)) | p(s(bool,X144)) | s(bool,X144) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) | $spl12),
  inference(resolution,[],[f906,f648])).
fof(f20168,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | s(bool,X25) != s(bool,X26)) ) | $spl12),
  inference(resolution,[],[f906,f484])).
fof(f20165,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) | $spl12),
  inference(resolution,[],[f906,f483])).
fof(f20164,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) | $spl12),
  inference(resolution,[],[f906,f482])).
fof(f20163,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | s(bool,X10) != s(bool,X11)) ) | $spl12),
  inference(resolution,[],[f906,f480])).
fof(f20161,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) | $spl12),
  inference(resolution,[],[f906,f478])).
fof(f19897,plain,(
  ( ! [X147,X148] : (p(s(bool,X147)) | ~p(s(bool,X148)) | s(bool,X148) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f788,f648])).
fof(f19828,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | s(bool,X22) != s(bool,X24)) )),
  inference(resolution,[],[f788,f484])).
fof(f19826,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f788,f483])).
fof(f19824,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f788,f482])).
fof(f19823,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | s(bool,X7) != s(bool,X9)) )),
  inference(resolution,[],[f788,f480])).
fof(f19821,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f788,f478])).
fof(f19565,plain,(
  ( ! [X147,X148] : (p(s(bool,X147)) | ~p(s(bool,X148)) | s(bool,X148) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f664,f648])).
fof(f19496,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | s(bool,X22) != s(bool,X24)) )),
  inference(resolution,[],[f664,f484])).
fof(f19494,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f664,f483])).
fof(f19492,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f664,f482])).
fof(f19491,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | s(bool,X7) != s(bool,X9)) )),
  inference(resolution,[],[f664,f480])).
fof(f19489,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f664,f478])).
fof(f19440,plain,(
  ( ! [X140,X141] : (p(s(bool,X140)) | p(s(bool,X141)) | s(bool,X141) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f657,f648])).
fof(f19374,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | s(bool,X22) != s(bool,X23)) )),
  inference(resolution,[],[f657,f484])).
fof(f19371,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f657,f483])).
fof(f19370,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f657,f482])).
fof(f19369,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | s(bool,X7) != s(bool,X8)) )),
  inference(resolution,[],[f657,f480])).
fof(f19367,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f657,f478])).
fof(f19323,plain,(
  ( ! [X146,X147] : (p(s(bool,X146)) | ~p(s(bool,X147)) | s(bool,X147) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f654,f648])).
fof(f19254,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | s(bool,X21) != s(bool,X23)) )),
  inference(resolution,[],[f654,f484])).
fof(f19252,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f654,f483])).
fof(f19250,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f654,f482])).
fof(f19249,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | s(bool,X6) != s(bool,X8)) )),
  inference(resolution,[],[f654,f480])).
fof(f19247,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f654,f478])).
fof(f19195,plain,(
  ( ! [X140,X141] : (p(s(bool,X140)) | p(s(bool,X141)) | s(bool,X141) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f458,f648])).
fof(f19129,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | s(bool,X22) != s(bool,X23)) )),
  inference(resolution,[],[f458,f484])).
fof(f19126,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f458,f483])).
fof(f19125,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f458,f482])).
fof(f19124,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | s(bool,X7) != s(bool,X8)) )),
  inference(resolution,[],[f458,f480])).
fof(f19122,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f458,f478])).
fof(f19077,plain,(
  ( ! [X144,X143] : (p(s(bool,X143)) | p(s(bool,X144)) | s(bool,X144) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f457,f648])).
fof(f19011,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | s(bool,X25) != s(bool,X26)) )),
  inference(resolution,[],[f457,f484])).
fof(f19008,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f457,f483])).
fof(f19007,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f457,f482])).
fof(f19006,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | s(bool,X10) != s(bool,X11)) )),
  inference(resolution,[],[f457,f480])).
fof(f19004,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f457,f478])).
fof(f18920,plain,(
  ( ! [X146,X147] : (p(s(bool,X146)) | ~p(s(bool,X147)) | s(bool,X147) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) ) | $spl54),
  inference(resolution,[],[f18167,f648])).
fof(f18167,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) | $spl54),
  inference(backward_demodulation,[],[f7058,f654])).
fof(f18851,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11))) | s(bool,X21) != s(bool,X23)) ) | $spl54),
  inference(resolution,[],[f18167,f484])).
fof(f18849,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) ) | $spl54),
  inference(resolution,[],[f18167,f483])).
fof(f18847,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) ) | $spl54),
  inference(resolution,[],[f18167,f482])).
fof(f18846,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11))) | s(bool,X6) != s(bool,X8)) ) | $spl54),
  inference(resolution,[],[f18167,f480])).
fof(f18844,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK11)))) ) | $spl54),
  inference(resolution,[],[f18167,f478])).
fof(f18793,plain,(
  ( ! [X146,X145] : (p(s(bool,X145)) | p(s(bool,X146)) | s(bool,X146) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | ($spl12 | $spl54)),
  inference(resolution,[],[f9846,f648])).
fof(f18727,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | s(bool,X27) != s(bool,X28)) ) | ($spl12 | $spl54)),
  inference(resolution,[],[f9846,f484])).
fof(f18724,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X19) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | ($spl12 | $spl54)),
  inference(resolution,[],[f9846,f483])).
fof(f18723,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | ($spl12 | $spl54)),
  inference(resolution,[],[f9846,f482])).
fof(f18722,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | s(bool,X12) != s(bool,X13)) ) | ($spl12 | $spl54)),
  inference(resolution,[],[f9846,f480])).
fof(f18720,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | ($spl12 | $spl54)),
  inference(resolution,[],[f9846,f478])).
fof(f18663,plain,(
  ( ! [X140,X141] : (p(s(bool,X140)) | p(s(bool,X141)) | s(bool,X141) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | $spl54),
  inference(resolution,[],[f18166,f648])).
fof(f18166,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) | $spl54),
  inference(backward_demodulation,[],[f7058,f458])).
fof(f18597,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | s(bool,X22) != s(bool,X23)) ) | $spl54),
  inference(resolution,[],[f18166,f484])).
fof(f18594,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | $spl54),
  inference(resolution,[],[f18166,f483])).
fof(f18593,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | $spl54),
  inference(resolution,[],[f18166,f482])).
fof(f18592,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | s(bool,X7) != s(bool,X8)) ) | $spl54),
  inference(resolution,[],[f18166,f480])).
fof(f18590,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | $spl54),
  inference(resolution,[],[f18166,f478])).
fof(f18115,plain,(
  ( ! [X147,X148] : (p(s(bool,X147)) | ~p(s(bool,X148)) | s(bool,X148) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl16),
  inference(resolution,[],[f6594,f648])).
fof(f18046,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | s(bool,X22) != s(bool,X24)) ) | $spl16),
  inference(resolution,[],[f6594,f484])).
fof(f18044,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl16),
  inference(resolution,[],[f6594,f483])).
fof(f18042,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl16),
  inference(resolution,[],[f6594,f482])).
fof(f18041,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | s(bool,X7) != s(bool,X9)) ) | $spl16),
  inference(resolution,[],[f6594,f480])).
fof(f18039,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl16),
  inference(resolution,[],[f6594,f478])).
fof(f17990,plain,(
  ( ! [X140,X141] : (p(s(bool,X140)) | p(s(bool,X141)) | s(bool,X141) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl16),
  inference(resolution,[],[f6587,f648])).
fof(f17924,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X22) != s(bool,X23)) ) | $spl16),
  inference(resolution,[],[f6587,f484])).
fof(f17921,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl16),
  inference(resolution,[],[f6587,f483])).
fof(f17920,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl16),
  inference(resolution,[],[f6587,f482])).
fof(f17919,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X7) != s(bool,X8)) ) | $spl16),
  inference(resolution,[],[f6587,f480])).
fof(f17917,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl16),
  inference(resolution,[],[f6587,f478])).
fof(f17873,plain,(
  ( ! [X146,X147] : (p(s(bool,X146)) | ~p(s(bool,X147)) | s(bool,X147) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl16),
  inference(resolution,[],[f6584,f648])).
fof(f17804,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | s(bool,X21) != s(bool,X23)) ) | $spl16),
  inference(resolution,[],[f6584,f484])).
fof(f17802,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl16),
  inference(resolution,[],[f6584,f483])).
fof(f17800,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl16),
  inference(resolution,[],[f6584,f482])).
fof(f17799,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | s(bool,X6) != s(bool,X8)) ) | $spl16),
  inference(resolution,[],[f6584,f480])).
fof(f17797,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl16),
  inference(resolution,[],[f6584,f478])).
fof(f17480,plain,(
  ( ! [X146,X147] : (p(s(bool,X146)) | ~p(s(bool,X147)) | s(bool,X147) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) | $spl12),
  inference(resolution,[],[f3242,f648])).
fof(f17411,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | s(bool,X21) != s(bool,X23)) ) | $spl12),
  inference(resolution,[],[f3242,f484])).
fof(f17409,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) | $spl12),
  inference(resolution,[],[f3242,f483])).
fof(f17407,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) | $spl12),
  inference(resolution,[],[f3242,f482])).
fof(f17406,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | s(bool,X6) != s(bool,X8)) ) | $spl12),
  inference(resolution,[],[f3242,f480])).
fof(f17404,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) | $spl12),
  inference(resolution,[],[f3242,f478])).
fof(f16358,plain,(
  ( ! [X144,X143] : (p(s(bool,X143)) | p(s(bool,X144)) | s(bool,X144) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl16),
  inference(resolution,[],[f1403,f648])).
fof(f16292,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X25) != s(bool,X26)) ) | $spl16),
  inference(resolution,[],[f1403,f484])).
fof(f16289,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl16),
  inference(resolution,[],[f1403,f483])).
fof(f16288,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl16),
  inference(resolution,[],[f1403,f482])).
fof(f16287,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X10) != s(bool,X11)) ) | $spl16),
  inference(resolution,[],[f1403,f480])).
fof(f16285,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl16),
  inference(resolution,[],[f1403,f478])).
fof(f16231,plain,(
  ( ! [X144,X143] : (p(s(bool,X143)) | p(s(bool,X144)) | s(bool,X144) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) | $spl12),
  inference(resolution,[],[f906,f648])).
fof(f16165,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | s(bool,X25) != s(bool,X26)) ) | $spl12),
  inference(resolution,[],[f906,f484])).
fof(f16162,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) | $spl12),
  inference(resolution,[],[f906,f483])).
fof(f16161,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) | $spl12),
  inference(resolution,[],[f906,f482])).
fof(f16160,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | s(bool,X10) != s(bool,X11)) ) | $spl12),
  inference(resolution,[],[f906,f480])).
fof(f16158,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) | $spl12),
  inference(resolution,[],[f906,f478])).
fof(f15894,plain,(
  ( ! [X147,X148] : (p(s(bool,X147)) | ~p(s(bool,X148)) | s(bool,X148) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f788,f648])).
fof(f15825,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | s(bool,X22) != s(bool,X24)) )),
  inference(resolution,[],[f788,f484])).
fof(f15823,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f788,f483])).
fof(f15821,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f788,f482])).
fof(f15820,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | s(bool,X7) != s(bool,X9)) )),
  inference(resolution,[],[f788,f480])).
fof(f15818,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f788,f478])).
fof(f15562,plain,(
  ( ! [X147,X148] : (p(s(bool,X147)) | ~p(s(bool,X148)) | s(bool,X148) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f664,f648])).
fof(f15493,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | s(bool,X22) != s(bool,X24)) )),
  inference(resolution,[],[f664,f484])).
fof(f15491,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f664,f483])).
fof(f15489,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f664,f482])).
fof(f15488,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | s(bool,X7) != s(bool,X9)) )),
  inference(resolution,[],[f664,f480])).
fof(f15486,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f664,f478])).
fof(f15437,plain,(
  ( ! [X140,X141] : (p(s(bool,X140)) | p(s(bool,X141)) | s(bool,X141) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f657,f648])).
fof(f15371,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | s(bool,X22) != s(bool,X23)) )),
  inference(resolution,[],[f657,f484])).
fof(f15368,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f657,f483])).
fof(f15367,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f657,f482])).
fof(f15366,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | s(bool,X7) != s(bool,X8)) )),
  inference(resolution,[],[f657,f480])).
fof(f15364,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f657,f478])).
fof(f15320,plain,(
  ( ! [X146,X147] : (p(s(bool,X146)) | ~p(s(bool,X147)) | s(bool,X147) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f654,f648])).
fof(f15251,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | s(bool,X21) != s(bool,X23)) )),
  inference(resolution,[],[f654,f484])).
fof(f15249,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f654,f483])).
fof(f15247,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f654,f482])).
fof(f15246,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | s(bool,X6) != s(bool,X8)) )),
  inference(resolution,[],[f654,f480])).
fof(f15244,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f654,f478])).
fof(f15192,plain,(
  ( ! [X140,X141] : (p(s(bool,X140)) | p(s(bool,X141)) | s(bool,X141) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f458,f648])).
fof(f15126,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | s(bool,X22) != s(bool,X23)) )),
  inference(resolution,[],[f458,f484])).
fof(f15123,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f458,f483])).
fof(f15122,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f458,f482])).
fof(f15121,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | s(bool,X7) != s(bool,X8)) )),
  inference(resolution,[],[f458,f480])).
fof(f15119,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f458,f478])).
fof(f15074,plain,(
  ( ! [X144,X143] : (p(s(bool,X143)) | p(s(bool,X144)) | s(bool,X144) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f457,f648])).
fof(f15008,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | s(bool,X25) != s(bool,X26)) )),
  inference(resolution,[],[f457,f484])).
fof(f15005,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f457,f483])).
fof(f15004,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f457,f482])).
fof(f15003,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | s(bool,X10) != s(bool,X11)) )),
  inference(resolution,[],[f457,f480])).
fof(f15001,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f457,f478])).
fof(f14916,plain,(
  ( ! [X146,X145] : (p(s(bool,X145)) | p(s(bool,X146)) | s(bool,X146) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | ($spl12 | $spl54)),
  inference(resolution,[],[f9846,f648])).
fof(f14850,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | s(bool,X27) != s(bool,X28)) ) | ($spl12 | $spl54)),
  inference(resolution,[],[f9846,f484])).
fof(f14847,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X19) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | ($spl12 | $spl54)),
  inference(resolution,[],[f9846,f483])).
fof(f14846,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | ($spl12 | $spl54)),
  inference(resolution,[],[f9846,f482])).
fof(f14845,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | s(bool,X12) != s(bool,X13)) ) | ($spl12 | $spl54)),
  inference(resolution,[],[f9846,f480])).
fof(f14843,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | ($spl12 | $spl54)),
  inference(resolution,[],[f9846,f478])).
fof(f14786,plain,(
  ( ! [X140,X141] : (p(s(bool,X140)) | p(s(bool,X141)) | s(bool,X141) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | $spl54),
  inference(resolution,[],[f14279,f648])).
fof(f14279,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) | $spl54),
  inference(backward_demodulation,[],[f7058,f458])).
fof(f14720,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | s(bool,X22) != s(bool,X23)) ) | $spl54),
  inference(resolution,[],[f14279,f484])).
fof(f14717,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | $spl54),
  inference(resolution,[],[f14279,f483])).
fof(f14716,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | $spl54),
  inference(resolution,[],[f14279,f482])).
fof(f14715,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | s(bool,X7) != s(bool,X8)) ) | $spl54),
  inference(resolution,[],[f14279,f480])).
fof(f14713,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | $spl54),
  inference(resolution,[],[f14279,f478])).
fof(f13691,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) | $spl17),
  inference(resolution,[],[f1402,f466])).
fof(f13692,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) | $spl17),
  inference(resolution,[],[f1402,f471])).
fof(f13696,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | $spl17),
  inference(resolution,[],[f1402,f629])).
fof(f13697,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | $spl17),
  inference(resolution,[],[f1402,f629])).
fof(f13817,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | $spl17),
  inference(resolution,[],[f6930,f471])).
fof(f13818,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) | $spl17),
  inference(resolution,[],[f6930,f466])).
fof(f14228,plain,(
  ( ! [X146,X147] : (p(s(bool,X146)) | ~p(s(bool,X147)) | s(bool,X147) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl53),
  inference(resolution,[],[f7056,f648])).
fof(f14159,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X21) != s(bool,X23)) ) | $spl53),
  inference(resolution,[],[f7056,f484])).
fof(f14157,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl53),
  inference(resolution,[],[f7056,f483])).
fof(f14155,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl53),
  inference(resolution,[],[f7056,f482])).
fof(f14154,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X6) != s(bool,X8)) ) | $spl53),
  inference(resolution,[],[f7056,f480])).
fof(f14152,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl53),
  inference(resolution,[],[f7056,f478])).
fof(f13895,plain,(
  ( ! [X140,X141] : (p(s(bool,X140)) | p(s(bool,X141)) | s(bool,X141) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl17),
  inference(resolution,[],[f6930,f648])).
fof(f13829,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | s(bool,X22) != s(bool,X23)) ) | $spl17),
  inference(resolution,[],[f6930,f484])).
fof(f13826,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl17),
  inference(resolution,[],[f6930,f483])).
fof(f13825,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl17),
  inference(resolution,[],[f6930,f482])).
fof(f13824,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | s(bool,X7) != s(bool,X8)) ) | $spl17),
  inference(resolution,[],[f6930,f480])).
fof(f13822,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl17),
  inference(resolution,[],[f6930,f478])).
fof(f13774,plain,(
  ( ! [X147,X148] : (p(s(bool,X147)) | ~p(s(bool,X148)) | s(bool,X148) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl17),
  inference(resolution,[],[f1402,f648])).
fof(f13705,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X22) != s(bool,X24)) ) | $spl17),
  inference(resolution,[],[f1402,f484])).
fof(f13703,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl17),
  inference(resolution,[],[f1402,f483])).
fof(f13701,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl17),
  inference(resolution,[],[f1402,f482])).
fof(f13700,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X7) != s(bool,X9)) ) | $spl17),
  inference(resolution,[],[f1402,f480])).
fof(f13698,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl17),
  inference(resolution,[],[f1402,f478])).
fof(f13647,plain,(
  ( ! [X140,X141] : (p(s(bool,X140)) | p(s(bool,X141)) | s(bool,X141) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl16),
  inference(resolution,[],[f6587,f648])).
fof(f13581,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X22) != s(bool,X23)) ) | $spl16),
  inference(resolution,[],[f6587,f484])).
fof(f13578,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl16),
  inference(resolution,[],[f6587,f483])).
fof(f13577,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl16),
  inference(resolution,[],[f6587,f482])).
fof(f13576,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X7) != s(bool,X8)) ) | $spl16),
  inference(resolution,[],[f6587,f480])).
fof(f13574,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl16),
  inference(resolution,[],[f6587,f478])).
fof(f13530,plain,(
  ( ! [X146,X147] : (p(s(bool,X146)) | ~p(s(bool,X147)) | s(bool,X147) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl16),
  inference(resolution,[],[f6584,f648])).
fof(f13461,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | s(bool,X21) != s(bool,X23)) ) | $spl16),
  inference(resolution,[],[f6584,f484])).
fof(f13459,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl16),
  inference(resolution,[],[f6584,f483])).
fof(f13457,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl16),
  inference(resolution,[],[f6584,f482])).
fof(f13456,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | s(bool,X6) != s(bool,X8)) ) | $spl16),
  inference(resolution,[],[f6584,f480])).
fof(f13454,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl16),
  inference(resolution,[],[f6584,f478])).
fof(f13403,plain,(
  ( ! [X144,X143] : (p(s(bool,X143)) | p(s(bool,X144)) | s(bool,X144) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl16),
  inference(resolution,[],[f1403,f648])).
fof(f13337,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X25) != s(bool,X26)) ) | $spl16),
  inference(resolution,[],[f1403,f484])).
fof(f13334,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl16),
  inference(resolution,[],[f1403,f483])).
fof(f13333,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl16),
  inference(resolution,[],[f1403,f482])).
fof(f13332,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X10) != s(bool,X11)) ) | $spl16),
  inference(resolution,[],[f1403,f480])).
fof(f13330,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl16),
  inference(resolution,[],[f1403,f478])).
fof(f13274,plain,(
  ( ! [X140,X141] : (p(s(bool,X140)) | p(s(bool,X141)) | s(bool,X141) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl17),
  inference(resolution,[],[f6930,f648])).
fof(f13208,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | s(bool,X22) != s(bool,X23)) ) | $spl17),
  inference(resolution,[],[f6930,f484])).
fof(f13205,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl17),
  inference(resolution,[],[f6930,f483])).
fof(f13204,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl17),
  inference(resolution,[],[f6930,f482])).
fof(f13203,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | s(bool,X7) != s(bool,X8)) ) | $spl17),
  inference(resolution,[],[f6930,f480])).
fof(f13201,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl17),
  inference(resolution,[],[f6930,f478])).
fof(f12890,plain,(
  ( ! [X146,X147] : (p(s(bool,X146)) | ~p(s(bool,X147)) | s(bool,X147) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) | $spl12),
  inference(resolution,[],[f3242,f648])).
fof(f12821,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | s(bool,X21) != s(bool,X23)) ) | $spl12),
  inference(resolution,[],[f3242,f484])).
fof(f12819,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) | $spl12),
  inference(resolution,[],[f3242,f483])).
fof(f12817,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) | $spl12),
  inference(resolution,[],[f3242,f482])).
fof(f12816,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | s(bool,X6) != s(bool,X8)) ) | $spl12),
  inference(resolution,[],[f3242,f480])).
fof(f12814,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) | $spl12),
  inference(resolution,[],[f3242,f478])).
fof(f11765,plain,(
  ( ! [X147,X148] : (p(s(bool,X147)) | ~p(s(bool,X148)) | s(bool,X148) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl17),
  inference(resolution,[],[f1402,f648])).
fof(f11696,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X22) != s(bool,X24)) ) | $spl17),
  inference(resolution,[],[f1402,f484])).
fof(f11694,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl17),
  inference(resolution,[],[f1402,f483])).
fof(f11692,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl17),
  inference(resolution,[],[f1402,f482])).
fof(f11691,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X7) != s(bool,X9)) ) | $spl17),
  inference(resolution,[],[f1402,f480])).
fof(f11689,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl17),
  inference(resolution,[],[f1402,f478])).
fof(f11638,plain,(
  ( ! [X144,X143] : (p(s(bool,X143)) | p(s(bool,X144)) | s(bool,X144) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) | $spl12),
  inference(resolution,[],[f906,f648])).
fof(f11572,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | s(bool,X25) != s(bool,X26)) ) | $spl12),
  inference(resolution,[],[f906,f484])).
fof(f11569,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) | $spl12),
  inference(resolution,[],[f906,f483])).
fof(f11568,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) | $spl12),
  inference(resolution,[],[f906,f482])).
fof(f11567,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | s(bool,X10) != s(bool,X11)) ) | $spl12),
  inference(resolution,[],[f906,f480])).
fof(f11565,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) | $spl12),
  inference(resolution,[],[f906,f478])).
fof(f11301,plain,(
  ( ! [X147,X148] : (p(s(bool,X147)) | ~p(s(bool,X148)) | s(bool,X148) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f788,f648])).
fof(f11232,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | s(bool,X22) != s(bool,X24)) )),
  inference(resolution,[],[f788,f484])).
fof(f11230,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f788,f483])).
fof(f11228,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f788,f482])).
fof(f11227,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | s(bool,X7) != s(bool,X9)) )),
  inference(resolution,[],[f788,f480])).
fof(f11225,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f788,f478])).
fof(f10969,plain,(
  ( ! [X147,X148] : (p(s(bool,X147)) | ~p(s(bool,X148)) | s(bool,X148) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f664,f648])).
fof(f10900,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | s(bool,X22) != s(bool,X24)) )),
  inference(resolution,[],[f664,f484])).
fof(f10898,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f664,f483])).
fof(f10896,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f664,f482])).
fof(f10895,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | s(bool,X7) != s(bool,X9)) )),
  inference(resolution,[],[f664,f480])).
fof(f10893,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f664,f478])).
fof(f10844,plain,(
  ( ! [X140,X141] : (p(s(bool,X140)) | p(s(bool,X141)) | s(bool,X141) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f657,f648])).
fof(f10778,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | s(bool,X22) != s(bool,X23)) )),
  inference(resolution,[],[f657,f484])).
fof(f10775,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f657,f483])).
fof(f10774,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f657,f482])).
fof(f10773,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | s(bool,X7) != s(bool,X8)) )),
  inference(resolution,[],[f657,f480])).
fof(f10771,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f657,f478])).
fof(f10727,plain,(
  ( ! [X146,X147] : (p(s(bool,X146)) | ~p(s(bool,X147)) | s(bool,X147) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f654,f648])).
fof(f10658,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | s(bool,X21) != s(bool,X23)) )),
  inference(resolution,[],[f654,f484])).
fof(f10656,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f654,f483])).
fof(f10654,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f654,f482])).
fof(f10653,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | s(bool,X6) != s(bool,X8)) )),
  inference(resolution,[],[f654,f480])).
fof(f10651,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f654,f478])).
fof(f10599,plain,(
  ( ! [X140,X141] : (p(s(bool,X140)) | p(s(bool,X141)) | s(bool,X141) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f458,f648])).
fof(f10533,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | s(bool,X22) != s(bool,X23)) )),
  inference(resolution,[],[f458,f484])).
fof(f10530,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f458,f483])).
fof(f10529,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f458,f482])).
fof(f10528,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | s(bool,X7) != s(bool,X8)) )),
  inference(resolution,[],[f458,f480])).
fof(f10526,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f458,f478])).
fof(f10471,plain,(
  ( ! [X144,X143] : (p(s(bool,X143)) | p(s(bool,X144)) | s(bool,X144) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f457,f648])).
fof(f10405,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | s(bool,X25) != s(bool,X26)) )),
  inference(resolution,[],[f457,f484])).
fof(f10402,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f457,f483])).
fof(f10401,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f457,f482])).
fof(f10400,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | s(bool,X10) != s(bool,X11)) )),
  inference(resolution,[],[f457,f480])).
fof(f10398,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f457,f478])).
fof(f10322,plain,(
  ( ! [X140,X141] : (p(s(bool,X140)) | p(s(bool,X141)) | s(bool,X141) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | $spl54),
  inference(resolution,[],[f9815,f648])).
fof(f9815,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) | $spl54),
  inference(backward_demodulation,[],[f7058,f458])).
fof(f10256,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | s(bool,X22) != s(bool,X23)) ) | $spl54),
  inference(resolution,[],[f9815,f484])).
fof(f10253,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | $spl54),
  inference(resolution,[],[f9815,f483])).
fof(f10252,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | $spl54),
  inference(resolution,[],[f9815,f482])).
fof(f10251,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | s(bool,X7) != s(bool,X8)) ) | $spl54),
  inference(resolution,[],[f9815,f480])).
fof(f10249,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | $spl54),
  inference(resolution,[],[f9815,f478])).
fof(f9774,plain,(
  ( ! [X146,X147] : (p(s(bool,X146)) | ~p(s(bool,X147)) | s(bool,X147) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl16),
  inference(resolution,[],[f6584,f648])).
fof(f9705,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | s(bool,X21) != s(bool,X23)) ) | $spl16),
  inference(resolution,[],[f6584,f484])).
fof(f9703,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl16),
  inference(resolution,[],[f6584,f483])).
fof(f9701,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl16),
  inference(resolution,[],[f6584,f482])).
fof(f9700,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | s(bool,X6) != s(bool,X8)) ) | $spl16),
  inference(resolution,[],[f6584,f480])).
fof(f9698,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl16),
  inference(resolution,[],[f6584,f478])).
fof(f8615,plain,(
  ( ! [X146,X147] : (p(s(bool,X146)) | ~p(s(bool,X147)) | s(bool,X147) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) | $spl12),
  inference(resolution,[],[f3242,f648])).
fof(f8546,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | s(bool,X21) != s(bool,X23)) ) | $spl12),
  inference(resolution,[],[f3242,f484])).
fof(f8544,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) | $spl12),
  inference(resolution,[],[f3242,f483])).
fof(f8542,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) | $spl12),
  inference(resolution,[],[f3242,f482])).
fof(f8541,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | s(bool,X6) != s(bool,X8)) ) | $spl12),
  inference(resolution,[],[f3242,f480])).
fof(f8539,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) ) | $spl12),
  inference(resolution,[],[f3242,f478])).
fof(f7555,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f7472,f7357])).
fof(f7357,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl4),
  inference(equality_resolution,[],[f783])).
fof(f7472,plain,(
  ( ! [X275,X276] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X275),s('type.integer.int',X276))))) ) | $spl4),
  inference(resolution,[],[f7357,f475])).
fof(f7554,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f7471,f7357])).
fof(f7471,plain,(
  ( ! [X273,X274] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X273),s('type.integer.int',X274))))) ) | $spl4),
  inference(resolution,[],[f7357,f475])).
fof(f7553,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f7470,f7357])).
fof(f7470,plain,(
  ( ! [X271,X272] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X271),s('type.integer.int',X272))))) ) | $spl4),
  inference(resolution,[],[f7357,f474])).
fof(f7552,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f7469,f7357])).
fof(f7469,plain,(
  ( ! [X269,X270] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X269),s('type.integer.int',X270))))) ) | $spl4),
  inference(resolution,[],[f7357,f473])).
fof(f7551,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f7464,f7357])).
fof(f7464,plain,(
  ( ! [X259,X260] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X259),s('type.integer.int',X260))))) ) | $spl4),
  inference(resolution,[],[f7357,f465])).
fof(f7550,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f7463,f7357])).
fof(f7463,plain,(
  ( ! [X257,X258] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X257),s('type.integer.int',X258))))) ) | $spl4),
  inference(resolution,[],[f7357,f462])).
fof(f7549,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f7462,f7357])).
fof(f7462,plain,(
  ( ! [X255,X256] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X255),s('type.integer.int',X256))))) ) | $spl4),
  inference(resolution,[],[f7357,f460])).
fof(f7546,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f7450,f7357])).
fof(f7450,plain,(
  ( ! [X241,X242] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X241),s('type.integer.int',X242))))) ) | $spl4),
  inference(resolution,[],[f7357,f474])).
fof(f7545,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f7449,f7357])).
fof(f7449,plain,(
  ( ! [X239,X240] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X239),s('type.integer.int',X240))))) ) | $spl4),
  inference(resolution,[],[f7357,f473])).
fof(f7544,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f7446,f7357])).
fof(f7446,plain,(
  ( ! [X233,X234] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X233),s('type.integer.int',X234))))) ) | $spl4),
  inference(resolution,[],[f7357,f465])).
fof(f7543,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f7445,f7357])).
fof(f7445,plain,(
  ( ! [X231,X232] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X231),s('type.integer.int',X232))))) ) | $spl4),
  inference(resolution,[],[f7357,f462])).
fof(f7542,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f7444,f7357])).
fof(f7444,plain,(
  ( ! [X229,X230] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X229),s('type.integer.int',X230))))) ) | $spl4),
  inference(resolution,[],[f7357,f460])).
fof(f7457,plain,(
  $false | $spl4),
  inference(resolution,[],[f7357,f457])).
fof(f7458,plain,(
  $false | ($spl4 | $spl28)),
  inference(resolution,[],[f7357,f2018])).
fof(f7459,plain,(
  $false | ($spl4 | $spl12)),
  inference(resolution,[],[f7357,f906])).
fof(f7460,plain,(
  $false | ($spl4 | $spl16)),
  inference(resolution,[],[f7357,f1403])).
fof(f7461,plain,(
  $false | ($spl4 | $spl24)),
  inference(resolution,[],[f7357,f1703])).
fof(f7473,plain,(
  $false | $spl4),
  inference(resolution,[],[f7357,f476])).
fof(f7474,plain,(
  $false | $spl4),
  inference(resolution,[],[f7357,f657])).
fof(f7475,plain,(
  $false | $spl4),
  inference(resolution,[],[f7357,f458])).
fof(f7476,plain,(
  $false | $spl4),
  inference(resolution,[],[f7357,f1573])).
fof(f7286,plain,(
  ( ! [X144,X143] : (p(s(bool,X143)) | p(s(bool,X144)) | s(bool,X144) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl28),
  inference(resolution,[],[f2018,f648])).
fof(f7220,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X25) != s(bool,X26)) ) | $spl28),
  inference(resolution,[],[f2018,f484])).
fof(f7217,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl28),
  inference(resolution,[],[f2018,f483])).
fof(f7216,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl28),
  inference(resolution,[],[f2018,f482])).
fof(f7215,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X10) != s(bool,X11)) ) | $spl28),
  inference(resolution,[],[f2018,f480])).
fof(f7213,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl28),
  inference(resolution,[],[f2018,f478])).
fof(f7159,plain,(
  ( ! [X144,X143] : (p(s(bool,X143)) | p(s(bool,X144)) | s(bool,X144) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl16),
  inference(resolution,[],[f1403,f648])).
fof(f7093,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X25) != s(bool,X26)) ) | $spl16),
  inference(resolution,[],[f1403,f484])).
fof(f7090,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl16),
  inference(resolution,[],[f1403,f483])).
fof(f7089,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl16),
  inference(resolution,[],[f1403,f482])).
fof(f7088,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X10) != s(bool,X11)) ) | $spl16),
  inference(resolution,[],[f1403,f480])).
fof(f7086,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl16),
  inference(resolution,[],[f1403,f478])).
fof(f7016,plain,(
  ( ! [X147,X148] : (p(s(bool,X147)) | ~p(s(bool,X148)) | s(bool,X148) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl17),
  inference(resolution,[],[f1402,f648])).
fof(f6947,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X22) != s(bool,X24)) ) | $spl17),
  inference(resolution,[],[f1402,f484])).
fof(f6945,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl17),
  inference(resolution,[],[f1402,f483])).
fof(f6943,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl17),
  inference(resolution,[],[f1402,f482])).
fof(f6942,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X7) != s(bool,X9)) ) | $spl17),
  inference(resolution,[],[f1402,f480])).
fof(f6940,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl17),
  inference(resolution,[],[f1402,f478])).
fof(f6939,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | $spl17),
  inference(resolution,[],[f1402,f629])).
fof(f6938,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | $spl17),
  inference(resolution,[],[f1402,f629])).
fof(f6934,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) | $spl17),
  inference(resolution,[],[f1402,f471])).
fof(f6933,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) | $spl17),
  inference(resolution,[],[f1402,f466])).
fof(f6922,plain,(
  $false | $spl2),
  inference(subsumption_resolution,[],[f6850,f6742])).
fof(f6742,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl2),
  inference(equality_resolution,[],[f782])).
fof(f6850,plain,(
  ( ! [X265,X264] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X264),s('type.integer.int',X265))))) ) | $spl2),
  inference(resolution,[],[f6742,f461])).
fof(f6921,plain,(
  $false | $spl2),
  inference(subsumption_resolution,[],[f6849,f6742])).
fof(f6849,plain,(
  ( ! [X263,X262] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X262),s('type.integer.int',X263))))) ) | $spl2),
  inference(resolution,[],[f6742,f463])).
fof(f463,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) )),
  inference(cnf_transformation,[],[f377])).
fof(f6920,plain,(
  $false | $spl2),
  inference(subsumption_resolution,[],[f6848,f6742])).
fof(f6848,plain,(
  ( ! [X261,X260] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X260),s('type.integer.int',X261))))) ) | $spl2),
  inference(resolution,[],[f6742,f464])).
fof(f6919,plain,(
  $false | $spl2),
  inference(subsumption_resolution,[],[f6839,f6742])).
fof(f6839,plain,(
  ( ! [X251,X250] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X250),s('type.integer.int',X251))))) ) | $spl2),
  inference(resolution,[],[f6742,f628])).
fof(f628,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f356])).
fof(f356,plain,(
  ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(ennf_transformation,[],[f316])).
fof(f316,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) => ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(flattening,[],[f315])).
fof(f315,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) => ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(rectify,[],[f51])).
fof(f51,axiom,(
  ! [X6,X7] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X6),s('type.integer.int',X7)))) => ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X7),s('type.integer.int',X6)))))),
  file('thm.integer.INT_LTE_TRANS','thm.integer.INT_LT_GT')).
fof(f6918,plain,(
  $false | $spl2),
  inference(subsumption_resolution,[],[f6838,f6742])).
fof(f6838,plain,(
  ( ! [X249,X248] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X248),s('type.integer.int',X249))))) ) | $spl2),
  inference(resolution,[],[f6742,f628])).
fof(f6917,plain,(
  $false | $spl2),
  inference(subsumption_resolution,[],[f6837,f6742])).
fof(f6837,plain,(
  ( ! [X246,X247] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X246),s('type.integer.int',X247))))) ) | $spl2),
  inference(resolution,[],[f6742,f626])).
fof(f6916,plain,(
  $false | $spl2),
  inference(subsumption_resolution,[],[f6836,f6742])).
fof(f6836,plain,(
  ( ! [X244,X245] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X244),s('type.integer.int',X245))))) ) | $spl2),
  inference(resolution,[],[f6742,f626])).
fof(f6915,plain,(
  $false | $spl2),
  inference(subsumption_resolution,[],[f6828,f6742])).
fof(f6828,plain,(
  ( ! [X225,X226] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X225),s('type.integer.int',X226))))) ) | $spl2),
  inference(resolution,[],[f6742,f464])).
fof(f6914,plain,(
  $false | $spl2),
  inference(subsumption_resolution,[],[f6827,f6742])).
fof(f6827,plain,(
  ( ! [X223,X224] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X223),s('type.integer.int',X224))))) ) | $spl2),
  inference(resolution,[],[f6742,f463])).
fof(f6913,plain,(
  $false | $spl2),
  inference(subsumption_resolution,[],[f6826,f6742])).
fof(f6826,plain,(
  ( ! [X222,X221] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X221),s('type.integer.int',X222))))) ) | $spl2),
  inference(resolution,[],[f6742,f461])).
fof(f6840,plain,(
  $false | $spl2),
  inference(resolution,[],[f6742,f631])).
fof(f6843,plain,(
  $false | $spl2),
  inference(resolution,[],[f6742,f459])).
fof(f6844,plain,(
  $false | $spl2),
  inference(resolution,[],[f6742,f664])).
fof(f6845,plain,(
  $false | $spl2),
  inference(resolution,[],[f6742,f788])).
fof(f6851,plain,(
  $false | $spl2),
  inference(resolution,[],[f6742,f654])).
fof(f6852,plain,(
  $false | ($spl2 | $spl21)),
  inference(resolution,[],[f6742,f1699])).
fof(f6671,plain,(
  ( ! [X144,X143] : (p(s(bool,X143)) | p(s(bool,X144)) | s(bool,X144) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl16),
  inference(resolution,[],[f1403,f648])).
fof(f6605,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X25) != s(bool,X26)) ) | $spl16),
  inference(resolution,[],[f1403,f484])).
fof(f6602,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl16),
  inference(resolution,[],[f1403,f483])).
fof(f6601,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl16),
  inference(resolution,[],[f1403,f482])).
fof(f6600,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X10) != s(bool,X11)) ) | $spl16),
  inference(resolution,[],[f1403,f480])).
fof(f6598,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl16),
  inference(resolution,[],[f1403,f478])).
fof(f6543,plain,(
  ( ! [X144,X143] : (p(s(bool,X143)) | p(s(bool,X144)) | s(bool,X144) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11)))) ) | $spl24),
  inference(resolution,[],[f1703,f648])).
fof(f6477,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | s(bool,X25) != s(bool,X26)) ) | $spl24),
  inference(resolution,[],[f1703,f484])).
fof(f6474,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11)))) ) | $spl24),
  inference(resolution,[],[f1703,f483])).
fof(f6473,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11)))) ) | $spl24),
  inference(resolution,[],[f1703,f482])).
fof(f6472,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | s(bool,X10) != s(bool,X11)) ) | $spl24),
  inference(resolution,[],[f1703,f480])).
fof(f6470,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11)))) ) | $spl24),
  inference(resolution,[],[f1703,f478])).
fof(f6417,plain,(
  ( ! [X146,X147] : (p(s(bool,X146)) | ~p(s(bool,X147)) | s(bool,X147) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK13)))) ) | $spl21),
  inference(resolution,[],[f1699,f648])).
fof(f6348,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | s(bool,X21) != s(bool,X23)) ) | $spl21),
  inference(resolution,[],[f1699,f484])).
fof(f6346,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK13)))) ) | $spl21),
  inference(resolution,[],[f1699,f483])).
fof(f6344,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK13)))) ) | $spl21),
  inference(resolution,[],[f1699,f482])).
fof(f6343,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | s(bool,X6) != s(bool,X8)) ) | $spl21),
  inference(resolution,[],[f1699,f480])).
fof(f6341,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK13)))) ) | $spl21),
  inference(resolution,[],[f1699,f478])).
fof(f6084,plain,(
  ( ! [X140,X141] : (p(s(bool,X140)) | p(s(bool,X141)) | s(bool,X141) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f1573,f648])).
fof(f6018,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | s(bool,X22) != s(bool,X23)) )),
  inference(resolution,[],[f1573,f484])).
fof(f6015,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f1573,f483])).
fof(f6014,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f1573,f482])).
fof(f6013,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | s(bool,X7) != s(bool,X8)) )),
  inference(resolution,[],[f1573,f480])).
fof(f6011,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f1573,f478])).
fof(f5966,plain,(
  ( ! [X144,X143] : (p(s(bool,X143)) | p(s(bool,X144)) | s(bool,X144) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) | $spl12),
  inference(resolution,[],[f906,f648])).
fof(f5900,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | s(bool,X25) != s(bool,X26)) ) | $spl12),
  inference(resolution,[],[f906,f484])).
fof(f5897,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) | $spl12),
  inference(resolution,[],[f906,f483])).
fof(f5896,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) | $spl12),
  inference(resolution,[],[f906,f482])).
fof(f5895,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | s(bool,X10) != s(bool,X11)) ) | $spl12),
  inference(resolution,[],[f906,f480])).
fof(f5893,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) | $spl12),
  inference(resolution,[],[f906,f478])).
fof(f5629,plain,(
  ( ! [X147,X148] : (p(s(bool,X147)) | ~p(s(bool,X148)) | s(bool,X148) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f788,f648])).
fof(f5560,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | s(bool,X22) != s(bool,X24)) )),
  inference(resolution,[],[f788,f484])).
fof(f5558,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f788,f483])).
fof(f5556,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f788,f482])).
fof(f5555,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | s(bool,X7) != s(bool,X9)) )),
  inference(resolution,[],[f788,f480])).
fof(f5553,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f788,f478])).
fof(f5500,plain,(
  ( ! [X147,X148] : (p(s(bool,X147)) | ~p(s(bool,X148)) | s(bool,X148) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f459,f648])).
fof(f5431,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | s(bool,X22) != s(bool,X24)) )),
  inference(resolution,[],[f459,f484])).
fof(f5429,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f459,f483])).
fof(f5427,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f459,f482])).
fof(f5426,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | s(bool,X7) != s(bool,X9)) )),
  inference(resolution,[],[f459,f480])).
fof(f5424,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f459,f478])).
fof(f5355,plain,(
  ( ! [X140,X141] : (p(s(bool,X140)) | p(s(bool,X141)) | s(bool,X141) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f458,f648])).
fof(f5289,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | s(bool,X22) != s(bool,X23)) )),
  inference(resolution,[],[f458,f484])).
fof(f5286,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f458,f483])).
fof(f5285,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f458,f482])).
fof(f5284,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | s(bool,X7) != s(bool,X8)) )),
  inference(resolution,[],[f458,f480])).
fof(f5282,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f458,f478])).
fof(f5278,plain,(
  s('type.integer.int',sK12) = s('type.integer.int',sK13) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))),
  inference(resolution,[],[f458,f466])).
fof(f5277,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | s('type.integer.int',sK12) = s('type.integer.int',sK13)),
  inference(resolution,[],[f458,f471])).
fof(f5103,plain,(
  $false | ($spl12 | $spl22)),
  inference(subsumption_resolution,[],[f4917,f664])).
fof(f4917,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) | ($spl12 | $spl22)),
  inference(backward_demodulation,[],[f1701,f906])).
fof(f5102,plain,(
  $false | $spl22),
  inference(subsumption_resolution,[],[f4907,f457])).
fof(f4907,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) | $spl22),
  inference(backward_demodulation,[],[f1701,f788])).
fof(f5101,plain,(
  $false | $spl22),
  inference(subsumption_resolution,[],[f4905,f654])).
fof(f4905,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) | $spl22),
  inference(backward_demodulation,[],[f1701,f458])).
fof(f1699,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK13)))) | $spl21),
  inference(cnf_transformation,[],[f1699_D])).
fof(f1699_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK13)))) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f1701,plain,(
  s('type.integer.int',sK11) = s('type.integer.int',sK13) | $spl22),
  inference(cnf_transformation,[],[f1701_D])).
fof(f1701_D,plain,(
  s('type.integer.int',sK11) = s('type.integer.int',sK13) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f2201,plain,(
  s('type.integer.int',sK11) = s('type.integer.int',sK13) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))))),
  inference(resolution,[],[f459,f466])).
fof(f2202,plain,(
  s('type.integer.int',sK11) = s('type.integer.int',sK13) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))))),
  inference(resolution,[],[f459,f471])).
fof(f1703,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11)))) | $spl24),
  inference(cnf_transformation,[],[f1703_D])).
fof(f1703_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11)))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f2206,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11)))) | s('type.integer.int',sK11) = s('type.integer.int',sK13)),
  inference(resolution,[],[f459,f629])).
fof(f2207,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11)))) | s('type.integer.int',sK11) = s('type.integer.int',sK13)),
  inference(resolution,[],[f459,f629])).
fof(f3864,plain,(
  s('type.integer.int',sK11) = s('type.integer.int',sK13) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))))),
  inference(resolution,[],[f459,f466])).
fof(f3865,plain,(
  s('type.integer.int',sK11) = s('type.integer.int',sK13) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))))),
  inference(resolution,[],[f459,f471])).
fof(f3869,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11)))) | s('type.integer.int',sK11) = s('type.integer.int',sK13)),
  inference(resolution,[],[f459,f629])).
fof(f3870,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11)))) | s('type.integer.int',sK11) = s('type.integer.int',sK13)),
  inference(resolution,[],[f459,f629])).
fof(f4453,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11)))) | s('type.integer.int',sK11) = s('type.integer.int',sK13)),
  inference(resolution,[],[f1573,f471])).
fof(f4454,plain,(
  s('type.integer.int',sK11) = s('type.integer.int',sK13) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))))),
  inference(resolution,[],[f1573,f466])).
fof(f4865,plain,(
  ( ! [X146,X147] : (p(s(bool,X146)) | ~p(s(bool,X147)) | s(bool,X147) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK13)))) ) | $spl21),
  inference(resolution,[],[f1699,f648])).
fof(f4796,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | s(bool,X21) != s(bool,X23)) ) | $spl21),
  inference(resolution,[],[f1699,f484])).
fof(f4794,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK13)))) ) | $spl21),
  inference(resolution,[],[f1699,f483])).
fof(f4792,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK13)))) ) | $spl21),
  inference(resolution,[],[f1699,f482])).
fof(f4791,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | s(bool,X6) != s(bool,X8)) ) | $spl21),
  inference(resolution,[],[f1699,f480])).
fof(f4789,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK13)))) ) | $spl21),
  inference(resolution,[],[f1699,f478])).
fof(f4531,plain,(
  ( ! [X140,X141] : (p(s(bool,X140)) | p(s(bool,X141)) | s(bool,X141) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f1573,f648])).
fof(f4465,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | s(bool,X22) != s(bool,X23)) )),
  inference(resolution,[],[f1573,f484])).
fof(f4462,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f1573,f483])).
fof(f4461,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f1573,f482])).
fof(f4460,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11))) | s(bool,X7) != s(bool,X8)) )),
  inference(resolution,[],[f1573,f480])).
fof(f4458,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f1573,f478])).
fof(f4413,plain,(
  ( ! [X144,X143] : (p(s(bool,X143)) | p(s(bool,X144)) | s(bool,X144) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) | $spl12),
  inference(resolution,[],[f906,f648])).
fof(f4347,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | s(bool,X25) != s(bool,X26)) ) | $spl12),
  inference(resolution,[],[f906,f484])).
fof(f4344,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) | $spl12),
  inference(resolution,[],[f906,f483])).
fof(f4343,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) | $spl12),
  inference(resolution,[],[f906,f482])).
fof(f4342,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | s(bool,X10) != s(bool,X11)) ) | $spl12),
  inference(resolution,[],[f906,f480])).
fof(f4340,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) | $spl12),
  inference(resolution,[],[f906,f478])).
fof(f4076,plain,(
  ( ! [X147,X148] : (p(s(bool,X147)) | ~p(s(bool,X148)) | s(bool,X148) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f788,f648])).
fof(f4007,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | s(bool,X22) != s(bool,X24)) )),
  inference(resolution,[],[f788,f484])).
fof(f4005,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f788,f483])).
fof(f4003,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f788,f482])).
fof(f4002,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | s(bool,X7) != s(bool,X9)) )),
  inference(resolution,[],[f788,f480])).
fof(f4000,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f788,f478])).
fof(f3947,plain,(
  ( ! [X147,X148] : (p(s(bool,X147)) | ~p(s(bool,X148)) | s(bool,X148) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f459,f648])).
fof(f3878,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | s(bool,X22) != s(bool,X24)) )),
  inference(resolution,[],[f459,f484])).
fof(f3876,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f459,f483])).
fof(f3874,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f459,f482])).
fof(f3873,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | s(bool,X7) != s(bool,X9)) )),
  inference(resolution,[],[f459,f480])).
fof(f3871,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f459,f478])).
fof(f3820,plain,(
  ( ! [X140,X141] : (p(s(bool,X140)) | p(s(bool,X141)) | s(bool,X141) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f458,f648])).
fof(f3754,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | s(bool,X22) != s(bool,X23)) )),
  inference(resolution,[],[f458,f484])).
fof(f3751,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f458,f483])).
fof(f3750,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f458,f482])).
fof(f3749,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | s(bool,X7) != s(bool,X8)) )),
  inference(resolution,[],[f458,f480])).
fof(f3747,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f458,f478])).
fof(f3558,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f3373,f457])).
fof(f3373,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) | $spl10),
  inference(backward_demodulation,[],[f904,f459])).
fof(f1454,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | s('type.integer.int',sK12) = s('type.integer.int',sK13)),
  inference(resolution,[],[f458,f471])).
fof(f1455,plain,(
  s('type.integer.int',sK12) = s('type.integer.int',sK13) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))),
  inference(resolution,[],[f458,f466])).
fof(f2079,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | s('type.integer.int',sK12) = s('type.integer.int',sK13)),
  inference(resolution,[],[f458,f471])).
fof(f2080,plain,(
  s('type.integer.int',sK12) = s('type.integer.int',sK13) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))),
  inference(resolution,[],[f458,f466])).
fof(f3371,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) | $spl43),
  inference(cnf_transformation,[],[f3371_D])).
fof(f3371_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) <=> ~$spl43),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl43])])).
fof(f2908,plain,(
  s('type.integer.int',sK12) = s('type.integer.int',sK13) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))))),
  inference(resolution,[],[f788,f466])).
fof(f2909,plain,(
  s('type.integer.int',sK12) = s('type.integer.int',sK13) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))))),
  inference(resolution,[],[f788,f471])).
fof(f2913,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | s('type.integer.int',sK12) = s('type.integer.int',sK13)),
  inference(resolution,[],[f788,f629])).
fof(f2914,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | s('type.integer.int',sK12) = s('type.integer.int',sK13)),
  inference(resolution,[],[f788,f629])).
fof(f3329,plain,(
  ( ! [X144,X143] : (p(s(bool,X143)) | p(s(bool,X144)) | s(bool,X144) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) | $spl12),
  inference(resolution,[],[f906,f648])).
fof(f3263,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | s(bool,X25) != s(bool,X26)) ) | $spl12),
  inference(resolution,[],[f906,f484])).
fof(f3260,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) | $spl12),
  inference(resolution,[],[f906,f483])).
fof(f3259,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) | $spl12),
  inference(resolution,[],[f906,f482])).
fof(f3258,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | s(bool,X10) != s(bool,X11)) ) | $spl12),
  inference(resolution,[],[f906,f480])).
fof(f3256,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) ) | $spl12),
  inference(resolution,[],[f906,f478])).
fof(f2991,plain,(
  ( ! [X147,X148] : (p(s(bool,X147)) | ~p(s(bool,X148)) | s(bool,X148) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f788,f648])).
fof(f2922,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | s(bool,X22) != s(bool,X24)) )),
  inference(resolution,[],[f788,f484])).
fof(f2920,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f788,f483])).
fof(f2918,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f788,f482])).
fof(f2917,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12))) | s(bool,X7) != s(bool,X9)) )),
  inference(resolution,[],[f788,f480])).
fof(f2915,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f788,f478])).
fof(f2658,plain,(
  ( ! [X147,X148] : (p(s(bool,X147)) | ~p(s(bool,X148)) | s(bool,X148) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f664,f648])).
fof(f2589,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | s(bool,X22) != s(bool,X24)) )),
  inference(resolution,[],[f664,f484])).
fof(f2587,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f664,f483])).
fof(f2585,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f664,f482])).
fof(f2584,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | s(bool,X7) != s(bool,X9)) )),
  inference(resolution,[],[f664,f480])).
fof(f2582,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f664,f478])).
fof(f2532,plain,(
  ( ! [X140,X141] : (p(s(bool,X140)) | p(s(bool,X141)) | s(bool,X141) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f657,f648])).
fof(f2466,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | s(bool,X22) != s(bool,X23)) )),
  inference(resolution,[],[f657,f484])).
fof(f2463,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f657,f483])).
fof(f2462,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f657,f482])).
fof(f2461,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | s(bool,X7) != s(bool,X8)) )),
  inference(resolution,[],[f657,f480])).
fof(f2459,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f657,f478])).
fof(f2414,plain,(
  ( ! [X146,X147] : (p(s(bool,X146)) | ~p(s(bool,X147)) | s(bool,X147) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f654,f648])).
fof(f2345,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | s(bool,X21) != s(bool,X23)) )),
  inference(resolution,[],[f654,f484])).
fof(f2343,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f654,f483])).
fof(f2341,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f654,f482])).
fof(f2340,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11))) | s(bool,X6) != s(bool,X8)) )),
  inference(resolution,[],[f654,f480])).
fof(f2338,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK11)))) )),
  inference(resolution,[],[f654,f478])).
fof(f2284,plain,(
  ( ! [X147,X148] : (p(s(bool,X147)) | ~p(s(bool,X148)) | s(bool,X148) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f459,f648])).
fof(f2215,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | s(bool,X22) != s(bool,X24)) )),
  inference(resolution,[],[f459,f484])).
fof(f2213,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f459,f483])).
fof(f2211,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f459,f482])).
fof(f2210,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | s(bool,X7) != s(bool,X9)) )),
  inference(resolution,[],[f459,f480])).
fof(f2208,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f459,f478])).
fof(f2157,plain,(
  ( ! [X140,X141] : (p(s(bool,X140)) | p(s(bool,X141)) | s(bool,X141) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f458,f648])).
fof(f2091,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | s(bool,X22) != s(bool,X23)) )),
  inference(resolution,[],[f458,f484])).
fof(f2088,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f458,f483])).
fof(f2087,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f458,f482])).
fof(f2086,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | s(bool,X7) != s(bool,X8)) )),
  inference(resolution,[],[f458,f480])).
fof(f2084,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f458,f478])).
fof(f2018,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) | $spl28),
  inference(cnf_transformation,[],[f2018_D])).
fof(f2018_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f2020,plain,(
  s('type.integer.int',sK11) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13)))) | $spl30),
  inference(cnf_transformation,[],[f2020_D])).
fof(f2020_D,plain,(
  s('type.integer.int',sK11) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13)))) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f2016,plain,(
  s('type.integer.int',sK11) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) | $spl22),
  inference(forward_demodulation,[],[f1825,f1701])).
fof(f1825,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int','const.integer.int_|48|_0')))) | s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13)))) | $spl22),
  inference(superposition,[],[f632,f1701])).
fof(f632,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(X0)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_|48|_0'))))) )),
  inference(cnf_transformation,[],[f451])).
fof(f451,plain,(
  ! [X0] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_|48|_0')))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK24])],[f357])).
fof(f357,plain,(
  ! [X0] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_|48|_0')))) | ? [X1] : s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))))),
  inference(ennf_transformation,[],[f322])).
fof(f322,plain,(
  ! [X0] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_|48|_0')))) => ? [X1] : s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))))),
  inference(flattening,[],[f321])).
fof(f321,plain,(
  ! [X0] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_|48|_0')))) => ? [X1] : s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))))),
  inference(rectify,[],[f58])).
fof(f58,axiom,(
  ! [X11] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X11),s('type.integer.int','const.integer.int_|48|_0')))) => ? [X43] : s('type.integer.int',X11) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X43))))),
  file('thm.integer.INT_LTE_TRANS','thm.integer.NUM_POSINT_EX')).
fof(f1659,plain,(
  ( ! [X147,X148] : (p(s(bool,X147)) | ~p(s(bool,X148)) | s(bool,X148) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f459,f648])).
fof(f1590,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | s(bool,X22) != s(bool,X24)) )),
  inference(resolution,[],[f459,f484])).
fof(f1588,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f459,f483])).
fof(f1586,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f459,f482])).
fof(f1585,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))) | s(bool,X7) != s(bool,X9)) )),
  inference(resolution,[],[f459,f480])).
fof(f1583,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f459,f478])).
fof(f1582,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11)))) | s('type.integer.int',sK11) = s('type.integer.int',sK13)),
  inference(resolution,[],[f459,f629])).
fof(f1581,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int',sK11)))) | s('type.integer.int',sK11) = s('type.integer.int',sK13)),
  inference(resolution,[],[f459,f629])).
fof(f1577,plain,(
  s('type.integer.int',sK11) = s('type.integer.int',sK13) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))))),
  inference(resolution,[],[f459,f471])).
fof(f1576,plain,(
  s('type.integer.int',sK11) = s('type.integer.int',sK13) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK11),s('type.integer.int',sK13))))),
  inference(resolution,[],[f459,f466])).
fof(f1532,plain,(
  ( ! [X140,X141] : (p(s(bool,X140)) | p(s(bool,X141)) | s(bool,X141) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f458,f648])).
fof(f1466,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | s(bool,X22) != s(bool,X23)) )),
  inference(resolution,[],[f458,f484])).
fof(f1463,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f458,f483])).
fof(f1462,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f458,f482])).
fof(f1461,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | s(bool,X7) != s(bool,X8)) )),
  inference(resolution,[],[f458,f480])).
fof(f1459,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f458,f478])).
fof(f1403,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) | $spl16),
  inference(cnf_transformation,[],[f1403_D])).
fof(f1403_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f1405,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13)))) | $spl18),
  inference(cnf_transformation,[],[f1405_D])).
fof(f1405_D,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13)))) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f1401,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) | $spl10),
  inference(forward_demodulation,[],[f1210,f904])).
fof(f1210,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) | s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13)))) | $spl10),
  inference(superposition,[],[f632,f904])).
fof(f1075,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f1005,f785])).
fof(f1005,plain,(
  ( ! [X265,X264] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X264),s('type.integer.int',X265))))) ) | $spl8),
  inference(resolution,[],[f785,f461])).
fof(f1074,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f1004,f785])).
fof(f1004,plain,(
  ( ! [X263,X262] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X262),s('type.integer.int',X263))))) ) | $spl8),
  inference(resolution,[],[f785,f463])).
fof(f1073,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f1003,f785])).
fof(f1003,plain,(
  ( ! [X261,X260] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X260),s('type.integer.int',X261))))) ) | $spl8),
  inference(resolution,[],[f785,f464])).
fof(f1072,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f997,f785])).
fof(f997,plain,(
  ( ! [X251,X250] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X250),s('type.integer.int',X251))))) ) | $spl8),
  inference(resolution,[],[f785,f628])).
fof(f1071,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f996,f785])).
fof(f996,plain,(
  ( ! [X249,X248] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X248),s('type.integer.int',X249))))) ) | $spl8),
  inference(resolution,[],[f785,f628])).
fof(f1070,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f995,f785])).
fof(f995,plain,(
  ( ! [X246,X247] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X246),s('type.integer.int',X247))))) ) | $spl8),
  inference(resolution,[],[f785,f626])).
fof(f1069,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f994,f785])).
fof(f994,plain,(
  ( ! [X244,X245] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X244),s('type.integer.int',X245))))) ) | $spl8),
  inference(resolution,[],[f785,f626])).
fof(f1068,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f986,f785])).
fof(f986,plain,(
  ( ! [X225,X226] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X225),s('type.integer.int',X226))))) ) | $spl8),
  inference(resolution,[],[f785,f464])).
fof(f1067,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f985,f785])).
fof(f985,plain,(
  ( ! [X223,X224] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X223),s('type.integer.int',X224))))) ) | $spl8),
  inference(resolution,[],[f785,f463])).
fof(f1066,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f984,f785])).
fof(f984,plain,(
  ( ! [X222,X221] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X221),s('type.integer.int',X222))))) ) | $spl8),
  inference(resolution,[],[f785,f461])).
fof(f998,plain,(
  $false | $spl8),
  inference(resolution,[],[f785,f631])).
fof(f864,plain,(
  ( ! [X140,X141] : (p(s(bool,X140)) | p(s(bool,X141)) | s(bool,X141) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f458,f648])).
fof(f798,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | s(bool,X22) != s(bool,X23)) )),
  inference(resolution,[],[f458,f484])).
fof(f795,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f458,f483])).
fof(f794,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f458,f482])).
fof(f793,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))) | s(bool,X7) != s(bool,X8)) )),
  inference(resolution,[],[f458,f480])).
fof(f791,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f458,f478])).
fof(f787,plain,(
  s('type.integer.int',sK12) = s('type.integer.int',sK13) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))),
  inference(resolution,[],[f458,f466])).
fof(f786,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))) | s('type.integer.int',sK12) = s('type.integer.int',sK13)),
  inference(resolution,[],[f458,f471])).
fof(f741,plain,(
  ( ! [X144,X143] : (p(s(bool,X143)) | p(s(bool,X144)) | s(bool,X144) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f457,f648])).
fof(f675,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | s(bool,X25) != s(bool,X26)) )),
  inference(resolution,[],[f457,f484])).
fof(f672,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f457,f483])).
fof(f671,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f457,f482])).
fof(f670,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12))) | s(bool,X10) != s(bool,X11)) )),
  inference(resolution,[],[f457,f480])).
fof(f668,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK11),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f457,f478])).
% SZS output end Proof for thm.integer
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 2720
% Passive clauses: 72791
% Generated clauses: 124565
% Final active clauses: 710
% Final passive clauses: 37982
% Input formulas: 261
% Initial clauses: 191
% 
% Selected by SInE selection: 52
% SInE iterations: 1
% 
% Duplicate literals: 3241
% Fw subsumption resolutions: 2265
% Fw demodulations: 3276
% Bw demodulations: 1078
% 
% Simple tautologies: 10741
% Equational tautologies: 78
% Forward subsumptions: 35702
% 
% Binary resolution: 52662
% Forward superposition: 53175
% Backward superposition: 2816
% Equality resolution: 49
% 
% Split clauses: 1550
% Split components: 3391
% SAT solver clauses: 1714
% SAT solver unit clauses: 64
% SAT solver binary clauses: 510
% 
% Sat splits: 1550
% Sat splitting refutations: 111
% 
% Memory used [KB]: 23283
% Time elapsed: 2.087 s
% ------------------------------
% ------------------------------
% Success in time 13.871 s
