Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.integral
% (87704)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 88
% Passive clauses: 2518
% Generated clauses: 14467
% Final active clauses: 72
% Final passive clauses: 2312
% Input formulas: 14
% Initial clauses: 36
% 
% Duplicate literals: 911
% Trivial inequalities: 144
% Fw subsumption resolutions: 1531
% Fw demodulations: 1224
% Bw demodulations: 12
% 
% Simple tautologies: 253
% Equational tautologies: 96
% Forward subsumptions: 5881
% Fw demodulations to eq. taut.: 2
% 
% Binary resolution: 74
% Forward superposition: 5652
% Backward superposition: 4116
% Self superposition: 353
% Equality factoring: 265
% Equality resolution: 11
% 
% Split clauses: 635
% Split components: 1285
% SAT solver clauses: 1679
% SAT solver unit clauses: 262
% SAT solver binary clauses: 437
% 
% Sat splits: 635
% Sat splitting refutations: 7
% 
% TWLsolver clauses: 1673
% TWLsolver calls for satisfiability: 88
% 
% Memory used [KB]: 9722
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.integral
% (87720)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 95
% Passive clauses: 705
% Generated clauses: 6761
% Final active clauses: 79
% Final passive clauses: 563
% Input formulas: 14
% Initial clauses: 36
% 
% Duplicate literals: 471
% Trivial inequalities: 23
% Fw subsumption resolutions: 428
% Fw demodulations: 706
% Bw demodulations: 58
% Condensations: 431
% Global subsumptions: 4
% 
% Simple tautologies: 83
% Equational tautologies: 25
% Forward subsumptions: 2553
% Fw demodulations to eq. taut.: 2
% Bw demodulations to eq. taut.: 3
% 
% Binary resolution: 47
% Forward superposition: 2165
% Backward superposition: 1432
% Self superposition: 120
% Equality factoring: 94
% 
% Split clauses: 208
% Split components: 417
% SAT solver clauses: 1333
% SAT solver unit clauses: 53
% SAT solver binary clauses: 56
% SAT solver learnt clauses: 8
% 
% Sat splits: 208
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 1317
% TWLsolver calls for satisfiability: 368
% 
% Memory used [KB]: 11641
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.integral
% (87737)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 257
% Passive clauses: 2593
% Generated clauses: 147684
% Final active clauses: 57
% Final passive clauses: 163
% Input formulas: 14
% Initial clauses: 36
% Discarded non-redundant clauses: 538
% 
% Duplicate literals: 7560
% Trivial inequalities: 1110
% Fw subsumption resolutions: 13278
% Fw demodulations: 13217
% Bw demodulations: 1434
% Condensations: 14907
% Global subsumptions: 19
% 
% Simple tautologies: 1145
% Equational tautologies: 684
% Forward subsumptions: 88201
% Fw demodulations to eq. taut.: 20
% Bw demodulations to eq. taut.: 37
% 
% Unit resulting resolution: 1901
% Forward superposition: 35310
% Backward superposition: 31620
% Self superposition: 764
% Equality factoring: 378
% Equality resolution: 42
% 
% Split clauses: 205
% Split components: 420
% SAT solver clauses: 3050
% SAT solver unit clauses: 24
% SAT solver binary clauses: 174
% 
% Sat splits: 205
% Sat splitting refutations: 26
% 
% Memory used [KB]: 58719
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.integral
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.integral
% SZS output start Proof for thm.integral
fof(f39336,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f932,f936,f934,f1077,f1078,f1122,f1080,f1123,f1171,f1173,f1176,f79,f1177,f1175,f2454,f2460,f2464,f2456,f3131,f3134,f3137,f3894,f3898,f3896,f39024,f39029,f39032,f39034,f39035,f39038,f39040,f39041,f39045,f39046,f39048,f39051,f39058,f39068,f39078,f39081,f39052,f39091,f39107,f39114,f39122,f39127,f39131,f39070,f39136,f39140,f39060,f39143,f39144,f31248,f39042,f39150,f39158,f39159,f39160,f25782,f39165,f39170,f39079,f39173,f21861,f21860,f39174,f39178,f39179,f39145,f17802,f39146,f17304,f39183,f39184,f39187,f2461,f39189,f2455,f39193,f39198,f39108,f39200,f39203,f10205,f39175,f10204,f39204,f39161,f7268,f39180,f39216,f2458,f39218,f39271,f39272,f39274,f39275,f39277,f39278,f39282,f39283,f39115,f39227,f39230,f39232,f39233,f39062,f39294,f39295,f39301,f39302,f39304,f39305,f39307,f39308,f39240,f3132,f39253,f39254,f39185,f39255,f39256,f39326,f39258,f39259,f39030,f39260,f39309,f39261,f39036,f39288,f39335])).
fof(f39335,plain,(
  $false | ($spl7 | $spl21)),
  inference(subsumption_resolution,[],[f39099,f2457])).
fof(f2457,plain,(
  ~sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | $spl21),
  inference(cnf_transformation,[],[f2457_D])).
fof(f2457_D,plain,(
  ~sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f39099,plain,(
  sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | $spl7),
  inference(subsumption_resolution,[],[f39098,f93])).
fof(f93,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f30])).
fof(f30,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.integral.BOLZANO_LEMMA_ALT',aHL_FALSITY)).
fof(f39098,plain,(
  p(s(bool,'F_0')) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | $spl7),
  inference(forward_demodulation,[],[f39097,f173])).
fof(f173,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',sK4)))),
  inference(subsumption_resolution,[],[f171,f94])).
fof(f94,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.integral.BOLZANO_LEMMA_ALT',aHL_TRUTH)).
fof(f171,plain,(
  ~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',sK4)))),
  inference(superposition,[],[f65,f96])).
fof(f96,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.integral.BOLZANO_LEMMA_ALT',aHL_BOOL_CASES)).
fof(f65,plain,(
  ~p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',sK4))))),
  inference(cnf_transformation,[],[f44])).
fof(f44,plain,(
  ! [X7,X8,X9] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real',X8)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X8),s('type.realax.real',X9)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X7))),s('type.realax.real',X8)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X8))),s('type.realax.real',X9)))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X7))),s('type.realax.real',X9))))) & ! [X3] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2(X3))))) & ! [X5,X6] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',X3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',X6)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X6),s('type.realax.real',X5))),s('type.realax.real',sK2(X3))))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X5))),s('type.realax.real',X6)))))) & (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))) & ~p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',sK4)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK1,sK2,sK3,sK4])],[f43])).
fof(f43,plain,(
  ? [X0] : (! [X7,X8,X9] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real',X8)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X8),s('type.realax.real',X9)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),X0),s('type.realax.real',X7))),s('type.realax.real',X8)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),X0),s('type.realax.real',X8))),s('type.realax.real',X9)))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),X0),s('type.realax.real',X7))),s('type.realax.real',X9))))) & ! [X3] : ? [X4] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X4)))) & ! [X5,X6] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',X3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',X6)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X6),s('type.realax.real',X5))),s('type.realax.real',X4)))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),X0),s('type.realax.real',X5))),s('type.realax.real',X6)))))) & ? [X1,X2] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) & ~p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),X0),s('type.realax.real',X1))),s('type.realax.real',X2))))))),
  inference(rectify,[],[f32])).
fof(f32,plain,(
  ? [X0] : (! [X5,X6,X7] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',X6)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X7)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),X0),s('type.realax.real',X5))),s('type.realax.real',X6)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),X0),s('type.realax.real',X6))),s('type.realax.real',X7)))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),X0),s('type.realax.real',X5))),s('type.realax.real',X7))))) & ! [X1] : ? [X2] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2)))) & ! [X3,X4] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',X1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X4)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X4),s('type.realax.real',X3))),s('type.realax.real',X2)))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),X0),s('type.realax.real',X3))),s('type.realax.real',X4)))))) & ? [X8,X9] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X8),s('type.realax.real',X9)))) & ~p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),X0),s('type.realax.real',X8))),s('type.realax.real',X9))))))),
  inference(flattening,[],[f31])).
fof(f31,plain,(
  ? [X0] : ((! [X5,X6,X7] : ((~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',X6)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X7)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),X0),s('type.realax.real',X5))),s('type.realax.real',X6)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),X0),s('type.realax.real',X6))),s('type.realax.real',X7))))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),X0),s('type.realax.real',X5))),s('type.realax.real',X7))))) & ! [X1] : ? [X2] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2)))) & ! [X3,X4] : ((~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',X1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X4)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X4),s('type.realax.real',X3))),s('type.realax.real',X2))))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),X0),s('type.realax.real',X3))),s('type.realax.real',X4))))))) & ? [X8,X9] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X8),s('type.realax.real',X9)))) & ~p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),X0),s('type.realax.real',X8))),s('type.realax.real',X9))))))),
  inference(ennf_transformation,[],[f16])).
fof(f16,plain,(
  ~! [X0] : ((! [X5,X6,X7] : ((p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',X6)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X7)))) & p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),X0),s('type.realax.real',X5))),s('type.realax.real',X6)))) & p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),X0),s('type.realax.real',X6))),s('type.realax.real',X7))))) => p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),X0),s('type.realax.real',X5))),s('type.realax.real',X7))))) & ! [X1] : ? [X2] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2)))) & ! [X3,X4] : ((p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',X1)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X4)))) & p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X4),s('type.realax.real',X3))),s('type.realax.real',X2))))) => p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),X0),s('type.realax.real',X3))),s('type.realax.real',X4))))))) => ! [X8,X9] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X8),s('type.realax.real',X9)))) => p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),X0),s('type.realax.real',X8))),s('type.realax.real',X9))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ((! [X7,X8,X9] : ((p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real',X8)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X8),s('type.realax.real',X9)))) & p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),X6),s('type.realax.real',X7))),s('type.realax.real',X8)))) & p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),X6),s('type.realax.real',X8))),s('type.realax.real',X9))))) => p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),X6),s('type.realax.real',X7))),s('type.realax.real',X9))))) & ! [X10] : ? [X11] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X11)))) & ! [X7,X8] : ((p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real',X10)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X10),s('type.realax.real',X8)))) & p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X8),s('type.realax.real',X7))),s('type.realax.real',X11))))) => p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),X6),s('type.realax.real',X7))),s('type.realax.real',X8))))))) => ! [X7,X8] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real',X8)))) => p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),X6),s('type.realax.real',X7))),s('type.realax.real',X8))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ((! [X7,X8,X9] : ((p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real',X8)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X8),s('type.realax.real',X9)))) & p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),X6),s('type.realax.real',X7))),s('type.realax.real',X8)))) & p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),X6),s('type.realax.real',X8))),s('type.realax.real',X9))))) => p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),X6),s('type.realax.real',X7))),s('type.realax.real',X9))))) & ! [X10] : ? [X11] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X11)))) & ! [X7,X8] : ((p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real',X10)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X10),s('type.realax.real',X8)))) & p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X8),s('type.realax.real',X7))),s('type.realax.real',X11))))) => p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),X6),s('type.realax.real',X7))),s('type.realax.real',X8))))))) => ! [X7,X8] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real',X8)))) => p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),X6),s('type.realax.real',X7))),s('type.realax.real',X8))))))),
  file('thm.integral.BOLZANO_LEMMA_ALT','thm.integral.BOLZANO_LEMMA_ALT')).
fof(f39097,plain,(
  p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',sK4)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | $spl7),
  inference(forward_demodulation,[],[f39096,f75])).
fof(f75,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X4))),s(X2,X5))) = s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X3),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X4),s(X2,X5)))))) )),
  inference(cnf_transformation,[],[f19])).
fof(f19,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X4))),s(X2,X5))) = s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X3),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X4),s(X2,X5)))))),
  inference(flattening,[],[f18])).
fof(f18,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X4))),s(X2,X5))) = s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X3),s('type.pair.prod'(X1,X2),'const.pair.,_2'(s(X1,X4),s(X2,X5)))))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X13,X14,X15] : ! [X16,X10,X17] : s(X13,happ(s(fun(X15,X13),happ(s(fun(X14,fun(X15,X13)),X16),s(X14,X10))),s(X15,X17))) = s(X13,'const.pair.UNCURRY_2'(s(fun(X14,fun(X15,X13)),X16),s('type.pair.prod'(X14,X15),'const.pair.,_2'(s(X14,X10),s(X15,X17)))))),
  file('thm.integral.BOLZANO_LEMMA_ALT','thm.pair.UNCURRY_DEF')).
fof(f39096,plain,(
  p(s(bool,'const.pair.UNCURRY_2'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | $spl7),
  inference(forward_demodulation,[],[f39095,f76])).
fof(f76,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X3),s('type.pair.prod'(X1,X2),X4))) = s(X0,happ(s(fun('type.pair.prod'(X1,X2),X0),'const.pair.UNCURRY_1'(s(fun(X1,fun(X2,X0)),X3))),s('type.pair.prod'(X1,X2),X4)))) )),
  inference(cnf_transformation,[],[f21])).
fof(f21,plain,(
  ! [X0,X1,X2,X3,X4] : s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X3),s('type.pair.prod'(X1,X2),X4))) = s(X0,happ(s(fun('type.pair.prod'(X1,X2),X0),'const.pair.UNCURRY_1'(s(fun(X1,fun(X2,X0)),X3))),s('type.pair.prod'(X1,X2),X4)))),
  inference(flattening,[],[f20])).
fof(f20,plain,(
  ! [X0,X1,X2] : ! [X3,X4] : s(X0,'const.pair.UNCURRY_2'(s(fun(X1,fun(X2,X0)),X3),s('type.pair.prod'(X1,X2),X4))) = s(X0,happ(s(fun('type.pair.prod'(X1,X2),X0),'const.pair.UNCURRY_1'(s(fun(X1,fun(X2,X0)),X3))),s('type.pair.prod'(X1,X2),X4)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X13,X14,X15] : ! [X18,X19] : s(X13,'const.pair.UNCURRY_2'(s(fun(X14,fun(X15,X13)),X18),s('type.pair.prod'(X14,X15),X19))) = s(X13,happ(s(fun('type.pair.prod'(X14,X15),X13),'const.pair.UNCURRY_1'(s(fun(X14,fun(X15,X13)),X18))),s('type.pair.prod'(X14,X15),X19)))),
  file('thm.integral.BOLZANO_LEMMA_ALT','ARITY_1_2_const.pair.UNCURRY')).
fof(f39095,plain,(
  sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | $spl7),
  inference(subsumption_resolution,[],[f39094,f1079])).
fof(f1079,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl7),
  inference(cnf_transformation,[],[f1079_D])).
fof(f1079_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f39094,plain,(
  s(bool,'T_0') = s(bool,'F_0') | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))))),
  inference(forward_demodulation,[],[f39093,f116])).
fof(f116,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2(X1))))) )),
  inference(subsumption_resolution,[],[f115,f93])).
fof(f115,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2(X1))))) )),
  inference(superposition,[],[f62,f96])).
fof(f62,plain,(
  ( ! [X3] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2(X3)))))) )),
  inference(cnf_transformation,[],[f44])).
fof(f39093,plain,(
  sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))))),
  inference(subsumption_resolution,[],[f39092,f93])).
fof(f39092,plain,(
  p(s(bool,'F_0')) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))))),
  inference(forward_demodulation,[],[f32700,f173])).
fof(f32700,plain,(
  p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',sK4)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))))),
  inference(duplicate_literal_removal,[],[f32662])).
fof(f32662,plain,(
  p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',sK4)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),
  inference(resolution,[],[f3417,f30421])).
fof(f30421,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK9(sK2(sK8(X0)),X0)))),s('type.realax.real',sK10(sK2(sK8(X0)),X0))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0(X0)) )),
  inference(subsumption_resolution,[],[f30416,f94])).
fof(f30416,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK9(sK2(sK8(X0)),X0)))),s('type.realax.real',sK10(sK2(sK8(X0)),X0))))) | sP0(X0)) )),
  inference(duplicate_literal_removal,[],[f30382])).
fof(f30382,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK9(sK2(sK8(X0)),X0)))),s('type.realax.real',sK10(sK2(sK8(X0)),X0))))) | sP0(X0) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))))) )),
  inference(superposition,[],[f1075,f106])).
fof(f106,plain,(
  s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))),
  inference(subsumption_resolution,[],[f105,f93])).
fof(f105,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))),
  inference(superposition,[],[f64,f96])).
fof(f64,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))),
  inference(cnf_transformation,[],[f44])).
fof(f1075,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X14),s('type.realax.real',X15)))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X13),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X14),s('type.realax.real',X15)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK9(sK2(sK8(X13)),X13)))),s('type.realax.real',sK10(sK2(sK8(X13)),X13))))) | sP0(X13) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X13),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))))) )),
  inference(subsumption_resolution,[],[f1068,f107])).
fof(f107,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2)))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X3),X0)),s('type.realax.real',sK8(X0))))) | sP0(X0)) )),
  inference(resolution,[],[f62,f71])).
fof(f71,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(X4,X0)),s('type.realax.real',sK8(X0))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X4)))) | sP0(X0)) )),
  inference(cnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0] : (sP0(X0) | ! [X4] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X4)))) | (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(X4,X0)),s('type.realax.real',sK8(X0))))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8(X0)),s('type.realax.real',sK10(X4,X0))))) & p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(X4,X0)),s('type.realax.real',sK9(X4,X0)))),s('type.realax.real',X4)))) & ~p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK9(X4,X0)),s('type.realax.real',sK10(X4,X0))))))))) | ! [X1,X2] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK8,sK9,sK10])],[f48])).
fof(f48,plain,(
  ! [X0] : (sP0(X0) | ? [X3] : ! [X4] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X4)))) | ? [X5,X6] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',X3)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',X6)))) & p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X6),s('type.realax.real',X5))),s('type.realax.real',X4)))) & ~p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X5),s('type.realax.real',X6)))))))) | ! [X1,X2] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2))))))))),
  inference(rectify,[],[f42])).
fof(f42,plain,(
  ! [X0] : (sP0(X0) | ? [X1] : ! [X2] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2)))) | ? [X3,X4] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',X1)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X4)))) & p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X4),s('type.realax.real',X3))),s('type.realax.real',X2)))) & ~p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X3),s('type.realax.real',X4)))))))) | ! [X8,X9] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X8),s('type.realax.real',X9)))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X8),s('type.realax.real',X9))))))))),
  inference(definition_folding,[],[f34,f41])).
fof(f41,plain,(
  ! [X0] : (~sP0(X0) | ? [X5,X6,X7] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',X6)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X7)))) & p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X5),s('type.realax.real',X6)))))) & p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X6),s('type.realax.real',X7)))))) & ~p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X5),s('type.realax.real',X7))))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f34,plain,(
  ! [X0] : (? [X5,X6,X7] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',X6)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X7)))) & p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X5),s('type.realax.real',X6)))))) & p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X6),s('type.realax.real',X7)))))) & ~p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X5),s('type.realax.real',X7))))))) | ? [X1] : ! [X2] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2)))) | ? [X3,X4] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',X1)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X4)))) & p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X4),s('type.realax.real',X3))),s('type.realax.real',X2)))) & ~p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X3),s('type.realax.real',X4)))))))) | ! [X8,X9] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X8),s('type.realax.real',X9)))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X8),s('type.realax.real',X9))))))))),
  inference(flattening,[],[f33])).
fof(f33,plain,(
  ! [X0] : ((? [X5,X6,X7] : ((p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',X6)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X7)))) & p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X5),s('type.realax.real',X6)))))) & p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X6),s('type.realax.real',X7))))))) & ~p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X5),s('type.realax.real',X7))))))) | ? [X1] : ! [X2] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2)))) | ? [X3,X4] : ((p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',X1)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X4)))) & p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X4),s('type.realax.real',X3))),s('type.realax.real',X2))))) & ~p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X3),s('type.realax.real',X4))))))))) | ! [X8,X9] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X8),s('type.realax.real',X9)))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X8),s('type.realax.real',X9))))))))),
  inference(ennf_transformation,[],[f17])).
fof(f17,plain,(
  ! [X0] : ((! [X5,X6,X7] : ((p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',X6)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X7)))) & p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X5),s('type.realax.real',X6)))))) & p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X6),s('type.realax.real',X7))))))) => p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X5),s('type.realax.real',X7))))))) & ! [X1] : ? [X2] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2)))) & ! [X3,X4] : ((p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',X1)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X4)))) & p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X4),s('type.realax.real',X3))),s('type.realax.real',X2))))) => p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X3),s('type.realax.real',X4))))))))) => ! [X8,X9] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X8),s('type.realax.real',X9)))) => p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X8),s('type.realax.real',X9))))))))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X6] : ((! [X7,X8,X9] : ((p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real',X8)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X8),s('type.realax.real',X9)))) & p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X6),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X7),s('type.realax.real',X8)))))) & p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X6),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X8),s('type.realax.real',X9))))))) => p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X6),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X7),s('type.realax.real',X9))))))) & ! [X10] : ? [X11] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X11)))) & ! [X7,X8] : ((p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real',X10)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X10),s('type.realax.real',X8)))) & p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X8),s('type.realax.real',X7))),s('type.realax.real',X11))))) => p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X6),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X7),s('type.realax.real',X8))))))))) => ! [X7,X8] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real',X8)))) => p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X6),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X7),s('type.realax.real',X8))))))))),
  file('thm.integral.BOLZANO_LEMMA_ALT','thm.seq.BOLZANO_LEMMA')).
fof(f1068,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK9(sK2(sK8(X13)),X13)))),s('type.realax.real',sK10(sK2(sK8(X13)),X13))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(sK8(X13)),X13)),s('type.realax.real',sK8(X13))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X13),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X14),s('type.realax.real',X15)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X14),s('type.realax.real',X15)))) | sP0(X13) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X13),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))))) )),
  inference(duplicate_literal_removal,[],[f1055])).
fof(f1055,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK9(sK2(sK8(X13)),X13)))),s('type.realax.real',sK10(sK2(sK8(X13)),X13))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(sK8(X13)),X13)),s('type.realax.real',sK8(X13))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X13),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X14),s('type.realax.real',X15)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X14),s('type.realax.real',X15)))) | sP0(X13) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X13),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0(X13)) )),
  inference(resolution,[],[f228,f365])).
fof(f365,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8(X0)),s('type.realax.real',sK10(sK2(X1),X0))))) | sP0(X0)) )),
  inference(subsumption_resolution,[],[f358,f94])).
fof(f358,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8(X0)),s('type.realax.real',sK10(sK2(X1),X0))))) | sP0(X0)) )),
  inference(superposition,[],[f108,f106])).
fof(f108,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',X6)))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X4),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X5),s('type.realax.real',X6)))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8(X4)),s('type.realax.real',sK10(sK2(X7),X4))))) | sP0(X4)) )),
  inference(resolution,[],[f62,f72])).
fof(f72,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8(X0)),s('type.realax.real',sK10(X4,X0))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X4)))) | sP0(X0)) )),
  inference(cnf_transformation,[],[f49])).
fof(f228,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',sK10(sK2(X0),X1))))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK9(sK2(X0),X1)))),s('type.realax.real',sK10(sK2(X0),X1))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X0),X1)),s('type.realax.real',X0)))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X1),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X2),s('type.realax.real',X3)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X3)))) | sP0(X1)) )),
  inference(subsumption_resolution,[],[f227,f82])).
fof(f82,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))))),
  inference(flattening,[],[f53])).
fof(f53,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f37])).
fof(f37,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f26])).
fof(f26,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X12] : ((p(s(bool,X12)) => p(s(bool,X12))) <=> p(s(bool,'T_0')))),
  file('thm.integral.BOLZANO_LEMMA_ALT','thm.bool.IMP_CLAUSES|split|3')).
fof(f227,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK9(sK2(X0),X1)))),s('type.realax.real',sK10(sK2(X0),X1))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',sK10(sK2(X0),X1))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X0),X1)),s('type.realax.real',X0)))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X1),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X2),s('type.realax.real',X3)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X3)))) | sP0(X1)) )),
  inference(forward_demodulation,[],[f216,f116])).
fof(f216,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK9(sK2(X0),X1)))),s('type.realax.real',sK10(sK2(X0),X1))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',sK10(sK2(X0),X1))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X0),X1)),s('type.realax.real',X0)))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X1),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X2),s('type.realax.real',X3)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X3)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2(X0))))) | sP0(X1)) )),
  inference(resolution,[],[f63,f73])).
fof(f73,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(X4,X0)),s('type.realax.real',sK9(X4,X0)))),s('type.realax.real',X4)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X4)))) | sP0(X0)) )),
  inference(cnf_transformation,[],[f49])).
fof(f63,plain,(
  ( ! [X6,X5,X3] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X6),s('type.realax.real',X5))),s('type.realax.real',sK2(X3))))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X5))),s('type.realax.real',X6)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',X6)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',X3))))) )),
  inference(cnf_transformation,[],[f44])).
fof(f3417,plain,(
  ( ! [X4,X5] : (~p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),X4),s('type.realax.real',sK9(X5,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4)))))),s('type.realax.real',sK10(X5,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4))))))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),X4),s('type.realax.real',sK3))),s('type.realax.real',sK4)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X5)))) )),
  inference(forward_demodulation,[],[f3416,f75])).
fof(f3416,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),X4),s('type.realax.real',sK9(X5,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4)))))),s('type.realax.real',sK10(X5,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4))))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X5)))) )),
  inference(forward_demodulation,[],[f3415,f76])).
fof(f3415,plain,(
  ( ! [X4,X5] : (~p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),X4),s('type.realax.real',sK9(X5,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4)))))),s('type.realax.real',sK10(X5,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4))))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X5)))) )),
  inference(forward_demodulation,[],[f3400,f75])).
fof(f3400,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK9(X5,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4)))),s('type.realax.real',sK10(X5,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4))))))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X5)))) )),
  inference(superposition,[],[f970,f76])).
fof(f970,plain,(
  ( ! [X2,X3] : (~p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X3),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK9(X2,X3)),s('type.realax.real',sK10(X2,X3))))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X3),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0(X3) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2)))) )),
  inference(subsumption_resolution,[],[f964,f82])).
fof(f964,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X3),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK9(X2,X3)),s('type.realax.real',sK10(X2,X3))))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X3),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0(X3) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2)))) )),
  inference(superposition,[],[f100,f96])).
fof(f100,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X7)))) | ~p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X6),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK9(X7,X6)),s('type.realax.real',sK10(X7,X6))))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X6),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0(X6)) )),
  inference(resolution,[],[f64,f74])).
fof(f74,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | ~p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK9(X4,X0)),s('type.realax.real',sK10(X4,X0))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X4)))) | sP0(X0)) )),
  inference(cnf_transformation,[],[f49])).
fof(f39288,plain,(
  ( ! [X21,X22] : (s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X21),s('type.realax.real',X22))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X21),s('type.realax.real',X22)))))) ) | $spl80),
  inference(cnf_transformation,[],[f39288_D])).
fof(f39288_D,plain,(
  ( ! [X21,X22] : (s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X21),s('type.realax.real',X22))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X21),s('type.realax.real',X22)))))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f39036,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(X0,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) ) | $spl32),
  inference(cnf_transformation,[],[f39036_D])).
fof(f39036_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(X0,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f39261,plain,(
  ( ! [X99,X100,X98] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(X98,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X98))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X99),s('type.realax.real',X100))))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X99),s('type.realax.real',X100)))) ) | $spl21),
  inference(resolution,[],[f2457,f8250])).
fof(f8250,plain,(
  ( ! [X30,X28,X29,X27] : (sP0(X30) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(X29,X30)),s('type.realax.real',sK8(X30)))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X29))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X30),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X27),s('type.realax.real',X28))))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X27),s('type.realax.real',X28)))) )),
  inference(subsumption_resolution,[],[f8237,f94])).
fof(f8237,plain,(
  ( ! [X30,X28,X29,X27] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(X29,X30)),s('type.realax.real',sK8(X30)))) | sP0(X30) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X29))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X30),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X27),s('type.realax.real',X28))))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X27),s('type.realax.real',X28)))) )),
  inference(superposition,[],[f5931,f96])).
fof(f5931,plain,(
  ( ! [X26,X24,X23,X25] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X23),s('type.realax.real',X24)))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(X25,X26)),s('type.realax.real',sK8(X26)))) | sP0(X26) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X25))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X26),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X23),s('type.realax.real',X24)))))) )),
  inference(resolution,[],[f2742,f1181])).
fof(f1181,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f93,f96])).
fof(f2742,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X5),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X6),s('type.realax.real',X7)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X7)))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(X4,X5)),s('type.realax.real',sK8(X5)))) | sP0(X5) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X4)))) )),
  inference(subsumption_resolution,[],[f2736,f82])).
fof(f2736,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X5),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X6),s('type.realax.real',X7)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X7)))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(X4,X5)),s('type.realax.real',sK8(X5)))) | sP0(X5) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X4)))) )),
  inference(superposition,[],[f1203,f96])).
fof(f1203,plain,(
  ( ! [X28,X26,X29,X27] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X26)))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X27),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X28),s('type.realax.real',X29)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X28),s('type.realax.real',X29)))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(X26,X27)),s('type.realax.real',sK8(X27)))) | sP0(X27)) )),
  inference(resolution,[],[f1181,f71])).
fof(f39309,plain,(
  ( ! [X46] : (s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK9(sK2(X46),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))))) ) | $spl90),
  inference(cnf_transformation,[],[f39309_D])).
fof(f39309_D,plain,(
  ( ! [X46] : (s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK9(sK2(X46),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f39260,plain,(
  ( ! [X97,X96] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(X96,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X96))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK9(sK2(X97),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))))) ) | ($spl18 | $spl21)),
  inference(resolution,[],[f2457,f8249])).
fof(f8249,plain,(
  ( ! [X26,X24,X25] : (sP0(X26) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(X25,X26)),s('type.realax.real',sK8(X26)))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X25))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X26),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK9(sK2(X24),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))))) ) | $spl18),
  inference(subsumption_resolution,[],[f8236,f94])).
fof(f8236,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(X25,X26)),s('type.realax.real',sK8(X26)))) | sP0(X26) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X25))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X26),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK9(sK2(X24),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))))) ) | $spl18),
  inference(superposition,[],[f5931,f2456])).
fof(f39030,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(X0,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) ) | $spl30),
  inference(cnf_transformation,[],[f39030_D])).
fof(f39030_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(X0,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f39259,plain,(
  ( ! [X94,X95,X93] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(X93,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X93))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X94),s('type.realax.real',X95))))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X94),s('type.realax.real',X95)))) ) | $spl21),
  inference(resolution,[],[f2457,f8118])).
fof(f8118,plain,(
  ( ! [X30,X28,X29,X27] : (sP0(X29) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8(X29)),s('type.realax.real',sK10(X30,X29)))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X30))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X29),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X27),s('type.realax.real',X28))))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X27),s('type.realax.real',X28)))) )),
  inference(subsumption_resolution,[],[f8105,f94])).
fof(f8105,plain,(
  ( ! [X30,X28,X29,X27] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8(X29)),s('type.realax.real',sK10(X30,X29)))) | sP0(X29) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X30))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X29),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X27),s('type.realax.real',X28))))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X27),s('type.realax.real',X28)))) )),
  inference(superposition,[],[f5884,f96])).
fof(f5884,plain,(
  ( ! [X26,X24,X23,X25] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X23),s('type.realax.real',X24)))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8(X25)),s('type.realax.real',sK10(X26,X25)))) | sP0(X25) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X26))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X25),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X23),s('type.realax.real',X24)))))) )),
  inference(resolution,[],[f2617,f1181])).
fof(f2617,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X5),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X6),s('type.realax.real',X7)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X7)))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8(X5)),s('type.realax.real',sK10(X4,X5)))) | sP0(X5) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X4)))) )),
  inference(subsumption_resolution,[],[f2611,f82])).
fof(f2611,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X5),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X6),s('type.realax.real',X7)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X7)))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8(X5)),s('type.realax.real',sK10(X4,X5)))) | sP0(X5) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X4)))) )),
  inference(superposition,[],[f1200,f96])).
fof(f1200,plain,(
  ( ! [X21,X19,X20,X18] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X19)))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X18),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X20),s('type.realax.real',X21)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X20),s('type.realax.real',X21)))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8(X18)),s('type.realax.real',sK10(X19,X18)))) | sP0(X18)) )),
  inference(resolution,[],[f1181,f72])).
fof(f39258,plain,(
  ( ! [X92,X91] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(X91,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X91))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK9(sK2(X92),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))))) ) | ($spl18 | $spl21)),
  inference(resolution,[],[f2457,f8117])).
fof(f8117,plain,(
  ( ! [X26,X24,X25] : (sP0(X25) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8(X25)),s('type.realax.real',sK10(X26,X25)))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X26))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X25),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK9(sK2(X24),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))))) ) | $spl18),
  inference(subsumption_resolution,[],[f8104,f94])).
fof(f8104,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8(X25)),s('type.realax.real',sK10(X26,X25)))) | sP0(X25) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X26))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X25),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK9(sK2(X24),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))))) ) | $spl18),
  inference(superposition,[],[f5884,f2456])).
fof(f39326,plain,(
  ( ! [X85] : (s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK9(sK2(X85),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(X85),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))))) ) | $spl92),
  inference(cnf_transformation,[],[f39326_D])).
fof(f39326_D,plain,(
  ( ! [X85] : (s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK9(sK2(X85),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(X85),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f39256,plain,(
  ( ! [X88,X87,X89] : (s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK9(sK2(X87),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(X87),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X88),s('type.realax.real',X89))))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X88),s('type.realax.real',X89)))) ) | $spl21),
  inference(resolution,[],[f2457,f7200])).
fof(f7200,plain,(
  ( ! [X30,X28,X29,X27] : (sP0(X29) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X29),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK9(sK2(X30),X29)),s('type.realax.real',sK10(sK2(X30),X29)))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X29),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X27),s('type.realax.real',X28))))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X27),s('type.realax.real',X28)))) )),
  inference(subsumption_resolution,[],[f7187,f94])).
fof(f7187,plain,(
  ( ! [X30,X28,X29,X27] : (~p(s(bool,'T_0')) | sP0(X29) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X29),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK9(sK2(X30),X29)),s('type.realax.real',sK10(sK2(X30),X29)))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X29),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X27),s('type.realax.real',X28))))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X27),s('type.realax.real',X28)))) )),
  inference(superposition,[],[f3067,f96])).
fof(f3067,plain,(
  ( ! [X17,X15,X18,X16] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X15),s('type.realax.real',X16)))) | sP0(X17) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X17),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK9(sK2(X18),X17)),s('type.realax.real',sK10(sK2(X18),X17)))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X17),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X15),s('type.realax.real',X16)))))) )),
  inference(resolution,[],[f766,f1181])).
fof(f766,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X2),s('type.realax.real',X3)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X3)))) | sP0(X0) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK9(sK2(X1),X0)),s('type.realax.real',sK10(sK2(X1),X0))))))) )),
  inference(subsumption_resolution,[],[f752,f82])).
fof(f752,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X3)))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X2),s('type.realax.real',X3)))))) | sP0(X0) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK9(sK2(X1),X0)),s('type.realax.real',sK10(sK2(X1),X0))))))) )),
  inference(superposition,[],[f110,f96])).
fof(f110,plain,(
  ( ! [X14,X12,X15,X13] : (~p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X12),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK9(sK2(X15),X12)),s('type.realax.real',sK10(sK2(X15),X12))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X13),s('type.realax.real',X14)))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X12),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X13),s('type.realax.real',X14)))))) | sP0(X12)) )),
  inference(resolution,[],[f62,f74])).
fof(f39255,plain,(
  ( ! [X85,X86] : (s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK9(sK2(X85),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(X85),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK9(sK2(X86),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))))) ) | ($spl18 | $spl21)),
  inference(resolution,[],[f2457,f7199])).
fof(f7199,plain,(
  ( ! [X26,X24,X25] : (sP0(X25) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X25),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK9(sK2(X26),X25)),s('type.realax.real',sK10(sK2(X26),X25)))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X25),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK9(sK2(X24),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))))) ) | $spl18),
  inference(subsumption_resolution,[],[f7186,f94])).
fof(f7186,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,'T_0')) | sP0(X25) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X25),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK9(sK2(X26),X25)),s('type.realax.real',sK10(sK2(X26),X25)))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X25),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK9(sK2(X24),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))))) ) | $spl18),
  inference(superposition,[],[f3067,f2456])).
fof(f39185,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(X0))))) ) | $spl66),
  inference(cnf_transformation,[],[f39185_D])).
fof(f39185_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(X0))))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f39254,plain,(
  ( ! [X83,X84] : (s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X83),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X83),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(X83)))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK9(sK2(X84),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))))) ) | ($spl18 | $spl21)),
  inference(resolution,[],[f2457,f5341])).
fof(f5341,plain,(
  ( ! [X26,X24,X25] : (sP0(X26) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X25),X26)),s('type.realax.real',sK9(sK2(X25),X26)))),s('type.realax.real',sK2(X25)))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X26),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK9(sK2(X24),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))))) ) | $spl18),
  inference(subsumption_resolution,[],[f5328,f94])).
fof(f5328,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X25),X26)),s('type.realax.real',sK9(sK2(X25),X26)))),s('type.realax.real',sK2(X25)))) | sP0(X26) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X26),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK9(sK2(X24),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))))) ) | $spl18),
  inference(superposition,[],[f4255,f2456])).
fof(f4255,plain,(
  ( ! [X21,X19,X20,X18] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X18),s('type.realax.real',X19)))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X20),X21)),s('type.realax.real',sK9(sK2(X20),X21)))),s('type.realax.real',sK2(X20)))) | sP0(X21) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X21),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X18),s('type.realax.real',X19)))))) )),
  inference(resolution,[],[f3602,f1181])).
fof(f3602,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X1),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X2),s('type.realax.real',X3)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X3)))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X0),X1)),s('type.realax.real',sK9(sK2(X0),X1)))),s('type.realax.real',sK2(X0)))) | sP0(X1)) )),
  inference(subsumption_resolution,[],[f3594,f82])).
fof(f3594,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X1),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X2),s('type.realax.real',X3)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X3)))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X0),X1)),s('type.realax.real',sK9(sK2(X0),X1)))),s('type.realax.real',sK2(X0)))) | sP0(X1)) )),
  inference(superposition,[],[f1217,f116])).
fof(f1217,plain,(
  ( ! [X64,X62,X65,X63] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X62)))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X63),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X64),s('type.realax.real',X65)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X64),s('type.realax.real',X65)))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(X62,X63)),s('type.realax.real',sK9(X62,X63)))),s('type.realax.real',X62))) | sP0(X63)) )),
  inference(resolution,[],[f1181,f73])).
fof(f39253,plain,(
  ( ! [X80,X81,X82] : (s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X80),s('type.realax.real',X81))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X80),s('type.realax.real',X81))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X82),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X82),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(X82))))) ) | $spl21),
  inference(resolution,[],[f2457,f5081])).
fof(f5081,plain,(
  ( ! [X6,X4,X7,X5] : (sP0(X4) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',X6))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X4),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X5),s('type.realax.real',X6))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X7),X4)),s('type.realax.real',sK9(sK2(X7),X4)))),s('type.realax.real',sK2(X7))))) )),
  inference(resolution,[],[f2329,f1181])).
fof(f2329,plain,(
  ( ! [X17,X15,X18,X16] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X15),X16)),s('type.realax.real',sK9(sK2(X15),X16)))),s('type.realax.real',sK2(X15))))) | sP0(X16) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X17),s('type.realax.real',X18))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X16),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X17),s('type.realax.real',X18)))))) )),
  inference(resolution,[],[f680,f1181])).
fof(f680,plain,(
  ( ! [X4,X2,X5,X3] : (p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X4),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X2),s('type.realax.real',X3)))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X5),X4)),s('type.realax.real',sK9(sK2(X5),X4)))),s('type.realax.real',sK2(X5))))) | sP0(X4) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X3)))) )),
  inference(subsumption_resolution,[],[f674,f94])).
fof(f674,plain,(
  ( ! [X4,X2,X5,X3] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X4),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X2),s('type.realax.real',X3)))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X5),X4)),s('type.realax.real',sK9(sK2(X5),X4)))),s('type.realax.real',sK2(X5))))) | sP0(X4) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X3)))) )),
  inference(superposition,[],[f109,f96])).
fof(f109,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real',X10)))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X8),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X9),s('type.realax.real',X10)))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X11),X8)),s('type.realax.real',sK9(sK2(X11),X8)))),s('type.realax.real',sK2(X11))))) | sP0(X8)) )),
  inference(resolution,[],[f62,f73])).
fof(f3132,plain,(
  ( ! [X8] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(X8),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) ) | $spl24),
  inference(cnf_transformation,[],[f3132_D])).
fof(f3132_D,plain,(
  ( ! [X8] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(X8),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f39240,plain,(
  ( ! [X45,X46] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(X45),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK9(sK2(X46),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))))) ) | ($spl18 | $spl21)),
  inference(resolution,[],[f2457,f2703])).
fof(f2703,plain,(
  ( ! [X12,X10,X11] : (sP0(X11) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8(X11)),s('type.realax.real',sK10(sK2(X12),X11)))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X11),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK9(sK2(X10),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))))) ) | $spl18),
  inference(subsumption_resolution,[],[f2692,f94])).
fof(f2692,plain,(
  ( ! [X12,X10,X11] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8(X11)),s('type.realax.real',sK10(sK2(X12),X11)))) | sP0(X11) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X11),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK9(sK2(X10),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))))) ) | $spl18),
  inference(superposition,[],[f2622,f2456])).
fof(f2622,plain,(
  ( ! [X17,X15,X18,X16] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X15),s('type.realax.real',X16)))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8(X17)),s('type.realax.real',sK10(sK2(X18),X17)))) | sP0(X17) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X17),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X15),s('type.realax.real',X16)))))) )),
  inference(resolution,[],[f2616,f1181])).
fof(f2616,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X1),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X2),s('type.realax.real',X3)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X3)))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8(X1)),s('type.realax.real',sK10(sK2(X0),X1)))) | sP0(X1)) )),
  inference(subsumption_resolution,[],[f2610,f82])).
fof(f2610,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X1),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X2),s('type.realax.real',X3)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X3)))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8(X1)),s('type.realax.real',sK10(sK2(X0),X1)))) | sP0(X1)) )),
  inference(superposition,[],[f1200,f116])).
fof(f39308,plain,(
  ( ! [X43,X42] : (s(bool,'T_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK9(sK2(X43),X42)))),s('type.realax.real',sK8(X42)))) | sP0(X42) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X42),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))))) ) | $spl88),
  inference(cnf_transformation,[],[f39308_D])).
fof(f39308_D,plain,(
  ( ! [X43,X42] : (s(bool,'T_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK9(sK2(X43),X42)))),s('type.realax.real',sK8(X42)))) | sP0(X42) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X42),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f39307,plain,(
  ( ! [X43,X44,X42] : (s(bool,'T_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK9(sK2(X43),X42)))),s('type.realax.real',sK8(X42)))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X42),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(X44),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | sP0(X42)) ) | $spl21),
  inference(forward_demodulation,[],[f39306,f75])).
fof(f39306,plain,(
  ( ! [X43,X44,X42] : (s(bool,'T_0') = s(bool,'const.pair.UNCURRY_2'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK9(sK2(X43),X42)),s('type.realax.real',sK8(X42)))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X42),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(X44),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | sP0(X42)) ) | $spl21),
  inference(forward_demodulation,[],[f39239,f76])).
fof(f39239,plain,(
  ( ! [X43,X44,X42] : (p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X42),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK9(sK2(X43),X42)),s('type.realax.real',sK8(X42)))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(X44),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | sP0(X42)) ) | $spl21),
  inference(resolution,[],[f2457,f2680])).
fof(f2680,plain,(
  ( ! [X107,X105,X108,X106] : (sP0(X105) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X108),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X105),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK9(sK2(X107),X108)),s('type.realax.real',sK8(X108)))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8(X105)),s('type.realax.real',sK10(sK2(X106),X105)))) | sP0(X108)) )),
  inference(resolution,[],[f2622,f212])).
fof(f212,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X1),X0)),s('type.realax.real',sK8(X0))))) | sP0(X0)) )),
  inference(subsumption_resolution,[],[f205,f94])).
fof(f205,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X1),X0)),s('type.realax.real',sK8(X0))))) | sP0(X0)) )),
  inference(superposition,[],[f107,f106])).
fof(f39305,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK9(sK2(X37),X36)))),s('type.realax.real',sK8(X36)))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X36),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))) | sP0(X36)) ) | $spl86),
  inference(cnf_transformation,[],[f39305_D])).
fof(f39305_D,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK9(sK2(X37),X36)))),s('type.realax.real',sK8(X36)))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X36),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))) | sP0(X36)) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f39304,plain,(
  ( ! [X37,X38,X36] : (s(bool,'T_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK9(sK2(X37),X36)))),s('type.realax.real',sK8(X36)))) | sP0(X36) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(X38),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X36),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) ) | $spl21),
  inference(forward_demodulation,[],[f39303,f75])).
fof(f39303,plain,(
  ( ! [X37,X38,X36] : (s(bool,'T_0') = s(bool,'const.pair.UNCURRY_2'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK9(sK2(X37),X36)),s('type.realax.real',sK8(X36)))))) | sP0(X36) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(X38),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X36),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) ) | $spl21),
  inference(forward_demodulation,[],[f39237,f76])).
fof(f39237,plain,(
  ( ! [X37,X38,X36] : (sP0(X36) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK9(sK2(X37),X36)),s('type.realax.real',sK8(X36)))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(X38),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X36),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) ) | $spl21),
  inference(resolution,[],[f2457,f2679])).
fof(f2679,plain,(
  ( ! [X103,X101,X102,X104] : (sP0(X101) | sP0(X104) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X101),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK9(sK2(X103),X104)),s('type.realax.real',sK8(X104)))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8(X101)),s('type.realax.real',sK10(sK2(X102),X101)))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X104),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) )),
  inference(resolution,[],[f2622,f267])).
fof(f267,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X3),X2)),s('type.realax.real',sK8(X2))))) | sP0(X2) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X2),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) )),
  inference(subsumption_resolution,[],[f261,f93])).
fof(f261,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X3),X2)),s('type.realax.real',sK8(X2))))) | sP0(X2) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X2),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) )),
  inference(superposition,[],[f212,f96])).
fof(f39302,plain,(
  ( ! [X33,X34] : (s(bool,'T_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK8(X33)))),s('type.realax.real',sK10(sK2(X34),X33)))) | sP0(X33) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X33),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))))) ) | $spl84),
  inference(cnf_transformation,[],[f39302_D])).
fof(f39302_D,plain,(
  ( ! [X33,X34] : (s(bool,'T_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK8(X33)))),s('type.realax.real',sK10(sK2(X34),X33)))) | sP0(X33) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X33),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f39301,plain,(
  ( ! [X35,X33,X34] : (s(bool,'T_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK8(X33)))),s('type.realax.real',sK10(sK2(X34),X33)))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X33),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(X35),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | sP0(X33)) ) | $spl21),
  inference(forward_demodulation,[],[f39300,f75])).
fof(f39300,plain,(
  ( ! [X35,X33,X34] : (s(bool,'T_0') = s(bool,'const.pair.UNCURRY_2'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK8(X33)),s('type.realax.real',sK10(sK2(X34),X33)))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X33),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(X35),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | sP0(X33)) ) | $spl21),
  inference(forward_demodulation,[],[f39236,f76])).
fof(f39236,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X33),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK8(X33)),s('type.realax.real',sK10(sK2(X34),X33)))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(X35),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | sP0(X33)) ) | $spl21),
  inference(resolution,[],[f2457,f2670])).
fof(f2670,plain,(
  ( ! [X59,X57,X60,X58] : (sP0(X57) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X59),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X57),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK8(X59)),s('type.realax.real',sK10(sK2(X60),X59)))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8(X57)),s('type.realax.real',sK10(sK2(X58),X57)))) | sP0(X59)) )),
  inference(resolution,[],[f2622,f365])).
fof(f39295,plain,(
  ( ! [X28,X27] : (s(bool,'T_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK8(X27)))),s('type.realax.real',sK10(sK2(X28),X27)))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X27),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))) | sP0(X27)) ) | $spl82),
  inference(cnf_transformation,[],[f39295_D])).
fof(f39295_D,plain,(
  ( ! [X28,X27] : (s(bool,'T_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK8(X27)))),s('type.realax.real',sK10(sK2(X28),X27)))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X27),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))) | sP0(X27)) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f39294,plain,(
  ( ! [X28,X29,X27] : (s(bool,'T_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK8(X27)))),s('type.realax.real',sK10(sK2(X28),X27)))) | sP0(X27) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(X29),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X27),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) ) | $spl21),
  inference(forward_demodulation,[],[f39293,f75])).
fof(f39293,plain,(
  ( ! [X28,X29,X27] : (s(bool,'T_0') = s(bool,'const.pair.UNCURRY_2'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK8(X27)),s('type.realax.real',sK10(sK2(X28),X27)))))) | sP0(X27) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(X29),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X27),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) ) | $spl21),
  inference(forward_demodulation,[],[f39234,f76])).
fof(f39234,plain,(
  ( ! [X28,X29,X27] : (sP0(X27) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK8(X27)),s('type.realax.real',sK10(sK2(X28),X27)))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(X29),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X27),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) ) | $spl21),
  inference(resolution,[],[f2457,f2669])).
fof(f2669,plain,(
  ( ! [X54,X56,X55,X53] : (sP0(X53) | sP0(X55) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X53),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK8(X55)),s('type.realax.real',sK10(sK2(X56),X55)))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8(X53)),s('type.realax.real',sK10(sK2(X54),X53)))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X55),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) )),
  inference(resolution,[],[f2622,f393])).
fof(f393,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8(X2)),s('type.realax.real',sK10(sK2(X3),X2))))) | sP0(X2) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X2),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) )),
  inference(subsumption_resolution,[],[f387,f93])).
fof(f387,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8(X2)),s('type.realax.real',sK10(sK2(X3),X2))))) | sP0(X2) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X2),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) )),
  inference(superposition,[],[f365,f96])).
fof(f39062,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))) | $spl42),
  inference(cnf_transformation,[],[f39062_D])).
fof(f39062_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f39233,plain,(
  ( ! [X26] : (s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X26))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(X26,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) ) | $spl21),
  inference(resolution,[],[f2457,f1104])).
fof(f1104,plain,(
  ( ! [X2,X3] : (sP0(X2) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X3))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8(X2)),s('type.realax.real',sK10(X3,X2)))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X2),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) )),
  inference(subsumption_resolution,[],[f1102,f93])).
fof(f1102,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | sP0(X2) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X3))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8(X2)),s('type.realax.real',sK10(X3,X2)))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X2),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) )),
  inference(superposition,[],[f639,f96])).
fof(f639,plain,(
  ( ! [X2,X3] : (p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X2),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0(X2) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X3))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8(X2)),s('type.realax.real',sK10(X3,X2))))) )),
  inference(subsumption_resolution,[],[f628,f93])).
fof(f628,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X2),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0(X2) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X3))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8(X2)),s('type.realax.real',sK10(X3,X2))))) )),
  inference(superposition,[],[f461,f96])).
fof(f461,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8(X3)),s('type.realax.real',sK10(X2,X3))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X3),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0(X3) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2)))) )),
  inference(subsumption_resolution,[],[f459,f94])).
fof(f459,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8(X3)),s('type.realax.real',sK10(X2,X3))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X3),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0(X3) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2)))) )),
  inference(superposition,[],[f98,f96])).
fof(f98,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X3)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8(X2)),s('type.realax.real',sK10(X3,X2))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X2),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0(X2)) )),
  inference(resolution,[],[f64,f72])).
fof(f39232,plain,(
  ( ! [X25] : (s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X25))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(X25,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) ) | $spl21),
  inference(resolution,[],[f2457,f1036])).
fof(f1036,plain,(
  ( ! [X2,X3] : (sP0(X2) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X3))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(X3,X2)),s('type.realax.real',sK8(X2)))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X2),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) )),
  inference(subsumption_resolution,[],[f1034,f93])).
fof(f1034,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | sP0(X2) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X3))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(X3,X2)),s('type.realax.real',sK8(X2)))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X2),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) )),
  inference(superposition,[],[f576,f96])).
fof(f576,plain,(
  ( ! [X2,X3] : (p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X3),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0(X3) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(X2,X3)),s('type.realax.real',sK8(X3))))) )),
  inference(subsumption_resolution,[],[f565,f93])).
fof(f565,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X3),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0(X3) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(X2,X3)),s('type.realax.real',sK8(X3))))) )),
  inference(superposition,[],[f417,f96])).
fof(f417,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(X2,X3)),s('type.realax.real',sK8(X3))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X3),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0(X3) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2)))) )),
  inference(subsumption_resolution,[],[f415,f94])).
fof(f415,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(X2,X3)),s('type.realax.real',sK8(X3))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X3),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0(X3) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2)))) )),
  inference(superposition,[],[f97,f96])).
fof(f97,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(X1,X0)),s('type.realax.real',sK8(X0))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0(X0)) )),
  inference(resolution,[],[f64,f71])).
fof(f39230,plain,(
  ( ! [X23,X21,X22] : (s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X21),s('type.realax.real',X22))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X21),s('type.realax.real',X22))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(X23),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) ) | $spl21),
  inference(resolution,[],[f2457,f782])).
fof(f782,plain,(
  ( ! [X6,X4,X7,X5] : (sP0(X4) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X7))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X4),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X6),s('type.realax.real',X7))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8(X4)),s('type.realax.real',sK10(sK2(X5),X4))))) )),
  inference(subsumption_resolution,[],[f781,f93])).
fof(f781,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,'F_0')) | sP0(X4) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X7))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X4),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X6),s('type.realax.real',X7))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8(X4)),s('type.realax.real',sK10(sK2(X5),X4))))) )),
  inference(superposition,[],[f541,f96])).
fof(f541,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8(X4)),s('type.realax.real',sK10(sK2(X7),X4))))) | sP0(X4) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',X6))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X4),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X5),s('type.realax.real',X6)))))) )),
  inference(subsumption_resolution,[],[f530,f93])).
fof(f530,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,'F_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8(X4)),s('type.realax.real',sK10(sK2(X7),X4))))) | sP0(X4) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',X6))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X4),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X5),s('type.realax.real',X6)))))) )),
  inference(superposition,[],[f366,f96])).
fof(f366,plain,(
  ( ! [X4,X2,X5,X3] : (p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X4),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X2),s('type.realax.real',X3)))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8(X4)),s('type.realax.real',sK10(sK2(X5),X4))))) | sP0(X4) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X3)))) )),
  inference(subsumption_resolution,[],[f359,f94])).
fof(f359,plain,(
  ( ! [X4,X2,X5,X3] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X4),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X2),s('type.realax.real',X3)))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8(X4)),s('type.realax.real',sK10(sK2(X5),X4))))) | sP0(X4) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X3)))) )),
  inference(superposition,[],[f108,f96])).
fof(f39227,plain,(
  ( ! [X16] : (s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(X16),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) ) | $spl21),
  inference(resolution,[],[f2457,f407])).
fof(f407,plain,(
  ( ! [X2,X3] : (sP0(X2) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X2),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8(X2)),s('type.realax.real',sK10(sK2(X3),X2))))) )),
  inference(subsumption_resolution,[],[f406,f93])).
fof(f406,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | sP0(X2) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X2),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8(X2)),s('type.realax.real',sK10(sK2(X3),X2))))) )),
  inference(superposition,[],[f393,f96])).
fof(f39115,plain,(
  ( ! [X2,X3] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X2))),s('type.realax.real',X3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X3))))) ) | $spl50),
  inference(cnf_transformation,[],[f39115_D])).
fof(f39115_D,plain,(
  ( ! [X2,X3] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X2))),s('type.realax.real',X3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X3))))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f39283,plain,(
  ( ! [X14] : (~p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK9(X14,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK10(X14,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X14))))) ) | $spl78),
  inference(cnf_transformation,[],[f39283_D])).
fof(f39283_D,plain,(
  ( ! [X14] : (~p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK9(X14,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK10(X14,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X14))))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f39282,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK9(X14,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK10(X14,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X12))),s('type.realax.real',X13)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X12),s('type.realax.real',X13)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X14))))) ) | $spl21),
  inference(forward_demodulation,[],[f39281,f75])).
fof(f39281,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK9(X14,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(X14,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X12))),s('type.realax.real',X13)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X12),s('type.realax.real',X13)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X14))))) ) | $spl21),
  inference(forward_demodulation,[],[f39280,f76])).
fof(f39280,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X12))),s('type.realax.real',X13)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X12),s('type.realax.real',X13)))) | ~p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK9(X14,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(X14,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X14))))) ) | $spl21),
  inference(forward_demodulation,[],[f39279,f75])).
fof(f39279,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X12),s('type.realax.real',X13)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X12),s('type.realax.real',X13)))) | ~p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK9(X14,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(X14,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X14))))) ) | $spl21),
  inference(forward_demodulation,[],[f39225,f76])).
fof(f39225,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X12),s('type.realax.real',X13)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X12),s('type.realax.real',X13)))) | ~p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK9(X14,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(X14,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X14))))) ) | $spl21),
  inference(resolution,[],[f2457,f74])).
fof(f39278,plain,(
  ( ! [X11] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(X11,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(X11,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',X11)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X11))))) ) | $spl76),
  inference(cnf_transformation,[],[f39278_D])).
fof(f39278_D,plain,(
  ( ! [X11] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(X11,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(X11,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',X11)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X11))))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f39277,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X9))),s('type.realax.real',X10)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real',X10)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(X11,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(X11,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',X11)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X11))))) ) | $spl21),
  inference(forward_demodulation,[],[f39276,f75])).
fof(f39276,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X9),s('type.realax.real',X10)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real',X10)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(X11,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(X11,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',X11)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X11))))) ) | $spl21),
  inference(forward_demodulation,[],[f39224,f76])).
fof(f39224,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X9),s('type.realax.real',X10)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real',X10)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(X11,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(X11,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',X11)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X11))))) ) | $spl21),
  inference(resolution,[],[f2457,f73])).
fof(f39275,plain,(
  ( ! [X8] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(X8,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X8))))) ) | $spl74),
  inference(cnf_transformation,[],[f39275_D])).
fof(f39275_D,plain,(
  ( ! [X8] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(X8,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X8))))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f39274,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X6))),s('type.realax.real',X7)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X7)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(X8,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X8))))) ) | $spl21),
  inference(forward_demodulation,[],[f39273,f75])).
fof(f39273,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X6),s('type.realax.real',X7)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X7)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(X8,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X8))))) ) | $spl21),
  inference(forward_demodulation,[],[f39223,f76])).
fof(f39223,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X6),s('type.realax.real',X7)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X7)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(X8,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X8))))) ) | $spl21),
  inference(resolution,[],[f2457,f72])).
fof(f39272,plain,(
  ( ! [X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(X5,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X5))))) ) | $spl72),
  inference(cnf_transformation,[],[f39272_D])).
fof(f39272_D,plain,(
  ( ! [X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(X5,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X5))))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f39271,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X3))),s('type.realax.real',X4)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',X4)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(X5,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X5))))) ) | $spl21),
  inference(forward_demodulation,[],[f39270,f75])).
fof(f39270,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X3),s('type.realax.real',X4)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',X4)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(X5,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X5))))) ) | $spl21),
  inference(forward_demodulation,[],[f39222,f76])).
fof(f39222,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X3),s('type.realax.real',X4)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',X4)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(X5,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X5))))) ) | $spl21),
  inference(resolution,[],[f2457,f71])).
fof(f39218,plain,(
  s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | $spl70),
  inference(cnf_transformation,[],[f39218_D])).
fof(f39218_D,plain,(
  s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f2458,plain,(
  sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | $spl20),
  inference(cnf_transformation,[],[f2458_D])).
fof(f2458_D,plain,(
  sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f39216,plain,(
  ( ! [X6,X7] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X6))),s('type.realax.real',X7)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X7)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) ) | $spl18),
  inference(subsumption_resolution,[],[f39215,f769])).
fof(f769,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),X8),s('type.realax.real',sK9(sK2(X9),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X8)))))),s('type.realax.real',sK10(sK2(X9),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X8))))))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),X8),s('type.realax.real',X10))),s('type.realax.real',X11)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X10),s('type.realax.real',X11)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X8)))) )),
  inference(forward_demodulation,[],[f768,f75])).
fof(f768,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun('type.realax.real',fun('type.realax.real',bool)),X8),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X10),s('type.realax.real',X11)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),X8),s('type.realax.real',sK9(sK2(X9),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X8)))))),s('type.realax.real',sK10(sK2(X9),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X8))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X10),s('type.realax.real',X11)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X8)))) )),
  inference(forward_demodulation,[],[f767,f76])).
fof(f767,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),X8),s('type.realax.real',sK9(sK2(X9),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X8)))))),s('type.realax.real',sK10(sK2(X9),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X8))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X10),s('type.realax.real',X11)))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X8))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X10),s('type.realax.real',X11)))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X8)))) )),
  inference(forward_demodulation,[],[f754,f75])).
fof(f754,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun('type.realax.real',fun('type.realax.real',bool)),X8),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK9(sK2(X9),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X8)))),s('type.realax.real',sK10(sK2(X9),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X8))))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X10),s('type.realax.real',X11)))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X8))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X10),s('type.realax.real',X11)))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X8)))) )),
  inference(superposition,[],[f110,f76])).
fof(f39215,plain,(
  ( ! [X6,X7] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X6))),s('type.realax.real',X7)))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK9(sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK10(sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X7)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) ) | $spl18),
  inference(forward_demodulation,[],[f39214,f75])).
fof(f39214,plain,(
  ( ! [X6,X7] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X6),s('type.realax.real',X7)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK9(sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK10(sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X7)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) ) | $spl18),
  inference(forward_demodulation,[],[f4887,f76])).
fof(f4887,plain,(
  ( ! [X6,X7] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK9(sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK10(sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X6),s('type.realax.real',X7)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X7)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) ) | $spl18),
  inference(subsumption_resolution,[],[f4853,f82])).
fof(f4853,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK9(sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK10(sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X6),s('type.realax.real',X7)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X7)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) ) | $spl18),
  inference(superposition,[],[f1092,f2456])).
fof(f1092,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X0),X1)),s('type.realax.real',X0)))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK9(sK2(X0),X1)))),s('type.realax.real',sK10(sK2(X0),X1))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X1),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X2),s('type.realax.real',X3)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X3)))) | sP0(X1) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',sK10(sK2(X0),X1))))) )),
  inference(subsumption_resolution,[],[f1064,f82])).
fof(f1064,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK9(sK2(X0),X1)))),s('type.realax.real',sK10(sK2(X0),X1))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X0),X1)),s('type.realax.real',X0)))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X1),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X2),s('type.realax.real',X3)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X3)))) | sP0(X1) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',sK10(sK2(X0),X1))))) )),
  inference(superposition,[],[f228,f96])).
fof(f39180,plain,(
  ( ! [X47,X50] : (s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X47),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X47),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(X50)))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X47),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',X50))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X50),s('type.realax.real',sK10(sK2(X47),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) ) | $spl64),
  inference(cnf_transformation,[],[f39180_D])).
fof(f39180_D,plain,(
  ( ! [X47,X50] : (s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X47),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X47),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(X50)))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X47),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',X50))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X50),s('type.realax.real',sK10(sK2(X47),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f7268,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X16),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X16),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(X17)))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X16),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',X17))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X17),s('type.realax.real',sK10(sK2(X16),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X18))),s('type.realax.real',X19)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X18),s('type.realax.real',X19)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))) )),
  inference(resolution,[],[f3246,f769])).
fof(f3246,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X17))),s('type.realax.real',X16)))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X16),s('type.realax.real',X17))),s('type.realax.real',sK2(X15)))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X17),s('type.realax.real',X15))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X15),s('type.realax.real',X16)))) )),
  inference(subsumption_resolution,[],[f3232,f94])).
fof(f3232,plain,(
  ( ! [X17,X15,X16] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X17))),s('type.realax.real',X16)))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X16),s('type.realax.real',X17))),s('type.realax.real',sK2(X15)))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X17),s('type.realax.real',X15))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X15),s('type.realax.real',X16)))) )),
  inference(superposition,[],[f886,f96])).
fof(f886,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X3)))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X1))),s('type.realax.real',X3)))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X3),s('type.realax.real',X1))),s('type.realax.real',sK2(X2)))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) )),
  inference(subsumption_resolution,[],[f874,f82])).
fof(f874,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X3)))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X1))),s('type.realax.real',X3)))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X3),s('type.realax.real',X1))),s('type.realax.real',sK2(X2)))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) )),
  inference(superposition,[],[f243,f96])).
fof(f243,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X0)))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X1))),s('type.realax.real',X0)))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real',sK2(X2))))) )),
  inference(subsumption_resolution,[],[f225,f82])).
fof(f225,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X1))),s('type.realax.real',X0)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X0)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real',sK2(X2))))) )),
  inference(superposition,[],[f63,f96])).
fof(f39161,plain,(
  ( ! [X0,X3,X1] : (s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1))) | s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X1))),s('type.realax.real',X3))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X3)))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X0))),s('type.realax.real',X3))))) ) | $spl58),
  inference(cnf_transformation,[],[f39161_D])).
fof(f39161_D,plain,(
  ( ! [X0,X3,X1] : (s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1))) | s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X1))),s('type.realax.real',X3))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X3)))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X0))),s('type.realax.real',X3))))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f39204,plain,(
  ( ! [X2] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(X2),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))))) ) | $spl68),
  inference(cnf_transformation,[],[f39204_D])).
fof(f39204_D,plain,(
  ( ! [X2] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(X2),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f10204,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X1))),s('type.realax.real',X3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X2))),s('type.realax.real',X3)))) )),
  inference(resolution,[],[f542,f336])).
fof(f336,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X3))),s('type.realax.real',X1)))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X3))),s('type.realax.real',X2)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',X1)))) | s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X1))),s('type.realax.real',X2)))) )),
  inference(subsumption_resolution,[],[f315,f82])).
fof(f315,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X3))),s('type.realax.real',X2)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X3))),s('type.realax.real',X1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',X1)))) | s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X1))),s('type.realax.real',X2)))) )),
  inference(superposition,[],[f61,f96])).
fof(f61,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X8))),s('type.realax.real',X9)))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X7))),s('type.realax.real',X9)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X7))),s('type.realax.real',X8)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X8),s('type.realax.real',X9)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real',X8))))) )),
  inference(cnf_transformation,[],[f44])).
fof(f542,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),X8),s('type.realax.real',X9))),s('type.realax.real',X10)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X8)))),s('type.realax.real',sK10(sK2(X11),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X8))))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X8))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real',X10)))) )),
  inference(forward_demodulation,[],[f531,f75])).
fof(f531,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun('type.realax.real',fun('type.realax.real',bool)),X8),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X9),s('type.realax.real',X10)))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X8)))),s('type.realax.real',sK10(sK2(X11),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X8))))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X8))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real',X10)))) )),
  inference(superposition,[],[f366,f76])).
fof(f39175,plain,(
  ( ! [X6,X7,X5] : (s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',X6))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real',X5)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',X6)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X7))),s('type.realax.real',X5)))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X7))),s('type.realax.real',X6))))) ) | $spl62),
  inference(cnf_transformation,[],[f39175_D])).
fof(f39175_D,plain,(
  ( ! [X6,X7,X5] : (s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',X6))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real',X5)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',X6)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X7))),s('type.realax.real',X5)))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X7))),s('type.realax.real',X6))))) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f10205,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(X4),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',X6))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X7))),s('type.realax.real',X6)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X7))),s('type.realax.real',X5)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',X6)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real',X5))))) )),
  inference(resolution,[],[f542,f61])).
fof(f39203,plain,(
  ( ! [X2] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(X2),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))) ) | $spl7),
  inference(subsumption_resolution,[],[f39202,f1079])).
fof(f39202,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,'F_0') | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(X2),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))) )),
  inference(forward_demodulation,[],[f39201,f106])).
fof(f39201,plain,(
  ( ! [X2] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(X2),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))) )),
  inference(subsumption_resolution,[],[f10256,f93])).
fof(f10256,plain,(
  ( ! [X2] : (p(s(bool,'F_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(X2),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))) )),
  inference(superposition,[],[f542,f173])).
fof(f39200,plain,(
  ( ! [X0,X1] : (sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(X1)))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',X1))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',sK10(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) )),
  inference(subsumption_resolution,[],[f39199,f93])).
fof(f39199,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(X1)))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',X1))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',sK10(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) )),
  inference(forward_demodulation,[],[f10859,f173])).
fof(f10859,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',sK4)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(X1)))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',X1))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',sK10(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) )),
  inference(resolution,[],[f1004,f3246])).
fof(f1004,plain,(
  ( ! [X4,X5] : (~p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),X4),s('type.realax.real',sK9(sK2(X5),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4)))))),s('type.realax.real',sK10(sK2(X5),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4))))))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),X4),s('type.realax.real',sK3))),s('type.realax.real',sK4)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4)))) )),
  inference(forward_demodulation,[],[f1003,f75])).
fof(f1003,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),X4),s('type.realax.real',sK9(sK2(X5),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4)))))),s('type.realax.real',sK10(sK2(X5),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4))))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4)))) )),
  inference(forward_demodulation,[],[f1002,f76])).
fof(f1002,plain,(
  ( ! [X4,X5] : (~p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),X4),s('type.realax.real',sK9(sK2(X5),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4)))))),s('type.realax.real',sK10(sK2(X5),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4))))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4)))) )),
  inference(forward_demodulation,[],[f990,f75])).
fof(f990,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.pair.UNCURRY_2'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK9(sK2(X5),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4)))),s('type.realax.real',sK10(sK2(X5),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4))))))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4)))) )),
  inference(superposition,[],[f969,f76])).
fof(f969,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X1),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK9(sK2(X0),X1)),s('type.realax.real',sK10(sK2(X0),X1))))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X1),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0(X1)) )),
  inference(subsumption_resolution,[],[f963,f82])).
fof(f963,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X1),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK9(sK2(X0),X1)),s('type.realax.real',sK10(sK2(X0),X1))))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X1),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0(X1)) )),
  inference(superposition,[],[f100,f116])).
fof(f39108,plain,(
  ( ! [X0] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',sK10(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',X0))))) ) | $spl48),
  inference(cnf_transformation,[],[f39108_D])).
fof(f39108_D,plain,(
  ( ! [X0] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',sK10(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',X0))))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f39198,plain,(
  ( ! [X0] : (sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',sK10(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',X0))))) )),
  inference(subsumption_resolution,[],[f39197,f93])).
fof(f39197,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',sK10(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',X0))))) )),
  inference(forward_demodulation,[],[f39196,f173])).
fof(f39196,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',sK4)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',sK10(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',X0))))) )),
  inference(duplicate_literal_removal,[],[f39195])).
fof(f39195,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',sK4)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',sK10(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',X0)))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',sK4))))) )),
  inference(forward_demodulation,[],[f39194,f75])).
fof(f39194,plain,(
  ( ! [X0] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',sK10(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',X0)))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',sK4))))) )),
  inference(forward_demodulation,[],[f16042,f76])).
fof(f16042,plain,(
  ( ! [X0] : (sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',sK10(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',X0)))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',sK4))))) )),
  inference(duplicate_literal_removal,[],[f16028])).
fof(f16028,plain,(
  ( ! [X0] : (sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',sK10(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',X0)))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',sK4)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))) )),
  inference(resolution,[],[f594,f1004])).
fof(f594,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK9(sK2(X1),X0)))),s('type.realax.real',sK10(sK2(X1),X0))))) | sP0(X0) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',sK10(sK2(X1),X0))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X1),X0)),s('type.realax.real',X1))))) )),
  inference(resolution,[],[f590,f63])).
fof(f590,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X0),X1)),s('type.realax.real',sK9(sK2(X0),X1)))),s('type.realax.real',sK2(X0))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X1),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0(X1)) )),
  inference(subsumption_resolution,[],[f585,f94])).
fof(f585,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X0),X1)),s('type.realax.real',sK9(sK2(X0),X1)))),s('type.realax.real',sK2(X0))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X1),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0(X1)) )),
  inference(superposition,[],[f99,f116])).
fof(f99,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X5)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(X5,X4)),s('type.realax.real',sK9(X5,X4)))),s('type.realax.real',X5)))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X4),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0(X4)) )),
  inference(resolution,[],[f64,f73])).
fof(f39193,plain,(
  ( ! [X2,X3,X1] : (sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',sK10(sK2(X1),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X1),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',X1)))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X2))),s('type.realax.real',X3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X3))))) )),
  inference(subsumption_resolution,[],[f39192,f93])).
fof(f39192,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,'F_0')) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',sK10(sK2(X1),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X1),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',X1)))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X2))),s('type.realax.real',X3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X3))))) )),
  inference(forward_demodulation,[],[f39191,f173])).
fof(f39191,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',sK4)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',sK10(sK2(X1),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X1),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',X1)))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X2))),s('type.realax.real',X3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X3))))) )),
  inference(forward_demodulation,[],[f39190,f75])).
fof(f39190,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',sK10(sK2(X1),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X1),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',X1)))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X2))),s('type.realax.real',X3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X3))))) )),
  inference(forward_demodulation,[],[f16041,f76])).
fof(f16041,plain,(
  ( ! [X2,X3,X1] : (sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',sK10(sK2(X1),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X1),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',X1)))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X2))),s('type.realax.real',X3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X3))))) )),
  inference(duplicate_literal_removal,[],[f16029])).
fof(f16029,plain,(
  ( ! [X2,X3,X1] : (sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',sK10(sK2(X1),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X1),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',X1)))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X2))),s('type.realax.real',X3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X3)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))) )),
  inference(resolution,[],[f594,f769])).
fof(f2455,plain,(
  ( ! [X9] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',X9)))) | s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK4))),s('type.realax.real',X9))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK4),s('type.realax.real',X9))))) ) | $spl16),
  inference(cnf_transformation,[],[f2455_D])).
fof(f2455_D,plain,(
  ( ! [X9] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',X9)))) | s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK4))),s('type.realax.real',X9))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK4),s('type.realax.real',X9))))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f39189,plain,(
  ( ! [X8,X9] : (sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X8),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X8),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(X8)))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',X9)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK4),s('type.realax.real',X9)))) | s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK4))),s('type.realax.real',X9)))) )),
  inference(subsumption_resolution,[],[f39188,f82])).
fof(f39188,plain,(
  ( ! [X8,X9] : (~p(s(bool,'T_0')) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X8),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X8),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(X8)))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',X9)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK4),s('type.realax.real',X9)))) | s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK4))),s('type.realax.real',X9)))) )),
  inference(forward_demodulation,[],[f17078,f106])).
fof(f17078,plain,(
  ( ! [X8,X9] : (sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X8),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X8),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(X8)))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',X9)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK4),s('type.realax.real',X9)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))) | s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK4))),s('type.realax.real',X9)))) )),
  inference(resolution,[],[f614,f336])).
fof(f614,plain,(
  ( ! [X4,X5] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),X4),s('type.realax.real',sK3))),s('type.realax.real',sK4)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X5),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4)))),s('type.realax.real',sK9(sK2(X5),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4)))))),s('type.realax.real',sK2(X5))))) )),
  inference(forward_demodulation,[],[f612,f75])).
fof(f612,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X5),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4)))),s('type.realax.real',sK9(sK2(X5),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4)))))),s('type.realax.real',sK2(X5))))) )),
  inference(superposition,[],[f601,f76])).
fof(f601,plain,(
  ( ! [X2,X3] : (p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X3),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0(X3) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X2),X3)),s('type.realax.real',sK9(sK2(X2),X3)))),s('type.realax.real',sK2(X2))))) )),
  inference(subsumption_resolution,[],[f599,f93])).
fof(f599,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X3),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0(X3) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X2),X3)),s('type.realax.real',sK9(sK2(X2),X3)))),s('type.realax.real',sK2(X2))))) )),
  inference(superposition,[],[f590,f96])).
fof(f2461,plain,(
  ( ! [X11] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X11))),s('type.realax.real',sK4)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real',sK3)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X11))),s('type.realax.real',sK3))))) ) | $spl22),
  inference(cnf_transformation,[],[f2461_D])).
fof(f2461_D,plain,(
  ( ! [X11] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X11))),s('type.realax.real',sK4)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real',sK3)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X11))),s('type.realax.real',sK3))))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f39187,plain,(
  ( ! [X10,X11] : (sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X10),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X10),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(X10)))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X11))),s('type.realax.real',sK4)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X11))),s('type.realax.real',sK3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real',sK3))))) )),
  inference(subsumption_resolution,[],[f39186,f82])).
fof(f39186,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X10),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X10),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(X10)))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X11))),s('type.realax.real',sK4)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X11))),s('type.realax.real',sK3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real',sK3))))) )),
  inference(forward_demodulation,[],[f17079,f106])).
fof(f17079,plain,(
  ( ! [X10,X11] : (sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X10),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X10),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(X10)))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X11))),s('type.realax.real',sK4)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X11))),s('type.realax.real',sK3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real',sK3))))) )),
  inference(resolution,[],[f614,f61])).
fof(f39184,plain,(
  ( ! [X0] : (sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(X0))))) )),
  inference(subsumption_resolution,[],[f17091,f93])).
fof(f17091,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(X0))))) )),
  inference(superposition,[],[f614,f173])).
fof(f39183,plain,(
  ( ! [X45,X46] : (sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X45),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X45),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(X46)))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X46),s('type.realax.real',sK10(sK2(X45),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X45),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',X46)))) )),
  inference(subsumption_resolution,[],[f39182,f93])).
fof(f39182,plain,(
  ( ! [X45,X46] : (p(s(bool,'F_0')) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X45),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X45),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(X46)))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X46),s('type.realax.real',sK10(sK2(X45),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X45),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',X46)))) )),
  inference(forward_demodulation,[],[f39181,f173])).
fof(f39181,plain,(
  ( ! [X45,X46] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',sK4)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X45),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X45),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(X46)))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X46),s('type.realax.real',sK10(sK2(X45),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X45),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',X46)))) )),
  inference(subsumption_resolution,[],[f17230,f94])).
fof(f17230,plain,(
  ( ! [X45,X46] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',sK4)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X45),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X45),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(X46)))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X46),s('type.realax.real',sK10(sK2(X45),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X45),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',X46)))) )),
  inference(superposition,[],[f1004,f7269])).
fof(f7269,plain,(
  ( ! [X21,X22,X20] : (s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X20),s('type.realax.real',X21))),s('type.realax.real',sK2(X22)))) | s(bool,'T_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X21))),s('type.realax.real',X20))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X22),s('type.realax.real',X20))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X21),s('type.realax.real',X22)))) )),
  inference(resolution,[],[f3246,f1181])).
fof(f17304,plain,(
  ( ! [X47,X50,X48,X49] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X48))),s('type.realax.real',X49)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X48),s('type.realax.real',X49)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X47),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X47),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(X50)))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X50),s('type.realax.real',sK10(sK2(X47),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X47),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',X50)))) )),
  inference(subsumption_resolution,[],[f17231,f82])).
fof(f17231,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X48))),s('type.realax.real',X49)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X48),s('type.realax.real',X49)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X47),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X47),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(X50)))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X50),s('type.realax.real',sK10(sK2(X47),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X47),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',X50)))) )),
  inference(superposition,[],[f769,f7269])).
fof(f39146,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X2))),s('type.realax.real',X3))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X3)))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X1))),s('type.realax.real',X3))))) ) | $spl54),
  inference(cnf_transformation,[],[f39146_D])).
fof(f39146_D,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X2))),s('type.realax.real',X3))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X3)))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X1))),s('type.realax.real',X3))))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f17802,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(X2),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X0))),s('type.realax.real',X3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X3)))) | s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X1))),s('type.realax.real',X3)))) )),
  inference(duplicate_literal_removal,[],[f17767])).
fof(f17767,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(X2),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X0))),s('type.realax.real',X3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X1))),s('type.realax.real',X3)))) )),
  inference(resolution,[],[f2656,f336])).
fof(f2656,plain,(
  ( ! [X23,X21,X22,X20] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),X20),s('type.realax.real',X21))),s('type.realax.real',X22)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X21),s('type.realax.real',X22)))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X20)))),s('type.realax.real',sK10(sK2(X23),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X20)))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X20)))) )),
  inference(forward_demodulation,[],[f2640,f75])).
fof(f2640,plain,(
  ( ! [X23,X21,X22,X20] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun('type.realax.real',fun('type.realax.real',bool)),X20),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X21),s('type.realax.real',X22)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X21),s('type.realax.real',X22)))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X20)))),s('type.realax.real',sK10(sK2(X23),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X20)))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X20)))) )),
  inference(superposition,[],[f2616,f76])).
fof(f39145,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',X6)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real',X5)))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X7))),s('type.realax.real',X6))))) ) | $spl52),
  inference(cnf_transformation,[],[f39145_D])).
fof(f39145_D,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',X6)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real',X5)))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X7))),s('type.realax.real',X6))))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f39179,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real',X5)))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(X6),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X7))),s('type.realax.real',X5)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real',X4))))) )),
  inference(subsumption_resolution,[],[f17801,f2656])).
fof(f17801,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real',X5)))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(X6),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X7))),s('type.realax.real',X5)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X7))),s('type.realax.real',X4)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real',X4))))) )),
  inference(duplicate_literal_removal,[],[f17768])).
fof(f17768,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real',X5)))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(X6),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X7))),s('type.realax.real',X5)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X7))),s('type.realax.real',X4)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real',X5)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real',X4))))) )),
  inference(resolution,[],[f2656,f61])).
fof(f39178,plain,(
  ( ! [X6] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(X6),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))) )),
  inference(subsumption_resolution,[],[f39177,f94])).
fof(f39177,plain,(
  ( ! [X6] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(X6),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))) )),
  inference(forward_demodulation,[],[f39176,f106])).
fof(f39176,plain,(
  ( ! [X6] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(X6),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))) )),
  inference(subsumption_resolution,[],[f17792,f93])).
fof(f17792,plain,(
  ( ! [X6] : (p(s(bool,'F_0')) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(X6),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))) )),
  inference(superposition,[],[f2656,f173])).
fof(f39174,plain,(
  ( ! [X6] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X6),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X6),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(X6)))))) ) | $spl60),
  inference(cnf_transformation,[],[f39174_D])).
fof(f39174_D,plain,(
  ( ! [X6] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X6),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X6),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(X6)))))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f21860,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(X0))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X1))),s('type.realax.real',X3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X2))),s('type.realax.real',X3)))) )),
  inference(resolution,[],[f2362,f336])).
fof(f2362,plain,(
  ( ! [X23,X21,X22,X20] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),X20),s('type.realax.real',X21))),s('type.realax.real',X22)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X23),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X20)))),s('type.realax.real',sK9(sK2(X23),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X20)))))),s('type.realax.real',sK2(X23))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X20))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X21),s('type.realax.real',X22)))) )),
  inference(forward_demodulation,[],[f2351,f75])).
fof(f2351,plain,(
  ( ! [X23,X21,X22,X20] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun('type.realax.real',fun('type.realax.real',bool)),X20),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X21),s('type.realax.real',X22)))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X23),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X20)))),s('type.realax.real',sK9(sK2(X23),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X20)))))),s('type.realax.real',sK2(X23))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X20))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X21),s('type.realax.real',X22)))) )),
  inference(superposition,[],[f680,f76])).
fof(f21861,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X4),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X4),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(X4))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',X6))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X7))),s('type.realax.real',X6)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X7))),s('type.realax.real',X5)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',X6)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real',X5))))) )),
  inference(resolution,[],[f2362,f61])).
fof(f39173,plain,(
  ( ! [X6] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X6),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X6),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(X6))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))) ) | $spl7),
  inference(subsumption_resolution,[],[f39172,f1079])).
fof(f39172,plain,(
  ( ! [X6] : (s(bool,'T_0') = s(bool,'F_0') | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X6),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X6),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(X6))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))) )),
  inference(forward_demodulation,[],[f39171,f106])).
fof(f39171,plain,(
  ( ! [X6] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X6),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X6),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(X6))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))) )),
  inference(subsumption_resolution,[],[f21889,f93])).
fof(f21889,plain,(
  ( ! [X6] : (p(s(bool,'F_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X6),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X6),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(X6))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))) )),
  inference(superposition,[],[f2362,f173])).
fof(f39079,plain,(
  ( ! [X28] : (s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X28),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X28),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))))) ) | $spl46),
  inference(cnf_transformation,[],[f39079_D])).
fof(f39079_D,plain,(
  ( ! [X28] : (s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X28),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X28),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f39170,plain,(
  ( ! [X16] : (s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X16),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X16),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))) )),
  inference(subsumption_resolution,[],[f39169,f93])).
fof(f39169,plain,(
  ( ! [X16] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X16),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X16),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))) )),
  inference(forward_demodulation,[],[f39168,f173])).
fof(f39168,plain,(
  ( ! [X16] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',sK4)))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X16),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X16),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))) )),
  inference(duplicate_literal_removal,[],[f39167])).
fof(f39167,plain,(
  ( ! [X16] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',sK4)))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X16),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X16),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',sK4))))) )),
  inference(forward_demodulation,[],[f39166,f75])).
fof(f39166,plain,(
  ( ! [X16] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X16),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X16),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',sK4))))) )),
  inference(forward_demodulation,[],[f24550,f76])).
fof(f24550,plain,(
  ( ! [X16] : (s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X16),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X16),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',sK4))))) )),
  inference(subsumption_resolution,[],[f24549,f365])).
fof(f24549,plain,(
  ( ! [X16] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(X16),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X16),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X16),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',sK4))))) )),
  inference(duplicate_literal_removal,[],[f24533])).
fof(f24533,plain,(
  ( ! [X16] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(X16),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X16),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X16),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',sK4)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))) )),
  inference(resolution,[],[f864,f1004])).
fof(f864,plain,(
  ( ! [X59,X60,X58] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK9(sK2(X60),X58)))),s('type.realax.real',X59)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8(X58)),s('type.realax.real',X59)))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X59),s('type.realax.real',sK9(sK2(X60),X58)))),s('type.realax.real',sK2(sK8(X58))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X58),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0(X58)) )),
  inference(resolution,[],[f243,f212])).
fof(f39165,plain,(
  ( ! [X19,X17,X18] : (s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X17),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X17),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X18))),s('type.realax.real',X19)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X18),s('type.realax.real',X19))))) )),
  inference(subsumption_resolution,[],[f39164,f93])).
fof(f39164,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X17),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X17),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X18))),s('type.realax.real',X19)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X18),s('type.realax.real',X19))))) )),
  inference(forward_demodulation,[],[f39163,f173])).
fof(f39163,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',sK4)))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X17),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X17),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X18))),s('type.realax.real',X19)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X18),s('type.realax.real',X19))))) )),
  inference(forward_demodulation,[],[f39162,f75])).
fof(f39162,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X17),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X17),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X18))),s('type.realax.real',X19)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X18),s('type.realax.real',X19))))) )),
  inference(forward_demodulation,[],[f24551,f76])).
fof(f24551,plain,(
  ( ! [X19,X17,X18] : (s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X17),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X17),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X18))),s('type.realax.real',X19)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X18),s('type.realax.real',X19))))) )),
  inference(subsumption_resolution,[],[f24548,f365])).
fof(f24548,plain,(
  ( ! [X19,X17,X18] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(X17),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X17),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X17),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X18))),s('type.realax.real',X19)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X18),s('type.realax.real',X19))))) )),
  inference(duplicate_literal_removal,[],[f24534])).
fof(f24534,plain,(
  ( ! [X19,X17,X18] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(X17),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X17),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X17),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X18))),s('type.realax.real',X19)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X18),s('type.realax.real',X19)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))) )),
  inference(resolution,[],[f864,f769])).
fof(f25782,plain,(
  ( ! [X2,X0,X3,X1] : (sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(X2),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X0))),s('type.realax.real',X3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X1))),s('type.realax.real',X3)))) )),
  inference(resolution,[],[f808,f336])).
fof(f808,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),X8),s('type.realax.real',X9))),s('type.realax.real',X10)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X8))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real',X10))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X8)))),s('type.realax.real',sK10(sK2(X11),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X8))))))) )),
  inference(forward_demodulation,[],[f800,f75])).
fof(f800,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun('type.realax.real',fun('type.realax.real',bool)),X8),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X9),s('type.realax.real',X10)))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X8))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real',X10))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X8)))),s('type.realax.real',sK10(sK2(X11),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X8))))))) )),
  inference(superposition,[],[f543,f76])).
fof(f543,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X4),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X6),s('type.realax.real',X7)))))) | sP0(X4) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X7))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8(X4)),s('type.realax.real',sK10(sK2(X5),X4))))) )),
  inference(subsumption_resolution,[],[f533,f93])).
fof(f533,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,'F_0')) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X4),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X6),s('type.realax.real',X7)))))) | sP0(X4) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X7))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8(X4)),s('type.realax.real',sK10(sK2(X5),X4))))) )),
  inference(superposition,[],[f366,f96])).
fof(f39160,plain,(
  ( ! [X4,X7,X5] : (s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real',X5))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real',X4)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real',X5)))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X7))),s('type.realax.real',X5))))) ) | $spl56),
  inference(cnf_transformation,[],[f39160_D])).
fof(f39160_D,plain,(
  ( ! [X4,X7,X5] : (s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real',X5))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real',X4)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real',X5)))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X7))),s('type.realax.real',X5))))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f39159,plain,(
  ( ! [X6,X4,X7,X5] : (sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real',X5))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(X6),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X7))),s('type.realax.real',X5)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real',X5)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real',X4))))) )),
  inference(subsumption_resolution,[],[f25783,f2656])).
fof(f25783,plain,(
  ( ! [X6,X4,X7,X5] : (sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real',X5))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(X6),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X7))),s('type.realax.real',X5)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X7))),s('type.realax.real',X4)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real',X5)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real',X4))))) )),
  inference(resolution,[],[f808,f61])).
fof(f39158,plain,(
  ( ! [X6] : (sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(X6),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) ) | $spl7),
  inference(subsumption_resolution,[],[f39157,f1079])).
fof(f39157,plain,(
  ( ! [X6] : (s(bool,'T_0') = s(bool,'F_0') | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(X6),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) )),
  inference(forward_demodulation,[],[f39156,f106])).
fof(f39156,plain,(
  ( ! [X6] : (sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(X6),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) )),
  inference(subsumption_resolution,[],[f25807,f93])).
fof(f25807,plain,(
  ( ! [X6] : (p(s(bool,'F_0')) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(X6),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) )),
  inference(superposition,[],[f808,f173])).
fof(f39150,plain,(
  ( ! [X0,X1] : (sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X0))),s('type.realax.real',X1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1))))) )),
  inference(subsumption_resolution,[],[f39149,f93])).
fof(f39149,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X0))),s('type.realax.real',X1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1))))) )),
  inference(forward_demodulation,[],[f39148,f173])).
fof(f39148,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',sK4)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X0))),s('type.realax.real',X1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1))))) )),
  inference(forward_demodulation,[],[f39147,f75])).
fof(f39147,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X0))),s('type.realax.real',X1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1))))) )),
  inference(forward_demodulation,[],[f30467,f76])).
fof(f30467,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X0))),s('type.realax.real',X1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1))))) )),
  inference(duplicate_literal_removal,[],[f30454])).
fof(f30454,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X0))),s('type.realax.real',X1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))) )),
  inference(resolution,[],[f30421,f769])).
fof(f39042,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK9(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK10(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) ) | $spl34),
  inference(cnf_transformation,[],[f39042_D])).
fof(f39042_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK9(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK10(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f31248,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK9(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK10(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X1))),s('type.realax.real',X3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X3)))) | s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X2))),s('type.realax.real',X3)))) )),
  inference(duplicate_literal_removal,[],[f31213])).
fof(f31213,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK9(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK10(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X1))),s('type.realax.real',X3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X2))),s('type.realax.real',X3)))) )),
  inference(resolution,[],[f3104,f336])).
fof(f3104,plain,(
  ( ! [X23,X21,X22,X20] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),X20),s('type.realax.real',X21))),s('type.realax.real',X22)))) | s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),X20),s('type.realax.real',sK9(sK2(X23),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X20)))))),s('type.realax.real',sK10(sK2(X23),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X20)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X21),s('type.realax.real',X22)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X20)))) )),
  inference(forward_demodulation,[],[f3103,f75])).
fof(f3103,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,'F_0') = s(bool,'const.pair.UNCURRY_2'(s(fun('type.realax.real',fun('type.realax.real',bool)),X20),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK9(sK2(X23),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X20)))),s('type.realax.real',sK10(sK2(X23),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X20)))))))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),X20),s('type.realax.real',X21))),s('type.realax.real',X22)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X21),s('type.realax.real',X22)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X20)))) )),
  inference(forward_demodulation,[],[f3102,f76])).
fof(f3102,plain,(
  ( ! [X23,X21,X22,X20] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),X20),s('type.realax.real',X21))),s('type.realax.real',X22)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X21),s('type.realax.real',X22)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X20))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X20))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK9(sK2(X23),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X20)))),s('type.realax.real',sK10(sK2(X23),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X20))))))))) )),
  inference(forward_demodulation,[],[f3085,f75])).
fof(f3085,plain,(
  ( ! [X23,X21,X22,X20] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun('type.realax.real',fun('type.realax.real',bool)),X20),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X21),s('type.realax.real',X22)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X21),s('type.realax.real',X22)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X20))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X20))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK9(sK2(X23),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X20)))),s('type.realax.real',sK10(sK2(X23),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X20))))))))) )),
  inference(superposition,[],[f766,f76])).
fof(f39144,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK9(sK2(X4),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK10(sK2(X4),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',X6)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X7))),s('type.realax.real',X6)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real',X5))))) )),
  inference(subsumption_resolution,[],[f31247,f3104])).
fof(f31247,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK9(sK2(X4),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK10(sK2(X4),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',X6)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X7))),s('type.realax.real',X6)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X7))),s('type.realax.real',X5)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real',X5))))) )),
  inference(duplicate_literal_removal,[],[f31214])).
fof(f31214,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK9(sK2(X4),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK10(sK2(X4),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',X6)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X7))),s('type.realax.real',X6)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X7))),s('type.realax.real',X5)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',X6)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real',X5))))) )),
  inference(resolution,[],[f3104,f61])).
fof(f39143,plain,(
  ( ! [X6] : (s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK9(sK2(X6),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK10(sK2(X6),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))) )),
  inference(subsumption_resolution,[],[f39142,f94])).
fof(f39142,plain,(
  ( ! [X6] : (~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK9(sK2(X6),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK10(sK2(X6),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))) )),
  inference(forward_demodulation,[],[f39141,f106])).
fof(f39141,plain,(
  ( ! [X6] : (s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK9(sK2(X6),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK10(sK2(X6),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))) )),
  inference(subsumption_resolution,[],[f31238,f93])).
fof(f31238,plain,(
  ( ! [X6] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK9(sK2(X6),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK10(sK2(X6),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))) )),
  inference(superposition,[],[f3104,f173])).
fof(f39060,plain,(
  s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK5('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK6('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | $spl40),
  inference(cnf_transformation,[],[f39060_D])).
fof(f39060_D,plain,(
  s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK5('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK6('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f39140,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK5('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK6('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))),
  inference(subsumption_resolution,[],[f39139,f93])).
fof(f39139,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK5('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK6('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))),
  inference(forward_demodulation,[],[f39138,f173])).
fof(f39138,plain,(
  p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',sK4)))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK5('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK6('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))),
  inference(subsumption_resolution,[],[f39137,f94])).
fof(f39137,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',sK4)))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK5('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK6('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))),
  inference(subsumption_resolution,[],[f31378,f1196])).
fof(f1196,plain,(
  ( ! [X12] : (~sP0(X12) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK5(X12)),s('type.realax.real',sK6(X12))))) )),
  inference(resolution,[],[f1181,f66])).
fof(f66,plain,(
  ( ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK5(X0)),s('type.realax.real',sK6(X0))))) | ~sP0(X0)) )),
  inference(cnf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0] : (~sP0(X0) | (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK5(X0)),s('type.realax.real',sK6(X0))))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK6(X0)),s('type.realax.real',sK7(X0))))) & p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK5(X0)),s('type.realax.real',sK6(X0))))))) & p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK6(X0)),s('type.realax.real',sK7(X0))))))) & ~p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK5(X0)),s('type.realax.real',sK7(X0)))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5,sK6,sK7])],[f46])).
fof(f46,plain,(
  ! [X0] : (~sP0(X0) | ? [X1,X2,X3] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X3)))) & p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2)))))) & p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X2),s('type.realax.real',X3)))))) & ~p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X3))))))))),
  inference(rectify,[],[f45])).
fof(f45,plain,(
  ! [X0] : (~sP0(X0) | ? [X5,X6,X7] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',X6)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X7)))) & p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X5),s('type.realax.real',X6)))))) & p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X6),s('type.realax.real',X7)))))) & ~p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X5),s('type.realax.real',X7))))))))),
  inference(nnf_transformation,[],[f41])).
fof(f31378,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',sK4)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK5('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK6('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))),
  inference(superposition,[],[f1004,f30507])).
fof(f30507,plain,(
  ( ! [X25] : (s(bool,'T_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK9(sK2(sK8(X25)),X25)))),s('type.realax.real',sK10(sK2(sK8(X25)),X25)))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X25),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK5(X25)),s('type.realax.real',sK6(X25))))) )),
  inference(resolution,[],[f30470,f1196])).
fof(f30470,plain,(
  ( ! [X0] : (sP0(X0) | s(bool,'T_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK9(sK2(sK8(X0)),X0)))),s('type.realax.real',sK10(sK2(sK8(X0)),X0)))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) )),
  inference(resolution,[],[f30457,f1181])).
fof(f30457,plain,(
  ( ! [X6] : (p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X6),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0(X6) | s(bool,'T_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK9(sK2(sK8(X6)),X6)))),s('type.realax.real',sK10(sK2(sK8(X6)),X6))))) )),
  inference(resolution,[],[f30421,f1181])).
fof(f39136,plain,(
  ( ! [X2,X1] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X1))),s('type.realax.real',X2)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK5('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK6('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) ) | $spl7),
  inference(subsumption_resolution,[],[f39135,f1079])).
fof(f39135,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') = s(bool,'F_0') | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X1))),s('type.realax.real',X2)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK5('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK6('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) )),
  inference(forward_demodulation,[],[f39134,f173])).
fof(f39134,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',sK4))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X1))),s('type.realax.real',X2)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK5('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK6('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) )),
  inference(forward_demodulation,[],[f39133,f75])).
fof(f39133,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') = s(bool,'const.pair.UNCURRY_2'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X1))),s('type.realax.real',X2)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK5('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK6('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) )),
  inference(forward_demodulation,[],[f39132,f76])).
fof(f39132,plain,(
  ( ! [X2,X1] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X1))),s('type.realax.real',X2)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK5('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK6('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) )),
  inference(subsumption_resolution,[],[f31415,f1196])).
fof(f31415,plain,(
  ( ! [X2,X1] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X1))),s('type.realax.real',X2)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK5('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK6('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) )),
  inference(subsumption_resolution,[],[f31379,f82])).
fof(f31379,plain,(
  ( ! [X2,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X1))),s('type.realax.real',X2)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK5('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK6('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) )),
  inference(superposition,[],[f769,f30507])).
fof(f39070,plain,(
  s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK6('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK7('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | $spl44),
  inference(cnf_transformation,[],[f39070_D])).
fof(f39070_D,plain,(
  s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK6('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK7('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f39131,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK6('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK7('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))),
  inference(subsumption_resolution,[],[f39130,f93])).
fof(f39130,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK6('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK7('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))),
  inference(forward_demodulation,[],[f39129,f173])).
fof(f39129,plain,(
  p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',sK4)))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK6('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK7('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))),
  inference(subsumption_resolution,[],[f39128,f94])).
fof(f39128,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',sK4)))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK6('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK7('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))),
  inference(subsumption_resolution,[],[f31430,f1199])).
fof(f1199,plain,(
  ( ! [X17] : (~sP0(X17) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK6(X17)),s('type.realax.real',sK7(X17))))) )),
  inference(resolution,[],[f1181,f67])).
fof(f67,plain,(
  ( ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK6(X0)),s('type.realax.real',sK7(X0))))) | ~sP0(X0)) )),
  inference(cnf_transformation,[],[f47])).
fof(f31430,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',sK4)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK6('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK7('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))),
  inference(superposition,[],[f1004,f30508])).
fof(f30508,plain,(
  ( ! [X26] : (s(bool,'T_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK9(sK2(sK8(X26)),X26)))),s('type.realax.real',sK10(sK2(sK8(X26)),X26)))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X26),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK6(X26)),s('type.realax.real',sK7(X26))))) )),
  inference(resolution,[],[f30470,f1199])).
fof(f39127,plain,(
  ( ! [X2,X3] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X2))),s('type.realax.real',X3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X3)))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK6('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK7('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) ) | $spl7),
  inference(subsumption_resolution,[],[f39126,f1079])).
fof(f39126,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,'F_0') | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X2))),s('type.realax.real',X3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X3)))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK6('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK7('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) )),
  inference(forward_demodulation,[],[f39125,f173])).
fof(f39125,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',sK4))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X2))),s('type.realax.real',X3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X3)))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK6('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK7('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) )),
  inference(forward_demodulation,[],[f39124,f75])).
fof(f39124,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,'const.pair.UNCURRY_2'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X2))),s('type.realax.real',X3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X3)))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK6('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK7('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) )),
  inference(forward_demodulation,[],[f39123,f76])).
fof(f39123,plain,(
  ( ! [X2,X3] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X2))),s('type.realax.real',X3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X3)))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK6('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK7('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) )),
  inference(subsumption_resolution,[],[f31467,f1199])).
fof(f31467,plain,(
  ( ! [X2,X3] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X2))),s('type.realax.real',X3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X3)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK6('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK7('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) )),
  inference(subsumption_resolution,[],[f31431,f82])).
fof(f31431,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X2))),s('type.realax.real',X3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X3)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK6('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK7('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) )),
  inference(superposition,[],[f769,f30508])).
fof(f39122,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))) ) | ($spl7 | $spl18)),
  inference(subsumption_resolution,[],[f39121,f93])).
fof(f39121,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))) ) | ($spl7 | $spl18)),
  inference(forward_demodulation,[],[f39120,f173])).
fof(f39120,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',sK4)))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))) ) | ($spl7 | $spl18)),
  inference(duplicate_literal_removal,[],[f39119])).
fof(f39119,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',sK4)))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',sK4))))) ) | ($spl7 | $spl18)),
  inference(forward_demodulation,[],[f39118,f75])).
fof(f39118,plain,(
  ( ! [X0] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',sK4))))) ) | ($spl7 | $spl18)),
  inference(forward_demodulation,[],[f39117,f76])).
fof(f39117,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',sK4))))) ) | ($spl7 | $spl18)),
  inference(subsumption_resolution,[],[f39116,f1079])).
fof(f39116,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',sK4))))) ) | $spl18),
  inference(forward_demodulation,[],[f32611,f2456])).
fof(f32611,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',sK4))))) )),
  inference(duplicate_literal_removal,[],[f32591])).
fof(f32591,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',sK4)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))) )),
  inference(resolution,[],[f3207,f1004])).
fof(f3207,plain,(
  ( ! [X50,X51,X49] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X49))),s('type.realax.real',sK10(sK2(X50),X51))))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X50),X51)),s('type.realax.real',X49))),s('type.realax.real',sK2(sK8(X51))))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X49),s('type.realax.real',sK8(X51)))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X51),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0(X51)) )),
  inference(resolution,[],[f886,f365])).
fof(f39114,plain,(
  ( ! [X2,X3,X1] : (s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X1),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X1),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X2))),s('type.realax.real',X3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X3))))) ) | ($spl7 | $spl18)),
  inference(subsumption_resolution,[],[f39113,f93])).
fof(f39113,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X1),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X1),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X2))),s('type.realax.real',X3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X3))))) ) | ($spl7 | $spl18)),
  inference(forward_demodulation,[],[f39112,f173])).
fof(f39112,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',sK4)))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X1),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X1),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X2))),s('type.realax.real',X3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X3))))) ) | ($spl7 | $spl18)),
  inference(forward_demodulation,[],[f39111,f75])).
fof(f39111,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X1),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X1),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X2))),s('type.realax.real',X3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X3))))) ) | ($spl7 | $spl18)),
  inference(forward_demodulation,[],[f39110,f76])).
fof(f39110,plain,(
  ( ! [X2,X3,X1] : (s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X1),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X1),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X2))),s('type.realax.real',X3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X3))))) ) | ($spl7 | $spl18)),
  inference(subsumption_resolution,[],[f39109,f1079])).
fof(f39109,plain,(
  ( ! [X2,X3,X1] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X1),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X1),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X2))),s('type.realax.real',X3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X3))))) ) | $spl18),
  inference(forward_demodulation,[],[f32610,f2456])).
fof(f32610,plain,(
  ( ! [X2,X3,X1] : (s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X1),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X1),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X1),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X2))),s('type.realax.real',X3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X3))))) )),
  inference(duplicate_literal_removal,[],[f32592])).
fof(f32592,plain,(
  ( ! [X2,X3,X1] : (s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X1),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X1),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X1),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X2))),s('type.realax.real',X3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X3)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))) )),
  inference(resolution,[],[f3207,f769])).
fof(f39107,plain,(
  ( ! [X0] : (sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',sK10(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',X0))))) ) | $spl7),
  inference(subsumption_resolution,[],[f39106,f93])).
fof(f39106,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',sK10(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',X0))))) ) | $spl7),
  inference(forward_demodulation,[],[f39105,f173])).
fof(f39105,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',sK4)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',sK10(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',X0))))) ) | $spl7),
  inference(forward_demodulation,[],[f39104,f75])).
fof(f39104,plain,(
  ( ! [X0] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',sK10(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',X0))))) ) | $spl7),
  inference(forward_demodulation,[],[f39103,f76])).
fof(f39103,plain,(
  ( ! [X0] : (sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',sK10(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',X0))))) ) | $spl7),
  inference(subsumption_resolution,[],[f39102,f1079])).
fof(f39102,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'F_0') | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',sK10(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',X0))))) )),
  inference(forward_demodulation,[],[f39101,f116])).
fof(f39101,plain,(
  ( ! [X0] : (sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2(X0)))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',sK10(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',X0))))) )),
  inference(subsumption_resolution,[],[f39100,f93])).
fof(f39100,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2(X0)))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',sK10(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',X0))))) )),
  inference(forward_demodulation,[],[f32701,f173])).
fof(f32701,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',sK4)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2(X0)))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',sK10(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',X0))))) )),
  inference(duplicate_literal_removal,[],[f32661])).
fof(f32661,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',sK4)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2(X0)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',sK10(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',X0))))) )),
  inference(resolution,[],[f3417,f594])).
fof(f39091,plain,(
  ( ! [X1] : (sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X1),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X1),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))))) ) | ($spl7 | $spl18)),
  inference(subsumption_resolution,[],[f39090,f93])).
fof(f39090,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X1),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X1),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))))) ) | ($spl7 | $spl18)),
  inference(forward_demodulation,[],[f39089,f173])).
fof(f39089,plain,(
  ( ! [X1] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',sK4)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X1),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X1),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))))) ) | ($spl7 | $spl18)),
  inference(forward_demodulation,[],[f39088,f75])).
fof(f39088,plain,(
  ( ! [X1] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X1),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X1),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))))) ) | ($spl7 | $spl18)),
  inference(forward_demodulation,[],[f39087,f76])).
fof(f39087,plain,(
  ( ! [X1] : (sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X1),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X1),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))))) ) | ($spl7 | $spl18)),
  inference(subsumption_resolution,[],[f39086,f1079])).
fof(f39086,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,'F_0') | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X1),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X1),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))))) ) | ($spl7 | $spl18)),
  inference(forward_demodulation,[],[f39085,f2456])).
fof(f39085,plain,(
  ( ! [X1] : (sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X1),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X1),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X1),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))))) ) | $spl7),
  inference(subsumption_resolution,[],[f39084,f1079])).
fof(f39084,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,'F_0') | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X1),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X1),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X1),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))))) )),
  inference(forward_demodulation,[],[f39083,f116])).
fof(f39083,plain,(
  ( ! [X1] : (sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2(X1)))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X1),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X1),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X1),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))))) )),
  inference(subsumption_resolution,[],[f39082,f93])).
fof(f39082,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2(X1)))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X1),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X1),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X1),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))))) )),
  inference(forward_demodulation,[],[f32699,f173])).
fof(f32699,plain,(
  ( ! [X1] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',sK4)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2(X1)))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X1),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X1),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X1),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))))) )),
  inference(duplicate_literal_removal,[],[f32663])).
fof(f32663,plain,(
  ( ! [X1] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',sK4)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2(X1)))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X1),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X1),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X1),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))) )),
  inference(resolution,[],[f3417,f3207])).
fof(f39052,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(X0,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',X1))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',sK10(X0,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(X0,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(X0,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(X1))))) ) | $spl38),
  inference(cnf_transformation,[],[f39052_D])).
fof(f39052_D,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(X0,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',X1))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',sK10(X0,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(X0,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(X0,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(X1))))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f39081,plain,(
  ( ! [X2,X3] : (sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(X2,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(X2,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(X3)))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(X2,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',X3))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',sK10(X2,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) )),
  inference(subsumption_resolution,[],[f39080,f93])).
fof(f39080,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(X2,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(X2,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(X3)))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(X2,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',X3))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',sK10(X2,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) )),
  inference(forward_demodulation,[],[f32664,f173])).
fof(f32664,plain,(
  ( ! [X2,X3] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',sK4)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(X2,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(X2,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(X3)))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(X2,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',X3))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',sK10(X2,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) )),
  inference(resolution,[],[f3417,f3246])).
fof(f39078,plain,(
  ( ! [X28] : (sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X28),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X28),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))))) ) | $spl7),
  inference(subsumption_resolution,[],[f39077,f93])).
fof(f39077,plain,(
  ( ! [X28] : (p(s(bool,'F_0')) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X28),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X28),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))))) ) | $spl7),
  inference(forward_demodulation,[],[f39076,f173])).
fof(f39076,plain,(
  ( ! [X28] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',sK4)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X28),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X28),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))))) ) | $spl7),
  inference(forward_demodulation,[],[f39075,f75])).
fof(f39075,plain,(
  ( ! [X28] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X28),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X28),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))))) ) | $spl7),
  inference(forward_demodulation,[],[f39074,f76])).
fof(f39074,plain,(
  ( ! [X28] : (sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X28),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X28),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))))) ) | $spl7),
  inference(subsumption_resolution,[],[f39073,f1079])).
fof(f39073,plain,(
  ( ! [X28] : (s(bool,'T_0') = s(bool,'F_0') | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X28),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X28),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))))) )),
  inference(forward_demodulation,[],[f39072,f116])).
fof(f39072,plain,(
  ( ! [X28] : (sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2(X28)))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X28),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X28),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))))) )),
  inference(subsumption_resolution,[],[f39071,f93])).
fof(f39071,plain,(
  ( ! [X28] : (p(s(bool,'F_0')) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2(X28)))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X28),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X28),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))))) )),
  inference(forward_demodulation,[],[f32702,f173])).
fof(f32702,plain,(
  ( ! [X28] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',sK4)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2(X28)))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X28),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X28),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))))) )),
  inference(subsumption_resolution,[],[f32690,f394])).
fof(f394,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4)))),s('type.realax.real',sK10(sK2(X5),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4))))))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),X4),s('type.realax.real',sK3))),s('type.realax.real',sK4)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4)))) )),
  inference(forward_demodulation,[],[f388,f75])).
fof(f388,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4)))),s('type.realax.real',sK10(sK2(X5),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4))))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4)))) )),
  inference(superposition,[],[f365,f76])).
fof(f32690,plain,(
  ( ! [X28] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',sK4)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2(X28)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(X28),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X28),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X28),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))))) )),
  inference(duplicate_literal_removal,[],[f32673])).
fof(f32673,plain,(
  ( ! [X28] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',sK4)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2(X28)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(X28),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(sK2(X28),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(sK2(X28),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))) )),
  inference(resolution,[],[f3417,f864])).
fof(f39068,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK6('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK7('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | $spl7),
  inference(subsumption_resolution,[],[f39067,f1079])).
fof(f39067,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK6('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK7('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))),
  inference(forward_demodulation,[],[f39066,f116])).
fof(f39066,plain,(
  s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK6('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK7('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))),
  inference(subsumption_resolution,[],[f39065,f93])).
fof(f39065,plain,(
  p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK6('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK7('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))),
  inference(forward_demodulation,[],[f39064,f173])).
fof(f39064,plain,(
  p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',sK4)))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK6('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK7('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))),
  inference(subsumption_resolution,[],[f39063,f94])).
fof(f39063,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',sK4)))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK6('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK7('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))),
  inference(subsumption_resolution,[],[f32685,f1199])).
fof(f32685,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',sK4)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK6('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK7('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))),
  inference(superposition,[],[f3417,f30508])).
fof(f39058,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK5('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK6('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | $spl7),
  inference(subsumption_resolution,[],[f39057,f1079])).
fof(f39057,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK5('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK6('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))),
  inference(forward_demodulation,[],[f39056,f116])).
fof(f39056,plain,(
  s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK5('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK6('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))),
  inference(subsumption_resolution,[],[f39055,f93])).
fof(f39055,plain,(
  p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK5('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK6('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))),
  inference(forward_demodulation,[],[f39054,f173])).
fof(f39054,plain,(
  p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',sK4)))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK5('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK6('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))),
  inference(subsumption_resolution,[],[f39053,f94])).
fof(f39053,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',sK4)))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK5('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK6('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))),
  inference(subsumption_resolution,[],[f32686,f1196])).
fof(f32686,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',sK4)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2(sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK5('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK6('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))),
  inference(superposition,[],[f3417,f30507])).
fof(f39051,plain,(
  ( ! [X0,X1] : (sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(X0,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(X0,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(X1)))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',sK10(X0,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(X0,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',X1)))) )),
  inference(subsumption_resolution,[],[f39050,f93])).
fof(f39050,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(X0,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(X0,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(X1)))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',sK10(X0,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(X0,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',X1)))) )),
  inference(forward_demodulation,[],[f39049,f173])).
fof(f39049,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',sK4)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(X0,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(X0,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(X1)))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',sK10(X0,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(X0,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',X1)))) )),
  inference(subsumption_resolution,[],[f32687,f94])).
fof(f32687,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',sK4)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10(X0,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK9(X0,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK2(X1)))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',sK10(X0,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(X0,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',X1)))) )),
  inference(superposition,[],[f3417,f7269])).
fof(f39048,plain,(
  ( ! [X8,X9] : (s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK9(sK2(X8),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK10(sK2(X8),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',X9)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK4),s('type.realax.real',X9)))) | s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK4))),s('type.realax.real',X9)))) )),
  inference(subsumption_resolution,[],[f39047,f82])).
fof(f39047,plain,(
  ( ! [X8,X9] : (~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK9(sK2(X8),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK10(sK2(X8),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',X9)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK4),s('type.realax.real',X9)))) | s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK4))),s('type.realax.real',X9)))) )),
  inference(forward_demodulation,[],[f33705,f106])).
fof(f33705,plain,(
  ( ! [X8,X9] : (s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK9(sK2(X8),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK10(sK2(X8),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',X9)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK4),s('type.realax.real',X9)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))) | s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK4))),s('type.realax.real',X9)))) )),
  inference(resolution,[],[f1020,f336])).
fof(f1020,plain,(
  ( ! [X4,X5] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),X4),s('type.realax.real',sK3))),s('type.realax.real',sK4)))) | s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),X4),s('type.realax.real',sK9(sK2(X5),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4)))))),s('type.realax.real',sK10(sK2(X5),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4)))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4)))) )),
  inference(forward_demodulation,[],[f1019,f75])).
fof(f1019,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') = s(bool,'const.pair.UNCURRY_2'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK9(sK2(X5),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4)))),s('type.realax.real',sK10(sK2(X5),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4)))))))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),X4),s('type.realax.real',sK3))),s('type.realax.real',sK4)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4)))) )),
  inference(forward_demodulation,[],[f1018,f76])).
fof(f1018,plain,(
  ( ! [X4,X5] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),X4),s('type.realax.real',sK3))),s('type.realax.real',sK4)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK9(sK2(X5),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4)))),s('type.realax.real',sK10(sK2(X5),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4))))))))) )),
  inference(forward_demodulation,[],[f1015,f75])).
fof(f1015,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4))) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK9(sK2(X5),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4)))),s('type.realax.real',sK10(sK2(X5),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4))))))))) )),
  inference(superposition,[],[f1000,f76])).
fof(f1000,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0(X0) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK9(sK2(X1),X0)),s('type.realax.real',sK10(sK2(X1),X0))))))) )),
  inference(subsumption_resolution,[],[f988,f94])).
fof(f988,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0(X0) | s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK9(sK2(X1),X0)),s('type.realax.real',sK10(sK2(X1),X0))))))) )),
  inference(superposition,[],[f969,f96])).
fof(f39046,plain,(
  ( ! [X11] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X11))),s('type.realax.real',sK4)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real',sK3))))) ) | $spl36),
  inference(cnf_transformation,[],[f39046_D])).
fof(f39046_D,plain,(
  ( ! [X11] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X11))),s('type.realax.real',sK4)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real',sK3))))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f39045,plain,(
  ( ! [X10,X11] : (s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK9(sK2(X10),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK10(sK2(X10),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X11))),s('type.realax.real',sK4)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real',sK3))))) )),
  inference(subsumption_resolution,[],[f39044,f82])).
fof(f39044,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK9(sK2(X10),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK10(sK2(X10),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X11))),s('type.realax.real',sK4)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real',sK3))))) )),
  inference(forward_demodulation,[],[f39043,f106])).
fof(f39043,plain,(
  ( ! [X10,X11] : (s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK9(sK2(X10),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK10(sK2(X10),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X11))),s('type.realax.real',sK4)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real',sK3))))) )),
  inference(subsumption_resolution,[],[f33706,f3104])).
fof(f33706,plain,(
  ( ! [X10,X11] : (s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK9(sK2(X10),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK10(sK2(X10),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X11))),s('type.realax.real',sK4)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X11))),s('type.realax.real',sK3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real',sK3))))) )),
  inference(resolution,[],[f1020,f61])).
fof(f39041,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK9(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK10(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))) )),
  inference(subsumption_resolution,[],[f33718,f93])).
fof(f33718,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK9(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK10(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))) )),
  inference(superposition,[],[f1020,f173])).
fof(f39040,plain,(
  ( ! [X8,X9] : (sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X8))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(X8,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',X9)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK4),s('type.realax.real',X9)))) | s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK4))),s('type.realax.real',X9)))) )),
  inference(subsumption_resolution,[],[f39039,f82])).
fof(f39039,plain,(
  ( ! [X8,X9] : (~p(s(bool,'T_0')) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X8))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(X8,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',X9)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK4),s('type.realax.real',X9)))) | s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK4))),s('type.realax.real',X9)))) )),
  inference(forward_demodulation,[],[f37934,f106])).
fof(f37934,plain,(
  ( ! [X8,X9] : (sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X8))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(X8,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',X9)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK4),s('type.realax.real',X9)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))) | s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK4))),s('type.realax.real',X9)))) )),
  inference(resolution,[],[f1037,f336])).
fof(f1037,plain,(
  ( ! [X4,X5] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),X4),s('type.realax.real',sK3))),s('type.realax.real',sK4)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X5))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(X5,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4)))),s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4))))))) )),
  inference(forward_demodulation,[],[f1035,f75])).
fof(f1035,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X5))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(X5,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4)))),s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4))))))) )),
  inference(superposition,[],[f576,f76])).
fof(f39038,plain,(
  ( ! [X10,X11] : (sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X10))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(X10,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X11))),s('type.realax.real',sK4)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X11))),s('type.realax.real',sK3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real',sK3))))) )),
  inference(subsumption_resolution,[],[f39037,f82])).
fof(f39037,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X10))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(X10,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X11))),s('type.realax.real',sK4)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X11))),s('type.realax.real',sK3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real',sK3))))) )),
  inference(forward_demodulation,[],[f37935,f106])).
fof(f37935,plain,(
  ( ! [X10,X11] : (sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X10))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(X10,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X11))),s('type.realax.real',sK4)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X11))),s('type.realax.real',sK3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real',sK3))))) )),
  inference(resolution,[],[f1037,f61])).
fof(f39035,plain,(
  ( ! [X0] : (sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(X0,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) )),
  inference(subsumption_resolution,[],[f37947,f93])).
fof(f37947,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(X0,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) )),
  inference(superposition,[],[f1037,f173])).
fof(f39034,plain,(
  ( ! [X8,X9] : (sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X8))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(X8,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',X9)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK4),s('type.realax.real',X9)))) | s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK4))),s('type.realax.real',X9)))) )),
  inference(subsumption_resolution,[],[f39033,f82])).
fof(f39033,plain,(
  ( ! [X8,X9] : (~p(s(bool,'T_0')) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X8))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(X8,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',X9)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK4),s('type.realax.real',X9)))) | s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK4))),s('type.realax.real',X9)))) )),
  inference(forward_demodulation,[],[f37960,f106])).
fof(f37960,plain,(
  ( ! [X8,X9] : (sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X8))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(X8,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',X9)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK4),s('type.realax.real',X9)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))) | s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK4))),s('type.realax.real',X9)))) )),
  inference(resolution,[],[f1105,f336])).
fof(f1105,plain,(
  ( ! [X4,X5] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),X4),s('type.realax.real',sK3))),s('type.realax.real',sK4)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X5))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4)))),s('type.realax.real',sK10(X5,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4))))))) )),
  inference(forward_demodulation,[],[f1103,f75])).
fof(f1103,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X5))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4)))),s('type.realax.real',sK10(X5,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4))))))) )),
  inference(superposition,[],[f639,f76])).
fof(f39032,plain,(
  ( ! [X10,X11] : (sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X10))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(X10,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X11))),s('type.realax.real',sK4)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X11))),s('type.realax.real',sK3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real',sK3))))) )),
  inference(subsumption_resolution,[],[f39031,f82])).
fof(f39031,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X10))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(X10,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X11))),s('type.realax.real',sK4)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X11))),s('type.realax.real',sK3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real',sK3))))) )),
  inference(forward_demodulation,[],[f37961,f106])).
fof(f37961,plain,(
  ( ! [X10,X11] : (sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X10))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(X10,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X11))),s('type.realax.real',sK4)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X11))),s('type.realax.real',sK3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real',sK3))))) )),
  inference(resolution,[],[f1105,f61])).
fof(f39029,plain,(
  ( ! [X0] : (sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(X0,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) )),
  inference(subsumption_resolution,[],[f37973,f93])).
fof(f37973,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(X0,'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) )),
  inference(superposition,[],[f1105,f173])).
fof(f39024,plain,(
  $false | ($spl7 | $spl20)),
  inference(subsumption_resolution,[],[f39023,f1079])).
fof(f39023,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl7 | $spl20)),
  inference(forward_demodulation,[],[f39022,f3364])).
fof(f3364,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK5('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK6('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | $spl20),
  inference(resolution,[],[f3167,f1181])).
fof(f3167,plain,(
  p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK5('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK6('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | $spl20),
  inference(forward_demodulation,[],[f3166,f75])).
fof(f3166,plain,(
  p(s(bool,'const.pair.UNCURRY_2'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK5('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK6('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))))) | $spl20),
  inference(forward_demodulation,[],[f3141,f76])).
fof(f3141,plain,(
  p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK5('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK6('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))))) | $spl20),
  inference(resolution,[],[f2458,f68])).
fof(f68,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK5(X0)),s('type.realax.real',sK6(X0))))))) | ~sP0(X0)) )),
  inference(cnf_transformation,[],[f47])).
fof(f39022,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK5('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK6('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | ($spl7 | $spl20)),
  inference(subsumption_resolution,[],[f39021,f1079])).
fof(f39021,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK5('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK6('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | $spl20),
  inference(forward_demodulation,[],[f39020,f3420])).
fof(f3420,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK6('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK7('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | $spl20),
  inference(resolution,[],[f3169,f1181])).
fof(f3169,plain,(
  p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK6('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK7('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | $spl20),
  inference(forward_demodulation,[],[f3168,f75])).
fof(f3168,plain,(
  p(s(bool,'const.pair.UNCURRY_2'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK6('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK7('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))))) | $spl20),
  inference(forward_demodulation,[],[f3142,f76])).
fof(f3142,plain,(
  p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK6('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK7('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))))) | $spl20),
  inference(resolution,[],[f2458,f69])).
fof(f69,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK6(X0)),s('type.realax.real',sK7(X0))))))) | ~sP0(X0)) )),
  inference(cnf_transformation,[],[f47])).
fof(f39020,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK6('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK7('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK5('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK6('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | $spl20),
  inference(subsumption_resolution,[],[f39019,f2458])).
fof(f39019,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK6('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK7('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK5('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK6('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | ~sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | $spl20),
  inference(subsumption_resolution,[],[f39018,f67])).
fof(f39018,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK6('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK7('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK6('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK7('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK5('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK6('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | ~sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | $spl20),
  inference(subsumption_resolution,[],[f39006,f93])).
fof(f39006,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK6('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK7('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK6('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK7('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK5('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK6('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | ~sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | $spl20),
  inference(superposition,[],[f4906,f3443])).
fof(f3443,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK5('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK7('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | $spl20),
  inference(subsumption_resolution,[],[f3441,f94])).
fof(f3441,plain,(
  ~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK5('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK7('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | $spl20),
  inference(superposition,[],[f3171,f96])).
fof(f3171,plain,(
  ~p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK5('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))),s('type.realax.real',sK7('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) | $spl20),
  inference(forward_demodulation,[],[f3170,f75])).
fof(f3170,plain,(
  ~p(s(bool,'const.pair.UNCURRY_2'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK5('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK7('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))))) | $spl20),
  inference(forward_demodulation,[],[f3143,f76])).
fof(f3143,plain,(
  ~p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK5('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK7('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))))) | $spl20),
  inference(resolution,[],[f2458,f70])).
fof(f70,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK5(X0)),s('type.realax.real',sK7(X0))))))) | ~sP0(X0)) )),
  inference(cnf_transformation,[],[f47])).
fof(f4906,plain,(
  ( ! [X6,X7] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK5(X6)))),s('type.realax.real',X7)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK6(X6)),s('type.realax.real',X7)))) | s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK6(X6)))),s('type.realax.real',X7))) | s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK5(X6)))),s('type.realax.real',sK6(X6)))) | ~sP0(X6)) )),
  inference(resolution,[],[f1449,f66])).
fof(f1449,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X3)))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X1))),s('type.realax.real',X3)))) | s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X2))),s('type.realax.real',X3))) | s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X1))),s('type.realax.real',X2)))) )),
  inference(subsumption_resolution,[],[f1431,f82])).
fof(f1431,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X1))),s('type.realax.real',X3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X2))),s('type.realax.real',X3))) | s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X1))),s('type.realax.real',X2)))) )),
  inference(superposition,[],[f336,f96])).
fof(f3896,plain,(
  s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK3),s('type.realax.real',sK3))) | $spl26),
  inference(cnf_transformation,[],[f3896_D])).
fof(f3896_D,plain,(
  s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK3),s('type.realax.real',sK3))) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f3898,plain,(
  s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK4),s('type.realax.real',sK3))),s('type.realax.real',sK2(sK3)))) | $spl28),
  inference(cnf_transformation,[],[f3898_D])).
fof(f3898_D,plain,(
  s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK4),s('type.realax.real',sK3))),s('type.realax.real',sK2(sK3)))) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f3894,plain,(
  s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK4),s('type.realax.real',sK3))),s('type.realax.real',sK2(sK3)))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK3),s('type.realax.real',sK3)))),
  inference(subsumption_resolution,[],[f3886,f93])).
fof(f3886,plain,(
  p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK4),s('type.realax.real',sK3))),s('type.realax.real',sK2(sK3)))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK3),s('type.realax.real',sK3)))),
  inference(superposition,[],[f3240,f173])).
fof(f3240,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X0))),s('type.realax.real',sK4)))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK4),s('type.realax.real',X0))),s('type.realax.real',sK2(sK3)))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',sK3)))) )),
  inference(subsumption_resolution,[],[f3226,f94])).
fof(f3226,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X0))),s('type.realax.real',sK4)))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK4),s('type.realax.real',X0))),s('type.realax.real',sK2(sK3)))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',sK3)))) )),
  inference(superposition,[],[f886,f106])).
fof(f3137,plain,(
  ( ! [X0] : (sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) )),
  inference(subsumption_resolution,[],[f3126,f93])).
fof(f3126,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) )),
  inference(superposition,[],[f432,f173])).
fof(f432,plain,(
  ( ! [X4,X5] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),X4),s('type.realax.real',sK3))),s('type.realax.real',sK4)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4)))),s('type.realax.real',sK10(sK2(X5),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4))))))) )),
  inference(forward_demodulation,[],[f430,f75])).
fof(f430,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4)))),s('type.realax.real',sK10(sK2(X5),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4))))))) )),
  inference(superposition,[],[f395,f76])).
fof(f395,plain,(
  ( ! [X2,X3] : (p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X2),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0(X2) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8(X2)),s('type.realax.real',sK10(sK2(X3),X2))))) )),
  inference(subsumption_resolution,[],[f390,f93])).
fof(f390,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X2),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0(X2) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8(X2)),s('type.realax.real',sK10(sK2(X3),X2))))) )),
  inference(superposition,[],[f365,f96])).
fof(f3134,plain,(
  ( ! [X10,X11] : (sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(X10),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X11))),s('type.realax.real',sK4)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X11))),s('type.realax.real',sK3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real',sK3))))) )),
  inference(subsumption_resolution,[],[f3133,f82])).
fof(f3133,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(X10),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X11))),s('type.realax.real',sK4)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X11))),s('type.realax.real',sK3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real',sK3))))) )),
  inference(forward_demodulation,[],[f3114,f106])).
fof(f3114,plain,(
  ( ! [X10,X11] : (sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(X10),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X11))),s('type.realax.real',sK4)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X11))),s('type.realax.real',sK3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real',sK3))))) )),
  inference(resolution,[],[f432,f61])).
fof(f3131,plain,(
  ( ! [X8,X9] : (sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(X8),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',X9)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK4),s('type.realax.real',X9)))) | s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK4))),s('type.realax.real',X9)))) )),
  inference(subsumption_resolution,[],[f3130,f82])).
fof(f3130,plain,(
  ( ! [X8,X9] : (~p(s(bool,'T_0')) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(X8),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',X9)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK4),s('type.realax.real',X9)))) | s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK4))),s('type.realax.real',X9)))) )),
  inference(forward_demodulation,[],[f3113,f106])).
fof(f3113,plain,(
  ( ! [X8,X9] : (sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK10(sK2(X8),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',X9)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK4),s('type.realax.real',X9)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))) | s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK4))),s('type.realax.real',X9)))) )),
  inference(resolution,[],[f432,f336])).
fof(f2456,plain,(
  ( ! [X8] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X8),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) ) | $spl18),
  inference(cnf_transformation,[],[f2456_D])).
fof(f2456_D,plain,(
  ( ! [X8] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X8),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f2464,plain,(
  ( ! [X0] : (sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) )),
  inference(subsumption_resolution,[],[f2449,f93])).
fof(f2449,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X0),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))))))) )),
  inference(superposition,[],[f293,f173])).
fof(f293,plain,(
  ( ! [X4,X5] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),X4),s('type.realax.real',sK3))),s('type.realax.real',sK4)))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X5),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4)))),s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4))))))) )),
  inference(forward_demodulation,[],[f291,f75])).
fof(f291,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.pair.UNCURRY_2'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X5),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4)))),s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),X4))))))) )),
  inference(superposition,[],[f269,f76])).
fof(f269,plain,(
  ( ! [X2,X3] : (p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X3),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0(X3) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X2),X3)),s('type.realax.real',sK8(X3))))) )),
  inference(subsumption_resolution,[],[f264,f93])).
fof(f264,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X3),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0(X3) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X2),X3)),s('type.realax.real',sK8(X3))))) )),
  inference(superposition,[],[f212,f96])).
fof(f2460,plain,(
  ( ! [X10,X11] : (sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X10),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X11))),s('type.realax.real',sK4)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X11))),s('type.realax.real',sK3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real',sK3))))) )),
  inference(subsumption_resolution,[],[f2459,f82])).
fof(f2459,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X10),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X11))),s('type.realax.real',sK4)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X11))),s('type.realax.real',sK3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real',sK3))))) )),
  inference(forward_demodulation,[],[f2437,f106])).
fof(f2437,plain,(
  ( ! [X10,X11] : (sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X10),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X11))),s('type.realax.real',sK4)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',X11))),s('type.realax.real',sK3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real',sK3))))) )),
  inference(resolution,[],[f293,f61])).
fof(f2454,plain,(
  ( ! [X8,X9] : (sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X8),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',X9)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK4),s('type.realax.real',X9)))) | s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK4))),s('type.realax.real',X9)))) )),
  inference(subsumption_resolution,[],[f2453,f82])).
fof(f2453,plain,(
  ( ! [X8,X9] : (~p(s(bool,'T_0')) | sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X8),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',X9)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK4),s('type.realax.real',X9)))) | s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK4))),s('type.realax.real',X9)))) )),
  inference(forward_demodulation,[],[f2436,f106])).
fof(f2436,plain,(
  ( ! [X8,X9] : (sP0('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(X8),'const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))),s('type.realax.real',sK8('const.pair.UNCURRY_1'(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',X9)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK4),s('type.realax.real',X9)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))) | s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK4))),s('type.realax.real',X9)))) )),
  inference(resolution,[],[f293,f336])).
fof(f1175,plain,(
  ~p(s(bool,'F_0')) | $spl11),
  inference(cnf_transformation,[],[f1175_D])).
fof(f1175_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f1177,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl14),
  inference(cnf_transformation,[],[f1177_D])).
fof(f1177_D,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f79,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f51])).
fof(f51,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f36])).
fof(f36,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f25])).
fof(f25,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f24])).
fof(f24,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X12] : ((p(s(bool,X12)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X12)))),
  file('thm.integral.BOLZANO_LEMMA_ALT','thm.bool.IMP_CLAUSES|split|4')).
fof(f1176,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl12),
  inference(cnf_transformation,[],[f1176_D])).
fof(f1176_D,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f1173,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(subsumption_resolution,[],[f86,f82])).
fof(f86,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f56])).
fof(f56,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f55])).
fof(f55,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f27])).
fof(f27,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X12] : ((p(s(bool,'F_0')) => p(s(bool,X12))) <=> p(s(bool,'T_0')))),
  file('thm.integral.BOLZANO_LEMMA_ALT','thm.bool.IMP_CLAUSES|split|2')).
fof(f1171,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f1128,f95])).
fof(f95,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  p(s(bool,'T_0'))),
  file('thm.integral.BOLZANO_LEMMA_ALT','thm.bool.TRUTH')).
fof(f1128,plain,(
  ~p(s(bool,'T_0')) | $spl6),
  inference(backward_demodulation,[],[f1080,f93])).
fof(f1123,plain,(
  ( ! [X2,X0,X1] : (sP0(X0) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK9(sK2(sK8(X0)),X0)))),s('type.realax.real',sK10(sK2(sK8(X0)),X0))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) ) | $spl8),
  inference(cnf_transformation,[],[f1123_D])).
fof(f1123_D,plain,(
  ( ! [X2,X0,X1] : (sP0(X0) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK9(sK2(sK8(X0)),X0)))),s('type.realax.real',sK10(sK2(sK8(X0)),X0))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f1080,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl6),
  inference(cnf_transformation,[],[f1080_D])).
fof(f1080_D,plain,(
  s(bool,'T_0') = s(bool,'F_0') <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f1122,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'F_0') | sP0(X0) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK9(sK2(sK8(X0)),X0)))),s('type.realax.real',sK10(sK2(sK8(X0)),X0))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2))))) )),
  inference(forward_demodulation,[],[f1121,f116])).
fof(f1121,plain,(
  ( ! [X2,X0,X1] : (sP0(X0) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2(sK8(X0))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK9(sK2(sK8(X0)),X0)))),s('type.realax.real',sK10(sK2(sK8(X0)),X0))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2))))) )),
  inference(subsumption_resolution,[],[f1120,f267])).
fof(f1120,plain,(
  ( ! [X2,X0,X1] : (sP0(X0) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2(sK8(X0))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK9(sK2(sK8(X0)),X0)))),s('type.realax.real',sK10(sK2(sK8(X0)),X0))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(sK8(X0)),X0)),s('type.realax.real',sK8(X0))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2))))) )),
  inference(duplicate_literal_removal,[],[f1106])).
fof(f1106,plain,(
  ( ! [X2,X0,X1] : (sP0(X0) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2(sK8(X0))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK9(sK2(sK8(X0)),X0)))),s('type.realax.real',sK10(sK2(sK8(X0)),X0))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(sK8(X0)),X0)),s('type.realax.real',sK8(X0))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | sP0(X0)) )),
  inference(resolution,[],[f640,f228])).
fof(f640,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8(X2)),s('type.realax.real',sK10(X3,X2))))) | sP0(X2) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X3))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X2),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) )),
  inference(subsumption_resolution,[],[f635,f93])).
fof(f635,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK8(X2)),s('type.realax.real',sK10(X3,X2))))) | sP0(X2) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X3))) | s(bool,'T_0') = s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X2),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) )),
  inference(superposition,[],[f461,f96])).
fof(f1078,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK9(sK2(sK8(X16)),X16)))),s('type.realax.real',sK10(sK2(sK8(X16)),X16))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X16),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0(X16) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X17),s('type.realax.real',X18)))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X16),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X17),s('type.realax.real',X18))))))) ) | $spl4),
  inference(cnf_transformation,[],[f1078_D])).
fof(f1078_D,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK9(sK2(sK8(X16)),X16)))),s('type.realax.real',sK10(sK2(sK8(X16)),X16))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X16),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0(X16) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X17),s('type.realax.real',X18)))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X16),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X17),s('type.realax.real',X18))))))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f1077,plain,(
  ( ! [X17,X18,X16] : (s(bool,'T_0') = s(bool,'F_0') | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK9(sK2(sK8(X16)),X16)))),s('type.realax.real',sK10(sK2(sK8(X16)),X16))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X16),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X17),s('type.realax.real',X18)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X17),s('type.realax.real',X18)))) | sP0(X16) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X16),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4))))))) )),
  inference(forward_demodulation,[],[f1076,f116])).
fof(f1076,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK9(sK2(sK8(X16)),X16)))),s('type.realax.real',sK10(sK2(sK8(X16)),X16))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X16),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X17),s('type.realax.real',X18)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X17),s('type.realax.real',X18)))) | sP0(X16) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X16),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2(sK8(X16)))))) )),
  inference(subsumption_resolution,[],[f1067,f107])).
fof(f1067,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK9(sK2(sK8(X16)),X16)))),s('type.realax.real',sK10(sK2(sK8(X16)),X16))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(sK8(X16)),X16)),s('type.realax.real',sK8(X16))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X16),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X17),s('type.realax.real',X18)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X17),s('type.realax.real',X18)))) | sP0(X16) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X16),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2(sK8(X16)))))) )),
  inference(duplicate_literal_removal,[],[f1056])).
fof(f1056,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK9(sK2(sK8(X16)),X16)))),s('type.realax.real',sK10(sK2(sK8(X16)),X16))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK9(sK2(sK8(X16)),X16)),s('type.realax.real',sK8(X16))))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X16),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X17),s('type.realax.real',X18)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X17),s('type.realax.real',X18)))) | sP0(X16) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X16),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK3),s('type.realax.real',sK4)))))) | sP0(X16) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2(sK8(X16)))))) )),
  inference(resolution,[],[f228,f461])).
fof(f934,plain,(
  s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK4),s('type.realax.real',sK4))) | $spl0),
  inference(cnf_transformation,[],[f934_D])).
fof(f934_D,plain,(
  s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK4),s('type.realax.real',sK4))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f936,plain,(
  s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK4),s('type.realax.real',sK3))),s('type.realax.real',sK2(sK4)))) | $spl2),
  inference(cnf_transformation,[],[f936_D])).
fof(f936_D,plain,(
  s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK4),s('type.realax.real',sK3))),s('type.realax.real',sK2(sK4)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f932,plain,(
  s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK4),s('type.realax.real',sK3))),s('type.realax.real',sK2(sK4)))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK4),s('type.realax.real',sK4)))),
  inference(subsumption_resolution,[],[f925,f93])).
fof(f925,plain,(
  p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK4),s('type.realax.real',sK3))),s('type.realax.real',sK2(sK4)))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK4),s('type.realax.real',sK4)))),
  inference(superposition,[],[f899,f173])).
fof(f899,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',X0)))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X0),s('type.realax.real',sK3))),s('type.realax.real',sK2(sK4)))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK4),s('type.realax.real',X0)))) )),
  inference(subsumption_resolution,[],[f893,f94])).
fof(f893,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',X0)))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X0),s('type.realax.real',sK3))),s('type.realax.real',sK2(sK4)))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK4),s('type.realax.real',X0)))) )),
  inference(superposition,[],[f885,f96])).
fof(f885,plain,(
  ( ! [X0] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK4),s('type.realax.real',X0)))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',X0)))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X0),s('type.realax.real',sK3))),s('type.realax.real',sK2(sK4))))) )),
  inference(subsumption_resolution,[],[f873,f82])).
fof(f873,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK4),s('type.realax.real',X0)))) | p(s(bool,happ(s(fun('type.realax.real',bool),happ(s(fun('type.realax.real',fun('type.realax.real',bool)),sK1),s('type.realax.real',sK3))),s('type.realax.real',X0)))) | s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X0),s('type.realax.real',sK3))),s('type.realax.real',sK2(sK4))))) )),
  inference(superposition,[],[f243,f106])).
% SZS output end Proof for thm.integral
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 850
% Passive clauses: 21957
% Generated clauses: 39296
% Final active clauses: 653
% Final passive clauses: 15192
% Input formulas: 14
% Initial clauses: 36
% 
% Selected by SInE selection: 14
% SInE iterations: 1
% 
% Duplicate literals: 1599
% Trivial inequalities: 44
% Fw subsumption resolutions: 3661
% Fw demodulations: 1723
% Bw demodulations: 34
% 
% Simple tautologies: 262
% Equational tautologies: 7
% Forward subsumptions: 10591
% 
% Binary resolution: 17653
% Factoring: 18
% Forward superposition: 8972
% Backward superposition: 5536
% Equality resolution: 8
% 
% Split clauses: 83
% Split components: 194
% SAT solver clauses: 87
% SAT solver unit clauses: 1
% SAT solver binary clauses: 25
% 
% Sat splits: 83
% Sat splitting refutations: 3
% 
% Memory used [KB]: 28272
% Time elapsed: 2.506 s
% ------------------------------
% ------------------------------
% Success in time 14.224 s
