Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.list
% (103761)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 142
% Passive clauses: 1640
% Generated clauses: 20405
% Final active clauses: 116
% Final passive clauses: 1317
% Input formulas: 26
% Initial clauses: 70
% 
% Duplicate literals: 1090
% Trivial inequalities: 312
% Fw subsumption resolutions: 3447
% Fw demodulations: 2257
% Bw demodulations: 25
% 
% Simple tautologies: 526
% Equational tautologies: 66
% Forward subsumptions: 9168
% 
% Binary resolution: 353
% Forward superposition: 4927
% Backward superposition: 7283
% Self superposition: 222
% Equality factoring: 181
% Equality resolution: 30
% 
% Split clauses: 1225
% Split components: 2537
% SAT solver clauses: 2041
% SAT solver unit clauses: 41
% SAT solver binary clauses: 230
% 
% Sat splits: 1225
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 2040
% TWLsolver calls for satisfiability: 202
% 
% Memory used [KB]: 9210
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.list
% (103774)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 118
% Passive clauses: 325
% Generated clauses: 9954
% Final active clauses: 81
% Final passive clauses: 148
% Input formulas: 26
% Initial clauses: 70
% 
% Duplicate literals: 975
% Trivial inequalities: 134
% Fw subsumption resolutions: 1255
% Fw demodulations: 575
% Bw demodulations: 77
% Condensations: 1783
% Global subsumptions: 8
% 
% Simple tautologies: 216
% Equational tautologies: 50
% Forward subsumptions: 3403
% Bw demodulations to eq. taut.: 6
% 
% Binary resolution: 160
% Forward superposition: 1935
% Backward superposition: 2219
% Self superposition: 142
% Equality factoring: 134
% Equality resolution: 39
% 
% Split clauses: 245
% Split components: 504
% SAT solver clauses: 1175
% SAT solver unit clauses: 75
% SAT solver binary clauses: 77
% SAT solver learnt clauses: 16
% 
% Sat splits: 245
% Sat splitting refutations: 20
% 
% TWLsolver clauses: 1143
% TWLsolver calls for satisfiability: 301
% 
% Memory used [KB]: 6524
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.list
% (103788)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 332
% Passive clauses: 3327
% Generated clauses: 168847
% Final active clauses: 51
% Final passive clauses: 369
% Input formulas: 26
% Initial clauses: 70
% Discarded non-redundant clauses: 254
% 
% Duplicate literals: 7543
% Trivial inequalities: 1263
% Fw subsumption resolutions: 14077
% Fw demodulations: 14354
% Bw demodulations: 1772
% Condensations: 17863
% Global subsumptions: 832
% 
% Simple tautologies: 1437
% Equational tautologies: 565
% Forward subsumptions: 105942
% Fw demodulations to eq. taut.: 9
% Bw demodulations to eq. taut.: 49
% 
% Unit resulting resolution: 1800
% Forward superposition: 33125
% Backward superposition: 22761
% Self superposition: 648
% Equality factoring: 379
% Equality resolution: 84
% 
% Split clauses: 401
% Split components: 809
% SAT solver clauses: 5744
% SAT solver unit clauses: 39
% SAT solver binary clauses: 218
% 
% Sat splits: 401
% Sat splitting refutations: 780
% 
% Memory used [KB]: 61534
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.list
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.list
% SZS output start Proof for thm.list
fof(f2998,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f185,f186,f187,f188,f189,f191,f213,f192,f202,f211,f221,f254,f222,f256,f252,f223,f224,f225,f226,f227,f214,f258,f228,f230,f241,f257,f290,f282,f307,f336,f338,f334,f333,f371,f372,f339,f368,f367,f337,f409,f404,f403,f402,f425,f426,f427,f428,f450,f452,f453,f455,f457,f468,f480,f530,f566,f568,f569,f570,f573,f583,f592,f838,f839,f1156,f1158,f1159,f1161,f1163,f1174,f1189,f1195,f1197,f1198,f1200,f1202,f1213,f1233,f1235,f1236,f1238,f1240,f1251,f1366,f1370,f1373,f1377,f1379,f1390,f1227,f1399,f1401,f1402,f1403,f1406,f1416,f1427,f1544,f1548,f1551,f1555,f1557,f1568,f1264,f1605,f1607,f1608,f1610,f1612,f1623,f1634,f1645,f1647,f1648,f1650,f1652,f1663,f1676,f1719,f1721,f1722,f1724,f1726,f1737,f1754,f1756,f1757,f1759,f1761,f1772,f1852,f1854,f1855,f1856,f1859,f1869,f1880,f1900,f1902,f1903,f1904,f1907,f1917,f1928,f2056,f2060,f2063,f2067,f2069,f2080,f1748,f2217,f2221,f2224,f2228,f2230,f2241,f1785,f2287,f2289,f2290,f2291,f2294,f2304,f2321,f2373,f2375,f2376,f212,f2378,f216,f2380,f215,f2391,f2404,f217,f2920,f2925])).
fof(f2925,plain,(
  $false | ($spl16 | $spl19)),
  inference(subsumption_resolution,[],[f2924,f175])).
fof(f175,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.list.LIST_REL_mono',aHL_TRUTH)).
fof(f2924,plain,(
  ~p(s(bool,'T_0')) | ($spl16 | $spl19)),
  inference(forward_demodulation,[],[f2923,f200])).
fof(f200,plain,(
  s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),sK2)))),
  inference(resolution,[],[f111,f170])).
fof(f170,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f106])).
fof(f106,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X14] : (s(bool,'T_0') = s(bool,X14) <=> p(s(bool,X14)))),
  file('thm.list.LIST_REL_mono','thm.bool.EQ_CLAUSES|split|0')).
fof(f111,plain,(
  p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),sK2))))),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X6,X7] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK5),s(sK0,X6))),s(sK1,X7)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,X6))),s(sK1,X7))))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),sK2)))) & ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),sK2))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5])],[f60])).
fof(f60,plain,(
  ? [X0,X1,X2,X3,X4,X5] : (! [X6,X7] : (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X5),s(X0,X6))),s(X1,X7)))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X4),s(X0,X6))),s(X1,X7))))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X5),s('type.list.list'(X0),X3),s('type.list.list'(X1),X2)))) & ~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X4),s('type.list.list'(X0),X3),s('type.list.list'(X1),X2)))))),
  inference(flattening,[],[f59])).
fof(f59,plain,(
  ? [X0,X1,X2,X3,X4,X5] : (! [X6,X7] : (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X5),s(X0,X6))),s(X1,X7)))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X4),s(X0,X6))),s(X1,X7))))) & (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X5),s('type.list.list'(X0),X3),s('type.list.list'(X1),X2)))) & ~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X4),s('type.list.list'(X0),X3),s('type.list.list'(X1),X2))))))),
  inference(ennf_transformation,[],[f29])).
fof(f29,plain,(
  ~! [X0,X1,X2,X3,X4,X5] : (! [X6,X7] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X5),s(X0,X6))),s(X1,X7)))) => p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X4),s(X0,X6))),s(X1,X7))))) => (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X5),s('type.list.list'(X0),X3),s('type.list.list'(X1),X2)))) => p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X4),s('type.list.list'(X0),X3),s('type.list.list'(X1),X2))))))),
  inference(flattening,[],[f28])).
fof(f28,plain,(
  ~! [X0,X1] : ! [X2,X3,X4,X5] : (! [X6,X7] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X5),s(X0,X6))),s(X1,X7)))) => p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X4),s(X0,X6))),s(X1,X7))))) => (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X5),s('type.list.list'(X0),X3),s('type.list.list'(X1),X2)))) => p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X4),s('type.list.list'(X0),X3),s('type.list.list'(X1),X2))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10,X11] : (! [X12,X13] : (p(s(bool,happ(s(fun(X7,bool),happ(s(fun(X6,fun(X7,bool)),X11),s(X6,X12))),s(X7,X13)))) => p(s(bool,happ(s(fun(X7,bool),happ(s(fun(X6,fun(X7,bool)),X10),s(X6,X12))),s(X7,X13))))) => (p(s(bool,'const.list.LIST_REL_3'(s(fun(X6,fun(X7,bool)),X11),s('type.list.list'(X6),X9),s('type.list.list'(X7),X8)))) => p(s(bool,'const.list.LIST_REL_3'(s(fun(X6,fun(X7,bool)),X10),s('type.list.list'(X6),X9),s('type.list.list'(X7),X8))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10,X11] : (! [X12,X13] : (p(s(bool,happ(s(fun(X7,bool),happ(s(fun(X6,fun(X7,bool)),X11),s(X6,X12))),s(X7,X13)))) => p(s(bool,happ(s(fun(X7,bool),happ(s(fun(X6,fun(X7,bool)),X10),s(X6,X12))),s(X7,X13))))) => (p(s(bool,'const.list.LIST_REL_3'(s(fun(X6,fun(X7,bool)),X11),s('type.list.list'(X6),X9),s('type.list.list'(X7),X8)))) => p(s(bool,'const.list.LIST_REL_3'(s(fun(X6,fun(X7,bool)),X10),s('type.list.list'(X6),X9),s('type.list.list'(X7),X8))))))),
  file('thm.list.LIST_REL_mono','thm.list.LIST_REL_mono')).
fof(f2923,plain,(
  ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),sK2)))) | ($spl16 | $spl19)),
  inference(subsumption_resolution,[],[f2922,f181])).
fof(f181,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f133])).
fof(f133,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f36])).
fof(f36,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f35])).
fof(f35,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X6] : ! [X12] : (s(X6,X12) = s(X6,X12) <=> p(s(bool,'T_0')))),
  file('thm.list.LIST_REL_mono','thm.bool.REFL_CLAUSE')).
fof(f2922,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),sK2)))) | ($spl16 | $spl19)),
  inference(forward_demodulation,[],[f2921,f581])).
fof(f581,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3))))) | $spl16),
  inference(resolution,[],[f254,f170])).
fof(f2921,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),sK2)))) | $spl19),
  inference(subsumption_resolution,[],[f2825,f136])).
fof(f136,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f85])).
fof(f85,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f84])).
fof(f84,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f66])).
fof(f66,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f37])).
fof(f37,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X14] : ((p(s(bool,X14)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X14)))),
  file('thm.list.LIST_REL_mono','thm.bool.IMP_CLAUSES|split|4')).
fof(f2825,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),sK2)))) | $spl19),
  inference(superposition,[],[f114,f1397])).
fof(f1397,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK5),s(sK0,'const.list.EL_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.list.list'(sK0),sK3))))),s(sK1,'const.list.EL_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.list.list'(sK1),sK2))))) | $spl19),
  inference(resolution,[],[f464,f256])).
fof(f464,plain,(
  ( ! [X56,X55] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,X55))),s(sK1,X56)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK5),s(sK0,X55))),s(sK1,X56)))) )),
  inference(resolution,[],[f110,f166])).
fof(f166,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f104])).
fof(f104,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f48])).
fof(f48,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X14] : (s(bool,'F_0') = s(bool,X14) <=> ~p(s(bool,X14)))),
  file('thm.list.LIST_REL_mono','thm.bool.EQ_CLAUSES|split|2')).
fof(f110,plain,(
  ( ! [X6,X7] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK5),s(sK0,X6))),s(sK1,X7)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,X6))),s(sK1,X7))))) )),
  inference(cnf_transformation,[],[f72])).
fof(f114,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.list.EL_2'(s('type.num.num',X6),s('type.list.list'(X0),X3))))),s(X1,'const.list.EL_2'(s('type.num.num',X6),s('type.list.list'(X1),X4)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4))))) )),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4)))) | (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4))) & ! [X6] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3)))))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.list.EL_2'(s('type.num.num',X6),s('type.list.list'(X0),X3))))),s(X1,'const.list.EL_2'(s('type.num.num',X6),s('type.list.list'(X1),X4))))))))) & (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4))) | (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(X4,X3,X2,X1,X0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3)))))) & ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.list.EL_2'(s('type.num.num',sK6(X4,X3,X2,X1,X0)),s('type.list.list'(X0),X3))))),s(X1,'const.list.EL_2'(s('type.num.num',sK6(X4,X3,X2,X1,X0)),s('type.list.list'(X1),X4))))))) | p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK6])],[f75])).
fof(f75,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4)))) | (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4))) & ! [X6] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3)))))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.list.EL_2'(s('type.num.num',X6),s('type.list.list'(X0),X3))))),s(X1,'const.list.EL_2'(s('type.num.num',X6),s('type.list.list'(X1),X4))))))))) & (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4))) | ? [X5] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3)))))) & ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X0),X3))))),s(X1,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X1),X4))))))) | p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4))))))),
  inference(rectify,[],[f74])).
fof(f74,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4)))) | (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4))) & ! [X5] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3)))))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X0),X3))))),s(X1,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X1),X4))))))))) & (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4))) | ? [X5] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3)))))) & ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X0),X3))))),s(X1,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X1),X4))))))) | p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4))))))),
  inference(flattening,[],[f73])).
fof(f73,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4)))) | (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4))) & ! [X5] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3)))))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X0),X3))))),s(X1,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X1),X4))))))))) & ((s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4))) | ? [X5] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3)))))) & ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X0),X3))))),s(X1,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X1),X4)))))))) | p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4))))))),
  inference(nnf_transformation,[],[f61])).
fof(f61,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4)))) <=> (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4))) & ! [X5] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3)))))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X0),X3))))),s(X1,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X1),X4)))))))))),
  inference(ennf_transformation,[],[f31])).
fof(f31,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4)))) <=> (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4))) & ! [X5] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3)))))) => p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X0),X3))))),s(X1,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X1),X4)))))))))),
  inference(flattening,[],[f30])).
fof(f30,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4)))) <=> (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4))) & ! [X5] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3)))))) => p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X0),X3))))),s(X1,'const.list.EL_2'(s('type.num.num',X5),s('type.list.list'(X1),X4)))))))))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X6,X7] : ! [X20,X9,X8] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X6,fun(X7,bool)),X20),s('type.list.list'(X6),X9),s('type.list.list'(X7),X8)))) <=> (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X9))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X7),X8))) & ! [X21] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X21),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X9)))))) => p(s(bool,happ(s(fun(X7,bool),happ(s(fun(X6,fun(X7,bool)),X20),s(X6,'const.list.EL_2'(s('type.num.num',X21),s('type.list.list'(X6),X9))))),s(X7,'const.list.EL_2'(s('type.num.num',X21),s('type.list.list'(X7),X8)))))))))),
  file('thm.list.LIST_REL_mono','thm.list.LIST_REL_EL_EQN')).
fof(f2920,plain,(
  $false | ($spl16 | $spl19)),
  inference(subsumption_resolution,[],[f2919,f175])).
fof(f2919,plain,(
  ~p(s(bool,'T_0')) | ($spl16 | $spl19)),
  inference(forward_demodulation,[],[f2918,f581])).
fof(f2918,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3)))))) | $spl19),
  inference(subsumption_resolution,[],[f2824,f136])).
fof(f2824,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3)))))) | $spl19),
  inference(superposition,[],[f184,f1397])).
fof(f184,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK5),s(sK0,'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(sK0),sK3))))),s(sK1,'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(sK1),sK2)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3))))))) )),
  inference(resolution,[],[f111,f114])).
fof(f217,plain,(
  ( ! [X13] : (p(s(bool,X13))) ) | $spl12),
  inference(cnf_transformation,[],[f217_D])).
fof(f217_D,plain,(
  ( ! [X13] : (p(s(bool,X13))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f2404,plain,(
  ( ! [X6,X10,X8,X7] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),X7)))))) | s(bool,X10) != s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.list.list'(sK0),X7),s('type.list.list'(sK1),X8))) | ~p(s(bool,X10)) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,'const.list.EL_2'(s('type.num.num',X6),s('type.list.list'(sK0),X7))))),s(sK1,'const.list.EL_2'(s('type.num.num',X6),s('type.list.list'(sK1),X8))))))) ) | $spl54),
  inference(cnf_transformation,[],[f2404_D])).
fof(f2404_D,plain,(
  ( ! [X6,X10,X8,X7] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),X7)))))) | s(bool,X10) != s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.list.list'(sK0),X7),s('type.list.list'(sK1),X8))) | ~p(s(bool,X10)) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,'const.list.EL_2'(s('type.num.num',X6),s('type.list.list'(sK0),X7))))),s(sK1,'const.list.EL_2'(s('type.num.num',X6),s('type.list.list'(sK1),X8))))))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f2391,plain,(
  ( ! [X88,X87,X85,X86,X84] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X84),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),X85)))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,'const.list.EL_2'(s('type.num.num',X84),s('type.list.list'(sK0),X85))))),s(sK1,'const.list.EL_2'(s('type.num.num',X84),s('type.list.list'(sK1),X86)))))) | p(s(bool,X87)) | ~p(s(bool,X88)) | s(bool,X88) != s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.list.list'(sK0),X85),s('type.list.list'(sK1),X86)))) )),
  inference(resolution,[],[f449,f180])).
fof(f180,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f121])).
fof(f121,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f77])).
fof(f77,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f63])).
fof(f63,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f62])).
fof(f62,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f32])).
fof(f32,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X12,X18,X13,X19] : ((s(bool,X12) = s(bool,X18) & (p(s(bool,X18)) => s(bool,X13) = s(bool,X19))) => ((p(s(bool,X12)) => p(s(bool,X13))) <=> (p(s(bool,X18)) => p(s(bool,X19)))))),
  file('thm.list.LIST_REL_mono','thm.bool.IMP_CONG')).
fof(f449,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.list.list'(sK0),X1),s('type.list.list'(sK1),X2)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),X1)))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(sK0),X1))))),s(sK1,'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(sK1),X2))))))) )),
  inference(resolution,[],[f110,f114])).
fof(f215,plain,(
  ( ! [X8,X9] : (s(bool,X8) != s(bool,X9) | ~p(s(bool,X8))) ) | $spl8),
  inference(cnf_transformation,[],[f215_D])).
fof(f215_D,plain,(
  ( ! [X8,X9] : (s(bool,X8) != s(bool,X9) | ~p(s(bool,X8))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f2380,plain,(
  ( ! [X52,X50,X48,X53,X51,X49] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X48),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),X49)))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,'const.list.EL_2'(s('type.num.num',X48),s('type.list.list'(sK0),X49))))),s(sK1,'const.list.EL_2'(s('type.num.num',X48),s('type.list.list'(sK1),X50)))))) | ~p(s(bool,X51)) | ~p(s(bool,X52)) | s(bool,X52) != s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.list.list'(sK0),X49),s('type.list.list'(sK1),X50))) | s(bool,X51) != s(bool,X53)) )),
  inference(resolution,[],[f449,f124])).
fof(f124,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f78])).
fof(f216,plain,(
  ( ! [X14] : (~p(s(bool,X14))) ) | $spl10),
  inference(cnf_transformation,[],[f216_D])).
fof(f216_D,plain,(
  ( ! [X14] : (~p(s(bool,X14))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f2378,plain,(
  ( ! [X39,X37,X41,X38,X36,X40] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X36),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),X37)))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,'const.list.EL_2'(s('type.num.num',X36),s('type.list.list'(sK0),X37))))),s(sK1,'const.list.EL_2'(s('type.num.num',X36),s('type.list.list'(sK1),X38)))))) | p(s(bool,X39)) | ~p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X40) != s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.list.list'(sK0),X37),s('type.list.list'(sK1),X38)))) )),
  inference(resolution,[],[f449,f123])).
fof(f123,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f78])).
fof(f212,plain,(
  ( ! [X3,X1] : (s(bool,X1) != s(bool,X3) | p(s(bool,X1))) ) | $spl2),
  inference(cnf_transformation,[],[f212_D])).
fof(f212_D,plain,(
  ( ! [X3,X1] : (s(bool,X1) != s(bool,X3) | p(s(bool,X1))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f2376,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X24),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),X25)))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,'const.list.EL_2'(s('type.num.num',X24),s('type.list.list'(sK0),X25))))),s(sK1,'const.list.EL_2'(s('type.num.num',X24),s('type.list.list'(sK1),X26)))))) | p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,X29) | s(bool,X28) != s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.list.list'(sK0),X25),s('type.list.list'(sK1),X26)))) )),
  inference(resolution,[],[f449,f122])).
fof(f122,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f78])).
fof(f2375,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X18),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),X19)))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,'const.list.EL_2'(s('type.num.num',X18),s('type.list.list'(sK0),X19))))),s(sK1,'const.list.EL_2'(s('type.num.num',X18),s('type.list.list'(sK1),X20)))))) | ~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.list.list'(sK0),X19),s('type.list.list'(sK1),X20))) | s(bool,X21) != s(bool,X23)) )),
  inference(resolution,[],[f449,f120])).
fof(f120,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f78])).
fof(f2373,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),X7)))))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,'const.list.EL_2'(s('type.num.num',X6),s('type.list.list'(sK0),X7))))),s(sK1,'const.list.EL_2'(s('type.num.num',X6),s('type.list.list'(sK1),X8)))))) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.list.list'(sK0),X7),s('type.list.list'(sK1),X8)))) )),
  inference(resolution,[],[f449,f118])).
fof(f118,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f78])).
fof(f2321,plain,(
  ( ! [X3,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3)))))) | s(bool,X3) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK5),s(sK0,'const.list.EL_2'(s('type.num.num',X1),s('type.list.list'(sK0),sK3))))),s(sK1,'const.list.EL_2'(s('type.num.num',X1),s('type.list.list'(sK1),sK2))))) | p(s(bool,X3))) ) | $spl52),
  inference(cnf_transformation,[],[f2321_D])).
fof(f2321_D,plain,(
  ( ! [X3,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3)))))) | s(bool,X3) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK5),s(sK0,'const.list.EL_2'(s('type.num.num',X1),s('type.list.list'(sK0),sK3))))),s(sK1,'const.list.EL_2'(s('type.num.num',X1),s('type.list.list'(sK1),sK2))))) | p(s(bool,X3))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f2304,plain,(
  ( ! [X43,X44,X42] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X42),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3)))))) | p(s(bool,X43)) | p(s(bool,X44)) | s(bool,X44) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK5),s(sK0,'const.list.EL_2'(s('type.num.num',X42),s('type.list.list'(sK0),sK3))))),s(sK1,'const.list.EL_2'(s('type.num.num',X42),s('type.list.list'(sK1),sK2)))))) )),
  inference(resolution,[],[f184,f180])).
fof(f2294,plain,(
  ( ! [X30,X31,X29,X32] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X29),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3)))))) | p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X30) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK5),s(sK0,'const.list.EL_2'(s('type.num.num',X29),s('type.list.list'(sK0),sK3))))),s(sK1,'const.list.EL_2'(s('type.num.num',X29),s('type.list.list'(sK1),sK2))))) | s(bool,X31) != s(bool,X32)) )),
  inference(resolution,[],[f184,f124])).
fof(f2291,plain,(
  ( ! [X19,X17,X20,X18] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X17),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3)))))) | p(s(bool,X18)) | ~p(s(bool,X19)) | p(s(bool,X20)) | s(bool,X20) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK5),s(sK0,'const.list.EL_2'(s('type.num.num',X17),s('type.list.list'(sK0),sK3))))),s(sK1,'const.list.EL_2'(s('type.num.num',X17),s('type.list.list'(sK1),sK2)))))) )),
  inference(resolution,[],[f184,f123])).
fof(f2290,plain,(
  ( ! [X14,X15,X13,X16] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X13),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3)))))) | p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK5),s(sK0,'const.list.EL_2'(s('type.num.num',X13),s('type.list.list'(sK0),sK3))))),s(sK1,'const.list.EL_2'(s('type.num.num',X13),s('type.list.list'(sK1),sK2)))))) )),
  inference(resolution,[],[f184,f122])).
fof(f2289,plain,(
  ( ! [X12,X10,X11,X9] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3)))))) | p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK5),s(sK0,'const.list.EL_2'(s('type.num.num',X9),s('type.list.list'(sK0),sK3))))),s(sK1,'const.list.EL_2'(s('type.num.num',X9),s('type.list.list'(sK1),sK2))))) | s(bool,X11) != s(bool,X12)) )),
  inference(resolution,[],[f184,f120])).
fof(f2287,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3)))))) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK5),s(sK0,'const.list.EL_2'(s('type.num.num',X1),s('type.list.list'(sK0),sK3))))),s(sK1,'const.list.EL_2'(s('type.num.num',X1),s('type.list.list'(sK1),sK2)))))) )),
  inference(resolution,[],[f184,f118])).
fof(f1785,plain,(
  ( ! [X6,X7,X11,X9] : (s(bool,X9) != s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(X6,bool)),X11),s('type.list.list'(sK0),sK3),s('type.list.list'(X6),X7))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X7))))) | ~p(s(bool,X9))) ) | $spl46),
  inference(cnf_transformation,[],[f1785_D])).
fof(f1785_D,plain,(
  ( ! [X6,X7,X11,X9] : (s(bool,X9) != s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(X6,bool)),X11),s('type.list.list'(sK0),sK3),s('type.list.list'(X6),X7))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X7))))) | ~p(s(bool,X9))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f2241,plain,(
  ( ! [X333,X335,X337,X334,X336] : (p(s(bool,X336)) | ~p(s(bool,X337)) | s(bool,X337) != s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(X333,bool)),X334),s('type.list.list'(sK0),sK3),s('type.list.list'(X333),X335))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X333),X335)))))) ) | $spl16),
  inference(subsumption_resolution,[],[f2188,f160])).
fof(f160,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0] : ((~p(s(bool,X0)) | ~p(s(bool,'F_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,X0)) & p(s(bool,'F_0')))))),
  inference(flattening,[],[f100])).
fof(f100,plain,(
  ! [X0] : (((~p(s(bool,X0)) | ~p(s(bool,'F_0'))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,X0)) & p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f46])).
fof(f46,plain,(
  ! [X0] : ((p(s(bool,X0)) & p(s(bool,'F_0'))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X14] : ((p(s(bool,X14)) & p(s(bool,'F_0'))) <=> p(s(bool,'F_0')))),
  file('thm.list.LIST_REL_mono','thm.bool.AND_CLAUSES|split|3')).
fof(f2188,plain,(
  ( ! [X333,X335,X337,X334,X336] : (p(s(bool,'F_0')) | p(s(bool,X336)) | ~p(s(bool,X337)) | s(bool,X337) != s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(X333,bool)),X334),s('type.list.list'(sK0),sK3),s('type.list.list'(X333),X335))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X333),X335)))))) ) | $spl16),
  inference(superposition,[],[f180,f1768])).
fof(f1768,plain,(
  ( ! [X66,X67,X65] : (s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(X65,bool)),X67),s('type.list.list'(sK0),sK3),s('type.list.list'(X65),X66))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X65),X66)))))) ) | $spl16),
  inference(resolution,[],[f1023,f166])).
fof(f1023,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(X3,bool)),X5),s('type.list.list'(sK0),sK3),s('type.list.list'(X3),X4)))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X3),X4)))))) ) | $spl16),
  inference(superposition,[],[f581,f113])).
fof(f113,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4))))) )),
  inference(cnf_transformation,[],[f76])).
fof(f2230,plain,(
  ( ! [X218,X222,X220,X219,X223,X221] : (~p(s(bool,X221)) | ~p(s(bool,X222)) | s(bool,X222) != s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(X218,bool)),X219),s('type.list.list'(sK0),sK3),s('type.list.list'(X218),X220))) | s(bool,X221) != s(bool,X223) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X218),X220)))))) ) | $spl16),
  inference(subsumption_resolution,[],[f2157,f136])).
fof(f2157,plain,(
  ( ! [X218,X222,X220,X219,X223,X221] : (p(s(bool,'F_0')) | ~p(s(bool,X221)) | ~p(s(bool,X222)) | s(bool,X222) != s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(X218,bool)),X219),s('type.list.list'(sK0),sK3),s('type.list.list'(X218),X220))) | s(bool,X221) != s(bool,X223) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X218),X220)))))) ) | $spl16),
  inference(superposition,[],[f124,f1768])).
fof(f2228,plain,(
  ( ! [X175,X173,X171,X174,X172,X170] : (p(s(bool,X173)) | ~p(s(bool,X174)) | ~p(s(bool,X175)) | s(bool,X174) != s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(X170,bool)),X171),s('type.list.list'(sK0),sK3),s('type.list.list'(X170),X172))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X170),X172)))))) ) | $spl16),
  inference(subsumption_resolution,[],[f2149,f160])).
fof(f2149,plain,(
  ( ! [X175,X173,X171,X174,X172,X170] : (p(s(bool,'F_0')) | p(s(bool,X173)) | ~p(s(bool,X174)) | ~p(s(bool,X175)) | s(bool,X174) != s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(X170,bool)),X171),s('type.list.list'(sK0),sK3),s('type.list.list'(X170),X172))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X170),X172)))))) ) | $spl16),
  inference(superposition,[],[f123,f1768])).
fof(f2224,plain,(
  ( ! [X144,X142,X140,X145,X143,X141] : (p(s(bool,X143)) | ~p(s(bool,X144)) | s(bool,X143) != s(bool,X145) | s(bool,X144) != s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(X140,bool)),X141),s('type.list.list'(sK0),sK3),s('type.list.list'(X140),X142))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X140),X142)))))) ) | $spl16),
  inference(subsumption_resolution,[],[f2144,f160])).
fof(f2144,plain,(
  ( ! [X144,X142,X140,X145,X143,X141] : (p(s(bool,'F_0')) | p(s(bool,X143)) | ~p(s(bool,X144)) | s(bool,X143) != s(bool,X145) | s(bool,X144) != s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(X140,bool)),X141),s('type.list.list'(sK0),sK3),s('type.list.list'(X140),X142))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X140),X142)))))) ) | $spl16),
  inference(superposition,[],[f122,f1768])).
fof(f2221,plain,(
  ( ! [X103,X101,X99,X102,X100,X98] : (~p(s(bool,X101)) | ~p(s(bool,X102)) | s(bool,X102) != s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(X98,bool)),X99),s('type.list.list'(sK0),sK3),s('type.list.list'(X98),X100))) | s(bool,X101) != s(bool,X103) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X98),X100)))))) ) | $spl16),
  inference(subsumption_resolution,[],[f2137,f136])).
fof(f2137,plain,(
  ( ! [X103,X101,X99,X102,X100,X98] : (p(s(bool,'F_0')) | ~p(s(bool,X101)) | ~p(s(bool,X102)) | s(bool,X102) != s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(X98,bool)),X99),s('type.list.list'(sK0),sK3),s('type.list.list'(X98),X100))) | s(bool,X101) != s(bool,X103) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X98),X100)))))) ) | $spl16),
  inference(superposition,[],[f120,f1768])).
fof(f2217,plain,(
  ( ! [X61,X59,X57,X60,X58,X56] : (p(s(bool,X59)) | ~p(s(bool,X60)) | s(bool,X59) != s(bool,X61) | s(bool,X60) != s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(X56,bool)),X57),s('type.list.list'(sK0),sK3),s('type.list.list'(X56),X58))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X56),X58)))))) ) | $spl16),
  inference(subsumption_resolution,[],[f2130,f160])).
fof(f2130,plain,(
  ( ! [X61,X59,X57,X60,X58,X56] : (p(s(bool,'F_0')) | p(s(bool,X59)) | ~p(s(bool,X60)) | s(bool,X59) != s(bool,X61) | s(bool,X60) != s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(X56,bool)),X57),s('type.list.list'(sK0),sK3),s('type.list.list'(X56),X58))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X56),X58)))))) ) | $spl16),
  inference(superposition,[],[f118,f1768])).
fof(f1748,plain,(
  ( ! [X6,X7,X11,X9] : (s(bool,X9) != s(bool,'const.list.LIST_REL_3'(s(fun(X6,fun(sK0,bool)),X11),s('type.list.list'(X6),X7),s('type.list.list'(sK0),sK3))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X7))))) | ~p(s(bool,X9))) ) | $spl44),
  inference(cnf_transformation,[],[f1748_D])).
fof(f1748_D,plain,(
  ( ! [X6,X7,X11,X9] : (s(bool,X9) != s(bool,'const.list.LIST_REL_3'(s(fun(X6,fun(sK0,bool)),X11),s('type.list.list'(X6),X7),s('type.list.list'(sK0),sK3))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X7))))) | ~p(s(bool,X9))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f2080,plain,(
  ( ! [X333,X335,X331,X332,X334] : (p(s(bool,X334)) | ~p(s(bool,X335)) | s(bool,X335) != s(bool,'const.list.LIST_REL_3'(s(fun(X331,fun(sK0,bool)),X332),s('type.list.list'(X331),X333),s('type.list.list'(sK0),sK3))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X331),X333)))))) ) | $spl16),
  inference(subsumption_resolution,[],[f2032,f160])).
fof(f2032,plain,(
  ( ! [X333,X335,X331,X332,X334] : (p(s(bool,'F_0')) | p(s(bool,X334)) | ~p(s(bool,X335)) | s(bool,X335) != s(bool,'const.list.LIST_REL_3'(s(fun(X331,fun(sK0,bool)),X332),s('type.list.list'(X331),X333),s('type.list.list'(sK0),sK3))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X331),X333)))))) ) | $spl16),
  inference(superposition,[],[f180,f1733])).
fof(f1733,plain,(
  ( ! [X66,X67,X65] : (s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X65,fun(sK0,bool)),X67),s('type.list.list'(X65),X66),s('type.list.list'(sK0),sK3))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X65),X66)))))) ) | $spl16),
  inference(resolution,[],[f1022,f166])).
fof(f1022,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(sK0,bool)),X2),s('type.list.list'(X0),X1),s('type.list.list'(sK0),sK3)))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1)))))) ) | $spl16),
  inference(superposition,[],[f581,f113])).
fof(f2069,plain,(
  ( ! [X218,X216,X220,X219,X217,X221] : (~p(s(bool,X219)) | ~p(s(bool,X220)) | s(bool,X220) != s(bool,'const.list.LIST_REL_3'(s(fun(X216,fun(sK0,bool)),X217),s('type.list.list'(X216),X218),s('type.list.list'(sK0),sK3))) | s(bool,X219) != s(bool,X221) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X216),X218)))))) ) | $spl16),
  inference(subsumption_resolution,[],[f2001,f136])).
fof(f2001,plain,(
  ( ! [X218,X216,X220,X219,X217,X221] : (p(s(bool,'F_0')) | ~p(s(bool,X219)) | ~p(s(bool,X220)) | s(bool,X220) != s(bool,'const.list.LIST_REL_3'(s(fun(X216,fun(sK0,bool)),X217),s('type.list.list'(X216),X218),s('type.list.list'(sK0),sK3))) | s(bool,X219) != s(bool,X221) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X216),X218)))))) ) | $spl16),
  inference(superposition,[],[f124,f1733])).
fof(f2067,plain,(
  ( ! [X173,X171,X169,X172,X170,X168] : (p(s(bool,X171)) | ~p(s(bool,X172)) | ~p(s(bool,X173)) | s(bool,X172) != s(bool,'const.list.LIST_REL_3'(s(fun(X168,fun(sK0,bool)),X169),s('type.list.list'(X168),X170),s('type.list.list'(sK0),sK3))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X168),X170)))))) ) | $spl16),
  inference(subsumption_resolution,[],[f1993,f160])).
fof(f1993,plain,(
  ( ! [X173,X171,X169,X172,X170,X168] : (p(s(bool,'F_0')) | p(s(bool,X171)) | ~p(s(bool,X172)) | ~p(s(bool,X173)) | s(bool,X172) != s(bool,'const.list.LIST_REL_3'(s(fun(X168,fun(sK0,bool)),X169),s('type.list.list'(X168),X170),s('type.list.list'(sK0),sK3))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X168),X170)))))) ) | $spl16),
  inference(superposition,[],[f123,f1733])).
fof(f2063,plain,(
  ( ! [X142,X140,X138,X143,X141,X139] : (p(s(bool,X141)) | ~p(s(bool,X142)) | s(bool,X141) != s(bool,X143) | s(bool,X142) != s(bool,'const.list.LIST_REL_3'(s(fun(X138,fun(sK0,bool)),X139),s('type.list.list'(X138),X140),s('type.list.list'(sK0),sK3))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X138),X140)))))) ) | $spl16),
  inference(subsumption_resolution,[],[f1988,f160])).
fof(f1988,plain,(
  ( ! [X142,X140,X138,X143,X141,X139] : (p(s(bool,'F_0')) | p(s(bool,X141)) | ~p(s(bool,X142)) | s(bool,X141) != s(bool,X143) | s(bool,X142) != s(bool,'const.list.LIST_REL_3'(s(fun(X138,fun(sK0,bool)),X139),s('type.list.list'(X138),X140),s('type.list.list'(sK0),sK3))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X138),X140)))))) ) | $spl16),
  inference(superposition,[],[f122,f1733])).
fof(f2060,plain,(
  ( ! [X101,X99,X97,X100,X98,X96] : (~p(s(bool,X99)) | ~p(s(bool,X100)) | s(bool,X100) != s(bool,'const.list.LIST_REL_3'(s(fun(X96,fun(sK0,bool)),X97),s('type.list.list'(X96),X98),s('type.list.list'(sK0),sK3))) | s(bool,X99) != s(bool,X101) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X96),X98)))))) ) | $spl16),
  inference(subsumption_resolution,[],[f1981,f136])).
fof(f1981,plain,(
  ( ! [X101,X99,X97,X100,X98,X96] : (p(s(bool,'F_0')) | ~p(s(bool,X99)) | ~p(s(bool,X100)) | s(bool,X100) != s(bool,'const.list.LIST_REL_3'(s(fun(X96,fun(sK0,bool)),X97),s('type.list.list'(X96),X98),s('type.list.list'(sK0),sK3))) | s(bool,X99) != s(bool,X101) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X96),X98)))))) ) | $spl16),
  inference(superposition,[],[f120,f1733])).
fof(f2056,plain,(
  ( ! [X59,X57,X54,X58,X56,X55] : (p(s(bool,X57)) | ~p(s(bool,X58)) | s(bool,X57) != s(bool,X59) | s(bool,X58) != s(bool,'const.list.LIST_REL_3'(s(fun(X54,fun(sK0,bool)),X55),s('type.list.list'(X54),X56),s('type.list.list'(sK0),sK3))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X54),X56)))))) ) | $spl16),
  inference(subsumption_resolution,[],[f1974,f160])).
fof(f1974,plain,(
  ( ! [X59,X57,X54,X58,X56,X55] : (p(s(bool,'F_0')) | p(s(bool,X57)) | ~p(s(bool,X58)) | s(bool,X57) != s(bool,X59) | s(bool,X58) != s(bool,'const.list.LIST_REL_3'(s(fun(X54,fun(sK0,bool)),X55),s('type.list.list'(X54),X56),s('type.list.list'(sK0),sK3))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X54),X56)))))) ) | $spl16),
  inference(superposition,[],[f118,f1733])).
fof(f1928,plain,(
  ( ! [X6,X10,X8,X7] : (s(bool,X10) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X8))))) | s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(X6,bool)),X7),s('type.list.list'(sK0),sK3),s('type.list.list'(X6),X8))) | p(s(bool,X10))) ) | $spl50),
  inference(cnf_transformation,[],[f1928_D])).
fof(f1928_D,plain,(
  ( ! [X6,X10,X8,X7] : (s(bool,X10) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X8))))) | s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(X6,bool)),X7),s('type.list.list'(sK0),sK3),s('type.list.list'(X6),X8))) | p(s(bool,X10))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f1917,plain,(
  ( ! [X85,X83,X81,X84,X82] : (s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(X81,bool)),X82),s('type.list.list'(sK0),sK3),s('type.list.list'(X81),X83))) | p(s(bool,X84)) | p(s(bool,X85)) | s(bool,X85) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X81),X83)))))) ) | $spl16),
  inference(resolution,[],[f1659,f180])).
fof(f1659,plain,(
  ( ! [X66,X67,X65] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X65),X66)))))) | s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(X65,bool)),X67),s('type.list.list'(sK0),sK3),s('type.list.list'(X65),X66)))) ) | $spl16),
  inference(resolution,[],[f585,f166])).
fof(f585,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(X3,bool)),X5),s('type.list.list'(sK0),sK3),s('type.list.list'(X3),X4)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X3),X4))))))) ) | $spl16),
  inference(superposition,[],[f254,f113])).
fof(f1907,plain,(
  ( ! [X52,X50,X48,X53,X51,X49] : (s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(X48,bool)),X49),s('type.list.list'(sK0),sK3),s('type.list.list'(X48),X50))) | p(s(bool,X51)) | ~p(s(bool,X52)) | s(bool,X51) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X48),X50))))) | s(bool,X52) != s(bool,X53)) ) | $spl16),
  inference(resolution,[],[f1659,f124])).
fof(f1904,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(X30,bool)),X31),s('type.list.list'(sK0),sK3),s('type.list.list'(X30),X32))) | p(s(bool,X33)) | ~p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X35) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X30),X32)))))) ) | $spl16),
  inference(resolution,[],[f1659,f123])).
fof(f1903,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(X24,bool)),X25),s('type.list.list'(sK0),sK3),s('type.list.list'(X24),X26))) | p(s(bool,X27)) | p(s(bool,X28)) | s(bool,X27) != s(bool,X29) | s(bool,X28) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X24),X26)))))) ) | $spl16),
  inference(resolution,[],[f1659,f122])).
fof(f1902,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(X18,bool)),X19),s('type.list.list'(sK0),sK3),s('type.list.list'(X18),X20))) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X18),X20))))) | s(bool,X22) != s(bool,X23)) ) | $spl16),
  inference(resolution,[],[f1659,f120])).
fof(f1900,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(X6,bool)),X7),s('type.list.list'(sK0),sK3),s('type.list.list'(X6),X8))) | p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X8)))))) ) | $spl16),
  inference(resolution,[],[f1659,f118])).
fof(f1880,plain,(
  ( ! [X6,X10,X8,X7] : (s(bool,X10) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X8))))) | s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X6,fun(sK0,bool)),X7),s('type.list.list'(X6),X8),s('type.list.list'(sK0),sK3))) | p(s(bool,X10))) ) | $spl48),
  inference(cnf_transformation,[],[f1880_D])).
fof(f1880_D,plain,(
  ( ! [X6,X10,X8,X7] : (s(bool,X10) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X8))))) | s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X6,fun(sK0,bool)),X7),s('type.list.list'(X6),X8),s('type.list.list'(sK0),sK3))) | p(s(bool,X10))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f1869,plain,(
  ( ! [X85,X83,X81,X84,X82] : (s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X81,fun(sK0,bool)),X82),s('type.list.list'(X81),X83),s('type.list.list'(sK0),sK3))) | p(s(bool,X84)) | p(s(bool,X85)) | s(bool,X85) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X81),X83)))))) ) | $spl16),
  inference(resolution,[],[f1619,f180])).
fof(f1619,plain,(
  ( ! [X66,X67,X65] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X65),X66)))))) | s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X65,fun(sK0,bool)),X67),s('type.list.list'(X65),X66),s('type.list.list'(sK0),sK3)))) ) | $spl16),
  inference(resolution,[],[f584,f166])).
fof(f584,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(sK0,bool)),X2),s('type.list.list'(X0),X1),s('type.list.list'(sK0),sK3)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))))))) ) | $spl16),
  inference(superposition,[],[f254,f113])).
fof(f1859,plain,(
  ( ! [X52,X50,X48,X53,X51,X49] : (s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X48,fun(sK0,bool)),X49),s('type.list.list'(X48),X50),s('type.list.list'(sK0),sK3))) | p(s(bool,X51)) | ~p(s(bool,X52)) | s(bool,X51) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X48),X50))))) | s(bool,X52) != s(bool,X53)) ) | $spl16),
  inference(resolution,[],[f1619,f124])).
fof(f1856,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X30,fun(sK0,bool)),X31),s('type.list.list'(X30),X32),s('type.list.list'(sK0),sK3))) | p(s(bool,X33)) | ~p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X35) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X30),X32)))))) ) | $spl16),
  inference(resolution,[],[f1619,f123])).
fof(f1855,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X24,fun(sK0,bool)),X25),s('type.list.list'(X24),X26),s('type.list.list'(sK0),sK3))) | p(s(bool,X27)) | p(s(bool,X28)) | s(bool,X27) != s(bool,X29) | s(bool,X28) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X24),X26)))))) ) | $spl16),
  inference(resolution,[],[f1619,f122])).
fof(f1854,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X18,fun(sK0,bool)),X19),s('type.list.list'(X18),X20),s('type.list.list'(sK0),sK3))) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X18),X20))))) | s(bool,X22) != s(bool,X23)) ) | $spl16),
  inference(resolution,[],[f1619,f120])).
fof(f1852,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X6,fun(sK0,bool)),X7),s('type.list.list'(X6),X8),s('type.list.list'(sK0),sK3))) | p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X8)))))) ) | $spl16),
  inference(resolution,[],[f1619,f118])).
fof(f1772,plain,(
  ( ! [X80,X78,X81,X79,X77] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X77),X78))))) | p(s(bool,X79)) | ~p(s(bool,X80)) | s(bool,X80) != s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(X77,bool)),X81),s('type.list.list'(sK0),sK3),s('type.list.list'(X77),X78)))) ) | $spl16),
  inference(resolution,[],[f1023,f180])).
fof(f1761,plain,(
  ( ! [X47,X52,X50,X48,X51,X49] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X47),X48))))) | ~p(s(bool,X49)) | ~p(s(bool,X50)) | s(bool,X50) != s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(X47,bool)),X51),s('type.list.list'(sK0),sK3),s('type.list.list'(X47),X48))) | s(bool,X49) != s(bool,X52)) ) | $spl16),
  inference(resolution,[],[f1023,f124])).
fof(f1759,plain,(
  ( ! [X39,X37,X41,X38,X36,X40] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X36),X37))))) | p(s(bool,X38)) | ~p(s(bool,X39)) | ~p(s(bool,X40)) | s(bool,X39) != s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(X36,bool)),X41),s('type.list.list'(sK0),sK3),s('type.list.list'(X36),X37)))) ) | $spl16),
  inference(resolution,[],[f1023,f123])).
fof(f1757,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X24),X25))))) | p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,X28) | s(bool,X27) != s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(X24,bool)),X29),s('type.list.list'(sK0),sK3),s('type.list.list'(X24),X25)))) ) | $spl16),
  inference(resolution,[],[f1023,f122])).
fof(f1756,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X18),X19))))) | ~p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X21) != s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(X18,bool)),X22),s('type.list.list'(sK0),sK3),s('type.list.list'(X18),X19))) | s(bool,X20) != s(bool,X23)) ) | $spl16),
  inference(resolution,[],[f1023,f120])).
fof(f1754,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X7))))) | p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(X6,bool)),X11),s('type.list.list'(sK0),sK3),s('type.list.list'(X6),X7)))) ) | $spl16),
  inference(resolution,[],[f1023,f118])).
fof(f1737,plain,(
  ( ! [X80,X78,X81,X79,X77] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X77),X78))))) | p(s(bool,X79)) | ~p(s(bool,X80)) | s(bool,X80) != s(bool,'const.list.LIST_REL_3'(s(fun(X77,fun(sK0,bool)),X81),s('type.list.list'(X77),X78),s('type.list.list'(sK0),sK3)))) ) | $spl16),
  inference(resolution,[],[f1022,f180])).
fof(f1726,plain,(
  ( ! [X47,X52,X50,X48,X51,X49] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X47),X48))))) | ~p(s(bool,X49)) | ~p(s(bool,X50)) | s(bool,X50) != s(bool,'const.list.LIST_REL_3'(s(fun(X47,fun(sK0,bool)),X51),s('type.list.list'(X47),X48),s('type.list.list'(sK0),sK3))) | s(bool,X49) != s(bool,X52)) ) | $spl16),
  inference(resolution,[],[f1022,f124])).
fof(f1724,plain,(
  ( ! [X39,X37,X41,X38,X36,X40] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X36),X37))))) | p(s(bool,X38)) | ~p(s(bool,X39)) | ~p(s(bool,X40)) | s(bool,X39) != s(bool,'const.list.LIST_REL_3'(s(fun(X36,fun(sK0,bool)),X41),s('type.list.list'(X36),X37),s('type.list.list'(sK0),sK3)))) ) | $spl16),
  inference(resolution,[],[f1022,f123])).
fof(f1722,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X24),X25))))) | p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,X28) | s(bool,X27) != s(bool,'const.list.LIST_REL_3'(s(fun(X24,fun(sK0,bool)),X29),s('type.list.list'(X24),X25),s('type.list.list'(sK0),sK3)))) ) | $spl16),
  inference(resolution,[],[f1022,f122])).
fof(f1721,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X18),X19))))) | ~p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X21) != s(bool,'const.list.LIST_REL_3'(s(fun(X18,fun(sK0,bool)),X22),s('type.list.list'(X18),X19),s('type.list.list'(sK0),sK3))) | s(bool,X20) != s(bool,X23)) ) | $spl16),
  inference(resolution,[],[f1022,f120])).
fof(f1719,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X7))))) | p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,'const.list.LIST_REL_3'(s(fun(X6,fun(sK0,bool)),X11),s('type.list.list'(X6),X7),s('type.list.list'(sK0),sK3)))) ) | $spl16),
  inference(resolution,[],[f1022,f118])).
fof(f1676,plain,(
  ( ! [X6,X7,X11,X9] : (s(bool,X9) != s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(X6,bool)),X11),s('type.list.list'(sK0),sK3),s('type.list.list'(X6),X7))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X7)))))) | ~p(s(bool,X9))) ) | $spl42),
  inference(cnf_transformation,[],[f1676_D])).
fof(f1676_D,plain,(
  ( ! [X6,X7,X11,X9] : (s(bool,X9) != s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(X6,bool)),X11),s('type.list.list'(sK0),sK3),s('type.list.list'(X6),X7))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X7)))))) | ~p(s(bool,X9))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f1663,plain,(
  ( ! [X80,X78,X81,X79,X77] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X77),X78)))))) | p(s(bool,X79)) | ~p(s(bool,X80)) | s(bool,X80) != s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(X77,bool)),X81),s('type.list.list'(sK0),sK3),s('type.list.list'(X77),X78)))) ) | $spl16),
  inference(resolution,[],[f585,f180])).
fof(f1652,plain,(
  ( ! [X47,X52,X50,X48,X51,X49] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X47),X48)))))) | ~p(s(bool,X49)) | ~p(s(bool,X50)) | s(bool,X50) != s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(X47,bool)),X51),s('type.list.list'(sK0),sK3),s('type.list.list'(X47),X48))) | s(bool,X49) != s(bool,X52)) ) | $spl16),
  inference(resolution,[],[f585,f124])).
fof(f1650,plain,(
  ( ! [X39,X37,X41,X38,X36,X40] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X36),X37)))))) | p(s(bool,X38)) | ~p(s(bool,X39)) | ~p(s(bool,X40)) | s(bool,X39) != s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(X36,bool)),X41),s('type.list.list'(sK0),sK3),s('type.list.list'(X36),X37)))) ) | $spl16),
  inference(resolution,[],[f585,f123])).
fof(f1648,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X24),X25)))))) | p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,X28) | s(bool,X27) != s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(X24,bool)),X29),s('type.list.list'(sK0),sK3),s('type.list.list'(X24),X25)))) ) | $spl16),
  inference(resolution,[],[f585,f122])).
fof(f1647,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X18),X19)))))) | ~p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X21) != s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(X18,bool)),X22),s('type.list.list'(sK0),sK3),s('type.list.list'(X18),X19))) | s(bool,X20) != s(bool,X23)) ) | $spl16),
  inference(resolution,[],[f585,f120])).
fof(f1645,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X7)))))) | p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(X6,bool)),X11),s('type.list.list'(sK0),sK3),s('type.list.list'(X6),X7)))) ) | $spl16),
  inference(resolution,[],[f585,f118])).
fof(f1634,plain,(
  ( ! [X6,X7,X11,X9] : (s(bool,X9) != s(bool,'const.list.LIST_REL_3'(s(fun(X6,fun(sK0,bool)),X11),s('type.list.list'(X6),X7),s('type.list.list'(sK0),sK3))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X7)))))) | ~p(s(bool,X9))) ) | $spl40),
  inference(cnf_transformation,[],[f1634_D])).
fof(f1634_D,plain,(
  ( ! [X6,X7,X11,X9] : (s(bool,X9) != s(bool,'const.list.LIST_REL_3'(s(fun(X6,fun(sK0,bool)),X11),s('type.list.list'(X6),X7),s('type.list.list'(sK0),sK3))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X7)))))) | ~p(s(bool,X9))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f1623,plain,(
  ( ! [X80,X78,X81,X79,X77] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X77),X78)))))) | p(s(bool,X79)) | ~p(s(bool,X80)) | s(bool,X80) != s(bool,'const.list.LIST_REL_3'(s(fun(X77,fun(sK0,bool)),X81),s('type.list.list'(X77),X78),s('type.list.list'(sK0),sK3)))) ) | $spl16),
  inference(resolution,[],[f584,f180])).
fof(f1612,plain,(
  ( ! [X47,X52,X50,X48,X51,X49] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X47),X48)))))) | ~p(s(bool,X49)) | ~p(s(bool,X50)) | s(bool,X50) != s(bool,'const.list.LIST_REL_3'(s(fun(X47,fun(sK0,bool)),X51),s('type.list.list'(X47),X48),s('type.list.list'(sK0),sK3))) | s(bool,X49) != s(bool,X52)) ) | $spl16),
  inference(resolution,[],[f584,f124])).
fof(f1610,plain,(
  ( ! [X39,X37,X41,X38,X36,X40] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X36),X37)))))) | p(s(bool,X38)) | ~p(s(bool,X39)) | ~p(s(bool,X40)) | s(bool,X39) != s(bool,'const.list.LIST_REL_3'(s(fun(X36,fun(sK0,bool)),X41),s('type.list.list'(X36),X37),s('type.list.list'(sK0),sK3)))) ) | $spl16),
  inference(resolution,[],[f584,f123])).
fof(f1608,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X24),X25)))))) | p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,X28) | s(bool,X27) != s(bool,'const.list.LIST_REL_3'(s(fun(X24,fun(sK0,bool)),X29),s('type.list.list'(X24),X25),s('type.list.list'(sK0),sK3)))) ) | $spl16),
  inference(resolution,[],[f584,f122])).
fof(f1607,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X18),X19)))))) | ~p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X21) != s(bool,'const.list.LIST_REL_3'(s(fun(X18,fun(sK0,bool)),X22),s('type.list.list'(X18),X19),s('type.list.list'(sK0),sK3))) | s(bool,X20) != s(bool,X23)) ) | $spl16),
  inference(resolution,[],[f584,f120])).
fof(f1605,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X7)))))) | p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,'const.list.LIST_REL_3'(s(fun(X6,fun(sK0,bool)),X11),s('type.list.list'(X6),X7),s('type.list.list'(sK0),sK3)))) ) | $spl16),
  inference(resolution,[],[f584,f118])).
fof(f1264,plain,(
  ( ! [X6,X7,X11,X9] : (s(bool,X9) != s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(X6,bool)),X11),s('type.list.list'(sK1),sK2),s('type.list.list'(X6),X7))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X7))) | ~p(s(bool,X9))) ) | $spl36),
  inference(cnf_transformation,[],[f1264_D])).
fof(f1264_D,plain,(
  ( ! [X6,X7,X11,X9] : (s(bool,X9) != s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(X6,bool)),X11),s('type.list.list'(sK1),sK2),s('type.list.list'(X6),X7))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X7))) | ~p(s(bool,X9))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f1568,plain,(
  ( ! [X321,X323,X324,X320,X322] : (p(s(bool,X323)) | ~p(s(bool,X324)) | s(bool,X324) != s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(X320,bool)),X321),s('type.list.list'(sK1),sK2),s('type.list.list'(X320),X322))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X320),X322)))) )),
  inference(subsumption_resolution,[],[f1520,f160])).
fof(f1520,plain,(
  ( ! [X321,X323,X324,X320,X322] : (p(s(bool,'F_0')) | p(s(bool,X323)) | ~p(s(bool,X324)) | s(bool,X324) != s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(X320,bool)),X321),s('type.list.list'(sK1),sK2),s('type.list.list'(X320),X322))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X320),X322)))) )),
  inference(superposition,[],[f180,f1247])).
fof(f1247,plain,(
  ( ! [X66,X67,X65] : (s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(X65,bool)),X67),s('type.list.list'(sK1),sK2),s('type.list.list'(X65),X66))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X65),X66)))) )),
  inference(resolution,[],[f434,f166])).
fof(f434,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(X3,bool)),X5),s('type.list.list'(sK1),sK2),s('type.list.list'(X3),X4)))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X3),X4)))) )),
  inference(superposition,[],[f183,f113])).
fof(f183,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK1),sK2)))),
  inference(resolution,[],[f111,f113])).
fof(f1557,plain,(
  ( ! [X210,X208,X206,X209,X207,X205] : (~p(s(bool,X208)) | ~p(s(bool,X209)) | s(bool,X209) != s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(X205,bool)),X206),s('type.list.list'(sK1),sK2),s('type.list.list'(X205),X207))) | s(bool,X208) != s(bool,X210) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X205),X207)))) )),
  inference(subsumption_resolution,[],[f1489,f136])).
fof(f1489,plain,(
  ( ! [X210,X208,X206,X209,X207,X205] : (p(s(bool,'F_0')) | ~p(s(bool,X208)) | ~p(s(bool,X209)) | s(bool,X209) != s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(X205,bool)),X206),s('type.list.list'(sK1),sK2),s('type.list.list'(X205),X207))) | s(bool,X208) != s(bool,X210) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X205),X207)))) )),
  inference(superposition,[],[f124,f1247])).
fof(f1555,plain,(
  ( ! [X158,X161,X159,X157,X162,X160] : (p(s(bool,X160)) | ~p(s(bool,X161)) | ~p(s(bool,X162)) | s(bool,X161) != s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(X157,bool)),X158),s('type.list.list'(sK1),sK2),s('type.list.list'(X157),X159))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X157),X159)))) )),
  inference(subsumption_resolution,[],[f1481,f160])).
fof(f1481,plain,(
  ( ! [X158,X161,X159,X157,X162,X160] : (p(s(bool,'F_0')) | p(s(bool,X160)) | ~p(s(bool,X161)) | ~p(s(bool,X162)) | s(bool,X161) != s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(X157,bool)),X158),s('type.list.list'(sK1),sK2),s('type.list.list'(X157),X159))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X157),X159)))) )),
  inference(superposition,[],[f123,f1247])).
fof(f1551,plain,(
  ( ! [X127,X132,X130,X128,X131,X129] : (p(s(bool,X130)) | ~p(s(bool,X131)) | s(bool,X130) != s(bool,X132) | s(bool,X131) != s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(X127,bool)),X128),s('type.list.list'(sK1),sK2),s('type.list.list'(X127),X129))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X127),X129)))) )),
  inference(subsumption_resolution,[],[f1476,f160])).
fof(f1476,plain,(
  ( ! [X127,X132,X130,X128,X131,X129] : (p(s(bool,'F_0')) | p(s(bool,X130)) | ~p(s(bool,X131)) | s(bool,X130) != s(bool,X132) | s(bool,X131) != s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(X127,bool)),X128),s('type.list.list'(sK1),sK2),s('type.list.list'(X127),X129))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X127),X129)))) )),
  inference(superposition,[],[f122,f1247])).
fof(f1548,plain,(
  ( ! [X90,X88,X87,X85,X89,X86] : (~p(s(bool,X88)) | ~p(s(bool,X89)) | s(bool,X89) != s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(X85,bool)),X86),s('type.list.list'(sK1),sK2),s('type.list.list'(X85),X87))) | s(bool,X88) != s(bool,X90) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X85),X87)))) )),
  inference(subsumption_resolution,[],[f1469,f136])).
fof(f1469,plain,(
  ( ! [X90,X88,X87,X85,X89,X86] : (p(s(bool,'F_0')) | ~p(s(bool,X88)) | ~p(s(bool,X89)) | s(bool,X89) != s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(X85,bool)),X86),s('type.list.list'(sK1),sK2),s('type.list.list'(X85),X87))) | s(bool,X88) != s(bool,X90) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X85),X87)))) )),
  inference(superposition,[],[f120,f1247])).
fof(f1544,plain,(
  ( ! [X47,X45,X43,X48,X46,X44] : (p(s(bool,X46)) | ~p(s(bool,X47)) | s(bool,X46) != s(bool,X48) | s(bool,X47) != s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(X43,bool)),X44),s('type.list.list'(sK1),sK2),s('type.list.list'(X43),X45))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X43),X45)))) )),
  inference(subsumption_resolution,[],[f1462,f160])).
fof(f1462,plain,(
  ( ! [X47,X45,X43,X48,X46,X44] : (p(s(bool,'F_0')) | p(s(bool,X46)) | ~p(s(bool,X47)) | s(bool,X46) != s(bool,X48) | s(bool,X47) != s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(X43,bool)),X44),s('type.list.list'(sK1),sK2),s('type.list.list'(X43),X45))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X43),X45)))) )),
  inference(superposition,[],[f118,f1247])).
fof(f1427,plain,(
  ( ! [X2,X5,X3] : (s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK5),s(sK0,X2))),s(sK1,X3))) | s(bool,X5) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,X2))),s(sK1,X3))) | p(s(bool,X5))) ) | $spl38),
  inference(cnf_transformation,[],[f1427_D])).
fof(f1427_D,plain,(
  ( ! [X2,X5,X3] : (s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK5),s(sK0,X2))),s(sK1,X3))) | s(bool,X5) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,X2))),s(sK1,X3))) | p(s(bool,X5))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f1416,plain,(
  ( ! [X61,X62,X60,X63] : (s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK5),s(sK0,X60))),s(sK1,X61))) | p(s(bool,X62)) | p(s(bool,X63)) | s(bool,X63) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,X60))),s(sK1,X61)))) )),
  inference(resolution,[],[f464,f180])).
fof(f1406,plain,(
  ( ! [X39,X37,X41,X38,X40] : (s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK5),s(sK0,X37))),s(sK1,X38))) | p(s(bool,X39)) | ~p(s(bool,X40)) | s(bool,X39) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,X37))),s(sK1,X38))) | s(bool,X40) != s(bool,X41)) )),
  inference(resolution,[],[f464,f124])).
fof(f1403,plain,(
  ( ! [X26,X24,X23,X25,X22] : (s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK5),s(sK0,X22))),s(sK1,X23))) | p(s(bool,X24)) | ~p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X26) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,X22))),s(sK1,X23)))) )),
  inference(resolution,[],[f464,f123])).
fof(f1402,plain,(
  ( ! [X21,X19,X17,X20,X18] : (s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK5),s(sK0,X17))),s(sK1,X18))) | p(s(bool,X19)) | p(s(bool,X20)) | s(bool,X19) != s(bool,X21) | s(bool,X20) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,X17))),s(sK1,X18)))) )),
  inference(resolution,[],[f464,f122])).
fof(f1401,plain,(
  ( ! [X14,X12,X15,X13,X16] : (s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK5),s(sK0,X12))),s(sK1,X13))) | p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X14) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,X12))),s(sK1,X13))) | s(bool,X15) != s(bool,X16)) )),
  inference(resolution,[],[f464,f120])).
fof(f1399,plain,(
  ( ! [X6,X4,X2,X5,X3] : (s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK5),s(sK0,X2))),s(sK1,X3))) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,X2))),s(sK1,X3)))) )),
  inference(resolution,[],[f464,f118])).
fof(f1227,plain,(
  ( ! [X6,X7,X11,X9] : (s(bool,X9) != s(bool,'const.list.LIST_REL_3'(s(fun(X6,fun(sK1,bool)),X11),s('type.list.list'(X6),X7),s('type.list.list'(sK1),sK2))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X7))) | ~p(s(bool,X9))) ) | $spl34),
  inference(cnf_transformation,[],[f1227_D])).
fof(f1227_D,plain,(
  ( ! [X6,X7,X11,X9] : (s(bool,X9) != s(bool,'const.list.LIST_REL_3'(s(fun(X6,fun(sK1,bool)),X11),s('type.list.list'(X6),X7),s('type.list.list'(sK1),sK2))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X7))) | ~p(s(bool,X9))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f1390,plain,(
  ( ! [X321,X318,X320,X322,X319] : (p(s(bool,X321)) | ~p(s(bool,X322)) | s(bool,X322) != s(bool,'const.list.LIST_REL_3'(s(fun(X318,fun(sK1,bool)),X319),s('type.list.list'(X318),X320),s('type.list.list'(sK1),sK2))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X318),X320)))) )),
  inference(subsumption_resolution,[],[f1344,f160])).
fof(f1344,plain,(
  ( ! [X321,X318,X320,X322,X319] : (p(s(bool,'F_0')) | p(s(bool,X321)) | ~p(s(bool,X322)) | s(bool,X322) != s(bool,'const.list.LIST_REL_3'(s(fun(X318,fun(sK1,bool)),X319),s('type.list.list'(X318),X320),s('type.list.list'(sK1),sK2))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X318),X320)))) )),
  inference(superposition,[],[f180,f1209])).
fof(f1209,plain,(
  ( ! [X66,X67,X65] : (s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X65,fun(sK1,bool)),X67),s('type.list.list'(X65),X66),s('type.list.list'(sK1),sK2))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X65),X66)))) )),
  inference(resolution,[],[f433,f166])).
fof(f433,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(sK1,bool)),X2),s('type.list.list'(X0),X1),s('type.list.list'(sK1),sK2)))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1)))) )),
  inference(superposition,[],[f183,f113])).
fof(f1379,plain,(
  ( ! [X208,X206,X204,X207,X205,X203] : (~p(s(bool,X206)) | ~p(s(bool,X207)) | s(bool,X207) != s(bool,'const.list.LIST_REL_3'(s(fun(X203,fun(sK1,bool)),X204),s('type.list.list'(X203),X205),s('type.list.list'(sK1),sK2))) | s(bool,X206) != s(bool,X208) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X203),X205)))) )),
  inference(subsumption_resolution,[],[f1313,f136])).
fof(f1313,plain,(
  ( ! [X208,X206,X204,X207,X205,X203] : (p(s(bool,'F_0')) | ~p(s(bool,X206)) | ~p(s(bool,X207)) | s(bool,X207) != s(bool,'const.list.LIST_REL_3'(s(fun(X203,fun(sK1,bool)),X204),s('type.list.list'(X203),X205),s('type.list.list'(sK1),sK2))) | s(bool,X206) != s(bool,X208) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X203),X205)))) )),
  inference(superposition,[],[f124,f1209])).
fof(f1377,plain,(
  ( ! [X158,X156,X159,X157,X155,X160] : (p(s(bool,X158)) | ~p(s(bool,X159)) | ~p(s(bool,X160)) | s(bool,X159) != s(bool,'const.list.LIST_REL_3'(s(fun(X155,fun(sK1,bool)),X156),s('type.list.list'(X155),X157),s('type.list.list'(sK1),sK2))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X155),X157)))) )),
  inference(subsumption_resolution,[],[f1305,f160])).
fof(f1305,plain,(
  ( ! [X158,X156,X159,X157,X155,X160] : (p(s(bool,'F_0')) | p(s(bool,X158)) | ~p(s(bool,X159)) | ~p(s(bool,X160)) | s(bool,X159) != s(bool,'const.list.LIST_REL_3'(s(fun(X155,fun(sK1,bool)),X156),s('type.list.list'(X155),X157),s('type.list.list'(sK1),sK2))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X155),X157)))) )),
  inference(superposition,[],[f123,f1209])).
fof(f1373,plain,(
  ( ! [X127,X125,X130,X128,X126,X129] : (p(s(bool,X128)) | ~p(s(bool,X129)) | s(bool,X128) != s(bool,X130) | s(bool,X129) != s(bool,'const.list.LIST_REL_3'(s(fun(X125,fun(sK1,bool)),X126),s('type.list.list'(X125),X127),s('type.list.list'(sK1),sK2))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X125),X127)))) )),
  inference(subsumption_resolution,[],[f1300,f160])).
fof(f1300,plain,(
  ( ! [X127,X125,X130,X128,X126,X129] : (p(s(bool,'F_0')) | p(s(bool,X128)) | ~p(s(bool,X129)) | s(bool,X128) != s(bool,X130) | s(bool,X129) != s(bool,'const.list.LIST_REL_3'(s(fun(X125,fun(sK1,bool)),X126),s('type.list.list'(X125),X127),s('type.list.list'(sK1),sK2))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X125),X127)))) )),
  inference(superposition,[],[f122,f1209])).
fof(f1370,plain,(
  ( ! [X88,X87,X85,X83,X86,X84] : (~p(s(bool,X86)) | ~p(s(bool,X87)) | s(bool,X87) != s(bool,'const.list.LIST_REL_3'(s(fun(X83,fun(sK1,bool)),X84),s('type.list.list'(X83),X85),s('type.list.list'(sK1),sK2))) | s(bool,X86) != s(bool,X88) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X83),X85)))) )),
  inference(subsumption_resolution,[],[f1293,f136])).
fof(f1293,plain,(
  ( ! [X88,X87,X85,X83,X86,X84] : (p(s(bool,'F_0')) | ~p(s(bool,X86)) | ~p(s(bool,X87)) | s(bool,X87) != s(bool,'const.list.LIST_REL_3'(s(fun(X83,fun(sK1,bool)),X84),s('type.list.list'(X83),X85),s('type.list.list'(sK1),sK2))) | s(bool,X86) != s(bool,X88) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X83),X85)))) )),
  inference(superposition,[],[f120,f1209])).
fof(f1366,plain,(
  ( ! [X45,X43,X41,X46,X44,X42] : (p(s(bool,X44)) | ~p(s(bool,X45)) | s(bool,X44) != s(bool,X46) | s(bool,X45) != s(bool,'const.list.LIST_REL_3'(s(fun(X41,fun(sK1,bool)),X42),s('type.list.list'(X41),X43),s('type.list.list'(sK1),sK2))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X41),X43)))) )),
  inference(subsumption_resolution,[],[f1286,f160])).
fof(f1286,plain,(
  ( ! [X45,X43,X41,X46,X44,X42] : (p(s(bool,'F_0')) | p(s(bool,X44)) | ~p(s(bool,X45)) | s(bool,X44) != s(bool,X46) | s(bool,X45) != s(bool,'const.list.LIST_REL_3'(s(fun(X41,fun(sK1,bool)),X42),s('type.list.list'(X41),X43),s('type.list.list'(sK1),sK2))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X41),X43)))) )),
  inference(superposition,[],[f118,f1209])).
fof(f1251,plain,(
  ( ! [X80,X78,X81,X79,X77] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X77),X78))) | p(s(bool,X79)) | ~p(s(bool,X80)) | s(bool,X80) != s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(X77,bool)),X81),s('type.list.list'(sK1),sK2),s('type.list.list'(X77),X78)))) )),
  inference(resolution,[],[f434,f180])).
fof(f1240,plain,(
  ( ! [X47,X52,X50,X48,X51,X49] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X47),X48))) | ~p(s(bool,X49)) | ~p(s(bool,X50)) | s(bool,X50) != s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(X47,bool)),X51),s('type.list.list'(sK1),sK2),s('type.list.list'(X47),X48))) | s(bool,X49) != s(bool,X52)) )),
  inference(resolution,[],[f434,f124])).
fof(f1238,plain,(
  ( ! [X39,X37,X41,X38,X36,X40] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X36),X37))) | p(s(bool,X38)) | ~p(s(bool,X39)) | ~p(s(bool,X40)) | s(bool,X39) != s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(X36,bool)),X41),s('type.list.list'(sK1),sK2),s('type.list.list'(X36),X37)))) )),
  inference(resolution,[],[f434,f123])).
fof(f1236,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X24),X25))) | p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,X28) | s(bool,X27) != s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(X24,bool)),X29),s('type.list.list'(sK1),sK2),s('type.list.list'(X24),X25)))) )),
  inference(resolution,[],[f434,f122])).
fof(f1235,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X18),X19))) | ~p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X21) != s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(X18,bool)),X22),s('type.list.list'(sK1),sK2),s('type.list.list'(X18),X19))) | s(bool,X20) != s(bool,X23)) )),
  inference(resolution,[],[f434,f120])).
fof(f1233,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X7))) | p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,'const.list.LIST_REL_3'(s(fun(sK1,fun(X6,bool)),X11),s('type.list.list'(sK1),sK2),s('type.list.list'(X6),X7)))) )),
  inference(resolution,[],[f434,f118])).
fof(f1213,plain,(
  ( ! [X80,X78,X81,X79,X77] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X77),X78))) | p(s(bool,X79)) | ~p(s(bool,X80)) | s(bool,X80) != s(bool,'const.list.LIST_REL_3'(s(fun(X77,fun(sK1,bool)),X81),s('type.list.list'(X77),X78),s('type.list.list'(sK1),sK2)))) )),
  inference(resolution,[],[f433,f180])).
fof(f1202,plain,(
  ( ! [X47,X52,X50,X48,X51,X49] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X47),X48))) | ~p(s(bool,X49)) | ~p(s(bool,X50)) | s(bool,X50) != s(bool,'const.list.LIST_REL_3'(s(fun(X47,fun(sK1,bool)),X51),s('type.list.list'(X47),X48),s('type.list.list'(sK1),sK2))) | s(bool,X49) != s(bool,X52)) )),
  inference(resolution,[],[f433,f124])).
fof(f1200,plain,(
  ( ! [X39,X37,X41,X38,X36,X40] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X36),X37))) | p(s(bool,X38)) | ~p(s(bool,X39)) | ~p(s(bool,X40)) | s(bool,X39) != s(bool,'const.list.LIST_REL_3'(s(fun(X36,fun(sK1,bool)),X41),s('type.list.list'(X36),X37),s('type.list.list'(sK1),sK2)))) )),
  inference(resolution,[],[f433,f123])).
fof(f1198,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X24),X25))) | p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,X28) | s(bool,X27) != s(bool,'const.list.LIST_REL_3'(s(fun(X24,fun(sK1,bool)),X29),s('type.list.list'(X24),X25),s('type.list.list'(sK1),sK2)))) )),
  inference(resolution,[],[f433,f122])).
fof(f1197,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X18),X19))) | ~p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X21) != s(bool,'const.list.LIST_REL_3'(s(fun(X18,fun(sK1,bool)),X22),s('type.list.list'(X18),X19),s('type.list.list'(sK1),sK2))) | s(bool,X20) != s(bool,X23)) )),
  inference(resolution,[],[f433,f120])).
fof(f1195,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X7))) | p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,'const.list.LIST_REL_3'(s(fun(X6,fun(sK1,bool)),X11),s('type.list.list'(X6),X7),s('type.list.list'(sK1),sK2)))) )),
  inference(resolution,[],[f433,f118])).
fof(f1189,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,'const.list.EL_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.list.list'(sK0),sK3))))),s(sK1,'const.list.EL_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.list.list'(sK1),sK2)))))) ) | $spl32),
  inference(cnf_transformation,[],[f1189_D])).
fof(f1189_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,'const.list.EL_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.list.list'(sK0),sK3))))),s(sK1,'const.list.EL_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.list.list'(sK1),sK2)))))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f1174,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,'const.list.EL_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.list.list'(sK0),sK3))))),s(sK1,'const.list.EL_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.list.list'(sK1),sK2)))))) ) | $spl19),
  inference(resolution,[],[f256,f180])).
fof(f1163,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,'const.list.EL_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.list.list'(sK0),sK3))))),s(sK1,'const.list.EL_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.list.list'(sK1),sK2))))) | s(bool,X21) != s(bool,X23)) ) | $spl19),
  inference(resolution,[],[f256,f124])).
fof(f1161,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,'const.list.EL_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.list.list'(sK0),sK3))))),s(sK1,'const.list.EL_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.list.list'(sK1),sK2)))))) ) | $spl19),
  inference(resolution,[],[f256,f123])).
fof(f1159,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,'const.list.EL_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.list.list'(sK0),sK3))))),s(sK1,'const.list.EL_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.list.list'(sK1),sK2)))))) ) | $spl19),
  inference(resolution,[],[f256,f122])).
fof(f1158,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,'const.list.EL_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.list.list'(sK0),sK3))))),s(sK1,'const.list.EL_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.list.list'(sK1),sK2))))) | s(bool,X6) != s(bool,X8)) ) | $spl19),
  inference(resolution,[],[f256,f120])).
fof(f1156,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,'const.list.EL_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.list.list'(sK0),sK3))))),s(sK1,'const.list.EL_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.list.list'(sK1),sK2)))))) ) | $spl19),
  inference(resolution,[],[f256,f118])).
fof(f839,plain,(
  $false | $spl4),
  inference(trivial_inequality_removal,[],[f828])).
fof(f828,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl4),
  inference(superposition,[],[f213,f200])).
fof(f838,plain,(
  $false | $spl4),
  inference(equality_resolution,[],[f213])).
fof(f592,plain,(
  ( ! [X4] : (p(s(bool,X4)) | s(bool,X4) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3)))))) ) | $spl30),
  inference(cnf_transformation,[],[f592_D])).
fof(f592_D,plain,(
  ( ! [X4] : (p(s(bool,X4)) | s(bool,X4) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3)))))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f583,plain,(
  ( ! [X28,X27] : (p(s(bool,X27)) | p(s(bool,X28)) | s(bool,X28) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3)))))) ) | $spl16),
  inference(resolution,[],[f254,f180])).
fof(f573,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3))))) | s(bool,X25) != s(bool,X26)) ) | $spl16),
  inference(resolution,[],[f254,f124])).
fof(f570,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X17) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3)))))) ) | $spl16),
  inference(resolution,[],[f254,f123])).
fof(f569,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3)))))) ) | $spl16),
  inference(resolution,[],[f254,f122])).
fof(f568,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3))))) | s(bool,X10) != s(bool,X11)) ) | $spl16),
  inference(resolution,[],[f254,f120])).
fof(f566,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3)))))) ) | $spl16),
  inference(resolution,[],[f254,f118])).
fof(f530,plain,(
  $false | $spl15),
  inference(subsumption_resolution,[],[f252,f183])).
fof(f480,plain,(
  ( ! [X6,X4,X3] : (s(bool,X6) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK5),s(sK0,X3))),s(sK1,X4))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,X3))),s(sK1,X4)))) | ~p(s(bool,X6))) ) | $spl28),
  inference(cnf_transformation,[],[f480_D])).
fof(f480_D,plain,(
  ( ! [X6,X4,X3] : (s(bool,X6) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK5),s(sK0,X3))),s(sK1,X4))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,X3))),s(sK1,X4)))) | ~p(s(bool,X6))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f468,plain,(
  ( ! [X66,X64,X65,X63] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,X63))),s(sK1,X64)))) | p(s(bool,X65)) | ~p(s(bool,X66)) | s(bool,X66) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK5),s(sK0,X63))),s(sK1,X64)))) )),
  inference(resolution,[],[f110,f180])).
fof(f457,plain,(
  ( ! [X39,X41,X38,X42,X40] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,X38))),s(sK1,X39)))) | ~p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X41) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK5),s(sK0,X38))),s(sK1,X39))) | s(bool,X40) != s(bool,X42)) )),
  inference(resolution,[],[f110,f124])).
fof(f455,plain,(
  ( ! [X30,X28,X31,X29,X32] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,X28))),s(sK1,X29)))) | p(s(bool,X30)) | ~p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK5),s(sK0,X28))),s(sK1,X29)))) )),
  inference(resolution,[],[f110,f123])).
fof(f453,plain,(
  ( ! [X21,X19,X22,X20,X18] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,X18))),s(sK1,X19)))) | p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X20) != s(bool,X22) | s(bool,X21) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK5),s(sK0,X18))),s(sK1,X19)))) )),
  inference(resolution,[],[f110,f122])).
fof(f452,plain,(
  ( ! [X14,X17,X15,X13,X16] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,X13))),s(sK1,X14)))) | ~p(s(bool,X15)) | ~p(s(bool,X16)) | s(bool,X16) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK5),s(sK0,X13))),s(sK1,X14))) | s(bool,X15) != s(bool,X17)) )),
  inference(resolution,[],[f110,f120])).
fof(f450,plain,(
  ( ! [X6,X4,X7,X5,X3] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,X3))),s(sK1,X4)))) | p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK5),s(sK0,X3))),s(sK1,X4)))) )),
  inference(resolution,[],[f110,f118])).
fof(f428,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f399,f216])).
fof(f399,plain,(
  ( ! [X44] : (p(s(bool,X44))) ) | $spl10),
  inference(resolution,[],[f216,f147])).
fof(f147,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f93])).
fof(f93,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,'T_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f92])).
fof(f92,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (~p(s(bool,'T_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f70])).
fof(f70,plain,(
  ! [X0] : ((~p(s(bool,'T_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(ennf_transformation,[],[f42])).
fof(f42,plain,(
  ! [X0] : ((p(s(bool,'T_0')) => p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X14] : ((p(s(bool,'T_0')) => p(s(bool,X14))) <=> p(s(bool,X14)))),
  file('thm.list.LIST_REL_mono','thm.bool.IMP_CLAUSES|split|0')).
fof(f427,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f397,f216])).
fof(f397,plain,(
  ( ! [X42] : (p(s(bool,X42))) ) | $spl10),
  inference(resolution,[],[f216,f144])).
fof(f144,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f91])).
fof(f91,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))))),
  inference(flattening,[],[f90])).
fof(f90,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f69])).
fof(f69,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f41])).
fof(f41,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X14] : ((p(s(bool,X14)) => p(s(bool,X14))) <=> p(s(bool,'T_0')))),
  file('thm.list.LIST_REL_mono','thm.bool.IMP_CLAUSES|split|3')).
fof(f426,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f395,f216])).
fof(f395,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(resolution,[],[f216,f141])).
fof(f141,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f89])).
fof(f89,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f88])).
fof(f88,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f68])).
fof(f68,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f40])).
fof(f40,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X14] : ((p(s(bool,'F_0')) => p(s(bool,X14))) <=> p(s(bool,'T_0')))),
  file('thm.list.LIST_REL_mono','thm.bool.IMP_CLAUSES|split|2')).
fof(f425,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f394,f216])).
fof(f394,plain,(
  ( ! [X40] : (p(s(bool,X40))) ) | $spl10),
  inference(resolution,[],[f216,f138])).
fof(f138,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f87])).
fof(f87,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,'T_0'))))),
  inference(flattening,[],[f86])).
fof(f86,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f67])).
fof(f67,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f39])).
fof(f39,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X14] : ((p(s(bool,X14)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  file('thm.list.LIST_REL_mono','thm.bool.IMP_CLAUSES|split|1')).
fof(f402,plain,(
  $false | $spl10),
  inference(resolution,[],[f216,f174])).
fof(f174,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  p(s(bool,'T_0'))),
  file('thm.list.LIST_REL_mono','thm.bool.TRUTH')).
fof(f403,plain,(
  $false | $spl10),
  inference(resolution,[],[f216,f175])).
fof(f404,plain,(
  $false | $spl10),
  inference(resolution,[],[f216,f181])).
fof(f409,plain,(
  $false | $spl10),
  inference(resolution,[],[f216,f111])).
fof(f337,plain,(
  ( ! [X7] : (s(bool,'T_0') != s(bool,X7)) ) | $spl24),
  inference(cnf_transformation,[],[f337_D])).
fof(f337_D,plain,(
  ( ! [X7] : (s(bool,'T_0') != s(bool,X7)) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f367,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') != s(bool,X13) | ~p(s(bool,X12))) ) | $spl8),
  inference(duplicate_literal_removal,[],[f348])).
fof(f348,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') != s(bool,X13) | ~p(s(bool,X12)) | ~p(s(bool,X12))) ) | $spl8),
  inference(superposition,[],[f215,f170])).
fof(f368,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') != s(bool,X11) | ~p(s(bool,X10))) ) | $spl8),
  inference(duplicate_literal_removal,[],[f347])).
fof(f347,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') != s(bool,X11) | ~p(s(bool,X10)) | ~p(s(bool,X10))) ) | $spl8),
  inference(superposition,[],[f215,f172])).
fof(f172,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f107])).
fof(f107,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X14] : (s(bool,'T_0') = s(bool,X14) <=> p(s(bool,X14)))),
  file('thm.list.LIST_REL_mono','thm.bool.EQ_CLAUSES|split|1')).
fof(f339,plain,(
  ( ! [X9] : (s(bool,'F_0') != s(bool,X9)) ) | $spl26),
  inference(cnf_transformation,[],[f339_D])).
fof(f339_D,plain,(
  ( ! [X9] : (s(bool,'F_0') != s(bool,X9)) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f372,plain,(
  ( ! [X8,X9] : (s(bool,'F_0') != s(bool,X9) | ~p(s(bool,X8))) ) | $spl8),
  inference(subsumption_resolution,[],[f346,f215])).
fof(f346,plain,(
  ( ! [X8,X9] : (s(bool,'F_0') != s(bool,X9) | ~p(s(bool,X8)) | s(bool,'T_0') = s(bool,X8)) ) | $spl8),
  inference(superposition,[],[f215,f178])).
fof(f178,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.list.LIST_REL_mono',aHL_BOOL_CASES)).
fof(f371,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X7) | ~p(s(bool,X6))) ) | $spl8),
  inference(subsumption_resolution,[],[f345,f215])).
fof(f345,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X7) | ~p(s(bool,X6)) | s(bool,'F_0') = s(bool,X6)) ) | $spl8),
  inference(superposition,[],[f215,f178])).
fof(f333,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,X17) | p(s(bool,X16))) ) | $spl2),
  inference(duplicate_literal_removal,[],[f316])).
fof(f316,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,X17) | p(s(bool,X16)) | p(s(bool,X16))) ) | $spl2),
  inference(superposition,[],[f212,f166])).
fof(f334,plain,(
  ( ! [X14,X15] : (s(bool,'F_0') != s(bool,X15) | p(s(bool,X14))) ) | $spl2),
  inference(duplicate_literal_removal,[],[f315])).
fof(f315,plain,(
  ( ! [X14,X15] : (s(bool,'F_0') != s(bool,X15) | p(s(bool,X14)) | p(s(bool,X14))) ) | $spl2),
  inference(superposition,[],[f212,f168])).
fof(f168,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f105])).
fof(f105,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f50])).
fof(f50,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X14] : (s(bool,'F_0') = s(bool,X14) <=> ~p(s(bool,X14)))),
  file('thm.list.LIST_REL_mono','thm.bool.EQ_CLAUSES|split|3')).
fof(f338,plain,(
  ( ! [X8,X9] : (s(bool,'F_0') != s(bool,X9) | p(s(bool,X8))) ) | $spl2),
  inference(subsumption_resolution,[],[f312,f212])).
fof(f312,plain,(
  ( ! [X8,X9] : (s(bool,'F_0') != s(bool,X9) | p(s(bool,X8)) | s(bool,'T_0') = s(bool,X8)) ) | $spl2),
  inference(superposition,[],[f212,f178])).
fof(f336,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X7) | p(s(bool,X6))) ) | $spl2),
  inference(subsumption_resolution,[],[f311,f212])).
fof(f311,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X7) | p(s(bool,X6)) | s(bool,'F_0') = s(bool,X6)) ) | $spl2),
  inference(superposition,[],[f212,f178])).
fof(f307,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f286,f217])).
fof(f286,plain,(
  ( ! [X42] : (~p(s(bool,X42))) ) | $spl12),
  inference(resolution,[],[f217,f136])).
fof(f282,plain,(
  $false | $spl12),
  inference(resolution,[],[f217,f173])).
fof(f173,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f54])).
fof(f54,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.list.LIST_REL_mono',aHL_FALSITY)).
fof(f290,plain,(
  $false | $spl12),
  inference(resolution,[],[f217,f112])).
fof(f112,plain,(
  ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),sK2))))),
  inference(cnf_transformation,[],[f72])).
fof(f257,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),sK2)))) ) | $spl20),
  inference(cnf_transformation,[],[f257_D])).
fof(f257_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),sK2)))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f241,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),sK2)))) )),
  inference(resolution,[],[f112,f180])).
fof(f230,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),sK2))) | s(bool,X21) != s(bool,X23)) )),
  inference(resolution,[],[f112,f124])).
fof(f228,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),sK2)))) )),
  inference(resolution,[],[f112,f123])).
fof(f258,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),sK2)))) ) | $spl22),
  inference(cnf_transformation,[],[f258_D])).
fof(f258_D,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),sK2)))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f214,plain,(
  ( ! [X4,X5] : (s(bool,X4) != s(bool,X5) | p(s(bool,X4)) | ~p(s(bool,X5))) ) | $spl6),
  inference(cnf_transformation,[],[f214_D])).
fof(f214_D,plain,(
  ( ! [X4,X5] : (s(bool,X4) != s(bool,X5) | p(s(bool,X4)) | ~p(s(bool,X5))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f227,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),sK2))) | s(bool,X12) != s(bool,X13)) )),
  inference(resolution,[],[f112,f122])).
fof(f226,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),sK2)))) )),
  inference(resolution,[],[f112,f122])).
fof(f225,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),sK2))) | s(bool,X6) != s(bool,X8)) )),
  inference(resolution,[],[f112,f120])).
fof(f224,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),sK2))) | s(bool,X3) != s(bool,X4)) )),
  inference(resolution,[],[f112,f118])).
fof(f223,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK4),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),sK2)))) )),
  inference(resolution,[],[f112,f118])).
fof(f252,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK1),sK2))) | $spl15),
  inference(cnf_transformation,[],[f252_D])).
fof(f252_D,plain,(
  s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK1),sK2))) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f256,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,'const.list.EL_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.list.list'(sK0),sK3))))),s(sK1,'const.list.EL_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.list.list'(sK1),sK2)))))) | $spl19),
  inference(cnf_transformation,[],[f256_D])).
fof(f256_D,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,'const.list.EL_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.list.list'(sK0),sK3))))),s(sK1,'const.list.EL_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.list.list'(sK1),sK2)))))) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f222,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK4),s(sK0,'const.list.EL_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.list.list'(sK0),sK3))))),s(sK1,'const.list.EL_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.list.list'(sK1),sK2)))))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK1),sK2)))),
  inference(resolution,[],[f112,f116])).
fof(f116,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4)))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,'const.list.EL_2'(s('type.num.num',sK6(X4,X3,X2,X1,X0)),s('type.list.list'(X0),X3))))),s(X1,'const.list.EL_2'(s('type.num.num',sK6(X4,X3,X2,X1,X0)),s('type.list.list'(X1),X4)))))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4)))) )),
  inference(cnf_transformation,[],[f76])).
fof(f254,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3)))))) | $spl16),
  inference(cnf_transformation,[],[f254_D])).
fof(f254_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3)))))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f221,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(sK2,sK3,sK4,sK1,sK0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3)))))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK0),sK3))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK1),sK2)))),
  inference(resolution,[],[f112,f115])).
fof(f115,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(X4,X3,X2,X1,X0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3)))))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X3))) != s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X4)))) )),
  inference(cnf_transformation,[],[f76])).
fof(f211,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),sK2)))) ) | $spl0),
  inference(cnf_transformation,[],[f211_D])).
fof(f211_D,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),sK2)))) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f202,plain,(
  ( ! [X26,X25] : (p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X26) != s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),sK2)))) )),
  inference(resolution,[],[f111,f180])).
fof(f192,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),sK2))) | s(bool,X23) != s(bool,X24)) )),
  inference(resolution,[],[f111,f124])).
fof(f213,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),sK2)))) ) | $spl4),
  inference(cnf_transformation,[],[f213_D])).
fof(f213_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),sK2)))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f191,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),sK2)))) )),
  inference(resolution,[],[f111,f124])).
fof(f189,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),sK2)))) )),
  inference(resolution,[],[f111,f123])).
fof(f188,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),sK2)))) )),
  inference(resolution,[],[f111,f122])).
fof(f187,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),sK2))) | s(bool,X8) != s(bool,X9)) )),
  inference(resolution,[],[f111,f120])).
fof(f186,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X5) | s(bool,X6) != s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),sK2)))) )),
  inference(resolution,[],[f111,f120])).
fof(f185,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK5),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),sK2)))) )),
  inference(resolution,[],[f111,f118])).
% SZS output end Proof for thm.list
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 102
% Passive clauses: 261
% Generated clauses: 2773
% Final active clauses: 91
% Final passive clauses: 124
% Input formulas: 26
% Initial clauses: 70
% 
% Selected by SInE selection: 26
% SInE iterations: 1
% 
% Duplicate literals: 77
% Trivial inequalities: 69
% Fw subsumption resolutions: 521
% Fw demodulations: 180
% Bw demodulations: 6
% 
% Simple tautologies: 136
% Equational tautologies: 1
% Forward subsumptions: 1460
% 
% Binary resolution: 401
% Forward superposition: 734
% Backward superposition: 743
% Equality resolution: 16
% 
% Split clauses: 134
% Split components: 288
% SAT solver clauses: 151
% SAT solver unit clauses: 14
% SAT solver binary clauses: 58
% 
% Sat splits: 134
% Sat splitting refutations: 16
% 
% Memory used [KB]: 2686
% Time elapsed: 0.074 s
% ------------------------------
% ------------------------------
% Success in time 11.796 s
