Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.quotient_sum
% (35039)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 167
% Passive clauses: 2146
% Generated clauses: 25779
% Final active clauses: 132
% Final passive clauses: 1688
% Input formulas: 37
% Initial clauses: 59
% 
% Duplicate literals: 1143
% Trivial inequalities: 517
% Fw subsumption resolutions: 3546
% Fw demodulations: 3678
% Bw demodulations: 30
% 
% Simple tautologies: 585
% Equational tautologies: 51
% Forward subsumptions: 13323
% 
% Binary resolution: 226
% Factoring: 2
% Forward superposition: 8790
% Backward superposition: 7318
% Self superposition: 205
% Equality factoring: 104
% Equality resolution: 47
% 
% Split clauses: 489
% Split components: 1005
% SAT solver clauses: 1032
% SAT solver unit clauses: 37
% SAT solver binary clauses: 264
% 
% Sat splits: 489
% Sat splitting refutations: 17
% 
% TWLsolver clauses: 1021
% TWLsolver calls for satisfiability: 149
% 
% Memory used [KB]: 11897
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 144 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.quotient_sum
% (35054)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 146
% Passive clauses: 1067
% Generated clauses: 20518
% Final active clauses: 117
% Final passive clauses: 699
% Input formulas: 37
% Initial clauses: 59
% 
% Duplicate literals: 572
% Trivial inequalities: 251
% Fw subsumption resolutions: 2099
% Fw demodulations: 3503
% Bw demodulations: 220
% Condensations: 747
% Global subsumptions: 9
% 
% Simple tautologies: 174
% Equational tautologies: 25
% Forward subsumptions: 10534
% Bw demodulations to eq. taut.: 16
% 
% Binary resolution: 277
% Factoring: 6
% Forward superposition: 6138
% Backward superposition: 5145
% Self superposition: 163
% Equality factoring: 53
% Equality resolution: 23
% 
% Split clauses: 70
% Split components: 149
% SAT solver clauses: 1613
% SAT solver unit clauses: 97
% SAT solver binary clauses: 42
% SAT solver learnt clauses: 18
% 
% Sat splits: 70
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 1577
% TWLsolver calls for satisfiability: 325
% 
% Memory used [KB]: 15735
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 137 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.quotient_sum
% (35064)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 508
% Passive clauses: 7201
% Generated clauses: 214128
% Final active clauses: 256
% Final passive clauses: 1054
% Input formulas: 37
% Initial clauses: 59
% Discarded non-redundant clauses: 26841
% 
% Duplicate literals: 5421
% Trivial inequalities: 2303
% Fw subsumption resolutions: 22255
% Fw demodulations: 45092
% Bw demodulations: 2329
% Condensations: 12471
% Global subsumptions: 93
% 
% Simple tautologies: 2201
% Equational tautologies: 591
% Forward subsumptions: 114611
% Fw demodulations to eq. taut.: 258
% Bw demodulations to eq. taut.: 330
% 
% Unit resulting resolution: 9490
% Factoring: 1
% Forward superposition: 54237
% Backward superposition: 55628
% Self superposition: 657
% Equality factoring: 293
% Equality resolution: 133
% 
% Split clauses: 281
% Split components: 572
% SAT solver clauses: 8021
% SAT solver unit clauses: 147
% SAT solver binary clauses: 999
% 
% Sat splits: 281
% Sat splitting refutations: 163
% 
% Memory used [KB]: 116160
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.quotient_sum
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.quotient_sum
% SZS output start Proof for thm.quotient_sum
fof(f35677,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f157,f158,f159,f160,f161,f162,f163,f164,f165,f166,f249,f253,f242,f243,f347,f348,f349,f345,f343,f250,f404,f634,f635,f649,f677,f715,f716,f743,f744,f1007,f2293,f2295,f2294,f2336,f2337,f2338,f2339,f2340,f2341,f2342,f2602,f2604,f2603,f2763,f2765,f2764,f3028,f3030,f3029,f4590,f4592,f4591,f4641,f4642,f4644,f4645,f4646,f4655,f8620,f8623,f8630,f8632,f8634,f8636,f8638,f8640,f8642,f8644,f8646,f8648,f8650,f8652,f8654,f8656,f8658,f8660,f8674,f8675,f8624,f8681,f8682,f8683,f8684,f8685,f8686,f8687,f8688,f8689,f8690,f8691,f8692,f8693,f8655,f8694,f8657,f8695,f8659,f8696,f8661,f9093,f9096,f9097,f9099,f9572,f9573,f9586,f9587,f10215,f10216,f10230,f10231,f10245,f10246,f10804,f10806,f10870,f10872,f11068,f11069,f11134,f11136,f11571,f11573,f11622,f11623,f11672,f11673,f11881,f11883,f11882,f11879,f11876,f11886,f11875,f11874,f11387,f11378,f11377,f11373,f11370,f11367,f11350,f11349,f11921,f11922,f11924,f11925,f11337,f11336,f11926,f11192,f11927,f11054,f11928,f10791,f11929,f10404,f10403,f10363,f10362,f11930,f9718,f9717,f9678,f9677,f11940,f8815,f8814,f11951,f8812,f11952,f11953,f11955,f8802,f11956,f8392,f8391,f11958,f8385,f8381,f8379,f8373,f8366,f8365,f8360,f8350,f8349,f8348,f8347,f8346,f8339,f8337,f8336,f11969,f11970,f11971,f11972,f8193,f8192,f11915,f11973,f11959,f8186,f8182,f11960,f8180,f11961,f8174,f11916,f8167,f8166,f11917,f8161,f11918,f8151,f11963,f8150,f11964,f8149,f11965,f8148,f11919,f8147,f11920,f8140,f8138,f8137,f11974,f11975,f11976,f11977,f11923,f7888,f7886,f7723,f7721,f7674,f7673,f7663,f7655,f7654,f7653,f7620,f11957,f7619,f8622,f7609,f11962,f7601,f11966,f7600,f11967,f7599,f254,f11968,f7454,f11979,f7439,f7436,f7423,f7421,f7417,f7405,f7402,f7344,f7341,f7057,f11986,f7042,f7039,f7026,f7024,f7020,f7008,f7005,f6948,f6945,f6670,f11987,f6655,f6652,f6639,f6637,f6633,f6621,f6618,f6562,f6559,f11988,f11989,f11954,f12003,f12015,f5477,f12018,f5467,f5465,f5450,f11981,f5448,f5446,f5437,f5435,f5408,f5406,f4567,f12020,f4566,f4565,f12021,f12023,f4553,f12024,f12027,f12029,f12030,f11982,f12031,f12032,f12019,f12034,f4502,f12038,f12040,f4489,f12041,f12042,f4381,f4380,f11984,f4320,f12047,f12048,f4337,f12052,f4334,f12053,f4332,f4330,f12062,f12063,f12064,f12065,f11985,f3719,f12046,f12067,f12068,f3693,f3602,f12070,f12071,f11983,f3022,f12076,f3018,f3017,f12077,f12079,f12080,f12081,f2753,f12035,f2749,f2748,f12082,f12084,f2596,f12085,f12086,f12087,f12054,f2286,f12069,f12096,f12098,f12097,f2058,f12100,f12101,f12103,f2062,f12104,f12105,f255,f12102,f12173,f12177,f10208,f1384,f429,f427,f426,f12183,f12209,f12210,f169,f12211,f12212,f170,f12213,f12214,f12215,f12216,f12217,f12218,f178,f12219,f12220,f172,f12223,f12221,f12222,f12257,f12258,f12278,f12279,f12394,f12395,f12543,f12544,f12545,f12546,f12547,f12548,f12549,f12550,f12551,f12552,f12573,f12574,f12575,f12576,f12577,f12578,f12579,f12175,f12642,f12606,f769,f12607,f771,f12615,f12616,f12617,f12618,f12640,f12620,f768,f12639,f12903,f12904,f12943,f12945,f12984,f12985,f13029,f13030,f13074,f13075,f13077,f13134,f13137,f13136,f13182,f13183,f773,f13482,f13700,f13701,f13740,f13742,f13782,f13784,f13823,f13825,f13869,f13870,f13894,f13913,f13915,f13917,f13918,f13919,f13921,f13924,f13925,f767,f13658,f13656,f13135,f13926,f14017,f14018,f14071,f14072,f171,f14073,f14521,f14565,f14567,f14601,f14603,f14650,f14652,f14710,f14712,f14724,f14726,f168,f14744,f14745,f14746,f14747,f14780,f14781,f14801,f14802,f251,f14844,f174,f14848,f14849,f12224,f176,f15237,f15239,f15238,f15246,f15498,f15501,f15511,f15514,f15517,f15521,f15524,f15527,f15532,f15535,f15538,f15541,f15592,f15595,f15615,f15624,f15633,f15636,f15645,f15652,f15655,f15658,f15757,f15760,f15770,f15773,f15776,f15780,f15783,f15786,f15788,f15793,f15796,f15799,f15802,f15826,f15868,f15870,f16226,f11887,f16593,f11888,f16638,f16640,f16673,f16669,f16664,f16663,f16728,f16729,f16730,f16731,f16732,f16778,f16780,f17165,f17166,f17210,f17212,f17368,f17372,f17383,f17386,f17388,f17389,f17391,f17395,f17415,f17934,f18005,f18007,f18101,f11885,f18144,f18146,f18197,f18199,f21365,f21367,f21371,f21377,f21379,f21383,f21410,f21411,f21414,f21419,f21420,f21423,f21743,f21744,f21745,f21747,f21748,f21750,f21751,f21786,f21788,f21792,f21793,f21794,f21795,f21798,f21806,f21813,f21867,f21869,f149,f21891,f12004,f138,f21890,f22049,f22077,f167,f22377,f22378,f22765,f22766,f8633,f22775,f22776,f8631,f22859,f22860,f11978,f22865,f22867,f22866,f23359,f23360,f8653,f23369,f23370,f8637,f23379,f23380,f8635,f23988,f23989,f24147,f24148,f21366,f24317,f24379,f24380,f21378,f24484,f24485,f21368,f24588,f24589,f21380,f25080,f25082,f25081,f25889,f25890,f24318,f26050,f26052,f26051,f26099,f26100,f11572,f26153,f26154,f8643,f26427,f26428,f26497,f26498,f10871,f26566,f26567,f21384,f26614,f26615,f8649,f26683,f26684,f10805,f26752,f26753,f11135,f26829,f26830,f21372,f26877,f26878,f8651,f27151,f27152,f12944,f27509,f27511,f27510,f27554,f27556,f252,f27555,f27566,f27568,f27567,f27951,f28107,f28109,f28011,f28493,f28652,f28654,f28555,f29076,f29077,f8639,f29195,f29197,f29304,f29306,f29305,f29293,f29292,f29443,f29444,f29196,f29495,f29496,f8641,f30058,f30059,f8645,f30180,f30182,f30297,f30299,f30298,f30286,f30285,f30442,f30443,f30181,f30495,f30496,f8647,f31089,f31090,f13741,f31432,f31433,f31494,f31495,f13824,f31556,f31557,f13783,f31682,f31684,f31825,f31827,f31904,f31905,f14566,f32043,f32045,f32044,f32829,f32831,f32830,f33073,f33074,f31683,f33273,f33274,f350,f31826,f33921,f33922,f34624,f34625,f17369,f34718,f34719,f14651,f34788,f34789,f15869,f34858,f34859,f16639,f34928,f34929,f11980,f16779,f34996,f34997,f17211,f35218,f35219,f18006,f35438,f35439,f18145,f35508,f35509,f18198,f8621,f35527,f35528,f35566,f35568,f35569,f35580,f35565,f35564,f35618,f35619,f35665,f35666])).
fof(f35666,plain,(
  $false | ($spl8 | $spl332)),
  inference(subsumption_resolution,[],[f35632,f173])).
fof(f173,plain,(
  s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | $spl8),
  inference(cnf_transformation,[],[f173_D])).
fof(f173_D,plain,(
  s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f35632,plain,(
  s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | $spl332),
  inference(superposition,[],[f5633,f35566])).
fof(f5633,plain,(
  ( ! [X76,X77,X75] : (s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X75,X76),'const.sum.INL_1'(s(X75,X77)))))) )),
  inference(subsumption_resolution,[],[f5575,f150])).
fof(f150,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.quotient_sum.ISR_RSP',aHL_FALSITY)).
fof(f5575,plain,(
  ( ! [X76,X77,X75] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X75,X76),'const.sum.INL_1'(s(X75,X77)))))) )),
  inference(superposition,[],[f122,f5433])).
fof(f5433,plain,(
  ( ! [X296,X298,X297] : (s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X296,X297),'const.sum.INL_1'(s(X296,X298)))))) )),
  inference(subsumption_resolution,[],[f5262,f129])).
fof(f129,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2))))))) )),
  inference(cnf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0,X1,X2] : p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2))))))),
  inference(flattening,[],[f56])).
fof(f56,plain,(
  ! [X0,X1] : ! [X2] : p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2))))))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X8,X9] : ! [X20] : p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X8,X9),'const.sum.INR_1'(s(X9,X20))))))),
  file('thm.quotient_sum.ISR_RSP','thm.sum.ISR|split|0')).
fof(f5262,plain,(
  ( ! [X296,X298,X297,X299] : (~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X299)))))) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X296,X297),'const.sum.INL_1'(s(X296,X298)))))) )),
  inference(superposition,[],[f128,f5092])).
fof(f5092,plain,(
  ( ! [X8,X7] : (s(bool,X7) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X8))))) | s(bool,'F_0') = s(bool,X7)) )),
  inference(resolution,[],[f4658,f133])).
fof(f133,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2))))))) )),
  inference(cnf_transformation,[],[f65])).
fof(f65,plain,(
  ! [X0,X1,X2] : ~p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2))))))),
  inference(flattening,[],[f64])).
fof(f64,plain,(
  ! [X0,X1] : ! [X2] : ~p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2))))))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X8,X9] : ! [X21] : ~p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X8,X9),'const.sum.INR_1'(s(X9,X21))))))),
  file('thm.quotient_sum.ISR_RSP','thm.sum.ISL|split|1')).
fof(f4658,plain,(
  ( ! [X6,X7] : (p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X7)))) | s(bool,X6) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X7))) | s(bool,'F_0') = s(bool,X6)) )),
  inference(resolution,[],[f4570,f119])).
fof(f119,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f92])).
fof(f92,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f46])).
fof(f46,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X22] : (s(bool,'F_0') = s(bool,X22) <=> ~p(s(bool,X22)))),
  file('thm.quotient_sum.ISR_RSP','thm.bool.EQ_CLAUSES|split|2')).
fof(f4570,plain,(
  ( ! [X45,X44] : (~p(s(bool,X45)) | s(bool,X45) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X44))) | p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X44))))) )),
  inference(trivial_inequality_removal,[],[f4564])).
fof(f4564,plain,(
  ( ! [X45,X44] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,X45) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X44))) | ~p(s(bool,X45)) | p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X44))))) )),
  inference(superposition,[],[f4132,f121])).
fof(f121,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f93])).
fof(f93,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f48])).
fof(f48,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X22] : (s(bool,'F_0') = s(bool,X22) <=> ~p(s(bool,X22)))),
  file('thm.quotient_sum.ISR_RSP','thm.bool.EQ_CLAUSES|split|3')).
fof(f4132,plain,(
  ( ! [X6,X5] : (s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X5))) | s(bool,X6) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X5))) | ~p(s(bool,X6))) )),
  inference(resolution,[],[f4105,f115])).
fof(f115,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f80])).
fof(f80,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f43])).
fof(f43,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X23,X24] : ((p(s(bool,X23)) => p(s(bool,X24))) => ((p(s(bool,X24)) => p(s(bool,X23))) => s(bool,X23) = s(bool,X24)))),
  file('thm.quotient_sum.ISR_RSP','thm.bool.IMP_ANTISYM_AX')).
fof(f4105,plain,(
  ( ! [X35] : (p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X35)))) | s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X35)))) )),
  inference(superposition,[],[f129,f3610])).
fof(f3610,plain,(
  ( ! [X11] : (s('type.sum.sum'(sK0,sK0),X11) = s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,sK13(X11,sK0,sK0)))) | s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X11)))) )),
  inference(subsumption_resolution,[],[f3587,f1592])).
fof(f1592,plain,(
  ( ! [X25] : (s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,X25)))))) )),
  inference(subsumption_resolution,[],[f1555,f128])).
fof(f1555,plain,(
  ( ! [X25] : (s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,X25))))) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK2),'const.sum.INL_1'(s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,X25))))))))) )),
  inference(superposition,[],[f122,f1442])).
fof(f1442,plain,(
  ( ! [X0] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,X0))))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK2),'const.sum.INL_1'(s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,X0)))))))) )),
  inference(superposition,[],[f1431,f131])).
fof(f131,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.sum.sum'(X2,X1),'const.sum.INL_1'(s(X2,happ(s(fun(X3,X2),X4),s(X3,X6))))) = s('type.sum.sum'(X2,X1),'const.sum.++_3'(s(fun(X3,X2),X4),s(fun(X0,X1),X5),s('type.sum.sum'(X3,X0),'const.sum.INL_1'(s(X3,X6)))))) )),
  inference(cnf_transformation,[],[f61])).
fof(f61,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s('type.sum.sum'(X2,X1),'const.sum.INL_1'(s(X2,happ(s(fun(X3,X2),X4),s(X3,X6))))) = s('type.sum.sum'(X2,X1),'const.sum.++_3'(s(fun(X3,X2),X4),s(fun(X0,X1),X5),s('type.sum.sum'(X3,X0),'const.sum.INL_1'(s(X3,X6)))))),
  inference(flattening,[],[f60])).
fof(f60,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6] : s('type.sum.sum'(X2,X1),'const.sum.INL_1'(s(X2,happ(s(fun(X3,X2),X4),s(X3,X6))))) = s('type.sum.sum'(X2,X1),'const.sum.++_3'(s(fun(X3,X2),X4),s(fun(X0,X1),X5),s('type.sum.sum'(X3,X0),'const.sum.INL_1'(s(X3,X6)))))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X9,X7,X6,X8] : ! [X25,X26,X18] : s('type.sum.sum'(X6,X7),'const.sum.INL_1'(s(X6,happ(s(fun(X8,X6),X25),s(X8,X18))))) = s('type.sum.sum'(X6,X7),'const.sum.++_3'(s(fun(X8,X6),X25),s(fun(X9,X7),X26),s('type.sum.sum'(X8,X9),'const.sum.INL_1'(s(X8,X18)))))),
  file('thm.quotient_sum.ISR_RSP','thm.sum.SUM_MAP_def|split|0')).
fof(f1431,plain,(
  ( ! [X0] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X0))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK2),'const.sum.++_3'(s(fun(sK0,sK2),sK6),s(fun(sK0,sK2),sK6),s('type.sum.sum'(sK0,sK0),X0)))))) )),
  inference(subsumption_resolution,[],[f1417,f155])).
fof(f155,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f19])).
fof(f19,axiom,(
  p(s(bool,'T_0'))),
  file('thm.quotient_sum.ISR_RSP','thm.bool.TRUTH')).
fof(f1417,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X0))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK2),'const.sum.++_3'(s(fun(sK0,sK2),sK6),s(fun(sK0,sK2),sK6),s('type.sum.sum'(sK0,sK0),X0)))))) )),
  inference(superposition,[],[f181,f187])).
fof(f187,plain,(
  s(bool,'T_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(sK2,sK0),sK5),s(fun(sK0,sK2),sK6)))),
  inference(resolution,[],[f107,f123])).
fof(f123,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f94])).
fof(f94,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X22] : (s(bool,'T_0') = s(bool,X22) <=> p(s(bool,X22)))),
  file('thm.quotient_sum.ISR_RSP','thm.bool.EQ_CLAUSES|split|0')).
fof(f107,plain,(
  p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(sK2,sK0),sK5),s(fun(sK0,sK2),sK6))))),
  inference(cnf_transformation,[],[f89])).
fof(f89,plain,(
  p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(sK2,sK0),sK5),s(fun(sK0,sK2),sK6)))) & (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK3,fun(sK3,bool)),sK7),s(fun(sK3,sK1),sK8),s(fun(sK1,sK3),sK9)))) & (p(s(bool,'const.quotient_sum.+++_4'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(sK3,fun(sK3,bool)),sK7),s('type.sum.sum'(sK2,sK3),sK10),s('type.sum.sum'(sK2,sK3),sK11)))) & s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK9,sK10,sK11])],[f79])).
fof(f79,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X2,fun(X2,bool)),X4),s(fun(X2,X0),X5),s(fun(X0,X2),X6)))) & ? [X7,X8,X9] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X3,fun(X3,bool)),X7),s(fun(X3,X1),X8),s(fun(X1,X3),X9)))) & ? [X10,X11] : (p(s(bool,'const.quotient_sum.+++_4'(s(fun(X2,fun(X2,bool)),X4),s(fun(X3,fun(X3,bool)),X7),s('type.sum.sum'(X2,X3),X10),s('type.sum.sum'(X2,X3),X11)))) & s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X2,X3),X10))) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X2,X3),X11))))))),
  inference(ennf_transformation,[],[f40])).
fof(f40,plain,(
  ~! [X0,X1,X2,X3,X4,X5,X6] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X2,fun(X2,bool)),X4),s(fun(X2,X0),X5),s(fun(X0,X2),X6)))) => ! [X7,X8,X9] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X3,fun(X3,bool)),X7),s(fun(X3,X1),X8),s(fun(X1,X3),X9)))) => ! [X10,X11] : (p(s(bool,'const.quotient_sum.+++_4'(s(fun(X2,fun(X2,bool)),X4),s(fun(X3,fun(X3,bool)),X7),s('type.sum.sum'(X2,X3),X10),s('type.sum.sum'(X2,X3),X11)))) => s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X2,X3),X10))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X2,X3),X11))))))),
  inference(flattening,[],[f39])).
fof(f39,plain,(
  ~! [X0,X1,X2,X3] : ! [X4,X5,X6] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X2,fun(X2,bool)),X4),s(fun(X2,X0),X5),s(fun(X0,X2),X6)))) => ! [X7,X8,X9] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X3,fun(X3,bool)),X7),s(fun(X3,X1),X8),s(fun(X1,X3),X9)))) => ! [X10,X11] : (p(s(bool,'const.quotient_sum.+++_4'(s(fun(X2,fun(X2,bool)),X4),s(fun(X3,fun(X3,bool)),X7),s('type.sum.sum'(X2,X3),X10),s('type.sum.sum'(X2,X3),X11)))) => s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X2,X3),X10))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X2,X3),X11))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8,X9] : ! [X10,X11,X12] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X8,fun(X8,bool)),X10),s(fun(X8,X6),X11),s(fun(X6,X8),X12)))) => ! [X13,X14,X15] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X9,fun(X9,bool)),X13),s(fun(X9,X7),X14),s(fun(X7,X9),X15)))) => ! [X16,X17] : (p(s(bool,'const.quotient_sum.+++_4'(s(fun(X8,fun(X8,bool)),X10),s(fun(X9,fun(X9,bool)),X13),s('type.sum.sum'(X8,X9),X16),s('type.sum.sum'(X8,X9),X17)))) => s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X8,X9),X16))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X8,X9),X17))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8,X9] : ! [X10,X11,X12] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X8,fun(X8,bool)),X10),s(fun(X8,X6),X11),s(fun(X6,X8),X12)))) => ! [X13,X14,X15] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X9,fun(X9,bool)),X13),s(fun(X9,X7),X14),s(fun(X7,X9),X15)))) => ! [X16,X17] : (p(s(bool,'const.quotient_sum.+++_4'(s(fun(X8,fun(X8,bool)),X10),s(fun(X9,fun(X9,bool)),X13),s('type.sum.sum'(X8,X9),X16),s('type.sum.sum'(X8,X9),X17)))) => s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X8,X9),X16))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X8,X9),X17))))))),
  file('thm.quotient_sum.ISR_RSP','thm.quotient_sum.ISR_RSP')).
fof(f181,plain,(
  ( ! [X14,X19,X17,X15,X18,X16] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X16,fun(X16,bool)),X18),s(fun(X16,X14),X19),s(fun(X14,X16),X17)))) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,X14),X15))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,X16),'const.sum.++_3'(s(fun(sK0,sK2),sK6),s(fun(X14,X16),X17),s('type.sum.sum'(sK0,X14),X15)))))) )),
  inference(resolution,[],[f107,f127])).
fof(f127,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X2,X3),X10))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X0,X1),'const.sum.++_3'(s(fun(X2,X0),X6),s(fun(X3,X1),X9),s('type.sum.sum'(X2,X3),X10))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X1,fun(X1,bool)),X7),s(fun(X1,X3),X8),s(fun(X3,X1),X9)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X0,fun(X0,bool)),X4),s(fun(X0,X2),X5),s(fun(X2,X0),X6))))) )),
  inference(cnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X0,fun(X0,bool)),X4),s(fun(X0,X2),X5),s(fun(X2,X0),X6)))) | ! [X7,X8,X9] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X1,fun(X1,bool)),X7),s(fun(X1,X3),X8),s(fun(X3,X1),X9)))) | ! [X10] : s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X2,X3),X10))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X0,X1),'const.sum.++_3'(s(fun(X2,X0),X6),s(fun(X3,X1),X9),s('type.sum.sum'(X2,X3),X10)))))))),
  inference(ennf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X0,fun(X0,bool)),X4),s(fun(X0,X2),X5),s(fun(X2,X0),X6)))) => ! [X7,X8,X9] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X1,fun(X1,bool)),X7),s(fun(X1,X3),X8),s(fun(X3,X1),X9)))) => ! [X10] : s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X2,X3),X10))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X0,X1),'const.sum.++_3'(s(fun(X2,X0),X6),s(fun(X3,X1),X9),s('type.sum.sum'(X2,X3),X10)))))))),
  inference(flattening,[],[f52])).
fof(f52,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X0,fun(X0,bool)),X4),s(fun(X0,X2),X5),s(fun(X2,X0),X6)))) => ! [X7,X8,X9] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X1,fun(X1,bool)),X7),s(fun(X1,X3),X8),s(fun(X3,X1),X9)))) => ! [X10] : s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X2,X3),X10))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X0,X1),'const.sum.++_3'(s(fun(X2,X0),X6),s(fun(X3,X1),X9),s('type.sum.sum'(X2,X3),X10)))))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X8,X9,X6,X7] : ! [X10,X11,X12] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X8,fun(X8,bool)),X10),s(fun(X8,X6),X11),s(fun(X6,X8),X12)))) => ! [X13,X14,X15] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X9,fun(X9,bool)),X13),s(fun(X9,X7),X14),s(fun(X7,X9),X15)))) => ! [X18] : s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X6,X7),X18))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X8,X9),'const.sum.++_3'(s(fun(X6,X8),X12),s(fun(X7,X9),X15),s('type.sum.sum'(X6,X7),X18)))))))),
  file('thm.quotient_sum.ISR_RSP','thm.quotient_sum.ISR_PRS')).
fof(f3587,plain,(
  ( ! [X11] : (s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X11))) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,sK12(X11,sK0,sK0)))))) | s('type.sum.sum'(sK0,sK0),X11) = s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,sK13(X11,sK0,sK0))))) )),
  inference(superposition,[],[f3541,f111])).
fof(f111,plain,(
  ( ! [X2,X0,X1] : (s('type.sum.sum'(X0,X1),X2) = s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,sK13(X2,X1,X0)))) | s('type.sum.sum'(X0,X1),X2) = s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,sK12(X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f90])).
fof(f90,plain,(
  ! [X0,X1,X2] : (s('type.sum.sum'(X0,X1),X2) = s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,sK12(X2,X1,X0)))) | s('type.sum.sum'(X0,X1),X2) = s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,sK13(X2,X1,X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13])],[f42])).
fof(f42,plain,(
  ! [X0,X1,X2] : (? [X4] : s('type.sum.sum'(X0,X1),X2) = s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,X4))) | ? [X3] : s('type.sum.sum'(X0,X1),X2) = s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X3))))),
  inference(flattening,[],[f41])).
fof(f41,plain,(
  ! [X0,X1] : ! [X2] : (? [X4] : s('type.sum.sum'(X0,X1),X2) = s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,X4))) | ? [X3] : s('type.sum.sum'(X0,X1),X2) = s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X3))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X8,X9] : ! [X19] : (? [X20] : s('type.sum.sum'(X8,X9),X19) = s('type.sum.sum'(X8,X9),'const.sum.INL_1'(s(X8,X20))) | ? [X21] : s('type.sum.sum'(X8,X9),X19) = s('type.sum.sum'(X8,X9),'const.sum.INR_1'(s(X9,X21))))),
  file('thm.quotient_sum.ISR_RSP','thm.sum.sum_CASES')).
fof(f3541,plain,(
  ( ! [X9] : (s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X9))) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X9)))) )),
  inference(resolution,[],[f3517,f118])).
fof(f118,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f92])).
fof(f3517,plain,(
  ( ! [X22] : (p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X22)))) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X22)))) )),
  inference(superposition,[],[f134,f3227])).
fof(f3227,plain,(
  ( ! [X2] : (s('type.sum.sum'(sK0,sK0),X2) = s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,sK12(X2,sK0,sK0)))) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X2)))) )),
  inference(superposition,[],[f3205,f111])).
fof(f3205,plain,(
  ( ! [X11] : (s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X11)))))) )),
  inference(subsumption_resolution,[],[f3157,f129])).
fof(f3157,plain,(
  ( ! [X11] : (s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X11))))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK2),'const.sum.INR_1'(s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,X11))))))))) )),
  inference(superposition,[],[f1443,f123])).
fof(f1443,plain,(
  ( ! [X1] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X1))))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK2),'const.sum.INR_1'(s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,X1)))))))) )),
  inference(superposition,[],[f1431,f132])).
fof(f132,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s('type.sum.sum'(X1,X2),'const.sum.INR_1'(s(X2,happ(s(fun(X3,X2),X5),s(X3,X6))))) = s('type.sum.sum'(X1,X2),'const.sum.++_3'(s(fun(X0,X1),X4),s(fun(X3,X2),X5),s('type.sum.sum'(X0,X3),'const.sum.INR_1'(s(X3,X6)))))) )),
  inference(cnf_transformation,[],[f63])).
fof(f63,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s('type.sum.sum'(X1,X2),'const.sum.INR_1'(s(X2,happ(s(fun(X3,X2),X5),s(X3,X6))))) = s('type.sum.sum'(X1,X2),'const.sum.++_3'(s(fun(X0,X1),X4),s(fun(X3,X2),X5),s('type.sum.sum'(X0,X3),'const.sum.INR_1'(s(X3,X6)))))),
  inference(flattening,[],[f62])).
fof(f62,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6] : s('type.sum.sum'(X1,X2),'const.sum.INR_1'(s(X2,happ(s(fun(X3,X2),X5),s(X3,X6))))) = s('type.sum.sum'(X1,X2),'const.sum.++_3'(s(fun(X0,X1),X4),s(fun(X3,X2),X5),s('type.sum.sum'(X0,X3),'const.sum.INR_1'(s(X3,X6)))))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X8,X6,X7,X9] : ! [X25,X26,X27] : s('type.sum.sum'(X6,X7),'const.sum.INR_1'(s(X7,happ(s(fun(X9,X7),X26),s(X9,X27))))) = s('type.sum.sum'(X6,X7),'const.sum.++_3'(s(fun(X8,X6),X25),s(fun(X9,X7),X26),s('type.sum.sum'(X8,X9),'const.sum.INR_1'(s(X9,X27)))))),
  file('thm.quotient_sum.ISR_RSP','thm.sum.SUM_MAP_def|split|1')).
fof(f134,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2))))))) )),
  inference(cnf_transformation,[],[f67])).
fof(f67,plain,(
  ! [X0,X1,X2] : p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2))))))),
  inference(flattening,[],[f66])).
fof(f66,plain,(
  ! [X0,X1] : ! [X2] : p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2))))))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X9,X8] : ! [X20] : p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X8,X9),'const.sum.INL_1'(s(X8,X20))))))),
  file('thm.quotient_sum.ISR_RSP','thm.sum.ISL|split|0')).
fof(f128,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2))))))) )),
  inference(cnf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0,X1,X2] : ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2))))))),
  inference(flattening,[],[f54])).
fof(f54,plain,(
  ! [X0,X1] : ! [X2] : ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2))))))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X9,X8] : ! [X21] : ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X8,X9),'const.sum.INL_1'(s(X8,X21))))))),
  file('thm.quotient_sum.ISR_RSP','thm.sum.ISR|split|1')).
fof(f122,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f94])).
fof(f35665,plain,(
  $false | ($spl8 | $spl332)),
  inference(subsumption_resolution,[],[f35664,f1128])).
fof(f1128,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(backward_demodulation,[],[f187,f185])).
fof(f185,plain,(
  s(bool,'F_0') != s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(sK2,sK0),sK5),s(fun(sK0,sK2),sK6)))),
  inference(resolution,[],[f107,f118])).
fof(f35664,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl8 | $spl332)),
  inference(forward_demodulation,[],[f35631,f173])).
fof(f35631,plain,(
  s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | $spl332),
  inference(superposition,[],[f5433,f35566])).
fof(f35619,plain,(
  $false | ($spl26 | $spl32 | $spl333)),
  inference(subsumption_resolution,[],[f35585,f1035])).
fof(f1035,plain,(
  s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) | ($spl26 | $spl32)),
  inference(forward_demodulation,[],[f429,f771])).
fof(f35585,plain,(
  s(bool,'T_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) | $spl333),
  inference(superposition,[],[f5815,f35582])).
fof(f35582,plain,(
  s('type.sum.sum'(sK2,sK3),sK10) = s('type.sum.sum'(sK2,sK3),'const.sum.INR_1'(s(sK3,sK13(sK10,sK3,sK2)))) | $spl333),
  inference(trivial_inequality_removal,[],[f35581])).
fof(f35581,plain,(
  s('type.sum.sum'(sK2,sK3),sK10) != s('type.sum.sum'(sK2,sK3),sK10) | s('type.sum.sum'(sK2,sK3),sK10) = s('type.sum.sum'(sK2,sK3),'const.sum.INR_1'(s(sK3,sK13(sK10,sK3,sK2)))) | $spl333),
  inference(superposition,[],[f35565,f111])).
fof(f5815,plain,(
  ( ! [X105,X106,X104] : (s(bool,'T_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X104,X105),'const.sum.INR_1'(s(X105,X106)))))) )),
  inference(subsumption_resolution,[],[f5757,f150])).
fof(f5757,plain,(
  ( ! [X105,X106,X104] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X104,X105),'const.sum.INR_1'(s(X105,X106)))))) )),
  inference(superposition,[],[f122,f5444])).
fof(f5444,plain,(
  ( ! [X377,X376,X375] : (s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X375,X376),'const.sum.INR_1'(s(X376,X377)))))) )),
  inference(subsumption_resolution,[],[f5286,f129])).
fof(f5286,plain,(
  ( ! [X377,X376,X378,X375] : (~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X378)))))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X375,X376),'const.sum.INR_1'(s(X376,X377)))))) )),
  inference(superposition,[],[f133,f5092])).
fof(f35618,plain,(
  $false | ($spl26 | $spl32 | $spl333)),
  inference(subsumption_resolution,[],[f35617,f1128])).
fof(f35617,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl26 | $spl32 | $spl333)),
  inference(forward_demodulation,[],[f35584,f1035])).
fof(f35584,plain,(
  s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) | $spl333),
  inference(superposition,[],[f5444,f35582])).
fof(f35564,plain,(
  s('type.sum.sum'(sK2,sK3),'const.sum.INL_1'(s(sK2,sK12(sK10,sK3,sK2)))) = s('type.sum.sum'(sK2,sK3),'const.sum.INL_1'(s(sK2,sK12('const.sum.INL_1'(s(sK2,sK12(sK10,sK3,sK2))),sK3,sK2)))) | $spl330),
  inference(cnf_transformation,[],[f35564_D])).
fof(f35564_D,plain,(
  s('type.sum.sum'(sK2,sK3),'const.sum.INL_1'(s(sK2,sK12(sK10,sK3,sK2)))) = s('type.sum.sum'(sK2,sK3),'const.sum.INL_1'(s(sK2,sK12('const.sum.INL_1'(s(sK2,sK12(sK10,sK3,sK2))),sK3,sK2)))) <=> ~$spl330),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl330])])).
fof(f35565,plain,(
  s('type.sum.sum'(sK2,sK3),sK10) != s('type.sum.sum'(sK2,sK3),'const.sum.INL_1'(s(sK2,sK12(sK10,sK3,sK2)))) | $spl333),
  inference(cnf_transformation,[],[f35565_D])).
fof(f35565_D,plain,(
  s('type.sum.sum'(sK2,sK3),sK10) != s('type.sum.sum'(sK2,sK3),'const.sum.INL_1'(s(sK2,sK12(sK10,sK3,sK2)))) <=> ~$spl333),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl333])])).
fof(f35580,plain,(
  s('type.sum.sum'(sK2,sK3),sK10) != s('type.sum.sum'(sK2,sK3),'const.sum.INL_1'(s(sK2,sK12(sK10,sK3,sK2)))) | s('type.sum.sum'(sK2,sK3),'const.sum.INL_1'(s(sK2,sK12(sK10,sK3,sK2)))) = s('type.sum.sum'(sK2,sK3),'const.sum.INL_1'(s(sK2,sK12('const.sum.INL_1'(s(sK2,sK12(sK10,sK3,sK2))),sK3,sK2)))) | $spl227),
  inference(superposition,[],[f12176,f111])).
fof(f12176,plain,(
  s('type.sum.sum'(sK2,sK3),sK10) != s('type.sum.sum'(sK2,sK3),'const.sum.INR_1'(s(sK3,sK13('const.sum.INL_1'(s(sK2,sK12(sK10,sK3,sK2))),sK3,sK2)))) | $spl227),
  inference(cnf_transformation,[],[f12176_D])).
fof(f12176_D,plain,(
  s('type.sum.sum'(sK2,sK3),sK10) != s('type.sum.sum'(sK2,sK3),'const.sum.INR_1'(s(sK3,sK13('const.sum.INL_1'(s(sK2,sK12(sK10,sK3,sK2))),sK3,sK2)))) <=> ~$spl227),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl227])])).
fof(f35569,plain,(
  $false | ($spl26 | $spl32 | $spl226)),
  inference(subsumption_resolution,[],[f35531,f1035])).
fof(f35531,plain,(
  s(bool,'T_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) | $spl226),
  inference(superposition,[],[f5815,f12177])).
fof(f35568,plain,(
  $false | ($spl26 | $spl32 | $spl226)),
  inference(subsumption_resolution,[],[f35567,f1128])).
fof(f35567,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl26 | $spl32 | $spl226)),
  inference(forward_demodulation,[],[f35530,f1035])).
fof(f35530,plain,(
  s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) | $spl226),
  inference(superposition,[],[f5444,f12177])).
fof(f35566,plain,(
  s('type.sum.sum'(sK2,sK3),sK10) = s('type.sum.sum'(sK2,sK3),'const.sum.INL_1'(s(sK2,sK12(sK10,sK3,sK2)))) | $spl332),
  inference(cnf_transformation,[],[f35566_D])).
fof(f35566_D,plain,(
  s('type.sum.sum'(sK2,sK3),sK10) = s('type.sum.sum'(sK2,sK3),'const.sum.INL_1'(s(sK2,sK12(sK10,sK3,sK2)))) <=> ~$spl332),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl332])])).
fof(f35528,plain,(
  s('type.sum.sum'(sK2,sK3),sK10) = s('type.sum.sum'(sK2,sK3),'const.sum.INL_1'(s(sK2,sK12(sK10,sK3,sK2)))) | s('type.sum.sum'(sK2,sK3),'const.sum.INL_1'(s(sK2,sK12(sK10,sK3,sK2)))) = s('type.sum.sum'(sK2,sK3),'const.sum.INL_1'(s(sK2,sK12('const.sum.INL_1'(s(sK2,sK12(sK10,sK3,sK2))),sK3,sK2)))) | $spl226),
  inference(superposition,[],[f111,f12177])).
fof(f35527,plain,(
  s('type.sum.sum'(sK2,sK3),sK10) = s('type.sum.sum'(sK2,sK3),'const.sum.INL_1'(s(sK2,sK12(sK10,sK3,sK2)))) | s('type.sum.sum'(sK2,sK3),'const.sum.INL_1'(s(sK2,sK12(sK10,sK3,sK2)))) = s('type.sum.sum'(sK2,sK3),'const.sum.INL_1'(s(sK2,sK12('const.sum.INL_1'(s(sK2,sK12(sK10,sK3,sK2))),sK3,sK2)))) | $spl226),
  inference(superposition,[],[f12177,f111])).
fof(f8621,plain,(
  ( ! [X81] : (s(bool,'F_0') = s(bool,X81)) ) | $spl46),
  inference(cnf_transformation,[],[f8621_D])).
fof(f8621_D,plain,(
  ( ! [X81] : (s(bool,'F_0') = s(bool,X81)) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f18198,plain,(
  ( ! [X165,X163,X173,X171,X169,X166,X164,X162,X172,X170,X168] : (s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X162,X168),'const.sum.++_3'(s(fun(X164,X162),X166),s(fun(X169,X168),X170),s('type.sum.sum'(X164,X169),X171))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X162,fun(X162,bool)),X163),s(fun(X162,X164),X165),s(fun(X164,X162),X166))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X164,X169),X171)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X168,fun(X168,bool)),X172),s(fun(X168,X169),X173),s(fun(X169,X168),X170))))) ) | $spl272),
  inference(cnf_transformation,[],[f18198_D])).
fof(f18198_D,plain,(
  ( ! [X165,X163,X173,X171,X169,X166,X164,X162,X172,X170,X168] : (s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X162,X168),'const.sum.++_3'(s(fun(X164,X162),X166),s(fun(X169,X168),X170),s('type.sum.sum'(X164,X169),X171))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X162,fun(X162,bool)),X163),s(fun(X162,X164),X165),s(fun(X164,X162),X166))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X164,X169),X171)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X168,fun(X168,bool)),X172),s(fun(X168,X169),X173),s(fun(X169,X168),X170))))) ) <=> ~$spl272),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl272])])).
fof(f35509,plain,(
  ( ! [X158,X156,X154,X152,X161,X151,X159,X157,X155,X153,X162,X160] : (s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X151,X157),'const.sum.++_3'(s(fun(X153,X151),X155),s(fun(X158,X157),X159),s('type.sum.sum'(X153,X158),X160))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X157,fun(X157,bool)),X161),s(fun(X157,X158),X162),s(fun(X158,X157),X159)))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X153,X158),X160)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X151,fun(X151,bool)),X152),s(fun(X151,X153),X154),s(fun(X153,X151),X155))) | s(bool,'F_0') = s(bool,X156)) )),
  inference(subsumption_resolution,[],[f35482,f119])).
fof(f35482,plain,(
  ( ! [X158,X156,X154,X152,X161,X151,X159,X157,X155,X153,X162,X160] : (~p(s(bool,X156)) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X151,X157),'const.sum.++_3'(s(fun(X153,X151),X155),s(fun(X158,X157),X159),s('type.sum.sum'(X153,X158),X160))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X157,fun(X157,bool)),X161),s(fun(X157,X158),X162),s(fun(X158,X157),X159)))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X153,X158),X160)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X151,fun(X151,bool)),X152),s(fun(X151,X153),X154),s(fun(X153,X151),X155))) | s(bool,'F_0') = s(bool,X156)) )),
  inference(superposition,[],[f9521,f5187])).
fof(f5187,plain,(
  ( ! [X2,X1] : (s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,'F_0') = s(bool,X1)) )),
  inference(superposition,[],[f5092,f5092])).
fof(f9521,plain,(
  ( ! [X231,X229,X227,X235,X233,X230,X228,X226,X236,X234,X232] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X226,fun(X226,bool)),X235),s(fun(X226,X228),X236),s(fun(X228,X226),X229)))) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X226,X227),'const.sum.++_3'(s(fun(X228,X226),X229),s(fun(X230,X227),X231),s('type.sum.sum'(X228,X230),X232))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X227,fun(X227,bool)),X233),s(fun(X227,X230),X234),s(fun(X230,X227),X231)))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X228,X230),X232))))) )),
  inference(superposition,[],[f123,f127])).
fof(f35508,plain,(
  ( ! [X146,X144,X142,X140,X149,X147,X145,X143,X141,X139,X150,X148] : (s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X139,X145),'const.sum.++_3'(s(fun(X141,X139),X143),s(fun(X146,X145),X147),s('type.sum.sum'(X141,X146),X148))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X145,fun(X145,bool)),X149),s(fun(X145,X146),X150),s(fun(X146,X145),X147)))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X141,X146),X148)))) | s(bool,'F_0') = s(bool,X144) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X139,fun(X139,bool)),X140),s(fun(X139,X141),X142),s(fun(X141,X139),X143)))) )),
  inference(subsumption_resolution,[],[f35481,f119])).
fof(f35481,plain,(
  ( ! [X146,X144,X142,X140,X149,X147,X145,X143,X141,X139,X150,X148] : (~p(s(bool,X144)) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X139,X145),'const.sum.++_3'(s(fun(X141,X139),X143),s(fun(X146,X145),X147),s('type.sum.sum'(X141,X146),X148))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X145,fun(X145,bool)),X149),s(fun(X145,X146),X150),s(fun(X146,X145),X147)))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X141,X146),X148)))) | s(bool,'F_0') = s(bool,X144) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X139,fun(X139,bool)),X140),s(fun(X139,X141),X142),s(fun(X141,X139),X143)))) )),
  inference(superposition,[],[f9521,f5187])).
fof(f18145,plain,(
  ( ! [X165,X163,X173,X171,X169,X166,X164,X162,X172,X170,X168] : (s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X162,X168),'const.sum.++_3'(s(fun(X164,X162),X166),s(fun(X169,X168),X170),s('type.sum.sum'(X164,X169),X171))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X162,fun(X162,bool)),X163),s(fun(X162,X164),X165),s(fun(X164,X162),X166))) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X164,X169),X171)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X168,fun(X168,bool)),X172),s(fun(X168,X169),X173),s(fun(X169,X168),X170))))) ) | $spl270),
  inference(cnf_transformation,[],[f18145_D])).
fof(f18145_D,plain,(
  ( ! [X165,X163,X173,X171,X169,X166,X164,X162,X172,X170,X168] : (s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X162,X168),'const.sum.++_3'(s(fun(X164,X162),X166),s(fun(X169,X168),X170),s('type.sum.sum'(X164,X169),X171))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X162,fun(X162,bool)),X163),s(fun(X162,X164),X165),s(fun(X164,X162),X166))) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X164,X169),X171)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X168,fun(X168,bool)),X172),s(fun(X168,X169),X173),s(fun(X169,X168),X170))))) ) <=> ~$spl270),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl270])])).
fof(f35439,plain,(
  ( ! [X158,X156,X154,X152,X161,X151,X159,X157,X155,X153,X162,X160] : (s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X151,X157),'const.sum.++_3'(s(fun(X153,X151),X155),s(fun(X158,X157),X159),s('type.sum.sum'(X153,X158),X160))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X157,fun(X157,bool)),X161),s(fun(X157,X158),X162),s(fun(X158,X157),X159)))) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X153,X158),X160)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X151,fun(X151,bool)),X152),s(fun(X151,X153),X154),s(fun(X153,X151),X155))) | s(bool,'F_0') = s(bool,X156)) )),
  inference(subsumption_resolution,[],[f35414,f119])).
fof(f35414,plain,(
  ( ! [X158,X156,X154,X152,X161,X151,X159,X157,X155,X153,X162,X160] : (~p(s(bool,X156)) | s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X151,X157),'const.sum.++_3'(s(fun(X153,X151),X155),s(fun(X158,X157),X159),s('type.sum.sum'(X153,X158),X160))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X157,fun(X157,bool)),X161),s(fun(X157,X158),X162),s(fun(X158,X157),X159)))) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X153,X158),X160)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X151,fun(X151,bool)),X152),s(fun(X151,X153),X154),s(fun(X153,X151),X155))) | s(bool,'F_0') = s(bool,X156)) )),
  inference(superposition,[],[f9519,f5187])).
fof(f9519,plain,(
  ( ! [X212,X210,X208,X206,X204,X213,X211,X209,X207,X205,X214] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X204,fun(X204,bool)),X213),s(fun(X204,X206),X214),s(fun(X206,X204),X207)))) | s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X204,X205),'const.sum.++_3'(s(fun(X206,X204),X207),s(fun(X208,X205),X209),s('type.sum.sum'(X206,X208),X210))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X205,fun(X205,bool)),X211),s(fun(X205,X208),X212),s(fun(X208,X205),X209)))) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X206,X208),X210))))) )),
  inference(superposition,[],[f122,f127])).
fof(f35438,plain,(
  ( ! [X146,X144,X142,X140,X149,X147,X145,X143,X141,X139,X150,X148] : (s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X139,X145),'const.sum.++_3'(s(fun(X141,X139),X143),s(fun(X146,X145),X147),s('type.sum.sum'(X141,X146),X148))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X145,fun(X145,bool)),X149),s(fun(X145,X146),X150),s(fun(X146,X145),X147)))) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X141,X146),X148)))) | s(bool,'F_0') = s(bool,X144) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X139,fun(X139,bool)),X140),s(fun(X139,X141),X142),s(fun(X141,X139),X143)))) )),
  inference(subsumption_resolution,[],[f35413,f119])).
fof(f35413,plain,(
  ( ! [X146,X144,X142,X140,X149,X147,X145,X143,X141,X139,X150,X148] : (~p(s(bool,X144)) | s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X139,X145),'const.sum.++_3'(s(fun(X141,X139),X143),s(fun(X146,X145),X147),s('type.sum.sum'(X141,X146),X148))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X145,fun(X145,bool)),X149),s(fun(X145,X146),X150),s(fun(X146,X145),X147)))) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X141,X146),X148)))) | s(bool,'F_0') = s(bool,X144) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X139,fun(X139,bool)),X140),s(fun(X139,X141),X142),s(fun(X141,X139),X143)))) )),
  inference(superposition,[],[f9519,f5187])).
fof(f18006,plain,(
  ( ! [X165,X163,X173,X171,X169,X166,X164,X162,X172,X170,X168] : (p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X162,X168),'const.sum.++_3'(s(fun(X164,X162),X166),s(fun(X169,X168),X170),s('type.sum.sum'(X164,X169),X171)))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X162,fun(X162,bool)),X163),s(fun(X162,X164),X165),s(fun(X164,X162),X166))) | s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X164,X169),X171))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X168,fun(X168,bool)),X172),s(fun(X168,X169),X173),s(fun(X169,X168),X170))))) ) | $spl268),
  inference(cnf_transformation,[],[f18006_D])).
fof(f18006_D,plain,(
  ( ! [X165,X163,X173,X171,X169,X166,X164,X162,X172,X170,X168] : (p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X162,X168),'const.sum.++_3'(s(fun(X164,X162),X166),s(fun(X169,X168),X170),s('type.sum.sum'(X164,X169),X171)))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X162,fun(X162,bool)),X163),s(fun(X162,X164),X165),s(fun(X164,X162),X166))) | s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X164,X169),X171))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X168,fun(X168,bool)),X172),s(fun(X168,X169),X173),s(fun(X169,X168),X170))))) ) <=> ~$spl268),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl268])])).
fof(f35219,plain,(
  ( ! [X158,X156,X154,X152,X161,X151,X159,X157,X155,X153,X162,X160] : (p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X151,X157),'const.sum.++_3'(s(fun(X153,X151),X155),s(fun(X158,X157),X159),s('type.sum.sum'(X153,X158),X160)))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X157,fun(X157,bool)),X161),s(fun(X157,X158),X162),s(fun(X158,X157),X159)))) | s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X153,X158),X160))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X151,fun(X151,bool)),X152),s(fun(X151,X153),X154),s(fun(X153,X151),X155))) | s(bool,'F_0') = s(bool,X156)) )),
  inference(subsumption_resolution,[],[f35194,f119])).
fof(f35194,plain,(
  ( ! [X158,X156,X154,X152,X161,X151,X159,X157,X155,X153,X162,X160] : (~p(s(bool,X156)) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X151,X157),'const.sum.++_3'(s(fun(X153,X151),X155),s(fun(X158,X157),X159),s('type.sum.sum'(X153,X158),X160)))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X157,fun(X157,bool)),X161),s(fun(X157,X158),X162),s(fun(X158,X157),X159)))) | s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X153,X158),X160))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X151,fun(X151,bool)),X152),s(fun(X151,X153),X154),s(fun(X153,X151),X155))) | s(bool,'F_0') = s(bool,X156)) )),
  inference(superposition,[],[f9518,f5187])).
fof(f9518,plain,(
  ( ! [X198,X196,X194,X202,X200,X199,X197,X195,X193,X203,X201] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X193,fun(X193,bool)),X202),s(fun(X193,X195),X203),s(fun(X195,X193),X196)))) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X193,X194),'const.sum.++_3'(s(fun(X195,X193),X196),s(fun(X197,X194),X198),s('type.sum.sum'(X195,X197),X199)))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X194,fun(X194,bool)),X200),s(fun(X194,X197),X201),s(fun(X197,X194),X198)))) | s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X195,X197),X199)))) )),
  inference(superposition,[],[f122,f127])).
fof(f35218,plain,(
  ( ! [X146,X144,X142,X140,X149,X147,X145,X143,X141,X139,X150,X148] : (p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X139,X145),'const.sum.++_3'(s(fun(X141,X139),X143),s(fun(X146,X145),X147),s('type.sum.sum'(X141,X146),X148)))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X145,fun(X145,bool)),X149),s(fun(X145,X146),X150),s(fun(X146,X145),X147)))) | s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X141,X146),X148))) | s(bool,'F_0') = s(bool,X144) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X139,fun(X139,bool)),X140),s(fun(X139,X141),X142),s(fun(X141,X139),X143)))) )),
  inference(subsumption_resolution,[],[f35193,f119])).
fof(f35193,plain,(
  ( ! [X146,X144,X142,X140,X149,X147,X145,X143,X141,X139,X150,X148] : (~p(s(bool,X144)) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X139,X145),'const.sum.++_3'(s(fun(X141,X139),X143),s(fun(X146,X145),X147),s('type.sum.sum'(X141,X146),X148)))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X145,fun(X145,bool)),X149),s(fun(X145,X146),X150),s(fun(X146,X145),X147)))) | s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X141,X146),X148))) | s(bool,'F_0') = s(bool,X144) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X139,fun(X139,bool)),X140),s(fun(X139,X141),X142),s(fun(X141,X139),X143)))) )),
  inference(superposition,[],[f9518,f5187])).
fof(f17211,plain,(
  ( ! [X165,X163,X173,X171,X169,X166,X164,X162,X172,X170,X168] : (s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X162,X168),'const.sum.++_3'(s(fun(X164,X162),X166),s(fun(X169,X168),X170),s('type.sum.sum'(X164,X169),X171))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X162,fun(X162,bool)),X163),s(fun(X162,X164),X165),s(fun(X164,X162),X166))) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X164,X169),X171)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X168,fun(X168,bool)),X172),s(fun(X168,X169),X173),s(fun(X169,X168),X170))))) ) | $spl264),
  inference(cnf_transformation,[],[f17211_D])).
fof(f17211_D,plain,(
  ( ! [X165,X163,X173,X171,X169,X166,X164,X162,X172,X170,X168] : (s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X162,X168),'const.sum.++_3'(s(fun(X164,X162),X166),s(fun(X169,X168),X170),s('type.sum.sum'(X164,X169),X171))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X162,fun(X162,bool)),X163),s(fun(X162,X164),X165),s(fun(X164,X162),X166))) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X164,X169),X171)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X168,fun(X168,bool)),X172),s(fun(X168,X169),X173),s(fun(X169,X168),X170))))) ) <=> ~$spl264),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl264])])).
fof(f34997,plain,(
  ( ! [X158,X156,X154,X152,X161,X151,X159,X157,X155,X153,X162,X160] : (s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X151,X157),'const.sum.++_3'(s(fun(X153,X151),X155),s(fun(X158,X157),X159),s('type.sum.sum'(X153,X158),X160))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X157,fun(X157,bool)),X161),s(fun(X157,X158),X162),s(fun(X158,X157),X159)))) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X153,X158),X160)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X151,fun(X151,bool)),X152),s(fun(X151,X153),X154),s(fun(X153,X151),X155))) | s(bool,'F_0') = s(bool,X156)) )),
  inference(subsumption_resolution,[],[f34972,f119])).
fof(f34972,plain,(
  ( ! [X158,X156,X154,X152,X161,X151,X159,X157,X155,X153,X162,X160] : (~p(s(bool,X156)) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X151,X157),'const.sum.++_3'(s(fun(X153,X151),X155),s(fun(X158,X157),X159),s('type.sum.sum'(X153,X158),X160))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X157,fun(X157,bool)),X161),s(fun(X157,X158),X162),s(fun(X158,X157),X159)))) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X153,X158),X160)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X151,fun(X151,bool)),X152),s(fun(X151,X153),X154),s(fun(X153,X151),X155))) | s(bool,'F_0') = s(bool,X156)) )),
  inference(superposition,[],[f9513,f5187])).
fof(f9513,plain,(
  ( ! [X146,X144,X142,X140,X138,X147,X145,X143,X141,X139,X148] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X138,fun(X138,bool)),X147),s(fun(X138,X140),X148),s(fun(X140,X138),X141)))) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X138,X139),'const.sum.++_3'(s(fun(X140,X138),X141),s(fun(X142,X139),X143),s('type.sum.sum'(X140,X142),X144))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X139,fun(X139,bool)),X145),s(fun(X139,X142),X146),s(fun(X142,X139),X143)))) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X140,X142),X144))))) )),
  inference(superposition,[],[f119,f127])).
fof(f34996,plain,(
  ( ! [X146,X144,X142,X140,X149,X147,X145,X143,X141,X139,X150,X148] : (s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X139,X145),'const.sum.++_3'(s(fun(X141,X139),X143),s(fun(X146,X145),X147),s('type.sum.sum'(X141,X146),X148))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X145,fun(X145,bool)),X149),s(fun(X145,X146),X150),s(fun(X146,X145),X147)))) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X141,X146),X148)))) | s(bool,'F_0') = s(bool,X144) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X139,fun(X139,bool)),X140),s(fun(X139,X141),X142),s(fun(X141,X139),X143)))) )),
  inference(subsumption_resolution,[],[f34971,f119])).
fof(f34971,plain,(
  ( ! [X146,X144,X142,X140,X149,X147,X145,X143,X141,X139,X150,X148] : (~p(s(bool,X144)) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X139,X145),'const.sum.++_3'(s(fun(X141,X139),X143),s(fun(X146,X145),X147),s('type.sum.sum'(X141,X146),X148))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X145,fun(X145,bool)),X149),s(fun(X145,X146),X150),s(fun(X146,X145),X147)))) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X141,X146),X148)))) | s(bool,'F_0') = s(bool,X144) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X139,fun(X139,bool)),X140),s(fun(X139,X141),X142),s(fun(X141,X139),X143)))) )),
  inference(superposition,[],[f9513,f5187])).
fof(f16779,plain,(
  ( ! [X165,X163,X173,X171,X169,X166,X164,X162,X172,X170,X168] : (s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X162,X168),'const.sum.++_3'(s(fun(X164,X162),X166),s(fun(X169,X168),X170),s('type.sum.sum'(X164,X169),X171))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X162,fun(X162,bool)),X163),s(fun(X162,X164),X165),s(fun(X164,X162),X166))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X164,X169),X171)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X168,fun(X168,bool)),X172),s(fun(X168,X169),X173),s(fun(X169,X168),X170))))) ) | $spl260),
  inference(cnf_transformation,[],[f16779_D])).
fof(f16779_D,plain,(
  ( ! [X165,X163,X173,X171,X169,X166,X164,X162,X172,X170,X168] : (s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X162,X168),'const.sum.++_3'(s(fun(X164,X162),X166),s(fun(X169,X168),X170),s('type.sum.sum'(X164,X169),X171))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X162,fun(X162,bool)),X163),s(fun(X162,X164),X165),s(fun(X164,X162),X166))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X164,X169),X171)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X168,fun(X168,bool)),X172),s(fun(X168,X169),X173),s(fun(X169,X168),X170))))) ) <=> ~$spl260),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl260])])).
fof(f11980,plain,(
  ( ! [X42] : (~p(s(bool,X42))) ) | $spl164),
  inference(cnf_transformation,[],[f11980_D])).
fof(f11980_D,plain,(
  ( ! [X42] : (~p(s(bool,X42))) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f34929,plain,(
  ( ! [X158,X156,X154,X152,X161,X151,X159,X157,X155,X153,X162,X160] : (~p(s(bool,X156)) | s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X151,X157),'const.sum.++_3'(s(fun(X153,X151),X155),s(fun(X158,X157),X159),s('type.sum.sum'(X153,X158),X160))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X157,fun(X157,bool)),X161),s(fun(X157,X158),X162),s(fun(X158,X157),X159)))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X153,X158),X160)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X151,fun(X151,bool)),X152),s(fun(X151,X153),X154),s(fun(X153,X151),X155)))) )),
  inference(subsumption_resolution,[],[f34902,f118])).
fof(f34902,plain,(
  ( ! [X158,X156,X154,X152,X161,X151,X159,X157,X155,X153,X162,X160] : (~p(s(bool,X156)) | s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X151,X157),'const.sum.++_3'(s(fun(X153,X151),X155),s(fun(X158,X157),X159),s('type.sum.sum'(X153,X158),X160))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X157,fun(X157,bool)),X161),s(fun(X157,X158),X162),s(fun(X158,X157),X159)))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X153,X158),X160)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X151,fun(X151,bool)),X152),s(fun(X151,X153),X154),s(fun(X153,X151),X155))) | s(bool,'F_0') = s(bool,X156)) )),
  inference(superposition,[],[f9511,f5187])).
fof(f9511,plain,(
  ( ! [X125,X123,X121,X118,X116,X126,X124,X122,X120,X119,X117] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X116,fun(X116,bool)),X125),s(fun(X116,X118),X126),s(fun(X118,X116),X119)))) | s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X116,X117),'const.sum.++_3'(s(fun(X118,X116),X119),s(fun(X120,X117),X121),s('type.sum.sum'(X118,X120),X122))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X117,fun(X117,bool)),X123),s(fun(X117,X120),X124),s(fun(X120,X117),X121)))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X118,X120),X122))))) )),
  inference(superposition,[],[f118,f127])).
fof(f34928,plain,(
  ( ! [X146,X144,X142,X140,X149,X147,X145,X143,X141,X139,X150,X148] : (~p(s(bool,X144)) | s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X139,X145),'const.sum.++_3'(s(fun(X141,X139),X143),s(fun(X146,X145),X147),s('type.sum.sum'(X141,X146),X148))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X145,fun(X145,bool)),X149),s(fun(X145,X146),X150),s(fun(X146,X145),X147)))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X141,X146),X148)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X139,fun(X139,bool)),X140),s(fun(X139,X141),X142),s(fun(X141,X139),X143)))) )),
  inference(subsumption_resolution,[],[f34901,f118])).
fof(f34901,plain,(
  ( ! [X146,X144,X142,X140,X149,X147,X145,X143,X141,X139,X150,X148] : (~p(s(bool,X144)) | s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X139,X145),'const.sum.++_3'(s(fun(X141,X139),X143),s(fun(X146,X145),X147),s('type.sum.sum'(X141,X146),X148))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X145,fun(X145,bool)),X149),s(fun(X145,X146),X150),s(fun(X146,X145),X147)))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X141,X146),X148)))) | s(bool,'F_0') = s(bool,X144) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X139,fun(X139,bool)),X140),s(fun(X139,X141),X142),s(fun(X141,X139),X143)))) )),
  inference(superposition,[],[f9511,f5187])).
fof(f16639,plain,(
  ( ! [X165,X163,X173,X171,X169,X166,X164,X162,X172,X170,X168] : (~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X162,X168),'const.sum.++_3'(s(fun(X164,X162),X166),s(fun(X169,X168),X170),s('type.sum.sum'(X164,X169),X171)))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X162,fun(X162,bool)),X163),s(fun(X162,X164),X165),s(fun(X164,X162),X166))) | s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X164,X169),X171))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X168,fun(X168,bool)),X172),s(fun(X168,X169),X173),s(fun(X169,X168),X170))))) ) | $spl258),
  inference(cnf_transformation,[],[f16639_D])).
fof(f16639_D,plain,(
  ( ! [X165,X163,X173,X171,X169,X166,X164,X162,X172,X170,X168] : (~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X162,X168),'const.sum.++_3'(s(fun(X164,X162),X166),s(fun(X169,X168),X170),s('type.sum.sum'(X164,X169),X171)))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X162,fun(X162,bool)),X163),s(fun(X162,X164),X165),s(fun(X164,X162),X166))) | s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X164,X169),X171))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X168,fun(X168,bool)),X172),s(fun(X168,X169),X173),s(fun(X169,X168),X170))))) ) <=> ~$spl258),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl258])])).
fof(f34859,plain,(
  ( ! [X158,X156,X154,X152,X161,X151,X159,X157,X155,X153,X162,X160] : (~p(s(bool,X156)) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X151,X157),'const.sum.++_3'(s(fun(X153,X151),X155),s(fun(X158,X157),X159),s('type.sum.sum'(X153,X158),X160)))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X157,fun(X157,bool)),X161),s(fun(X157,X158),X162),s(fun(X158,X157),X159)))) | s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X153,X158),X160))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X151,fun(X151,bool)),X152),s(fun(X151,X153),X154),s(fun(X153,X151),X155)))) )),
  inference(subsumption_resolution,[],[f34832,f118])).
fof(f34832,plain,(
  ( ! [X158,X156,X154,X152,X161,X151,X159,X157,X155,X153,X162,X160] : (~p(s(bool,X156)) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X151,X157),'const.sum.++_3'(s(fun(X153,X151),X155),s(fun(X158,X157),X159),s('type.sum.sum'(X153,X158),X160)))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X157,fun(X157,bool)),X161),s(fun(X157,X158),X162),s(fun(X158,X157),X159)))) | s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X153,X158),X160))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X151,fun(X151,bool)),X152),s(fun(X151,X153),X154),s(fun(X153,X151),X155))) | s(bool,'F_0') = s(bool,X156)) )),
  inference(superposition,[],[f9510,f5187])).
fof(f9510,plain,(
  ( ! [X111,X109,X107,X105,X114,X112,X110,X108,X106,X115,X113] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X105,fun(X105,bool)),X114),s(fun(X105,X107),X115),s(fun(X107,X105),X108)))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X105,X106),'const.sum.++_3'(s(fun(X107,X105),X108),s(fun(X109,X106),X110),s('type.sum.sum'(X107,X109),X111)))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X106,fun(X106,bool)),X112),s(fun(X106,X109),X113),s(fun(X109,X106),X110)))) | s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X107,X109),X111)))) )),
  inference(superposition,[],[f118,f127])).
fof(f34858,plain,(
  ( ! [X146,X144,X142,X140,X149,X147,X145,X143,X141,X139,X150,X148] : (~p(s(bool,X144)) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X139,X145),'const.sum.++_3'(s(fun(X141,X139),X143),s(fun(X146,X145),X147),s('type.sum.sum'(X141,X146),X148)))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X145,fun(X145,bool)),X149),s(fun(X145,X146),X150),s(fun(X146,X145),X147)))) | s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X141,X146),X148))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X139,fun(X139,bool)),X140),s(fun(X139,X141),X142),s(fun(X141,X139),X143)))) )),
  inference(subsumption_resolution,[],[f34831,f118])).
fof(f34831,plain,(
  ( ! [X146,X144,X142,X140,X149,X147,X145,X143,X141,X139,X150,X148] : (~p(s(bool,X144)) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X139,X145),'const.sum.++_3'(s(fun(X141,X139),X143),s(fun(X146,X145),X147),s('type.sum.sum'(X141,X146),X148)))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X145,fun(X145,bool)),X149),s(fun(X145,X146),X150),s(fun(X146,X145),X147)))) | s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X141,X146),X148))) | s(bool,'F_0') = s(bool,X144) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X139,fun(X139,bool)),X140),s(fun(X139,X141),X142),s(fun(X141,X139),X143)))) )),
  inference(superposition,[],[f9510,f5187])).
fof(f15869,plain,(
  ( ! [X165,X163,X173,X171,X169,X166,X164,X162,X172,X170,X168] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X168,fun(X168,bool)),X169),s(fun(X168,X170),X171),s(fun(X170,X168),X172)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X162,fun(X162,bool)),X163),s(fun(X162,X164),X165),s(fun(X164,X162),X166))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X162,X168),'const.sum.++_3'(s(fun(X164,X162),X166),s(fun(X170,X168),X172),s('type.sum.sum'(X164,X170),X173)))))) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X164,X170),X173)))) ) | $spl256),
  inference(cnf_transformation,[],[f15869_D])).
fof(f15869_D,plain,(
  ( ! [X165,X163,X173,X171,X169,X166,X164,X162,X172,X170,X168] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X168,fun(X168,bool)),X169),s(fun(X168,X170),X171),s(fun(X170,X168),X172)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X162,fun(X162,bool)),X163),s(fun(X162,X164),X165),s(fun(X164,X162),X166))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X162,X168),'const.sum.++_3'(s(fun(X164,X162),X166),s(fun(X170,X168),X172),s('type.sum.sum'(X164,X170),X173)))))) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X164,X170),X173)))) ) <=> ~$spl256),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl256])])).
fof(f34789,plain,(
  ( ! [X158,X156,X154,X152,X161,X151,X159,X157,X155,X153,X162,X160] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X157,fun(X157,bool)),X158),s(fun(X157,X159),X160),s(fun(X159,X157),X161)))) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X153,X159),X162))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X151,X157),'const.sum.++_3'(s(fun(X153,X151),X155),s(fun(X159,X157),X161),s('type.sum.sum'(X153,X159),X162)))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X151,fun(X151,bool)),X152),s(fun(X151,X153),X154),s(fun(X153,X151),X155))) | s(bool,'F_0') = s(bool,X156)) )),
  inference(subsumption_resolution,[],[f34762,f119])).
fof(f34762,plain,(
  ( ! [X158,X156,X154,X152,X161,X151,X159,X157,X155,X153,X162,X160] : (~p(s(bool,X156)) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X157,fun(X157,bool)),X158),s(fun(X157,X159),X160),s(fun(X159,X157),X161)))) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X153,X159),X162))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X151,X157),'const.sum.++_3'(s(fun(X153,X151),X155),s(fun(X159,X157),X161),s('type.sum.sum'(X153,X159),X162)))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X151,fun(X151,bool)),X152),s(fun(X151,X153),X154),s(fun(X153,X151),X155))) | s(bool,'F_0') = s(bool,X156)) )),
  inference(superposition,[],[f9449,f5187])).
fof(f9449,plain,(
  ( ! [X80,X78,X87,X85,X83,X81,X79,X77,X86,X84,X82] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X77,fun(X77,bool)),X86),s(fun(X77,X79),X87),s(fun(X79,X77),X80)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X78,fun(X78,bool)),X84),s(fun(X78,X81),X85),s(fun(X81,X78),X82)))) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X79,X81),X83))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X77,X78),'const.sum.++_3'(s(fun(X79,X77),X80),s(fun(X81,X78),X82),s('type.sum.sum'(X79,X81),X83))))))) )),
  inference(superposition,[],[f127,f123])).
fof(f34788,plain,(
  ( ! [X146,X144,X142,X140,X149,X147,X145,X143,X141,X139,X150,X148] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X145,fun(X145,bool)),X146),s(fun(X145,X147),X148),s(fun(X147,X145),X149)))) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X141,X147),X150))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X139,X145),'const.sum.++_3'(s(fun(X141,X139),X143),s(fun(X147,X145),X149),s('type.sum.sum'(X141,X147),X150)))))) | s(bool,'F_0') = s(bool,X144) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X139,fun(X139,bool)),X140),s(fun(X139,X141),X142),s(fun(X141,X139),X143)))) )),
  inference(subsumption_resolution,[],[f34761,f119])).
fof(f34761,plain,(
  ( ! [X146,X144,X142,X140,X149,X147,X145,X143,X141,X139,X150,X148] : (~p(s(bool,X144)) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X145,fun(X145,bool)),X146),s(fun(X145,X147),X148),s(fun(X147,X145),X149)))) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X141,X147),X150))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X139,X145),'const.sum.++_3'(s(fun(X141,X139),X143),s(fun(X147,X145),X149),s('type.sum.sum'(X141,X147),X150)))))) | s(bool,'F_0') = s(bool,X144) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X139,fun(X139,bool)),X140),s(fun(X139,X141),X142),s(fun(X141,X139),X143)))) )),
  inference(superposition,[],[f9449,f5187])).
fof(f14651,plain,(
  ( ! [X165,X163,X173,X171,X169,X166,X164,X162,X172,X170,X168] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X168,fun(X168,bool)),X169),s(fun(X168,X170),X171),s(fun(X170,X168),X172)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X162,fun(X162,bool)),X163),s(fun(X162,X164),X165),s(fun(X164,X162),X166))) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X162,X168),'const.sum.++_3'(s(fun(X164,X162),X166),s(fun(X170,X168),X172),s('type.sum.sum'(X164,X170),X173)))))) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X164,X170),X173)))) ) | $spl250),
  inference(cnf_transformation,[],[f14651_D])).
fof(f14651_D,plain,(
  ( ! [X165,X163,X173,X171,X169,X166,X164,X162,X172,X170,X168] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X168,fun(X168,bool)),X169),s(fun(X168,X170),X171),s(fun(X170,X168),X172)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X162,fun(X162,bool)),X163),s(fun(X162,X164),X165),s(fun(X164,X162),X166))) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X162,X168),'const.sum.++_3'(s(fun(X164,X162),X166),s(fun(X170,X168),X172),s('type.sum.sum'(X164,X170),X173)))))) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X164,X170),X173)))) ) <=> ~$spl250),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl250])])).
fof(f34719,plain,(
  ( ! [X158,X156,X154,X152,X161,X151,X159,X157,X155,X153,X162,X160] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X157,fun(X157,bool)),X158),s(fun(X157,X159),X160),s(fun(X159,X157),X161)))) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X153,X159),X162))) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X151,X157),'const.sum.++_3'(s(fun(X153,X151),X155),s(fun(X159,X157),X161),s('type.sum.sum'(X153,X159),X162)))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X151,fun(X151,bool)),X152),s(fun(X151,X153),X154),s(fun(X153,X151),X155))) | s(bool,'F_0') = s(bool,X156)) )),
  inference(subsumption_resolution,[],[f34694,f119])).
fof(f34694,plain,(
  ( ! [X158,X156,X154,X152,X161,X151,X159,X157,X155,X153,X162,X160] : (~p(s(bool,X156)) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X157,fun(X157,bool)),X158),s(fun(X157,X159),X160),s(fun(X159,X157),X161)))) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X153,X159),X162))) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X151,X157),'const.sum.++_3'(s(fun(X153,X151),X155),s(fun(X159,X157),X161),s('type.sum.sum'(X153,X159),X162)))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X151,fun(X151,bool)),X152),s(fun(X151,X153),X154),s(fun(X153,X151),X155))) | s(bool,'F_0') = s(bool,X156)) )),
  inference(superposition,[],[f9447,f5187])).
fof(f9447,plain,(
  ( ! [X61,X59,X57,X64,X62,X60,X58,X56,X65,X55,X63] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X55,fun(X55,bool)),X64),s(fun(X55,X57),X65),s(fun(X57,X55),X58)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X56,fun(X56,bool)),X62),s(fun(X56,X59),X63),s(fun(X59,X56),X60)))) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X57,X59),X61))) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X55,X56),'const.sum.++_3'(s(fun(X57,X55),X58),s(fun(X59,X56),X60),s('type.sum.sum'(X57,X59),X61))))))) )),
  inference(superposition,[],[f127,f119])).
fof(f34718,plain,(
  ( ! [X146,X144,X142,X140,X149,X147,X145,X143,X141,X139,X150,X148] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X145,fun(X145,bool)),X146),s(fun(X145,X147),X148),s(fun(X147,X145),X149)))) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X141,X147),X150))) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X139,X145),'const.sum.++_3'(s(fun(X141,X139),X143),s(fun(X147,X145),X149),s('type.sum.sum'(X141,X147),X150)))))) | s(bool,'F_0') = s(bool,X144) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X139,fun(X139,bool)),X140),s(fun(X139,X141),X142),s(fun(X141,X139),X143)))) )),
  inference(subsumption_resolution,[],[f34693,f119])).
fof(f34693,plain,(
  ( ! [X146,X144,X142,X140,X149,X147,X145,X143,X141,X139,X150,X148] : (~p(s(bool,X144)) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X145,fun(X145,bool)),X146),s(fun(X145,X147),X148),s(fun(X147,X145),X149)))) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X141,X147),X150))) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X139,X145),'const.sum.++_3'(s(fun(X141,X139),X143),s(fun(X147,X145),X149),s('type.sum.sum'(X141,X147),X150)))))) | s(bool,'F_0') = s(bool,X144) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X139,fun(X139,bool)),X140),s(fun(X139,X141),X142),s(fun(X141,X139),X143)))) )),
  inference(superposition,[],[f9447,f5187])).
fof(f17369,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl266),
  inference(cnf_transformation,[],[f17369_D])).
fof(f17369_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl266),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl266])])).
fof(f34625,plain,(
  ( ! [X14,X15,X13] : (s('type.sum.sum'(X13,X14),'const.sum.INR_1'(s(X14,X15))) = s('type.sum.sum'(X13,X14),'const.sum.INR_1'(s(X14,sK13('const.sum.INR_1'(s(X14,X15)),X14,X13))))) ) | $spl328),
  inference(cnf_transformation,[],[f34625_D])).
fof(f34625_D,plain,(
  ( ! [X14,X15,X13] : (s('type.sum.sum'(X13,X14),'const.sum.INR_1'(s(X14,X15))) = s('type.sum.sum'(X13,X14),'const.sum.INR_1'(s(X14,sK13('const.sum.INR_1'(s(X14,X15)),X14,X13))))) ) <=> ~$spl328),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl328])])).
fof(f34624,plain,(
  ( ! [X14,X15,X13,X16] : (s('type.sum.sum'(X13,X14),'const.sum.INR_1'(s(X14,X15))) = s('type.sum.sum'(X13,X14),'const.sum.INR_1'(s(X14,sK13('const.sum.INR_1'(s(X14,X15)),X14,X13)))) | s(bool,'T_0') = s(bool,X16)) )),
  inference(subsumption_resolution,[],[f34319,f123])).
fof(f34319,plain,(
  ( ! [X14,X15,X13,X16] : (p(s(bool,X16)) | s('type.sum.sum'(X13,X14),'const.sum.INR_1'(s(X14,X15))) = s('type.sum.sum'(X13,X14),'const.sum.INR_1'(s(X14,sK13('const.sum.INR_1'(s(X14,X15)),X14,X13)))) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f9288,f31809])).
fof(f31809,plain,(
  ( ! [X39,X37,X38,X36] : (s(bool,X36) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X37,X38),'const.sum.INR_1'(s(X38,X39))))) | s(bool,'T_0') = s(bool,X36)) )),
  inference(trivial_inequality_removal,[],[f31745])).
fof(f31745,plain,(
  ( ! [X39,X37,X38,X36] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,X36) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X37,X38),'const.sum.INR_1'(s(X38,X39))))) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f23990,f126])).
fof(f126,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.quotient_sum.ISR_RSP',aHL_BOOL_CASES)).
fof(f23990,plain,(
  ( ! [X138,X136,X139,X137] : (s(bool,'F_0') != s(bool,X139) | s(bool,X139) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X136,X137),'const.sum.INR_1'(s(X137,X138)))))) )),
  inference(subsumption_resolution,[],[f23969,f150])).
fof(f23969,plain,(
  ( ! [X138,X136,X139,X137] : (p(s(bool,'F_0')) | s(bool,X139) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X136,X137),'const.sum.INR_1'(s(X137,X138))))) | s(bool,'F_0') != s(bool,X139)) )),
  inference(superposition,[],[f21924,f5444])).
fof(f21924,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | s(bool,X0) = s(bool,X1) | s(bool,'F_0') != s(bool,X0)) )),
  inference(resolution,[],[f118,f112])).
fof(f112,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f81])).
fof(f9288,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X0,X1),X2)))) | s('type.sum.sum'(X0,X1),X2) = s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,sK13(X2,X1,X0))))) )),
  inference(superposition,[],[f134,f111])).
fof(f33922,plain,(
  ( ! [X14,X15,X13] : (s('type.sum.sum'(X13,X14),'const.sum.INL_1'(s(X13,X15))) = s('type.sum.sum'(X13,X14),'const.sum.INL_1'(s(X13,sK12('const.sum.INL_1'(s(X13,X15)),X14,X13))))) ) | $spl326),
  inference(cnf_transformation,[],[f33922_D])).
fof(f33922_D,plain,(
  ( ! [X14,X15,X13] : (s('type.sum.sum'(X13,X14),'const.sum.INL_1'(s(X13,X15))) = s('type.sum.sum'(X13,X14),'const.sum.INL_1'(s(X13,sK12('const.sum.INL_1'(s(X13,X15)),X14,X13))))) ) <=> ~$spl326),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl326])])).
fof(f33921,plain,(
  ( ! [X14,X15,X13,X16] : (s('type.sum.sum'(X13,X14),'const.sum.INL_1'(s(X13,X15))) = s('type.sum.sum'(X13,X14),'const.sum.INL_1'(s(X13,sK12('const.sum.INL_1'(s(X13,X15)),X14,X13)))) | s(bool,'T_0') = s(bool,X16)) )),
  inference(subsumption_resolution,[],[f33623,f123])).
fof(f33623,plain,(
  ( ! [X14,X15,X13,X16] : (p(s(bool,X16)) | s('type.sum.sum'(X13,X14),'const.sum.INL_1'(s(X13,X15))) = s('type.sum.sum'(X13,X14),'const.sum.INL_1'(s(X13,sK12('const.sum.INL_1'(s(X13,X15)),X14,X13)))) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f9189,f31666])).
fof(f31666,plain,(
  ( ! [X39,X37,X38,X36] : (s(bool,X36) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X37,X38),'const.sum.INL_1'(s(X37,X39))))) | s(bool,'T_0') = s(bool,X36)) )),
  inference(trivial_inequality_removal,[],[f31602])).
fof(f31602,plain,(
  ( ! [X39,X37,X38,X36] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,X36) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X37,X38),'const.sum.INL_1'(s(X37,X39))))) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f23984,f126])).
fof(f23984,plain,(
  ( ! [X92,X90,X91,X89] : (s(bool,'F_0') != s(bool,X92) | s(bool,X92) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X89,X90),'const.sum.INL_1'(s(X89,X91)))))) )),
  inference(subsumption_resolution,[],[f23958,f150])).
fof(f23958,plain,(
  ( ! [X92,X90,X91,X89] : (p(s(bool,'F_0')) | s(bool,X92) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X89,X90),'const.sum.INL_1'(s(X89,X91))))) | s(bool,'F_0') != s(bool,X92)) )),
  inference(superposition,[],[f21924,f5433])).
fof(f9189,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X0,X1),X2)))) | s('type.sum.sum'(X0,X1),X2) = s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,sK12(X2,X1,X0))))) )),
  inference(superposition,[],[f129,f111])).
fof(f31826,plain,(
  ( ! [X70,X74,X72,X73] : (s(bool,X70) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X72,X73),'const.sum.INR_1'(s(X73,X74))))) | s(bool,'T_0') = s(bool,X70)) ) | $spl320),
  inference(cnf_transformation,[],[f31826_D])).
fof(f31826_D,plain,(
  ( ! [X70,X74,X72,X73] : (s(bool,X70) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X72,X73),'const.sum.INR_1'(s(X73,X74))))) | s(bool,'T_0') = s(bool,X70)) ) <=> ~$spl320),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl320])])).
fof(f350,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3)) ) | $spl20),
  inference(cnf_transformation,[],[f350_D])).
fof(f350_D,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3)) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f33274,plain,(
  ( ! [X103,X107,X105,X106,X104] : (s(bool,'T_0') != s(bool,X106) | s(bool,'T_0') = s(bool,X107) | s(bool,X107) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X103,X104),'const.sum.INR_1'(s(X104,X105)))))) )),
  inference(subsumption_resolution,[],[f33225,f21926])).
fof(f21926,plain,(
  ( ! [X4] : (s(bool,'F_0') != s(bool,X4) | s(bool,'T_0') != s(bool,X4)) )),
  inference(resolution,[],[f118,f122])).
fof(f33225,plain,(
  ( ! [X103,X107,X105,X106,X104] : (s(bool,'T_0') != s(bool,X106) | s(bool,'F_0') = s(bool,X106) | s(bool,'T_0') = s(bool,X107) | s(bool,X107) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X103,X104),'const.sum.INR_1'(s(X104,X105)))))) )),
  inference(superposition,[],[f24319,f19171])).
fof(f19171,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') = s(bool,X10) | s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X11)) )),
  inference(resolution,[],[f19124,f123])).
fof(f19124,plain,(
  ( ! [X6,X5] : (p(s(bool,X6)) | s(bool,'T_0') = s(bool,X5) | s(bool,X5) = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f19123,f149])).
fof(f19123,plain,(
  ( ! [X6,X5] : (p(s(bool,'F_0')) | s(bool,X5) = s(bool,X6) | s(bool,'T_0') = s(bool,X5) | p(s(bool,X6))) )),
  inference(forward_demodulation,[],[f19012,f5433])).
fof(f19012,plain,(
  ( ! [X6,X4,X5] : (s(bool,X5) = s(bool,X6) | s(bool,'T_0') = s(bool,X5) | p(s(bool,X6)) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(bool,X4),'const.sum.INL_1'(s(bool,X5))))))) )),
  inference(superposition,[],[f5507,f112])).
fof(f5507,plain,(
  ( ! [X90,X89] : (s(bool,X89) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(bool,X90),'const.sum.INL_1'(s(bool,X89))))) | s(bool,'T_0') = s(bool,X89)) )),
  inference(superposition,[],[f5433,f126])).
fof(f24319,plain,(
  ( ! [X138,X136,X139,X137] : (s(bool,X139) != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X136,X137),'const.sum.INR_1'(s(X137,X138))))) | s(bool,'F_0') = s(bool,X139)) )),
  inference(subsumption_resolution,[],[f24298,f150])).
fof(f24298,plain,(
  ( ! [X138,X136,X139,X137] : (p(s(bool,'F_0')) | s(bool,X139) != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X136,X137),'const.sum.INR_1'(s(X137,X138))))) | s(bool,'F_0') = s(bool,X139)) )),
  inference(superposition,[],[f22799,f5444])).
fof(f22799,plain,(
  ( ! [X6,X7] : (p(s(bool,X6)) | s(bool,X6) != s(bool,X7) | s(bool,'F_0') = s(bool,X7)) )),
  inference(resolution,[],[f22072,f119])).
fof(f22072,plain,(
  ( ! [X8,X9] : (~p(s(bool,X9)) | s(bool,X8) != s(bool,X9) | p(s(bool,X8))) )),
  inference(subsumption_resolution,[],[f21978,f150])).
fof(f21978,plain,(
  ( ! [X8,X9] : (s(bool,X8) != s(bool,X9) | ~p(s(bool,X9)) | p(s(bool,X8)) | p(s(bool,'F_0'))) )),
  inference(superposition,[],[f118,f112])).
fof(f33273,plain,(
  ( ! [X101,X99,X102,X100,X98] : (s(bool,'T_0') != s(bool,X101) | s(bool,X102) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X98,X99),'const.sum.INR_1'(s(X99,X100))))) | s(bool,'T_0') = s(bool,X102)) )),
  inference(subsumption_resolution,[],[f33224,f21926])).
fof(f33224,plain,(
  ( ! [X101,X99,X102,X100,X98] : (s(bool,'T_0') != s(bool,X101) | s(bool,'F_0') = s(bool,X101) | s(bool,X102) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X98,X99),'const.sum.INR_1'(s(X99,X100))))) | s(bool,'T_0') = s(bool,X102)) )),
  inference(superposition,[],[f24319,f19171])).
fof(f31683,plain,(
  ( ! [X70,X74,X72,X73] : (s(bool,X70) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X72,X73),'const.sum.INL_1'(s(X72,X74))))) | s(bool,'T_0') = s(bool,X70)) ) | $spl318),
  inference(cnf_transformation,[],[f31683_D])).
fof(f31683_D,plain,(
  ( ! [X70,X74,X72,X73] : (s(bool,X70) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X72,X73),'const.sum.INL_1'(s(X72,X74))))) | s(bool,'T_0') = s(bool,X70)) ) <=> ~$spl318),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl318])])).
fof(f33074,plain,(
  ( ! [X80,X78,X81,X79,X77] : (s(bool,'T_0') != s(bool,X80) | s(bool,'T_0') = s(bool,X81) | s(bool,X81) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X77,X78),'const.sum.INL_1'(s(X77,X79)))))) )),
  inference(subsumption_resolution,[],[f33026,f21926])).
fof(f33026,plain,(
  ( ! [X80,X78,X81,X79,X77] : (s(bool,'T_0') != s(bool,X80) | s(bool,'F_0') = s(bool,X80) | s(bool,'T_0') = s(bool,X81) | s(bool,X81) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X77,X78),'const.sum.INL_1'(s(X77,X79)))))) )),
  inference(superposition,[],[f24313,f19171])).
fof(f24313,plain,(
  ( ! [X92,X90,X91,X89] : (s(bool,X92) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X89,X90),'const.sum.INL_1'(s(X89,X91))))) | s(bool,'F_0') = s(bool,X92)) )),
  inference(subsumption_resolution,[],[f24287,f150])).
fof(f24287,plain,(
  ( ! [X92,X90,X91,X89] : (p(s(bool,'F_0')) | s(bool,X92) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X89,X90),'const.sum.INL_1'(s(X89,X91))))) | s(bool,'F_0') = s(bool,X92)) )),
  inference(superposition,[],[f22799,f5433])).
fof(f33073,plain,(
  ( ! [X76,X74,X72,X75,X73] : (s(bool,'T_0') != s(bool,X75) | s(bool,X76) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X72,X73),'const.sum.INL_1'(s(X72,X74))))) | s(bool,'T_0') = s(bool,X76)) )),
  inference(subsumption_resolution,[],[f33025,f21926])).
fof(f33025,plain,(
  ( ! [X76,X74,X72,X75,X73] : (s(bool,'T_0') != s(bool,X75) | s(bool,'F_0') = s(bool,X75) | s(bool,X76) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X72,X73),'const.sum.INL_1'(s(X72,X74))))) | s(bool,'T_0') = s(bool,X76)) )),
  inference(superposition,[],[f24313,f19171])).
fof(f32830,plain,(
  ( ! [X54,X52,X50,X53] : (s(bool,X50) != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X52,X53),'const.sum.INR_1'(s(X53,X54))))) | s(bool,'F_0') = s(bool,X50)) ) | $spl324),
  inference(cnf_transformation,[],[f32830_D])).
fof(f32830_D,plain,(
  ( ! [X54,X52,X50,X53] : (s(bool,X50) != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X52,X53),'const.sum.INR_1'(s(X53,X54))))) | s(bool,'F_0') = s(bool,X50)) ) <=> ~$spl324),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl324])])).
fof(f32831,plain,(
  ( ! [X59,X57,X58,X56,X55] : (s(bool,'T_0') != s(bool,X56) | s(bool,X55) != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X57,X58),'const.sum.INR_1'(s(X58,X59))))) | s(bool,'F_0') = s(bool,X55)) )),
  inference(subsumption_resolution,[],[f32737,f21926])).
fof(f32737,plain,(
  ( ! [X59,X57,X58,X56,X55] : (s(bool,'T_0') != s(bool,X56) | s(bool,X55) != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X57,X58),'const.sum.INR_1'(s(X58,X59))))) | s(bool,'F_0') = s(bool,X55) | s(bool,'F_0') = s(bool,X56)) )),
  inference(superposition,[],[f24237,f5187])).
fof(f24237,plain,(
  ( ! [X138,X136,X139,X137] : (s(bool,'T_0') != s(bool,X139) | s(bool,X139) != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X136,X137),'const.sum.INR_1'(s(X137,X138)))))) )),
  inference(subsumption_resolution,[],[f24210,f150])).
fof(f24210,plain,(
  ( ! [X138,X136,X139,X137] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,X139) | s(bool,X139) != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X136,X137),'const.sum.INR_1'(s(X137,X138)))))) )),
  inference(superposition,[],[f22526,f5444])).
fof(f22526,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X1) | s(bool,X0) != s(bool,X1)) )),
  inference(superposition,[],[f21926,f119])).
fof(f32829,plain,(
  ( ! [X54,X52,X50,X53,X51] : (s(bool,'T_0') != s(bool,X51) | s(bool,X50) != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X52,X53),'const.sum.INR_1'(s(X53,X54))))) | s(bool,'F_0') = s(bool,X50)) )),
  inference(subsumption_resolution,[],[f32736,f21926])).
fof(f32736,plain,(
  ( ! [X54,X52,X50,X53,X51] : (s(bool,'T_0') != s(bool,X51) | s(bool,X50) != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X52,X53),'const.sum.INR_1'(s(X53,X54))))) | s(bool,'F_0') = s(bool,X51) | s(bool,'F_0') = s(bool,X50)) )),
  inference(superposition,[],[f24237,f5187])).
fof(f32044,plain,(
  ( ! [X54,X52,X50,X53] : (s(bool,X50) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X52,X53),'const.sum.INL_1'(s(X52,X54))))) | s(bool,'F_0') = s(bool,X50)) ) | $spl322),
  inference(cnf_transformation,[],[f32044_D])).
fof(f32044_D,plain,(
  ( ! [X54,X52,X50,X53] : (s(bool,X50) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X52,X53),'const.sum.INL_1'(s(X52,X54))))) | s(bool,'F_0') = s(bool,X50)) ) <=> ~$spl322),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl322])])).
fof(f32045,plain,(
  ( ! [X59,X57,X58,X56,X55] : (s(bool,'T_0') != s(bool,X56) | s(bool,X55) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X57,X58),'const.sum.INL_1'(s(X57,X59))))) | s(bool,'F_0') = s(bool,X55)) )),
  inference(subsumption_resolution,[],[f31951,f21926])).
fof(f31951,plain,(
  ( ! [X59,X57,X58,X56,X55] : (s(bool,'T_0') != s(bool,X56) | s(bool,X55) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X57,X58),'const.sum.INL_1'(s(X57,X59))))) | s(bool,'F_0') = s(bool,X55) | s(bool,'F_0') = s(bool,X56)) )),
  inference(superposition,[],[f24230,f5187])).
fof(f24230,plain,(
  ( ! [X92,X90,X91,X89] : (s(bool,'T_0') != s(bool,X92) | s(bool,X92) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X89,X90),'const.sum.INL_1'(s(X89,X91)))))) )),
  inference(subsumption_resolution,[],[f24199,f150])).
fof(f24199,plain,(
  ( ! [X92,X90,X91,X89] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,X92) | s(bool,X92) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X89,X90),'const.sum.INL_1'(s(X89,X91)))))) )),
  inference(superposition,[],[f22526,f5433])).
fof(f32043,plain,(
  ( ! [X54,X52,X50,X53,X51] : (s(bool,'T_0') != s(bool,X51) | s(bool,X50) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X52,X53),'const.sum.INL_1'(s(X52,X54))))) | s(bool,'F_0') = s(bool,X50)) )),
  inference(subsumption_resolution,[],[f31950,f21926])).
fof(f31950,plain,(
  ( ! [X54,X52,X50,X53,X51] : (s(bool,'T_0') != s(bool,X51) | s(bool,X50) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X52,X53),'const.sum.INL_1'(s(X52,X54))))) | s(bool,'F_0') = s(bool,X51) | s(bool,'F_0') = s(bool,X50)) )),
  inference(superposition,[],[f24230,f5187])).
fof(f14566,plain,(
  ( ! [X165,X163,X173,X171,X169,X166,X164,X162,X172,X170,X168] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X168,X164),X169))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X170,X162),'const.sum.++_3'(s(fun(X168,X170),X171),s(fun(X164,X162),X166),s('type.sum.sum'(X168,X164),X169))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X162,fun(X162,bool)),X163),s(fun(X162,X164),X165),s(fun(X164,X162),X166))) | s(bool,'T_0') != s(bool,'const.quotient.QUOTIENT_3'(s(fun(X170,fun(X170,bool)),X172),s(fun(X170,X168),X173),s(fun(X168,X170),X171)))) ) | $spl248),
  inference(cnf_transformation,[],[f14566_D])).
fof(f14566_D,plain,(
  ( ! [X165,X163,X173,X171,X169,X166,X164,X162,X172,X170,X168] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X168,X164),X169))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X170,X162),'const.sum.++_3'(s(fun(X168,X170),X171),s(fun(X164,X162),X166),s('type.sum.sum'(X168,X164),X169))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X162,fun(X162,bool)),X163),s(fun(X162,X164),X165),s(fun(X164,X162),X166))) | s(bool,'T_0') != s(bool,'const.quotient.QUOTIENT_3'(s(fun(X170,fun(X170,bool)),X172),s(fun(X170,X168),X173),s(fun(X168,X170),X171)))) ) <=> ~$spl248),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl248])])).
fof(f31905,plain,(
  ( ! [X158,X156,X154,X152,X161,X151,X159,X157,X155,X153,X162,X160] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X157,X153),X158))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X159,X151),'const.sum.++_3'(s(fun(X157,X159),X160),s(fun(X153,X151),X155),s('type.sum.sum'(X157,X153),X158))))) | s(bool,'T_0') != s(bool,'const.quotient.QUOTIENT_3'(s(fun(X159,fun(X159,bool)),X161),s(fun(X159,X157),X162),s(fun(X157,X159),X160))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X151,fun(X151,bool)),X152),s(fun(X151,X153),X154),s(fun(X153,X151),X155))) | s(bool,'F_0') = s(bool,X156)) )),
  inference(subsumption_resolution,[],[f31884,f119])).
fof(f31884,plain,(
  ( ! [X158,X156,X154,X152,X161,X151,X159,X157,X155,X153,X162,X160] : (~p(s(bool,X156)) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X157,X153),X158))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X159,X151),'const.sum.++_3'(s(fun(X157,X159),X160),s(fun(X153,X151),X155),s('type.sum.sum'(X157,X153),X158))))) | s(bool,'T_0') != s(bool,'const.quotient.QUOTIENT_3'(s(fun(X159,fun(X159,bool)),X161),s(fun(X159,X157),X162),s(fun(X157,X159),X160))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X151,fun(X151,bool)),X152),s(fun(X151,X153),X154),s(fun(X153,X151),X155))) | s(bool,'F_0') = s(bool,X156)) )),
  inference(superposition,[],[f9426,f5187])).
fof(f9426,plain,(
  ( ! [X47,X54,X52,X50,X48,X56,X46,X55,X53,X51,X49] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X50,fun(X50,bool)),X53),s(fun(X50,X47),X54),s(fun(X47,X50),X52)))) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X46,X47),X48))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X49,X50),'const.sum.++_3'(s(fun(X46,X49),X51),s(fun(X47,X50),X52),s('type.sum.sum'(X46,X47),X48))))) | s(bool,'T_0') != s(bool,'const.quotient.QUOTIENT_3'(s(fun(X49,fun(X49,bool)),X55),s(fun(X49,X46),X56),s(fun(X46,X49),X51)))) )),
  inference(resolution,[],[f127,f122])).
fof(f31904,plain,(
  ( ! [X146,X144,X142,X140,X149,X147,X145,X143,X141,X139,X150,X148] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X145,X141),X146))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X147,X139),'const.sum.++_3'(s(fun(X145,X147),X148),s(fun(X141,X139),X143),s('type.sum.sum'(X145,X141),X146))))) | s(bool,'T_0') != s(bool,'const.quotient.QUOTIENT_3'(s(fun(X147,fun(X147,bool)),X149),s(fun(X147,X145),X150),s(fun(X145,X147),X148))) | s(bool,'F_0') = s(bool,X144) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X139,fun(X139,bool)),X140),s(fun(X139,X141),X142),s(fun(X141,X139),X143)))) )),
  inference(subsumption_resolution,[],[f31883,f119])).
fof(f31883,plain,(
  ( ! [X146,X144,X142,X140,X149,X147,X145,X143,X141,X139,X150,X148] : (~p(s(bool,X144)) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X145,X141),X146))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X147,X139),'const.sum.++_3'(s(fun(X145,X147),X148),s(fun(X141,X139),X143),s('type.sum.sum'(X145,X141),X146))))) | s(bool,'T_0') != s(bool,'const.quotient.QUOTIENT_3'(s(fun(X147,fun(X147,bool)),X149),s(fun(X147,X145),X150),s(fun(X145,X147),X148))) | s(bool,'F_0') = s(bool,X144) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X139,fun(X139,bool)),X140),s(fun(X139,X141),X142),s(fun(X141,X139),X143)))) )),
  inference(superposition,[],[f9426,f5187])).
fof(f31827,plain,(
  ( ! [X78,X76,X79,X77,X75] : (s(bool,X75) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X77,X78),'const.sum.INR_1'(s(X78,X79))))) | s(bool,'T_0') = s(bool,X76) | s(bool,'T_0') = s(bool,X75)) )),
  inference(subsumption_resolution,[],[f31753,f126])).
fof(f31753,plain,(
  ( ! [X78,X76,X79,X77,X75] : (s(bool,'F_0') != s(bool,X76) | s(bool,X75) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X77,X78),'const.sum.INR_1'(s(X78,X79))))) | s(bool,'T_0') = s(bool,X76) | s(bool,'T_0') = s(bool,X75)) )),
  inference(superposition,[],[f23990,f19171])).
fof(f31825,plain,(
  ( ! [X70,X74,X72,X71,X73] : (s(bool,X70) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X72,X73),'const.sum.INR_1'(s(X73,X74))))) | s(bool,'T_0') = s(bool,X70) | s(bool,'T_0') = s(bool,X71)) )),
  inference(subsumption_resolution,[],[f31752,f126])).
fof(f31752,plain,(
  ( ! [X70,X74,X72,X71,X73] : (s(bool,'F_0') != s(bool,X71) | s(bool,X70) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X72,X73),'const.sum.INR_1'(s(X73,X74))))) | s(bool,'T_0') = s(bool,X70) | s(bool,'T_0') = s(bool,X71)) )),
  inference(superposition,[],[f23990,f19171])).
fof(f31684,plain,(
  ( ! [X78,X76,X79,X77,X75] : (s(bool,X75) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X77,X78),'const.sum.INL_1'(s(X77,X79))))) | s(bool,'T_0') = s(bool,X76) | s(bool,'T_0') = s(bool,X75)) )),
  inference(subsumption_resolution,[],[f31610,f126])).
fof(f31610,plain,(
  ( ! [X78,X76,X79,X77,X75] : (s(bool,'F_0') != s(bool,X76) | s(bool,X75) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X77,X78),'const.sum.INL_1'(s(X77,X79))))) | s(bool,'T_0') = s(bool,X76) | s(bool,'T_0') = s(bool,X75)) )),
  inference(superposition,[],[f23984,f19171])).
fof(f31682,plain,(
  ( ! [X70,X74,X72,X71,X73] : (s(bool,X70) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X72,X73),'const.sum.INL_1'(s(X72,X74))))) | s(bool,'T_0') = s(bool,X70) | s(bool,'T_0') = s(bool,X71)) )),
  inference(subsumption_resolution,[],[f31609,f126])).
fof(f31609,plain,(
  ( ! [X70,X74,X72,X71,X73] : (s(bool,'F_0') != s(bool,X71) | s(bool,X70) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X72,X73),'const.sum.INL_1'(s(X72,X74))))) | s(bool,'T_0') = s(bool,X70) | s(bool,'T_0') = s(bool,X71)) )),
  inference(superposition,[],[f23984,f19171])).
fof(f13783,plain,(
  ( ! [X165,X163,X173,X171,X169,X166,X164,X162,X172,X170,X168] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X164,X168),X169))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X162,X170),'const.sum.++_3'(s(fun(X164,X162),X166),s(fun(X168,X170),X171),s('type.sum.sum'(X164,X168),X169))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X162,fun(X162,bool)),X163),s(fun(X162,X164),X165),s(fun(X164,X162),X166))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X170,fun(X170,bool)),X172),s(fun(X170,X168),X173),s(fun(X168,X170),X171)))) ) | $spl244),
  inference(cnf_transformation,[],[f13783_D])).
fof(f13783_D,plain,(
  ( ! [X165,X163,X173,X171,X169,X166,X164,X162,X172,X170,X168] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X164,X168),X169))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X162,X170),'const.sum.++_3'(s(fun(X164,X162),X166),s(fun(X168,X170),X171),s('type.sum.sum'(X164,X168),X169))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X162,fun(X162,bool)),X163),s(fun(X162,X164),X165),s(fun(X164,X162),X166))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X170,fun(X170,bool)),X172),s(fun(X170,X168),X173),s(fun(X168,X170),X171)))) ) <=> ~$spl244),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl244])])).
fof(f31557,plain,(
  ( ! [X158,X156,X154,X152,X161,X151,X159,X157,X155,X153,X162,X160] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X157,X153),X158))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X159,X151),'const.sum.++_3'(s(fun(X157,X159),X160),s(fun(X153,X151),X155),s('type.sum.sum'(X157,X153),X158))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X159,fun(X159,bool)),X161),s(fun(X159,X157),X162),s(fun(X157,X159),X160))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X151,fun(X151,bool)),X152),s(fun(X151,X153),X154),s(fun(X153,X151),X155))) | s(bool,'F_0') = s(bool,X156)) )),
  inference(subsumption_resolution,[],[f31536,f119])).
fof(f31536,plain,(
  ( ! [X158,X156,X154,X152,X161,X151,X159,X157,X155,X153,X162,X160] : (~p(s(bool,X156)) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X157,X153),X158))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X159,X151),'const.sum.++_3'(s(fun(X157,X159),X160),s(fun(X153,X151),X155),s('type.sum.sum'(X157,X153),X158))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X159,fun(X159,bool)),X161),s(fun(X159,X157),X162),s(fun(X157,X159),X160))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X151,fun(X151,bool)),X152),s(fun(X151,X153),X154),s(fun(X153,X151),X155))) | s(bool,'F_0') = s(bool,X156)) )),
  inference(superposition,[],[f9424,f5187])).
fof(f9424,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X34,X32] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X28,fun(X28,bool)),X31),s(fun(X28,X25),X32),s(fun(X25,X28),X30)))) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X24,X25),X26))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X27,X28),'const.sum.++_3'(s(fun(X24,X27),X29),s(fun(X25,X28),X30),s('type.sum.sum'(X24,X25),X26))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X27,fun(X27,bool)),X33),s(fun(X27,X24),X34),s(fun(X24,X27),X29)))) )),
  inference(resolution,[],[f127,f119])).
fof(f31556,plain,(
  ( ! [X146,X144,X142,X140,X149,X147,X145,X143,X141,X139,X150,X148] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X145,X141),X146))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X147,X139),'const.sum.++_3'(s(fun(X145,X147),X148),s(fun(X141,X139),X143),s('type.sum.sum'(X145,X141),X146))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X147,fun(X147,bool)),X149),s(fun(X147,X145),X150),s(fun(X145,X147),X148))) | s(bool,'F_0') = s(bool,X144) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X139,fun(X139,bool)),X140),s(fun(X139,X141),X142),s(fun(X141,X139),X143)))) )),
  inference(subsumption_resolution,[],[f31535,f119])).
fof(f31535,plain,(
  ( ! [X146,X144,X142,X140,X149,X147,X145,X143,X141,X139,X150,X148] : (~p(s(bool,X144)) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X145,X141),X146))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X147,X139),'const.sum.++_3'(s(fun(X145,X147),X148),s(fun(X141,X139),X143),s('type.sum.sum'(X145,X141),X146))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X147,fun(X147,bool)),X149),s(fun(X147,X145),X150),s(fun(X145,X147),X148))) | s(bool,'F_0') = s(bool,X144) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X139,fun(X139,bool)),X140),s(fun(X139,X141),X142),s(fun(X141,X139),X143)))) )),
  inference(superposition,[],[f9424,f5187])).
fof(f13824,plain,(
  ( ! [X165,X163,X173,X171,X169,X166,X164,X162,X172,X170,X168] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X164,X168),X169))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X162,X170),'const.sum.++_3'(s(fun(X164,X162),X166),s(fun(X168,X170),X171),s('type.sum.sum'(X164,X168),X169))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X162,fun(X162,bool)),X163),s(fun(X162,X164),X165),s(fun(X164,X162),X166))) | s(bool,'T_0') != s(bool,'const.quotient.QUOTIENT_3'(s(fun(X170,fun(X170,bool)),X172),s(fun(X170,X168),X173),s(fun(X168,X170),X171)))) ) | $spl246),
  inference(cnf_transformation,[],[f13824_D])).
fof(f13824_D,plain,(
  ( ! [X165,X163,X173,X171,X169,X166,X164,X162,X172,X170,X168] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X164,X168),X169))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X162,X170),'const.sum.++_3'(s(fun(X164,X162),X166),s(fun(X168,X170),X171),s('type.sum.sum'(X164,X168),X169))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X162,fun(X162,bool)),X163),s(fun(X162,X164),X165),s(fun(X164,X162),X166))) | s(bool,'T_0') != s(bool,'const.quotient.QUOTIENT_3'(s(fun(X170,fun(X170,bool)),X172),s(fun(X170,X168),X173),s(fun(X168,X170),X171)))) ) <=> ~$spl246),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl246])])).
fof(f31495,plain,(
  ( ! [X158,X156,X154,X152,X161,X151,X159,X157,X155,X153,X162,X160] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X153,X157),X158))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X151,X159),'const.sum.++_3'(s(fun(X153,X151),X155),s(fun(X157,X159),X160),s('type.sum.sum'(X153,X157),X158))))) | s(bool,'T_0') != s(bool,'const.quotient.QUOTIENT_3'(s(fun(X159,fun(X159,bool)),X161),s(fun(X159,X157),X162),s(fun(X157,X159),X160))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X151,fun(X151,bool)),X152),s(fun(X151,X153),X154),s(fun(X153,X151),X155))) | s(bool,'F_0') = s(bool,X156)) )),
  inference(subsumption_resolution,[],[f31474,f119])).
fof(f31474,plain,(
  ( ! [X158,X156,X154,X152,X161,X151,X159,X157,X155,X153,X162,X160] : (~p(s(bool,X156)) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X153,X157),X158))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X151,X159),'const.sum.++_3'(s(fun(X153,X151),X155),s(fun(X157,X159),X160),s('type.sum.sum'(X153,X157),X158))))) | s(bool,'T_0') != s(bool,'const.quotient.QUOTIENT_3'(s(fun(X159,fun(X159,bool)),X161),s(fun(X159,X157),X162),s(fun(X157,X159),X160))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X151,fun(X151,bool)),X152),s(fun(X151,X153),X154),s(fun(X153,X151),X155))) | s(bool,'F_0') = s(bool,X156)) )),
  inference(superposition,[],[f9416,f5187])).
fof(f9416,plain,(
  ( ! [X47,X54,X52,X50,X48,X56,X46,X55,X53,X51,X49] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X49,fun(X49,bool)),X53),s(fun(X49,X46),X54),s(fun(X46,X49),X51)))) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X46,X47),X48))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X49,X50),'const.sum.++_3'(s(fun(X46,X49),X51),s(fun(X47,X50),X52),s('type.sum.sum'(X46,X47),X48))))) | s(bool,'T_0') != s(bool,'const.quotient.QUOTIENT_3'(s(fun(X50,fun(X50,bool)),X55),s(fun(X50,X47),X56),s(fun(X47,X50),X52)))) )),
  inference(resolution,[],[f127,f122])).
fof(f31494,plain,(
  ( ! [X146,X144,X142,X140,X149,X147,X145,X143,X141,X139,X150,X148] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X141,X145),X146))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X139,X147),'const.sum.++_3'(s(fun(X141,X139),X143),s(fun(X145,X147),X148),s('type.sum.sum'(X141,X145),X146))))) | s(bool,'T_0') != s(bool,'const.quotient.QUOTIENT_3'(s(fun(X147,fun(X147,bool)),X149),s(fun(X147,X145),X150),s(fun(X145,X147),X148))) | s(bool,'F_0') = s(bool,X144) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X139,fun(X139,bool)),X140),s(fun(X139,X141),X142),s(fun(X141,X139),X143)))) )),
  inference(subsumption_resolution,[],[f31473,f119])).
fof(f31473,plain,(
  ( ! [X146,X144,X142,X140,X149,X147,X145,X143,X141,X139,X150,X148] : (~p(s(bool,X144)) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X141,X145),X146))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X139,X147),'const.sum.++_3'(s(fun(X141,X139),X143),s(fun(X145,X147),X148),s('type.sum.sum'(X141,X145),X146))))) | s(bool,'T_0') != s(bool,'const.quotient.QUOTIENT_3'(s(fun(X147,fun(X147,bool)),X149),s(fun(X147,X145),X150),s(fun(X145,X147),X148))) | s(bool,'F_0') = s(bool,X144) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X139,fun(X139,bool)),X140),s(fun(X139,X141),X142),s(fun(X141,X139),X143)))) )),
  inference(superposition,[],[f9416,f5187])).
fof(f31433,plain,(
  ( ! [X158,X156,X154,X152,X161,X151,X159,X157,X155,X153,X162,X160] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X153,X157),X158))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X151,X159),'const.sum.++_3'(s(fun(X153,X151),X155),s(fun(X157,X159),X160),s('type.sum.sum'(X153,X157),X158))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X159,fun(X159,bool)),X161),s(fun(X159,X157),X162),s(fun(X157,X159),X160))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X151,fun(X151,bool)),X152),s(fun(X151,X153),X154),s(fun(X153,X151),X155))) | s(bool,'F_0') = s(bool,X156)) )),
  inference(subsumption_resolution,[],[f31412,f119])).
fof(f31412,plain,(
  ( ! [X158,X156,X154,X152,X161,X151,X159,X157,X155,X153,X162,X160] : (~p(s(bool,X156)) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X153,X157),X158))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X151,X159),'const.sum.++_3'(s(fun(X153,X151),X155),s(fun(X157,X159),X160),s('type.sum.sum'(X153,X157),X158))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X159,fun(X159,bool)),X161),s(fun(X159,X157),X162),s(fun(X157,X159),X160))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X151,fun(X151,bool)),X152),s(fun(X151,X153),X154),s(fun(X153,X151),X155))) | s(bool,'F_0') = s(bool,X156)) )),
  inference(superposition,[],[f9414,f5187])).
fof(f9414,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X34,X32] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X27,fun(X27,bool)),X31),s(fun(X27,X24),X32),s(fun(X24,X27),X29)))) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X24,X25),X26))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X27,X28),'const.sum.++_3'(s(fun(X24,X27),X29),s(fun(X25,X28),X30),s('type.sum.sum'(X24,X25),X26))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X28,fun(X28,bool)),X33),s(fun(X28,X25),X34),s(fun(X25,X28),X30)))) )),
  inference(resolution,[],[f127,f119])).
fof(f31432,plain,(
  ( ! [X146,X144,X142,X140,X149,X147,X145,X143,X141,X139,X150,X148] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X141,X145),X146))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X139,X147),'const.sum.++_3'(s(fun(X141,X139),X143),s(fun(X145,X147),X148),s('type.sum.sum'(X141,X145),X146))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X147,fun(X147,bool)),X149),s(fun(X147,X145),X150),s(fun(X145,X147),X148))) | s(bool,'F_0') = s(bool,X144) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X139,fun(X139,bool)),X140),s(fun(X139,X141),X142),s(fun(X141,X139),X143)))) )),
  inference(subsumption_resolution,[],[f31411,f119])).
fof(f31411,plain,(
  ( ! [X146,X144,X142,X140,X149,X147,X145,X143,X141,X139,X150,X148] : (~p(s(bool,X144)) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X141,X145),X146))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X139,X147),'const.sum.++_3'(s(fun(X141,X139),X143),s(fun(X145,X147),X148),s('type.sum.sum'(X141,X145),X146))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X147,fun(X147,bool)),X149),s(fun(X147,X145),X150),s(fun(X145,X147),X148))) | s(bool,'F_0') = s(bool,X144) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X139,fun(X139,bool)),X140),s(fun(X139,X141),X142),s(fun(X141,X139),X143)))) )),
  inference(superposition,[],[f9414,f5187])).
fof(f13741,plain,(
  ( ! [X103,X109,X105,X102,X108,X106,X104] : (~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X102,fun(X102,bool)),X103),s(fun(X102,fun(X102,bool)),X103),s('type.sum.sum'(X102,X102),X108),s('type.sum.sum'(X102,X102),X109)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X102,fun(X102,bool)),X103),s(fun(X102,X104),X105),s(fun(X104,X102),X106))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X102,X102),X108))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X102,X102),X109)))) ) | $spl242),
  inference(cnf_transformation,[],[f13741_D])).
fof(f13741_D,plain,(
  ( ! [X103,X109,X105,X102,X108,X106,X104] : (~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X102,fun(X102,bool)),X103),s(fun(X102,fun(X102,bool)),X103),s('type.sum.sum'(X102,X102),X108),s('type.sum.sum'(X102,X102),X109)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X102,fun(X102,bool)),X103),s(fun(X102,X104),X105),s(fun(X104,X102),X106))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X102,X102),X108))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X102,X102),X109)))) ) <=> ~$spl242),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl242])])).
fof(f31090,plain,(
  ( ! [X101,X99,X97,X95,X102,X100,X98,X96] : (~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X95,fun(X95,bool)),X96),s(fun(X95,fun(X95,bool)),X96),s('type.sum.sum'(X95,X95),X101),s('type.sum.sum'(X95,X95),X102)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X95,X95),X101))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X95,X95),X102))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X95,fun(X95,bool)),X96),s(fun(X95,X97),X98),s(fun(X97,X95),X99))) | s(bool,'F_0') = s(bool,X100)) )),
  inference(subsumption_resolution,[],[f31067,f119])).
fof(f31067,plain,(
  ( ! [X101,X99,X97,X95,X102,X100,X98,X96] : (~p(s(bool,X100)) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X95,fun(X95,bool)),X96),s(fun(X95,fun(X95,bool)),X96),s('type.sum.sum'(X95,X95),X101),s('type.sum.sum'(X95,X95),X102)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X95,X95),X101))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X95,X95),X102))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X95,fun(X95,bool)),X96),s(fun(X95,X97),X98),s(fun(X97,X95),X99))) | s(bool,'F_0') = s(bool,X100)) )),
  inference(superposition,[],[f9892,f5187])).
fof(f9892,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X0,fun(X0,bool)),X3),s(fun(X0,X4),X5),s(fun(X4,X0),X6)))) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X0,fun(X0,bool)),X3),s(fun(X0,fun(X0,bool)),X3),s('type.sum.sum'(X0,X0),X1),s('type.sum.sum'(X0,X0),X2)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X0,X0),X1))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X0,X0),X2)))) )),
  inference(factoring,[],[f130])).
fof(f130,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X2,X3),X10))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X2,X3),X11))) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X2,fun(X2,bool)),X4),s(fun(X3,fun(X3,bool)),X7),s('type.sum.sum'(X2,X3),X10),s('type.sum.sum'(X2,X3),X11)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X3,fun(X3,bool)),X7),s(fun(X3,X1),X8),s(fun(X1,X3),X9)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X2,fun(X2,bool)),X4),s(fun(X2,X0),X5),s(fun(X0,X2),X6))))) )),
  inference(cnf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X2,fun(X2,bool)),X4),s(fun(X2,X0),X5),s(fun(X0,X2),X6)))) | ! [X7,X8,X9] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X3,fun(X3,bool)),X7),s(fun(X3,X1),X8),s(fun(X1,X3),X9)))) | ! [X10,X11] : (~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X2,fun(X2,bool)),X4),s(fun(X3,fun(X3,bool)),X7),s('type.sum.sum'(X2,X3),X10),s('type.sum.sum'(X2,X3),X11)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X2,X3),X10))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X2,X3),X11))))))),
  inference(ennf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X2,fun(X2,bool)),X4),s(fun(X2,X0),X5),s(fun(X0,X2),X6)))) => ! [X7,X8,X9] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X3,fun(X3,bool)),X7),s(fun(X3,X1),X8),s(fun(X1,X3),X9)))) => ! [X10,X11] : (p(s(bool,'const.quotient_sum.+++_4'(s(fun(X2,fun(X2,bool)),X4),s(fun(X3,fun(X3,bool)),X7),s('type.sum.sum'(X2,X3),X10),s('type.sum.sum'(X2,X3),X11)))) => s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X2,X3),X10))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X2,X3),X11))))))),
  inference(flattening,[],[f58])).
fof(f58,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X2,fun(X2,bool)),X4),s(fun(X2,X0),X5),s(fun(X0,X2),X6)))) => ! [X7,X8,X9] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X3,fun(X3,bool)),X7),s(fun(X3,X1),X8),s(fun(X1,X3),X9)))) => ! [X10,X11] : (p(s(bool,'const.quotient_sum.+++_4'(s(fun(X2,fun(X2,bool)),X4),s(fun(X3,fun(X3,bool)),X7),s('type.sum.sum'(X2,X3),X10),s('type.sum.sum'(X2,X3),X11)))) => s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X2,X3),X10))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X2,X3),X11))))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6,X7,X8,X9] : ! [X10,X11,X12] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X8,fun(X8,bool)),X10),s(fun(X8,X6),X11),s(fun(X6,X8),X12)))) => ! [X13,X14,X15] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X9,fun(X9,bool)),X13),s(fun(X9,X7),X14),s(fun(X7,X9),X15)))) => ! [X16,X17] : (p(s(bool,'const.quotient_sum.+++_4'(s(fun(X8,fun(X8,bool)),X10),s(fun(X9,fun(X9,bool)),X13),s('type.sum.sum'(X8,X9),X16),s('type.sum.sum'(X8,X9),X17)))) => s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X8,X9),X16))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X8,X9),X17))))))),
  file('thm.quotient_sum.ISR_RSP','thm.quotient_sum.ISL_RSP')).
fof(f31089,plain,(
  ( ! [X94,X92,X90,X88,X87,X93,X91,X89] : (~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X87,fun(X87,bool)),X88),s(fun(X87,fun(X87,bool)),X88),s('type.sum.sum'(X87,X87),X93),s('type.sum.sum'(X87,X87),X94)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X87,X87),X93))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X87,X87),X94))) | s(bool,'F_0') = s(bool,X92) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X87,fun(X87,bool)),X88),s(fun(X87,X89),X90),s(fun(X89,X87),X91)))) )),
  inference(subsumption_resolution,[],[f31066,f119])).
fof(f31066,plain,(
  ( ! [X94,X92,X90,X88,X87,X93,X91,X89] : (~p(s(bool,X92)) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X87,fun(X87,bool)),X88),s(fun(X87,fun(X87,bool)),X88),s('type.sum.sum'(X87,X87),X93),s('type.sum.sum'(X87,X87),X94)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X87,X87),X93))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X87,X87),X94))) | s(bool,'F_0') = s(bool,X92) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X87,fun(X87,bool)),X88),s(fun(X87,X89),X90),s(fun(X89,X87),X91)))) )),
  inference(superposition,[],[f9892,f5187])).
fof(f8647,plain,(
  ( ! [X292,X288,X290,X287,X293,X289,X286] : (~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X286,fun(X286,bool)),X287),s(fun(sK3,fun(sK3,bool)),sK7),s('type.sum.sum'(X286,sK3),X292),s('type.sum.sum'(X286,sK3),X293)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X286,fun(X286,bool)),X287),s(fun(X286,X288),X289),s(fun(X288,X286),X290))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X286,sK3),X292))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X286,sK3),X293)))) ) | $spl68),
  inference(cnf_transformation,[],[f8647_D])).
fof(f8647_D,plain,(
  ( ! [X292,X288,X290,X287,X293,X289,X286] : (~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X286,fun(X286,bool)),X287),s(fun(sK3,fun(sK3,bool)),sK7),s('type.sum.sum'(X286,sK3),X292),s('type.sum.sum'(X286,sK3),X293)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X286,fun(X286,bool)),X287),s(fun(X286,X288),X289),s(fun(X288,X286),X290))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X286,sK3),X292))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X286,sK3),X293)))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f30496,plain,(
  ( ! [X101,X99,X97,X95,X102,X100,X98,X96] : (~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X95,fun(X95,bool)),X96),s(fun(sK3,fun(sK3,bool)),sK7),s('type.sum.sum'(X95,sK3),X101),s('type.sum.sum'(X95,sK3),X102)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X95,sK3),X101))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X95,sK3),X102))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X95,fun(X95,bool)),X96),s(fun(X95,X97),X98),s(fun(X97,X95),X99))) | s(bool,'F_0') = s(bool,X100)) )),
  inference(subsumption_resolution,[],[f30474,f119])).
fof(f30474,plain,(
  ( ! [X101,X99,X97,X95,X102,X100,X98,X96] : (~p(s(bool,X100)) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X95,fun(X95,bool)),X96),s(fun(sK3,fun(sK3,bool)),sK7),s('type.sum.sum'(X95,sK3),X101),s('type.sum.sum'(X95,sK3),X102)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X95,sK3),X101))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X95,sK3),X102))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X95,fun(X95,bool)),X96),s(fun(X95,X97),X98),s(fun(X97,X95),X99))) | s(bool,'F_0') = s(bool,X100)) )),
  inference(superposition,[],[f301,f5187])).
fof(f301,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X7,fun(X7,bool)),X10),s(fun(X7,X11),X12),s(fun(X11,X7),X13)))) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X7,fun(X7,bool)),X10),s(fun(sK3,fun(sK3,bool)),sK7),s('type.sum.sum'(X7,sK3),X8),s('type.sum.sum'(X7,sK3),X9)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X7,sK3),X8))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X7,sK3),X9)))) )),
  inference(resolution,[],[f108,f130])).
fof(f108,plain,(
  p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK3,fun(sK3,bool)),sK7),s(fun(sK3,sK1),sK8),s(fun(sK1,sK3),sK9))))),
  inference(cnf_transformation,[],[f89])).
fof(f30495,plain,(
  ( ! [X94,X92,X90,X88,X87,X93,X91,X89] : (~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X87,fun(X87,bool)),X88),s(fun(sK3,fun(sK3,bool)),sK7),s('type.sum.sum'(X87,sK3),X93),s('type.sum.sum'(X87,sK3),X94)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X87,sK3),X93))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X87,sK3),X94))) | s(bool,'F_0') = s(bool,X92) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X87,fun(X87,bool)),X88),s(fun(X87,X89),X90),s(fun(X89,X87),X91)))) )),
  inference(subsumption_resolution,[],[f30473,f119])).
fof(f30473,plain,(
  ( ! [X94,X92,X90,X88,X87,X93,X91,X89] : (~p(s(bool,X92)) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X87,fun(X87,bool)),X88),s(fun(sK3,fun(sK3,bool)),sK7),s('type.sum.sum'(X87,sK3),X93),s('type.sum.sum'(X87,sK3),X94)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X87,sK3),X93))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X87,sK3),X94))) | s(bool,'F_0') = s(bool,X92) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X87,fun(X87,bool)),X88),s(fun(X87,X89),X90),s(fun(X89,X87),X91)))) )),
  inference(superposition,[],[f301,f5187])).
fof(f30181,plain,(
  ( ! [X42,X40] : (s(bool,X40) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X42))))) | s(bool,'T_0') = s(bool,X40)) ) | $spl314),
  inference(cnf_transformation,[],[f30181_D])).
fof(f30181_D,plain,(
  ( ! [X42,X40] : (s(bool,X40) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X42))))) | s(bool,'T_0') = s(bool,X40)) ) <=> ~$spl314),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl314])])).
fof(f30443,plain,(
  ( ! [X64,X65,X63] : (s(bool,'T_0') != s(bool,X64) | s(bool,'T_0') = s(bool,X65) | s(bool,X65) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X63)))))) )),
  inference(subsumption_resolution,[],[f30397,f21926])).
fof(f30397,plain,(
  ( ! [X64,X65,X63] : (s(bool,'T_0') != s(bool,X64) | s(bool,'F_0') = s(bool,X64) | s(bool,'T_0') = s(bool,X65) | s(bool,X65) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X63)))))) )),
  inference(superposition,[],[f24320,f19171])).
fof(f24320,plain,(
  ( ! [X142,X143] : (s(bool,X143) != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X142))))) | s(bool,'F_0') = s(bool,X143)) )),
  inference(subsumption_resolution,[],[f24300,f150])).
fof(f24300,plain,(
  ( ! [X142,X143] : (p(s(bool,'F_0')) | s(bool,X143) != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X142))))) | s(bool,'F_0') = s(bool,X143)) )),
  inference(superposition,[],[f22799,f2351])).
fof(f2351,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X0)))))) )),
  inference(resolution,[],[f2288,f129])).
fof(f2288,plain,(
  ( ! [X21] : (~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X21)))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X21)))) )),
  inference(trivial_inequality_removal,[],[f2284])).
fof(f2284,plain,(
  ( ! [X21] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X21))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X21))))) )),
  inference(superposition,[],[f2232,f125])).
fof(f125,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f95])).
fof(f95,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X22] : (s(bool,'T_0') = s(bool,X22) <=> p(s(bool,X22)))),
  file('thm.quotient_sum.ISR_RSP','thm.bool.EQ_CLAUSES|split|1')).
fof(f2232,plain,(
  ( ! [X5] : (s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X5))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X5)))) )),
  inference(resolution,[],[f2215,f119])).
fof(f2215,plain,(
  ( ! [X6] : (~p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X6)))) | s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X6)))) )),
  inference(superposition,[],[f133,f1692])).
fof(f1692,plain,(
  ( ! [X0] : (s('type.sum.sum'(sK0,sK0),X0) = s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,sK13(X0,sK0,sK0)))) | s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X0)))) )),
  inference(superposition,[],[f1592,f111])).
fof(f30442,plain,(
  ( ! [X61,X62,X60] : (s(bool,'T_0') != s(bool,X61) | s(bool,X62) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X60))))) | s(bool,'T_0') = s(bool,X62)) )),
  inference(subsumption_resolution,[],[f30396,f21926])).
fof(f30396,plain,(
  ( ! [X61,X62,X60] : (s(bool,'T_0') != s(bool,X61) | s(bool,'F_0') = s(bool,X61) | s(bool,X62) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X60))))) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f24320,f19171])).
fof(f30285,plain,(
  ( ! [X64,X65,X63] : (s(bool,'T_0') != s(bool,X64) | s(bool,'T_0') = s(bool,X65) | s(bool,X65) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X63)))))) )),
  inference(duplicate_literal_removal,[],[f30274])).
fof(f30274,plain,(
  ( ! [X64,X65,X63] : (s(bool,'T_0') != s(bool,X64) | s(bool,'T_0') != s(bool,X64) | s(bool,'T_0') = s(bool,X65) | s(bool,X65) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X63)))))) )),
  inference(superposition,[],[f24238,f19171])).
fof(f24238,plain,(
  ( ! [X142,X143] : (s(bool,X143) != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X142))))) | s(bool,'T_0') != s(bool,X143)) )),
  inference(subsumption_resolution,[],[f24212,f150])).
fof(f24212,plain,(
  ( ! [X142,X143] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,X143) | s(bool,X143) != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X142)))))) )),
  inference(superposition,[],[f22526,f2351])).
fof(f30286,plain,(
  ( ! [X61,X62,X60] : (s(bool,'T_0') != s(bool,X61) | s(bool,X62) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X60))))) | s(bool,'T_0') = s(bool,X62)) )),
  inference(duplicate_literal_removal,[],[f30273])).
fof(f30273,plain,(
  ( ! [X61,X62,X60] : (s(bool,'T_0') != s(bool,X61) | s(bool,'T_0') != s(bool,X61) | s(bool,X62) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X60))))) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f24238,f19171])).
fof(f30298,plain,(
  ( ! [X30,X29] : (s(bool,X29) != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X30))))) | s(bool,'F_0') = s(bool,X29)) ) | $spl316),
  inference(cnf_transformation,[],[f30298_D])).
fof(f30298_D,plain,(
  ( ! [X30,X29] : (s(bool,X29) != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X30))))) | s(bool,'F_0') = s(bool,X29)) ) <=> ~$spl316),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl316])])).
fof(f30299,plain,(
  ( ! [X33,X31,X32] : (s(bool,X32) != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X33))))) | s(bool,'T_0') != s(bool,X31) | s(bool,'F_0') = s(bool,X32)) )),
  inference(subsumption_resolution,[],[f30221,f21926])).
fof(f30221,plain,(
  ( ! [X33,X31,X32] : (s(bool,X32) != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X33))))) | s(bool,'T_0') != s(bool,X31) | s(bool,'F_0') = s(bool,X31) | s(bool,'F_0') = s(bool,X32)) )),
  inference(superposition,[],[f24238,f5187])).
fof(f30297,plain,(
  ( ! [X30,X28,X29] : (s(bool,X29) != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X30))))) | s(bool,'T_0') != s(bool,X28) | s(bool,'F_0') = s(bool,X29)) )),
  inference(subsumption_resolution,[],[f30220,f21926])).
fof(f30220,plain,(
  ( ! [X30,X28,X29] : (s(bool,X29) != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X30))))) | s(bool,'T_0') != s(bool,X28) | s(bool,'F_0') = s(bool,X29) | s(bool,'F_0') = s(bool,X28)) )),
  inference(superposition,[],[f24238,f5187])).
fof(f30182,plain,(
  ( ! [X45,X43,X44] : (s(bool,X43) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X45))))) | s(bool,'T_0') = s(bool,X44) | s(bool,'T_0') = s(bool,X43)) )),
  inference(subsumption_resolution,[],[f30111,f126])).
fof(f30111,plain,(
  ( ! [X45,X43,X44] : (s(bool,'F_0') != s(bool,X44) | s(bool,X43) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X45))))) | s(bool,'T_0') = s(bool,X44) | s(bool,'T_0') = s(bool,X43)) )),
  inference(superposition,[],[f23991,f19171])).
fof(f23991,plain,(
  ( ! [X142,X143] : (s(bool,'F_0') != s(bool,X143) | s(bool,X143) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X142)))))) )),
  inference(subsumption_resolution,[],[f23971,f150])).
fof(f23971,plain,(
  ( ! [X142,X143] : (p(s(bool,'F_0')) | s(bool,X143) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X142))))) | s(bool,'F_0') != s(bool,X143)) )),
  inference(superposition,[],[f21924,f2351])).
fof(f30180,plain,(
  ( ! [X41,X42,X40] : (s(bool,X40) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X42))))) | s(bool,'T_0') = s(bool,X40) | s(bool,'T_0') = s(bool,X41)) )),
  inference(subsumption_resolution,[],[f30110,f126])).
fof(f30110,plain,(
  ( ! [X41,X42,X40] : (s(bool,'F_0') != s(bool,X41) | s(bool,X40) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X42))))) | s(bool,'T_0') = s(bool,X40) | s(bool,'T_0') = s(bool,X41)) )),
  inference(superposition,[],[f23991,f19171])).
fof(f8645,plain,(
  ( ! [X280,X282,X279,X285,X281,X278,X284] : (~p(s(bool,'const.quotient_sum.+++_4'(s(fun(sK3,fun(sK3,bool)),sK7),s(fun(X278,fun(X278,bool)),X279),s('type.sum.sum'(sK3,X278),X284),s('type.sum.sum'(sK3,X278),X285)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X278,fun(X278,bool)),X279),s(fun(X278,X280),X281),s(fun(X280,X278),X282))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK3,X278),X284))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK3,X278),X285)))) ) | $spl66),
  inference(cnf_transformation,[],[f8645_D])).
fof(f8645_D,plain,(
  ( ! [X280,X282,X279,X285,X281,X278,X284] : (~p(s(bool,'const.quotient_sum.+++_4'(s(fun(sK3,fun(sK3,bool)),sK7),s(fun(X278,fun(X278,bool)),X279),s('type.sum.sum'(sK3,X278),X284),s('type.sum.sum'(sK3,X278),X285)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X278,fun(X278,bool)),X279),s(fun(X278,X280),X281),s(fun(X280,X278),X282))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK3,X278),X284))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK3,X278),X285)))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f30059,plain,(
  ( ! [X101,X99,X97,X95,X102,X100,X98,X96] : (~p(s(bool,'const.quotient_sum.+++_4'(s(fun(sK3,fun(sK3,bool)),sK7),s(fun(X95,fun(X95,bool)),X96),s('type.sum.sum'(sK3,X95),X101),s('type.sum.sum'(sK3,X95),X102)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK3,X95),X101))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK3,X95),X102))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X95,fun(X95,bool)),X96),s(fun(X95,X97),X98),s(fun(X97,X95),X99))) | s(bool,'F_0') = s(bool,X100)) )),
  inference(subsumption_resolution,[],[f30037,f119])).
fof(f30037,plain,(
  ( ! [X101,X99,X97,X95,X102,X100,X98,X96] : (~p(s(bool,X100)) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(sK3,fun(sK3,bool)),sK7),s(fun(X95,fun(X95,bool)),X96),s('type.sum.sum'(sK3,X95),X101),s('type.sum.sum'(sK3,X95),X102)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK3,X95),X101))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK3,X95),X102))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X95,fun(X95,bool)),X96),s(fun(X95,X97),X98),s(fun(X97,X95),X99))) | s(bool,'F_0') = s(bool,X100)) )),
  inference(superposition,[],[f300,f5187])).
fof(f300,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X0,fun(X0,bool)),X3),s(fun(X0,X4),X5),s(fun(X4,X0),X6)))) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(sK3,fun(sK3,bool)),sK7),s(fun(X0,fun(X0,bool)),X3),s('type.sum.sum'(sK3,X0),X1),s('type.sum.sum'(sK3,X0),X2)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK3,X0),X1))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK3,X0),X2)))) )),
  inference(resolution,[],[f108,f130])).
fof(f30058,plain,(
  ( ! [X94,X92,X90,X88,X87,X93,X91,X89] : (~p(s(bool,'const.quotient_sum.+++_4'(s(fun(sK3,fun(sK3,bool)),sK7),s(fun(X87,fun(X87,bool)),X88),s('type.sum.sum'(sK3,X87),X93),s('type.sum.sum'(sK3,X87),X94)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK3,X87),X93))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK3,X87),X94))) | s(bool,'F_0') = s(bool,X92) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X87,fun(X87,bool)),X88),s(fun(X87,X89),X90),s(fun(X89,X87),X91)))) )),
  inference(subsumption_resolution,[],[f30036,f119])).
fof(f30036,plain,(
  ( ! [X94,X92,X90,X88,X87,X93,X91,X89] : (~p(s(bool,X92)) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(sK3,fun(sK3,bool)),sK7),s(fun(X87,fun(X87,bool)),X88),s('type.sum.sum'(sK3,X87),X93),s('type.sum.sum'(sK3,X87),X94)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK3,X87),X93))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK3,X87),X94))) | s(bool,'F_0') = s(bool,X92) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X87,fun(X87,bool)),X88),s(fun(X87,X89),X90),s(fun(X89,X87),X91)))) )),
  inference(superposition,[],[f300,f5187])).
fof(f8641,plain,(
  ( ! [X263,X257,X259,X260,X262,X256,X258] : (~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X256,fun(X256,bool)),X257),s(fun(sK2,fun(sK2,bool)),sK4),s('type.sum.sum'(X256,sK2),X262),s('type.sum.sum'(X256,sK2),X263)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X256,fun(X256,bool)),X257),s(fun(X256,X258),X259),s(fun(X258,X256),X260))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X256,sK2),X262))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X256,sK2),X263)))) ) | $spl62),
  inference(cnf_transformation,[],[f8641_D])).
fof(f8641_D,plain,(
  ( ! [X263,X257,X259,X260,X262,X256,X258] : (~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X256,fun(X256,bool)),X257),s(fun(sK2,fun(sK2,bool)),sK4),s('type.sum.sum'(X256,sK2),X262),s('type.sum.sum'(X256,sK2),X263)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X256,fun(X256,bool)),X257),s(fun(X256,X258),X259),s(fun(X258,X256),X260))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X256,sK2),X262))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X256,sK2),X263)))) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f29496,plain,(
  ( ! [X101,X99,X97,X95,X102,X100,X98,X96] : (~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X95,fun(X95,bool)),X96),s(fun(sK2,fun(sK2,bool)),sK4),s('type.sum.sum'(X95,sK2),X101),s('type.sum.sum'(X95,sK2),X102)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X95,sK2),X101))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X95,sK2),X102))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X95,fun(X95,bool)),X96),s(fun(X95,X97),X98),s(fun(X97,X95),X99))) | s(bool,'F_0') = s(bool,X100)) )),
  inference(subsumption_resolution,[],[f29475,f119])).
fof(f29475,plain,(
  ( ! [X101,X99,X97,X95,X102,X100,X98,X96] : (~p(s(bool,X100)) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X95,fun(X95,bool)),X96),s(fun(sK2,fun(sK2,bool)),sK4),s('type.sum.sum'(X95,sK2),X101),s('type.sum.sum'(X95,sK2),X102)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X95,sK2),X101))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X95,sK2),X102))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X95,fun(X95,bool)),X96),s(fun(X95,X97),X98),s(fun(X97,X95),X99))) | s(bool,'F_0') = s(bool,X100)) )),
  inference(superposition,[],[f180,f5187])).
fof(f180,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X7,fun(X7,bool)),X10),s(fun(X7,X11),X12),s(fun(X11,X7),X13)))) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X7,fun(X7,bool)),X10),s(fun(sK2,fun(sK2,bool)),sK4),s('type.sum.sum'(X7,sK2),X8),s('type.sum.sum'(X7,sK2),X9)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X7,sK2),X8))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X7,sK2),X9)))) )),
  inference(resolution,[],[f107,f130])).
fof(f29495,plain,(
  ( ! [X94,X92,X90,X88,X87,X93,X91,X89] : (~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X87,fun(X87,bool)),X88),s(fun(sK2,fun(sK2,bool)),sK4),s('type.sum.sum'(X87,sK2),X93),s('type.sum.sum'(X87,sK2),X94)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X87,sK2),X93))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X87,sK2),X94))) | s(bool,'F_0') = s(bool,X92) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X87,fun(X87,bool)),X88),s(fun(X87,X89),X90),s(fun(X89,X87),X91)))) )),
  inference(subsumption_resolution,[],[f29474,f119])).
fof(f29474,plain,(
  ( ! [X94,X92,X90,X88,X87,X93,X91,X89] : (~p(s(bool,X92)) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X87,fun(X87,bool)),X88),s(fun(sK2,fun(sK2,bool)),sK4),s('type.sum.sum'(X87,sK2),X93),s('type.sum.sum'(X87,sK2),X94)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X87,sK2),X93))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X87,sK2),X94))) | s(bool,'F_0') = s(bool,X92) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X87,fun(X87,bool)),X88),s(fun(X87,X89),X90),s(fun(X89,X87),X91)))) )),
  inference(superposition,[],[f180,f5187])).
fof(f29196,plain,(
  ( ! [X42,X40] : (s(bool,X40) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,X42))))) | s(bool,'T_0') = s(bool,X40)) ) | $spl310),
  inference(cnf_transformation,[],[f29196_D])).
fof(f29196_D,plain,(
  ( ! [X42,X40] : (s(bool,X40) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,X42))))) | s(bool,'T_0') = s(bool,X40)) ) <=> ~$spl310),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl310])])).
fof(f29444,plain,(
  ( ! [X43,X41,X42] : (s(bool,'T_0') != s(bool,X42) | s(bool,'T_0') = s(bool,X43) | s(bool,X43) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,X41)))))) )),
  inference(subsumption_resolution,[],[f29399,f21926])).
fof(f29399,plain,(
  ( ! [X43,X41,X42] : (s(bool,'T_0') != s(bool,X42) | s(bool,'F_0') = s(bool,X42) | s(bool,'T_0') = s(bool,X43) | s(bool,X43) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,X41)))))) )),
  inference(superposition,[],[f24315,f19171])).
fof(f24315,plain,(
  ( ! [X99,X98] : (s(bool,X99) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,X98))))) | s(bool,'F_0') = s(bool,X99)) )),
  inference(subsumption_resolution,[],[f24290,f150])).
fof(f24290,plain,(
  ( ! [X99,X98] : (p(s(bool,'F_0')) | s(bool,X99) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,X98))))) | s(bool,'F_0') = s(bool,X99)) )),
  inference(superposition,[],[f22799,f1580])).
fof(f1580,plain,(
  ( ! [X9] : (s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,X9)))))) )),
  inference(subsumption_resolution,[],[f1532,f128])).
fof(f1532,plain,(
  ( ! [X9] : (s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,X9))))) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK2),'const.sum.INL_1'(s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,X9))))))))) )),
  inference(superposition,[],[f1442,f119])).
fof(f29443,plain,(
  ( ! [X39,X38,X40] : (s(bool,'T_0') != s(bool,X39) | s(bool,X40) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,X38))))) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f29398,f21926])).
fof(f29398,plain,(
  ( ! [X39,X38,X40] : (s(bool,'T_0') != s(bool,X39) | s(bool,'F_0') = s(bool,X39) | s(bool,X40) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,X38))))) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f24315,f19171])).
fof(f29292,plain,(
  ( ! [X43,X41,X42] : (s(bool,'T_0') != s(bool,X42) | s(bool,'T_0') = s(bool,X43) | s(bool,X43) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,X41)))))) )),
  inference(duplicate_literal_removal,[],[f29282])).
fof(f29282,plain,(
  ( ! [X43,X41,X42] : (s(bool,'T_0') != s(bool,X42) | s(bool,'T_0') != s(bool,X42) | s(bool,'T_0') = s(bool,X43) | s(bool,X43) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,X41)))))) )),
  inference(superposition,[],[f24233,f19171])).
fof(f24233,plain,(
  ( ! [X99,X98] : (s(bool,X99) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,X98))))) | s(bool,'T_0') != s(bool,X99)) )),
  inference(subsumption_resolution,[],[f24202,f150])).
fof(f24202,plain,(
  ( ! [X99,X98] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,X99) | s(bool,X99) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,X98)))))) )),
  inference(superposition,[],[f22526,f1580])).
fof(f29293,plain,(
  ( ! [X39,X38,X40] : (s(bool,'T_0') != s(bool,X39) | s(bool,X40) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,X38))))) | s(bool,'T_0') = s(bool,X40)) )),
  inference(duplicate_literal_removal,[],[f29281])).
fof(f29281,plain,(
  ( ! [X39,X38,X40] : (s(bool,'T_0') != s(bool,X39) | s(bool,'T_0') != s(bool,X39) | s(bool,X40) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,X38))))) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f24233,f19171])).
fof(f29305,plain,(
  ( ! [X30,X29] : (s(bool,X29) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,X30))))) | s(bool,'F_0') = s(bool,X29)) ) | $spl312),
  inference(cnf_transformation,[],[f29305_D])).
fof(f29305_D,plain,(
  ( ! [X30,X29] : (s(bool,X29) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,X30))))) | s(bool,'F_0') = s(bool,X29)) ) <=> ~$spl312),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl312])])).
fof(f29306,plain,(
  ( ! [X33,X31,X32] : (s(bool,X32) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,X33))))) | s(bool,'T_0') != s(bool,X31) | s(bool,'F_0') = s(bool,X32)) )),
  inference(subsumption_resolution,[],[f29233,f21926])).
fof(f29233,plain,(
  ( ! [X33,X31,X32] : (s(bool,X32) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,X33))))) | s(bool,'T_0') != s(bool,X31) | s(bool,'F_0') = s(bool,X31) | s(bool,'F_0') = s(bool,X32)) )),
  inference(superposition,[],[f24233,f5187])).
fof(f29304,plain,(
  ( ! [X30,X28,X29] : (s(bool,X29) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,X30))))) | s(bool,'T_0') != s(bool,X28) | s(bool,'F_0') = s(bool,X29)) )),
  inference(subsumption_resolution,[],[f29232,f21926])).
fof(f29232,plain,(
  ( ! [X30,X28,X29] : (s(bool,X29) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,X30))))) | s(bool,'T_0') != s(bool,X28) | s(bool,'F_0') = s(bool,X29) | s(bool,'F_0') = s(bool,X28)) )),
  inference(superposition,[],[f24233,f5187])).
fof(f29197,plain,(
  ( ! [X45,X43,X44] : (s(bool,X43) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,X45))))) | s(bool,'T_0') = s(bool,X44) | s(bool,'T_0') = s(bool,X43)) )),
  inference(subsumption_resolution,[],[f29128,f126])).
fof(f29128,plain,(
  ( ! [X45,X43,X44] : (s(bool,'F_0') != s(bool,X44) | s(bool,X43) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,X45))))) | s(bool,'T_0') = s(bool,X44) | s(bool,'T_0') = s(bool,X43)) )),
  inference(superposition,[],[f23986,f19171])).
fof(f23986,plain,(
  ( ! [X99,X98] : (s(bool,'F_0') != s(bool,X99) | s(bool,X99) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,X98)))))) )),
  inference(subsumption_resolution,[],[f23961,f150])).
fof(f23961,plain,(
  ( ! [X99,X98] : (p(s(bool,'F_0')) | s(bool,X99) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,X98))))) | s(bool,'F_0') != s(bool,X99)) )),
  inference(superposition,[],[f21924,f1580])).
fof(f29195,plain,(
  ( ! [X41,X42,X40] : (s(bool,X40) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,X42))))) | s(bool,'T_0') = s(bool,X40) | s(bool,'T_0') = s(bool,X41)) )),
  inference(subsumption_resolution,[],[f29127,f126])).
fof(f29127,plain,(
  ( ! [X41,X42,X40] : (s(bool,'F_0') != s(bool,X41) | s(bool,X40) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,X42))))) | s(bool,'T_0') = s(bool,X40) | s(bool,'T_0') = s(bool,X41)) )),
  inference(superposition,[],[f23986,f19171])).
fof(f8639,plain,(
  ( ! [X255,X251,X249,X254,X252,X250,X248] : (~p(s(bool,'const.quotient_sum.+++_4'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(X248,fun(X248,bool)),X249),s('type.sum.sum'(sK2,X248),X254),s('type.sum.sum'(sK2,X248),X255)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X248,fun(X248,bool)),X249),s(fun(X248,X250),X251),s(fun(X250,X248),X252))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,X248),X254))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,X248),X255)))) ) | $spl60),
  inference(cnf_transformation,[],[f8639_D])).
fof(f8639_D,plain,(
  ( ! [X255,X251,X249,X254,X252,X250,X248] : (~p(s(bool,'const.quotient_sum.+++_4'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(X248,fun(X248,bool)),X249),s('type.sum.sum'(sK2,X248),X254),s('type.sum.sum'(sK2,X248),X255)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X248,fun(X248,bool)),X249),s(fun(X248,X250),X251),s(fun(X250,X248),X252))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,X248),X254))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,X248),X255)))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f29077,plain,(
  ( ! [X101,X99,X97,X95,X102,X100,X98,X96] : (~p(s(bool,'const.quotient_sum.+++_4'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(X95,fun(X95,bool)),X96),s('type.sum.sum'(sK2,X95),X101),s('type.sum.sum'(sK2,X95),X102)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,X95),X101))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,X95),X102))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X95,fun(X95,bool)),X96),s(fun(X95,X97),X98),s(fun(X97,X95),X99))) | s(bool,'F_0') = s(bool,X100)) )),
  inference(subsumption_resolution,[],[f29056,f119])).
fof(f29056,plain,(
  ( ! [X101,X99,X97,X95,X102,X100,X98,X96] : (~p(s(bool,X100)) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(X95,fun(X95,bool)),X96),s('type.sum.sum'(sK2,X95),X101),s('type.sum.sum'(sK2,X95),X102)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,X95),X101))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,X95),X102))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X95,fun(X95,bool)),X96),s(fun(X95,X97),X98),s(fun(X97,X95),X99))) | s(bool,'F_0') = s(bool,X100)) )),
  inference(superposition,[],[f179,f5187])).
fof(f179,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X0,fun(X0,bool)),X3),s(fun(X0,X4),X5),s(fun(X4,X0),X6)))) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(X0,fun(X0,bool)),X3),s('type.sum.sum'(sK2,X0),X1),s('type.sum.sum'(sK2,X0),X2)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,X0),X1))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,X0),X2)))) )),
  inference(resolution,[],[f107,f130])).
fof(f29076,plain,(
  ( ! [X94,X92,X90,X88,X87,X93,X91,X89] : (~p(s(bool,'const.quotient_sum.+++_4'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(X87,fun(X87,bool)),X88),s('type.sum.sum'(sK2,X87),X93),s('type.sum.sum'(sK2,X87),X94)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,X87),X93))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,X87),X94))) | s(bool,'F_0') = s(bool,X92) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X87,fun(X87,bool)),X88),s(fun(X87,X89),X90),s(fun(X89,X87),X91)))) )),
  inference(subsumption_resolution,[],[f29055,f119])).
fof(f29055,plain,(
  ( ! [X94,X92,X90,X88,X87,X93,X91,X89] : (~p(s(bool,X92)) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(X87,fun(X87,bool)),X88),s('type.sum.sum'(sK2,X87),X93),s('type.sum.sum'(sK2,X87),X94)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,X87),X93))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,X87),X94))) | s(bool,'F_0') = s(bool,X92) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X87,fun(X87,bool)),X88),s(fun(X87,X89),X90),s(fun(X89,X87),X91)))) )),
  inference(superposition,[],[f179,f5187])).
fof(f28555,plain,(
  ( ! [X182,X184,X183,X181] : (s(bool,X184) != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X181,X182),'const.sum.INL_1'(s(X181,X183))))) | s(bool,'T_0') = s(bool,X184)) ) | $spl308),
  inference(cnf_transformation,[],[f28555_D])).
fof(f28555_D,plain,(
  ( ! [X182,X184,X183,X181] : (s(bool,X184) != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X181,X182),'const.sum.INL_1'(s(X181,X183))))) | s(bool,'T_0') = s(bool,X184)) ) <=> ~$spl308),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl308])])).
fof(f28654,plain,(
  ( ! [X241,X239,X242,X240,X238] : (s(bool,'T_0') != s(bool,X241) | s(bool,'T_0') = s(bool,X242) | s(bool,X242) != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X238,X239),'const.sum.INL_1'(s(X238,X240)))))) ) | $spl92),
  inference(subsumption_resolution,[],[f28444,f21926])).
fof(f28444,plain,(
  ( ! [X241,X239,X242,X240,X238] : (s(bool,'T_0') != s(bool,X241) | s(bool,'T_0') = s(bool,X242) | s(bool,X242) != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X238,X239),'const.sum.INL_1'(s(X238,X240))))) | s(bool,'F_0') = s(bool,X241)) ) | $spl92),
  inference(superposition,[],[f15157,f24080])).
fof(f24080,plain,(
  ( ! [X109,X107,X108,X106] : (s(bool,X109) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X106,X107),'const.sum.INL_1'(s(X106,X108))))) | s(bool,'F_0') = s(bool,X109)) )),
  inference(subsumption_resolution,[],[f24047,f155])).
fof(f24047,plain,(
  ( ! [X109,X107,X108,X106] : (~p(s(bool,'T_0')) | s(bool,X109) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X106,X107),'const.sum.INL_1'(s(X106,X108))))) | s(bool,'F_0') = s(bool,X109)) )),
  inference(superposition,[],[f22078,f9229])).
fof(f9229,plain,(
  ( ! [X14,X15,X16] : (s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X14,X15),'const.sum.INL_1'(s(X14,X16)))))) )),
  inference(resolution,[],[f134,f123])).
fof(f22078,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X0)) )),
  inference(resolution,[],[f119,f115])).
fof(f15157,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') != s(bool,X10) | s(bool,'T_0') = s(bool,X11) | s(bool,X10) != s(bool,X11)) ) | $spl92),
  inference(resolution,[],[f11882,f122])).
fof(f28652,plain,(
  ( ! [X212,X210,X213,X211,X214] : (s(bool,X214) != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X210,X211),'const.sum.INL_1'(s(X210,X212))))) | s(bool,'T_0') = s(bool,X214) | s(bool,'F_0') = s(bool,X213)) ) | $spl92),
  inference(subsumption_resolution,[],[f28438,f119])).
fof(f28438,plain,(
  ( ! [X212,X210,X213,X211,X214] : (~p(s(bool,X213)) | s(bool,X214) != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X210,X211),'const.sum.INL_1'(s(X210,X212))))) | s(bool,'T_0') = s(bool,X214) | s(bool,'F_0') = s(bool,X213)) ) | $spl92),
  inference(superposition,[],[f11882,f24080])).
fof(f28493,plain,(
  ( ! [X182,X180,X184,X183,X181] : (s(bool,X184) != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X181,X182),'const.sum.INL_1'(s(X181,X183))))) | s(bool,'F_0') = s(bool,X180) | s(bool,'T_0') = s(bool,X184)) ) | $spl160),
  inference(duplicate_literal_removal,[],[f28301])).
fof(f28301,plain,(
  ( ! [X182,X180,X184,X183,X181] : (s(bool,X184) != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X181,X182),'const.sum.INL_1'(s(X181,X183))))) | s(bool,'F_0') = s(bool,X180) | s(bool,'T_0') = s(bool,X184) | s(bool,'F_0') = s(bool,X180)) ) | $spl160),
  inference(superposition,[],[f15050,f24080])).
fof(f15050,plain,(
  ( ! [X10,X11] : (s(bool,X10) != s(bool,X11) | s(bool,'F_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X11)) ) | $spl160),
  inference(resolution,[],[f11978,f123])).
fof(f28011,plain,(
  ( ! [X182,X184,X183,X181] : (s(bool,X184) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X181,X182),'const.sum.INR_1'(s(X182,X183))))) | s(bool,'T_0') = s(bool,X184)) ) | $spl306),
  inference(cnf_transformation,[],[f28011_D])).
fof(f28011_D,plain,(
  ( ! [X182,X184,X183,X181] : (s(bool,X184) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X181,X182),'const.sum.INR_1'(s(X182,X183))))) | s(bool,'T_0') = s(bool,X184)) ) <=> ~$spl306),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl306])])).
fof(f28109,plain,(
  ( ! [X218,X220,X219,X217,X221] : (s(bool,'T_0') != s(bool,X220) | s(bool,'T_0') = s(bool,X221) | s(bool,X221) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X217,X218),'const.sum.INR_1'(s(X218,X219)))))) ) | $spl92),
  inference(subsumption_resolution,[],[f27905,f21926])).
fof(f27905,plain,(
  ( ! [X218,X220,X219,X217,X221] : (s(bool,'T_0') != s(bool,X220) | s(bool,'T_0') = s(bool,X221) | s(bool,X221) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X217,X218),'const.sum.INR_1'(s(X218,X219))))) | s(bool,'F_0') = s(bool,X220)) ) | $spl92),
  inference(superposition,[],[f15157,f24076])).
fof(f24076,plain,(
  ( ! [X76,X74,X75,X73] : (s(bool,X76) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X73,X74),'const.sum.INR_1'(s(X74,X75))))) | s(bool,'F_0') = s(bool,X76)) )),
  inference(subsumption_resolution,[],[f24037,f155])).
fof(f24037,plain,(
  ( ! [X76,X74,X75,X73] : (~p(s(bool,'T_0')) | s(bool,X76) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X73,X74),'const.sum.INR_1'(s(X74,X75))))) | s(bool,'F_0') = s(bool,X76)) )),
  inference(superposition,[],[f22078,f9130])).
fof(f9130,plain,(
  ( ! [X14,X15,X16] : (s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X14,X15),'const.sum.INR_1'(s(X15,X16)))))) )),
  inference(resolution,[],[f129,f123])).
fof(f28107,plain,(
  ( ! [X191,X189,X192,X190,X193] : (s(bool,X193) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X189,X190),'const.sum.INR_1'(s(X190,X191))))) | s(bool,'T_0') = s(bool,X193) | s(bool,'F_0') = s(bool,X192)) ) | $spl92),
  inference(subsumption_resolution,[],[f27899,f119])).
fof(f27899,plain,(
  ( ! [X191,X189,X192,X190,X193] : (~p(s(bool,X192)) | s(bool,X193) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X189,X190),'const.sum.INR_1'(s(X190,X191))))) | s(bool,'T_0') = s(bool,X193) | s(bool,'F_0') = s(bool,X192)) ) | $spl92),
  inference(superposition,[],[f11882,f24076])).
fof(f27951,plain,(
  ( ! [X182,X180,X184,X183,X181] : (s(bool,X184) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X181,X182),'const.sum.INR_1'(s(X182,X183))))) | s(bool,'F_0') = s(bool,X180) | s(bool,'T_0') = s(bool,X184)) ) | $spl160),
  inference(duplicate_literal_removal,[],[f27765])).
fof(f27765,plain,(
  ( ! [X182,X180,X184,X183,X181] : (s(bool,X184) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X181,X182),'const.sum.INR_1'(s(X182,X183))))) | s(bool,'F_0') = s(bool,X180) | s(bool,'T_0') = s(bool,X184) | s(bool,'F_0') = s(bool,X180)) ) | $spl160),
  inference(superposition,[],[f15050,f24076])).
fof(f27567,plain,(
  ( ! [X78,X76,X75] : (s(bool,X75) != s(bool,X76) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78) | p(s(bool,X76))) ) | $spl304),
  inference(cnf_transformation,[],[f27567_D])).
fof(f27567_D,plain,(
  ( ! [X78,X76,X75] : (s(bool,X75) != s(bool,X76) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78) | p(s(bool,X76))) ) <=> ~$spl304),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl304])])).
fof(f27568,plain,(
  ( ! [X80,X81,X79,X82] : (s(bool,X79) != s(bool,X80) | p(s(bool,X80)) | s(bool,'T_0') = s(bool,X81) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(subsumption_resolution,[],[f27468,f123])).
fof(f27468,plain,(
  ( ! [X80,X81,X79,X82] : (s(bool,X79) != s(bool,X80) | p(s(bool,X80)) | s(bool,'T_0') = s(bool,X81) | p(s(bool,X81)) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(superposition,[],[f22797,f19171])).
fof(f22797,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | p(s(bool,X0)) | s(bool,X1) = s(bool,X2) | p(s(bool,X2))) )),
  inference(resolution,[],[f22072,f112])).
fof(f27566,plain,(
  ( ! [X78,X76,X77,X75] : (s(bool,X75) != s(bool,X76) | p(s(bool,X76)) | s(bool,'T_0') = s(bool,X77) | s(bool,'T_0') = s(bool,X78) | s(bool,X75) = s(bool,X78)) )),
  inference(subsumption_resolution,[],[f27467,f123])).
fof(f27467,plain,(
  ( ! [X78,X76,X77,X75] : (s(bool,X75) != s(bool,X76) | p(s(bool,X76)) | s(bool,'T_0') = s(bool,X77) | p(s(bool,X77)) | s(bool,'T_0') = s(bool,X78) | s(bool,X75) = s(bool,X78)) )),
  inference(superposition,[],[f22797,f19171])).
fof(f27555,plain,(
  ( ! [X33,X31,X34] : (s(bool,X31) = s(bool,X33) | s(bool,'F_0') = s(bool,X34) | s(bool,X31) = s(bool,X34) | p(s(bool,X33))) ) | $spl302),
  inference(cnf_transformation,[],[f27555_D])).
fof(f27555_D,plain,(
  ( ! [X33,X31,X34] : (s(bool,X31) = s(bool,X33) | s(bool,'F_0') = s(bool,X34) | s(bool,X31) = s(bool,X34) | p(s(bool,X33))) ) <=> ~$spl302),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl302])])).
fof(f252,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) | $spl16),
  inference(cnf_transformation,[],[f252_D])).
fof(f252_D,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f27556,plain,(
  ( ! [X37,X35,X38,X36] : (s(bool,'F_0') != s(bool,X36) | s(bool,X35) = s(bool,X37) | p(s(bool,X37)) | s(bool,X35) = s(bool,X38) | s(bool,'F_0') = s(bool,X38)) )),
  inference(subsumption_resolution,[],[f27456,f118])).
fof(f27456,plain,(
  ( ! [X37,X35,X38,X36] : (s(bool,'F_0') != s(bool,X36) | p(s(bool,X36)) | s(bool,X35) = s(bool,X37) | p(s(bool,X37)) | s(bool,X35) = s(bool,X38) | s(bool,'F_0') = s(bool,X38)) )),
  inference(superposition,[],[f22797,f5187])).
fof(f27554,plain,(
  ( ! [X33,X31,X34,X32] : (s(bool,'F_0') != s(bool,X32) | s(bool,X31) = s(bool,X33) | p(s(bool,X33)) | s(bool,X31) = s(bool,X34) | s(bool,'F_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f27455,f118])).
fof(f27455,plain,(
  ( ! [X33,X31,X34,X32] : (s(bool,'F_0') != s(bool,X32) | p(s(bool,X32)) | s(bool,X31) = s(bool,X33) | p(s(bool,X33)) | s(bool,X31) = s(bool,X34) | s(bool,'F_0') = s(bool,X34)) )),
  inference(superposition,[],[f22797,f5187])).
fof(f27510,plain,(
  ( ! [X57,X58,X56] : (s(bool,X56) != s(bool,X57) | s(bool,'T_0') = s(bool,X56) | p(s(bool,X58)) | s(bool,X57) = s(bool,X58)) ) | $spl300),
  inference(cnf_transformation,[],[f27510_D])).
fof(f27510_D,plain,(
  ( ! [X57,X58,X56] : (s(bool,X56) != s(bool,X57) | s(bool,'T_0') = s(bool,X56) | p(s(bool,X58)) | s(bool,X57) = s(bool,X58)) ) <=> ~$spl300),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl300])])).
fof(f27511,plain,(
  ( ! [X61,X59,X62,X60] : (s(bool,X60) != s(bool,X61) | s(bool,X61) = s(bool,X62) | p(s(bool,X62)) | s(bool,'T_0') = s(bool,X60) | s(bool,'T_0') = s(bool,X59)) )),
  inference(subsumption_resolution,[],[f27406,f123])).
fof(f27406,plain,(
  ( ! [X61,X59,X62,X60] : (s(bool,X60) != s(bool,X61) | p(s(bool,X59)) | s(bool,X61) = s(bool,X62) | p(s(bool,X62)) | s(bool,'T_0') = s(bool,X60) | s(bool,'T_0') = s(bool,X59)) )),
  inference(superposition,[],[f22797,f19171])).
fof(f27509,plain,(
  ( ! [X57,X58,X56,X55] : (s(bool,X56) != s(bool,X57) | s(bool,X57) = s(bool,X58) | p(s(bool,X58)) | s(bool,'T_0') = s(bool,X55) | s(bool,'T_0') = s(bool,X56)) )),
  inference(subsumption_resolution,[],[f27405,f123])).
fof(f27405,plain,(
  ( ! [X57,X58,X56,X55] : (s(bool,X56) != s(bool,X57) | p(s(bool,X55)) | s(bool,X57) = s(bool,X58) | p(s(bool,X58)) | s(bool,'T_0') = s(bool,X55) | s(bool,'T_0') = s(bool,X56)) )),
  inference(superposition,[],[f22797,f19171])).
fof(f12944,plain,(
  ( ! [X90,X88,X87,X93,X91,X89] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X89,X89),X93))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X87,X87),'const.sum.++_3'(s(fun(X89,X87),X91),s(fun(X89,X87),X91),s('type.sum.sum'(X89,X89),X93))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X87,fun(X87,bool)),X88),s(fun(X87,X89),X90),s(fun(X89,X87),X91)))) ) | $spl238),
  inference(cnf_transformation,[],[f12944_D])).
fof(f12944_D,plain,(
  ( ! [X90,X88,X87,X93,X91,X89] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X89,X89),X93))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X87,X87),'const.sum.++_3'(s(fun(X89,X87),X91),s(fun(X89,X87),X91),s('type.sum.sum'(X89,X89),X93))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X87,fun(X87,bool)),X88),s(fun(X87,X89),X90),s(fun(X89,X87),X91)))) ) <=> ~$spl238),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl238])])).
fof(f27152,plain,(
  ( ! [X87,X85,X83,X81,X86,X84,X82] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X83,X83),X87))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X81,X81),'const.sum.++_3'(s(fun(X83,X81),X85),s(fun(X83,X81),X85),s('type.sum.sum'(X83,X83),X87))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X81,fun(X81,bool)),X82),s(fun(X81,X83),X84),s(fun(X83,X81),X85))) | s(bool,'F_0') = s(bool,X86)) )),
  inference(subsumption_resolution,[],[f27135,f119])).
fof(f27135,plain,(
  ( ! [X87,X85,X83,X81,X86,X84,X82] : (~p(s(bool,X86)) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X83,X83),X87))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X81,X81),'const.sum.++_3'(s(fun(X83,X81),X85),s(fun(X83,X81),X85),s('type.sum.sum'(X83,X83),X87))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X81,fun(X81,bool)),X82),s(fun(X81,X83),X84),s(fun(X83,X81),X85))) | s(bool,'F_0') = s(bool,X86)) )),
  inference(superposition,[],[f9432,f5187])).
fof(f9432,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X2,fun(X2,bool)),X4),s(fun(X2,X0),X5),s(fun(X0,X2),X3)))) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X0,X0),X1))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X2,X2),'const.sum.++_3'(s(fun(X0,X2),X3),s(fun(X0,X2),X3),s('type.sum.sum'(X0,X0),X1)))))) )),
  inference(factoring,[],[f127])).
fof(f27151,plain,(
  ( ! [X80,X78,X76,X74,X79,X77,X75] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X76,X76),X80))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X74,X74),'const.sum.++_3'(s(fun(X76,X74),X78),s(fun(X76,X74),X78),s('type.sum.sum'(X76,X76),X80))))) | s(bool,'F_0') = s(bool,X79) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X74,fun(X74,bool)),X75),s(fun(X74,X76),X77),s(fun(X76,X74),X78)))) )),
  inference(subsumption_resolution,[],[f27134,f119])).
fof(f27134,plain,(
  ( ! [X80,X78,X76,X74,X79,X77,X75] : (~p(s(bool,X79)) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X76,X76),X80))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X74,X74),'const.sum.++_3'(s(fun(X76,X74),X78),s(fun(X76,X74),X78),s('type.sum.sum'(X76,X76),X80))))) | s(bool,'F_0') = s(bool,X79) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X74,fun(X74,bool)),X75),s(fun(X74,X76),X77),s(fun(X76,X74),X78)))) )),
  inference(superposition,[],[f9432,f5187])).
fof(f8651,plain,(
  ( ! [X302,X304,X301,X303,X305,X307] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X303,sK1),X307))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X301,sK3),'const.sum.++_3'(s(fun(X303,X301),X305),s(fun(sK1,sK3),sK9),s('type.sum.sum'(X303,sK1),X307))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X301,fun(X301,bool)),X302),s(fun(X301,X303),X304),s(fun(X303,X301),X305)))) ) | $spl72),
  inference(cnf_transformation,[],[f8651_D])).
fof(f8651_D,plain,(
  ( ! [X302,X304,X301,X303,X305,X307] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X303,sK1),X307))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X301,sK3),'const.sum.++_3'(s(fun(X303,X301),X305),s(fun(sK1,sK3),sK9),s('type.sum.sum'(X303,sK1),X307))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X301,fun(X301,bool)),X302),s(fun(X301,X303),X304),s(fun(X303,X301),X305)))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f26878,plain,(
  ( ! [X87,X85,X83,X81,X86,X84,X82] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X83,sK1),X87))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X81,sK3),'const.sum.++_3'(s(fun(X83,X81),X85),s(fun(sK1,sK3),sK9),s('type.sum.sum'(X83,sK1),X87))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X81,fun(X81,bool)),X82),s(fun(X81,X83),X84),s(fun(X83,X81),X85))) | s(bool,'F_0') = s(bool,X86)) )),
  inference(subsumption_resolution,[],[f26861,f119])).
fof(f26861,plain,(
  ( ! [X87,X85,X83,X81,X86,X84,X82] : (~p(s(bool,X86)) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X83,sK1),X87))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X81,sK3),'const.sum.++_3'(s(fun(X83,X81),X85),s(fun(sK1,sK3),sK9),s('type.sum.sum'(X83,sK1),X87))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X81,fun(X81,bool)),X82),s(fun(X81,X83),X84),s(fun(X83,X81),X85))) | s(bool,'F_0') = s(bool,X86)) )),
  inference(superposition,[],[f303,f5187])).
fof(f303,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X22,fun(X22,bool)),X24),s(fun(X22,X20),X25),s(fun(X20,X22),X23)))) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X20,sK1),X21))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X22,sK3),'const.sum.++_3'(s(fun(X20,X22),X23),s(fun(sK1,sK3),sK9),s('type.sum.sum'(X20,sK1),X21)))))) )),
  inference(resolution,[],[f108,f127])).
fof(f26877,plain,(
  ( ! [X80,X78,X76,X74,X79,X77,X75] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X76,sK1),X80))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X74,sK3),'const.sum.++_3'(s(fun(X76,X74),X78),s(fun(sK1,sK3),sK9),s('type.sum.sum'(X76,sK1),X80))))) | s(bool,'F_0') = s(bool,X79) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X74,fun(X74,bool)),X75),s(fun(X74,X76),X77),s(fun(X76,X74),X78)))) )),
  inference(subsumption_resolution,[],[f26860,f119])).
fof(f26860,plain,(
  ( ! [X80,X78,X76,X74,X79,X77,X75] : (~p(s(bool,X79)) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X76,sK1),X80))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X74,sK3),'const.sum.++_3'(s(fun(X76,X74),X78),s(fun(sK1,sK3),sK9),s('type.sum.sum'(X76,sK1),X80))))) | s(bool,'F_0') = s(bool,X79) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X74,fun(X74,bool)),X75),s(fun(X74,X76),X77),s(fun(X76,X74),X78)))) )),
  inference(superposition,[],[f303,f5187])).
fof(f21372,plain,(
  ( ! [X579] : (s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X579))) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X579)))) ) | $spl278),
  inference(cnf_transformation,[],[f21372_D])).
fof(f21372_D,plain,(
  ( ! [X579] : (s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X579))) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X579)))) ) <=> ~$spl278),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl278])])).
fof(f26830,plain,(
  ( ! [X41,X40] : (s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X40))) | s(bool,'T_0') = s(bool,X41) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X40)))) )),
  inference(subsumption_resolution,[],[f26802,f123])).
fof(f26802,plain,(
  ( ! [X41,X40] : (p(s(bool,X41)) | s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X40))) | s(bool,'T_0') = s(bool,X41) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X40)))) )),
  inference(superposition,[],[f4105,f19171])).
fof(f26829,plain,(
  ( ! [X39,X38] : (s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X38))) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X38))) | s(bool,'T_0') = s(bool,X39)) )),
  inference(subsumption_resolution,[],[f26801,f123])).
fof(f26801,plain,(
  ( ! [X39,X38] : (p(s(bool,X39)) | s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X38))) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X38))) | s(bool,'T_0') = s(bool,X39)) )),
  inference(superposition,[],[f4105,f19171])).
fof(f11135,plain,(
  ( ! [X35] : (s(bool,'T_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X35))) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X35)))) ) | $spl88),
  inference(cnf_transformation,[],[f11135_D])).
fof(f11135_D,plain,(
  ( ! [X35] : (s(bool,'T_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X35))) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X35)))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f26753,plain,(
  ( ! [X33,X32] : (s(bool,'T_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X32))) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X32))) | s(bool,'F_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f26730,f119])).
fof(f26730,plain,(
  ( ! [X33,X32] : (~p(s(bool,X33)) | s(bool,'T_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X32))) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X32))) | s(bool,'F_0') = s(bool,X33)) )),
  inference(superposition,[],[f2598,f5187])).
fof(f2598,plain,(
  ( ! [X21] : (~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X21)))) | s(bool,'T_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X21)))) )),
  inference(trivial_inequality_removal,[],[f2594])).
fof(f2594,plain,(
  ( ! [X21] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'T_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X21))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X21))))) )),
  inference(superposition,[],[f2234,f125])).
fof(f2234,plain,(
  ( ! [X7] : (s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X7))) | s(bool,'T_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X7)))) )),
  inference(resolution,[],[f2215,f122])).
fof(f26752,plain,(
  ( ! [X30,X31] : (s(bool,'T_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X30))) | s(bool,'F_0') = s(bool,X31) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X30)))) )),
  inference(subsumption_resolution,[],[f26729,f119])).
fof(f26729,plain,(
  ( ! [X30,X31] : (~p(s(bool,X31)) | s(bool,'T_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X30))) | s(bool,'F_0') = s(bool,X31) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X30)))) )),
  inference(superposition,[],[f2598,f5187])).
fof(f10805,plain,(
  ( ! [X44] : (s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X44))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X44)))) ) | $spl84),
  inference(cnf_transformation,[],[f10805_D])).
fof(f10805_D,plain,(
  ( ! [X44] : (s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X44))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X44)))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f26684,plain,(
  ( ! [X33,X32] : (s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X32))) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X32))) | s(bool,'F_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f26661,f119])).
fof(f26661,plain,(
  ( ! [X33,X32] : (~p(s(bool,X33)) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X32))) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X32))) | s(bool,'F_0') = s(bool,X33)) )),
  inference(superposition,[],[f2288,f5187])).
fof(f26683,plain,(
  ( ! [X30,X31] : (s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X30))) | s(bool,'F_0') = s(bool,X31) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X30)))) )),
  inference(subsumption_resolution,[],[f26660,f119])).
fof(f26660,plain,(
  ( ! [X30,X31] : (~p(s(bool,X31)) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X30))) | s(bool,'F_0') = s(bool,X31) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X30)))) )),
  inference(superposition,[],[f2288,f5187])).
fof(f8649,plain,(
  ( ! [X294,X300,X296,X298,X295,X297] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK1,X296),X300))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK3,X294),'const.sum.++_3'(s(fun(sK1,sK3),sK9),s(fun(X296,X294),X298),s('type.sum.sum'(sK1,X296),X300))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X294,fun(X294,bool)),X295),s(fun(X294,X296),X297),s(fun(X296,X294),X298)))) ) | $spl70),
  inference(cnf_transformation,[],[f8649_D])).
fof(f8649_D,plain,(
  ( ! [X294,X300,X296,X298,X295,X297] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK1,X296),X300))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK3,X294),'const.sum.++_3'(s(fun(sK1,sK3),sK9),s(fun(X296,X294),X298),s('type.sum.sum'(sK1,X296),X300))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X294,fun(X294,bool)),X295),s(fun(X294,X296),X297),s(fun(X296,X294),X298)))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f26615,plain,(
  ( ! [X87,X85,X83,X81,X86,X84,X82] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK1,X83),X87))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK3,X81),'const.sum.++_3'(s(fun(sK1,sK3),sK9),s(fun(X83,X81),X85),s('type.sum.sum'(sK1,X83),X87))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X81,fun(X81,bool)),X82),s(fun(X81,X83),X84),s(fun(X83,X81),X85))) | s(bool,'F_0') = s(bool,X86)) )),
  inference(subsumption_resolution,[],[f26598,f119])).
fof(f26598,plain,(
  ( ! [X87,X85,X83,X81,X86,X84,X82] : (~p(s(bool,X86)) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK1,X83),X87))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK3,X81),'const.sum.++_3'(s(fun(sK1,sK3),sK9),s(fun(X83,X81),X85),s('type.sum.sum'(sK1,X83),X87))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X81,fun(X81,bool)),X82),s(fun(X81,X83),X84),s(fun(X83,X81),X85))) | s(bool,'F_0') = s(bool,X86)) )),
  inference(superposition,[],[f302,f5187])).
fof(f302,plain,(
  ( ! [X14,X19,X17,X15,X18,X16] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X16,fun(X16,bool)),X18),s(fun(X16,X14),X19),s(fun(X14,X16),X17)))) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK1,X14),X15))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK3,X16),'const.sum.++_3'(s(fun(sK1,sK3),sK9),s(fun(X14,X16),X17),s('type.sum.sum'(sK1,X14),X15)))))) )),
  inference(resolution,[],[f108,f127])).
fof(f26614,plain,(
  ( ! [X80,X78,X76,X74,X79,X77,X75] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK1,X76),X80))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK3,X74),'const.sum.++_3'(s(fun(sK1,sK3),sK9),s(fun(X76,X74),X78),s('type.sum.sum'(sK1,X76),X80))))) | s(bool,'F_0') = s(bool,X79) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X74,fun(X74,bool)),X75),s(fun(X74,X76),X77),s(fun(X76,X74),X78)))) )),
  inference(subsumption_resolution,[],[f26597,f119])).
fof(f26597,plain,(
  ( ! [X80,X78,X76,X74,X79,X77,X75] : (~p(s(bool,X79)) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK1,X76),X80))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK3,X74),'const.sum.++_3'(s(fun(sK1,sK3),sK9),s(fun(X76,X74),X78),s('type.sum.sum'(sK1,X76),X80))))) | s(bool,'F_0') = s(bool,X79) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X74,fun(X74,bool)),X75),s(fun(X74,X76),X77),s(fun(X76,X74),X78)))) )),
  inference(superposition,[],[f302,f5187])).
fof(f21384,plain,(
  ( ! [X668] : (s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X668))) | s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X668)))) ) | $spl284),
  inference(cnf_transformation,[],[f21384_D])).
fof(f21384_D,plain,(
  ( ! [X668] : (s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X668))) | s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X668)))) ) <=> ~$spl284),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl284])])).
fof(f26567,plain,(
  ( ! [X52,X51] : (s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X51))) | s(bool,'T_0') = s(bool,X52) | s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X51)))) )),
  inference(subsumption_resolution,[],[f26548,f123])).
fof(f26548,plain,(
  ( ! [X52,X51] : (p(s(bool,X52)) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X51))) | s(bool,'T_0') = s(bool,X52) | s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X51)))) )),
  inference(superposition,[],[f3517,f19171])).
fof(f26566,plain,(
  ( ! [X50,X49] : (s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X49))) | s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X49))) | s(bool,'T_0') = s(bool,X50)) )),
  inference(subsumption_resolution,[],[f26547,f123])).
fof(f26547,plain,(
  ( ! [X50,X49] : (p(s(bool,X50)) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X49))) | s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X49))) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f3517,f19171])).
fof(f10871,plain,(
  ( ! [X44] : (s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X44))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X44)))) ) | $spl86),
  inference(cnf_transformation,[],[f10871_D])).
fof(f10871_D,plain,(
  ( ! [X44] : (s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X44))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X44)))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f26498,plain,(
  ( ! [X43,X44] : (s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X43))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X43))) | s(bool,'F_0') = s(bool,X44)) )),
  inference(subsumption_resolution,[],[f26475,f119])).
fof(f26475,plain,(
  ( ! [X43,X44] : (~p(s(bool,X44)) | s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X43))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X43))) | s(bool,'F_0') = s(bool,X44)) )),
  inference(superposition,[],[f2215,f5187])).
fof(f26497,plain,(
  ( ! [X41,X42] : (s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X41))) | s(bool,'F_0') = s(bool,X42) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X41)))) )),
  inference(subsumption_resolution,[],[f26474,f119])).
fof(f26474,plain,(
  ( ! [X41,X42] : (~p(s(bool,X42)) | s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X41))) | s(bool,'F_0') = s(bool,X42) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X41)))) )),
  inference(superposition,[],[f2215,f5187])).
fof(f26428,plain,(
  ( ! [X43,X44] : (s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X43))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X43))) | s(bool,'F_0') = s(bool,X44)) )),
  inference(subsumption_resolution,[],[f26405,f119])).
fof(f26405,plain,(
  ( ! [X43,X44] : (~p(s(bool,X44)) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X43))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X43))) | s(bool,'F_0') = s(bool,X44)) )),
  inference(superposition,[],[f1886,f5187])).
fof(f1886,plain,(
  ( ! [X2] : (~p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X2)))) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X2)))) )),
  inference(superposition,[],[f133,f1600])).
fof(f1600,plain,(
  ( ! [X0] : (s('type.sum.sum'(sK0,sK0),X0) = s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,sK13(X0,sK0,sK0)))) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X0)))) )),
  inference(superposition,[],[f1580,f111])).
fof(f26427,plain,(
  ( ! [X41,X42] : (s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X41))) | s(bool,'F_0') = s(bool,X42) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X41)))) )),
  inference(subsumption_resolution,[],[f26404,f119])).
fof(f26404,plain,(
  ( ! [X41,X42] : (~p(s(bool,X42)) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X41))) | s(bool,'F_0') = s(bool,X42) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X41)))) )),
  inference(superposition,[],[f1886,f5187])).
fof(f8643,plain,(
  ( ! [X271,X277,X273,X275,X272,X274] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X273,sK0),X277))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X271,sK2),'const.sum.++_3'(s(fun(X273,X271),X275),s(fun(sK0,sK2),sK6),s('type.sum.sum'(X273,sK0),X277))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X271,fun(X271,bool)),X272),s(fun(X271,X273),X274),s(fun(X273,X271),X275)))) ) | $spl64),
  inference(cnf_transformation,[],[f8643_D])).
fof(f8643_D,plain,(
  ( ! [X271,X277,X273,X275,X272,X274] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X273,sK0),X277))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X271,sK2),'const.sum.++_3'(s(fun(X273,X271),X275),s(fun(sK0,sK2),sK6),s('type.sum.sum'(X273,sK0),X277))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X271,fun(X271,bool)),X272),s(fun(X271,X273),X274),s(fun(X273,X271),X275)))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f26154,plain,(
  ( ! [X87,X85,X83,X81,X86,X84,X82] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X83,sK0),X87))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X81,sK2),'const.sum.++_3'(s(fun(X83,X81),X85),s(fun(sK0,sK2),sK6),s('type.sum.sum'(X83,sK0),X87))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X81,fun(X81,bool)),X82),s(fun(X81,X83),X84),s(fun(X83,X81),X85))) | s(bool,'F_0') = s(bool,X86)) )),
  inference(subsumption_resolution,[],[f26137,f119])).
fof(f26137,plain,(
  ( ! [X87,X85,X83,X81,X86,X84,X82] : (~p(s(bool,X86)) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X83,sK0),X87))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X81,sK2),'const.sum.++_3'(s(fun(X83,X81),X85),s(fun(sK0,sK2),sK6),s('type.sum.sum'(X83,sK0),X87))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X81,fun(X81,bool)),X82),s(fun(X81,X83),X84),s(fun(X83,X81),X85))) | s(bool,'F_0') = s(bool,X86)) )),
  inference(superposition,[],[f182,f5187])).
fof(f182,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X22,fun(X22,bool)),X24),s(fun(X22,X20),X25),s(fun(X20,X22),X23)))) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X20,sK0),X21))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X22,sK2),'const.sum.++_3'(s(fun(X20,X22),X23),s(fun(sK0,sK2),sK6),s('type.sum.sum'(X20,sK0),X21)))))) )),
  inference(resolution,[],[f107,f127])).
fof(f26153,plain,(
  ( ! [X80,X78,X76,X74,X79,X77,X75] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X76,sK0),X80))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X74,sK2),'const.sum.++_3'(s(fun(X76,X74),X78),s(fun(sK0,sK2),sK6),s('type.sum.sum'(X76,sK0),X80))))) | s(bool,'F_0') = s(bool,X79) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X74,fun(X74,bool)),X75),s(fun(X74,X76),X77),s(fun(X76,X74),X78)))) )),
  inference(subsumption_resolution,[],[f26136,f119])).
fof(f26136,plain,(
  ( ! [X80,X78,X76,X74,X79,X77,X75] : (~p(s(bool,X79)) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X76,sK0),X80))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X74,sK2),'const.sum.++_3'(s(fun(X76,X74),X78),s(fun(sK0,sK2),sK6),s('type.sum.sum'(X76,sK0),X80))))) | s(bool,'F_0') = s(bool,X79) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X74,fun(X74,bool)),X75),s(fun(X74,X76),X77),s(fun(X76,X74),X78)))) )),
  inference(superposition,[],[f182,f5187])).
fof(f11572,plain,(
  ( ! [X90,X88,X87,X93,X91,X89] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,X89),X93))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,X87),'const.sum.++_3'(s(fun(sK0,sK2),sK6),s(fun(X89,X87),X91),s('type.sum.sum'(sK0,X89),X93))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X87,fun(X87,bool)),X88),s(fun(X87,X89),X90),s(fun(X89,X87),X91)))) ) | $spl90),
  inference(cnf_transformation,[],[f11572_D])).
fof(f11572_D,plain,(
  ( ! [X90,X88,X87,X93,X91,X89] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,X89),X93))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,X87),'const.sum.++_3'(s(fun(sK0,sK2),sK6),s(fun(X89,X87),X91),s('type.sum.sum'(sK0,X89),X93))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X87,fun(X87,bool)),X88),s(fun(X87,X89),X90),s(fun(X89,X87),X91)))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f26100,plain,(
  ( ! [X87,X85,X83,X81,X86,X84,X82] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,X83),X87))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,X81),'const.sum.++_3'(s(fun(sK0,sK2),sK6),s(fun(X83,X81),X85),s('type.sum.sum'(sK0,X83),X87))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X81,fun(X81,bool)),X82),s(fun(X81,X83),X84),s(fun(X83,X81),X85))) | s(bool,'F_0') = s(bool,X86)) )),
  inference(subsumption_resolution,[],[f26083,f119])).
fof(f26083,plain,(
  ( ! [X87,X85,X83,X81,X86,X84,X82] : (~p(s(bool,X86)) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,X83),X87))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,X81),'const.sum.++_3'(s(fun(sK0,sK2),sK6),s(fun(X83,X81),X85),s('type.sum.sum'(sK0,X83),X87))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X81,fun(X81,bool)),X82),s(fun(X81,X83),X84),s(fun(X83,X81),X85))) | s(bool,'F_0') = s(bool,X86)) )),
  inference(superposition,[],[f181,f5187])).
fof(f26099,plain,(
  ( ! [X80,X78,X76,X74,X79,X77,X75] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,X76),X80))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,X74),'const.sum.++_3'(s(fun(sK0,sK2),sK6),s(fun(X76,X74),X78),s('type.sum.sum'(sK0,X76),X80))))) | s(bool,'F_0') = s(bool,X79) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X74,fun(X74,bool)),X75),s(fun(X74,X76),X77),s(fun(X76,X74),X78)))) )),
  inference(subsumption_resolution,[],[f26082,f119])).
fof(f26082,plain,(
  ( ! [X80,X78,X76,X74,X79,X77,X75] : (~p(s(bool,X79)) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,X76),X80))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,X74),'const.sum.++_3'(s(fun(sK0,sK2),sK6),s(fun(X76,X74),X78),s('type.sum.sum'(sK0,X76),X80))))) | s(bool,'F_0') = s(bool,X79) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X74,fun(X74,bool)),X75),s(fun(X74,X76),X77),s(fun(X76,X74),X78)))) )),
  inference(superposition,[],[f181,f5187])).
fof(f26051,plain,(
  ( ! [X41,X40] : (s('type.sum.sum'(X41,bool),X40) = s('type.sum.sum'(X41,bool),'const.sum.INL_1'(s(X41,sK12(X40,bool,X41)))) | s(bool,'T_0') = s(bool,sK13(X40,bool,X41)) | s('type.sum.sum'(X41,bool),X40) = s('type.sum.sum'(X41,bool),'const.sum.INR_1'(s(bool,'F_0')))) ) | $spl298),
  inference(cnf_transformation,[],[f26051_D])).
fof(f26051_D,plain,(
  ( ! [X41,X40] : (s('type.sum.sum'(X41,bool),X40) = s('type.sum.sum'(X41,bool),'const.sum.INL_1'(s(X41,sK12(X40,bool,X41)))) | s(bool,'T_0') = s(bool,sK13(X40,bool,X41)) | s('type.sum.sum'(X41,bool),X40) = s('type.sum.sum'(X41,bool),'const.sum.INR_1'(s(bool,'F_0')))) ) <=> ~$spl298),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl298])])).
fof(f26052,plain,(
  ( ! [X45,X43,X44] : (s('type.sum.sum'(X44,bool),X43) = s('type.sum.sum'(X44,bool),'const.sum.INL_1'(s(X44,sK12(X43,bool,X44)))) | s('type.sum.sum'(X44,bool),X43) = s('type.sum.sum'(X44,bool),'const.sum.INR_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X45) | s(bool,'T_0') = s(bool,sK13(X43,bool,X44))) )),
  inference(subsumption_resolution,[],[f26038,f123])).
fof(f26038,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,X45)) | s('type.sum.sum'(X44,bool),X43) = s('type.sum.sum'(X44,bool),'const.sum.INL_1'(s(X44,sK12(X43,bool,X44)))) | s('type.sum.sum'(X44,bool),X43) = s('type.sum.sum'(X44,bool),'const.sum.INR_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X45) | s(bool,'T_0') = s(bool,sK13(X43,bool,X44))) )),
  inference(superposition,[],[f22314,f19171])).
fof(f22314,plain,(
  ( ! [X149,X148] : (p(s(bool,sK13(X148,bool,X149))) | s('type.sum.sum'(X149,bool),X148) = s('type.sum.sum'(X149,bool),'const.sum.INL_1'(s(X149,sK12(X148,bool,X149)))) | s('type.sum.sum'(X149,bool),X148) = s('type.sum.sum'(X149,bool),'const.sum.INR_1'(s(bool,'F_0')))) )),
  inference(superposition,[],[f111,f119])).
fof(f26050,plain,(
  ( ! [X41,X42,X40] : (s('type.sum.sum'(X41,bool),X40) = s('type.sum.sum'(X41,bool),'const.sum.INL_1'(s(X41,sK12(X40,bool,X41)))) | s('type.sum.sum'(X41,bool),X40) = s('type.sum.sum'(X41,bool),'const.sum.INR_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK13(X40,bool,X41)) | s(bool,'T_0') = s(bool,X42)) )),
  inference(subsumption_resolution,[],[f26037,f123])).
fof(f26037,plain,(
  ( ! [X41,X42,X40] : (p(s(bool,X42)) | s('type.sum.sum'(X41,bool),X40) = s('type.sum.sum'(X41,bool),'const.sum.INL_1'(s(X41,sK12(X40,bool,X41)))) | s('type.sum.sum'(X41,bool),X40) = s('type.sum.sum'(X41,bool),'const.sum.INR_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK13(X40,bool,X41)) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f22314,f19171])).
fof(f24318,plain,(
  ( ! [X105] : (s(bool,X105) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,X105)) ) | $spl294),
  inference(cnf_transformation,[],[f24318_D])).
fof(f24318_D,plain,(
  ( ! [X105] : (s(bool,X105) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,X105)) ) <=> ~$spl294),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl294])])).
fof(f25890,plain,(
  ( ! [X19,X20] : (s(bool,X20) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'T_0') != s(bool,X19) | s(bool,'F_0') = s(bool,X20)) ) | $spl10),
  inference(subsumption_resolution,[],[f25822,f21926])).
fof(f25822,plain,(
  ( ! [X19,X20] : (s(bool,X20) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'T_0') != s(bool,X19) | s(bool,'F_0') = s(bool,X19) | s(bool,'F_0') = s(bool,X20)) ) | $spl10),
  inference(superposition,[],[f24234,f5187])).
fof(f24234,plain,(
  ( ! [X103] : (s(bool,X103) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'T_0') != s(bool,X103)) ) | $spl10),
  inference(subsumption_resolution,[],[f24205,f150])).
fof(f24205,plain,(
  ( ! [X103] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,X103) | s(bool,X103) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11)))) ) | $spl10),
  inference(superposition,[],[f22526,f176])).
fof(f25889,plain,(
  ( ! [X17,X18] : (s(bool,X18) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'T_0') != s(bool,X17) | s(bool,'F_0') = s(bool,X18)) ) | $spl10),
  inference(subsumption_resolution,[],[f25821,f21926])).
fof(f25821,plain,(
  ( ! [X17,X18] : (s(bool,X18) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'T_0') != s(bool,X17) | s(bool,'F_0') = s(bool,X18) | s(bool,'F_0') = s(bool,X17)) ) | $spl10),
  inference(superposition,[],[f24234,f5187])).
fof(f25081,plain,(
  ( ! [X41,X40] : (s('type.sum.sum'(bool,X41),X40) = s('type.sum.sum'(bool,X41),'const.sum.INR_1'(s(X41,sK13(X40,X41,bool)))) | s(bool,'T_0') = s(bool,sK12(X40,X41,bool)) | s('type.sum.sum'(bool,X41),X40) = s('type.sum.sum'(bool,X41),'const.sum.INL_1'(s(bool,'F_0')))) ) | $spl296),
  inference(cnf_transformation,[],[f25081_D])).
fof(f25081_D,plain,(
  ( ! [X41,X40] : (s('type.sum.sum'(bool,X41),X40) = s('type.sum.sum'(bool,X41),'const.sum.INR_1'(s(X41,sK13(X40,X41,bool)))) | s(bool,'T_0') = s(bool,sK12(X40,X41,bool)) | s('type.sum.sum'(bool,X41),X40) = s('type.sum.sum'(bool,X41),'const.sum.INL_1'(s(bool,'F_0')))) ) <=> ~$spl296),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl296])])).
fof(f25082,plain,(
  ( ! [X45,X43,X44] : (s('type.sum.sum'(bool,X44),X43) = s('type.sum.sum'(bool,X44),'const.sum.INR_1'(s(X44,sK13(X43,X44,bool)))) | s('type.sum.sum'(bool,X44),X43) = s('type.sum.sum'(bool,X44),'const.sum.INL_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X45) | s(bool,'T_0') = s(bool,sK12(X43,X44,bool))) )),
  inference(subsumption_resolution,[],[f25070,f123])).
fof(f25070,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,X45)) | s('type.sum.sum'(bool,X44),X43) = s('type.sum.sum'(bool,X44),'const.sum.INR_1'(s(X44,sK13(X43,X44,bool)))) | s('type.sum.sum'(bool,X44),X43) = s('type.sum.sum'(bool,X44),'const.sum.INL_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X45) | s(bool,'T_0') = s(bool,sK12(X43,X44,bool))) )),
  inference(superposition,[],[f22313,f19171])).
fof(f22313,plain,(
  ( ! [X146,X147] : (p(s(bool,sK12(X146,X147,bool))) | s('type.sum.sum'(bool,X147),X146) = s('type.sum.sum'(bool,X147),'const.sum.INR_1'(s(X147,sK13(X146,X147,bool)))) | s('type.sum.sum'(bool,X147),X146) = s('type.sum.sum'(bool,X147),'const.sum.INL_1'(s(bool,'F_0')))) )),
  inference(superposition,[],[f111,f119])).
fof(f25080,plain,(
  ( ! [X41,X42,X40] : (s('type.sum.sum'(bool,X41),X40) = s('type.sum.sum'(bool,X41),'const.sum.INR_1'(s(X41,sK13(X40,X41,bool)))) | s('type.sum.sum'(bool,X41),X40) = s('type.sum.sum'(bool,X41),'const.sum.INL_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK12(X40,X41,bool)) | s(bool,'T_0') = s(bool,X42)) )),
  inference(subsumption_resolution,[],[f25069,f123])).
fof(f25069,plain,(
  ( ! [X41,X42,X40] : (p(s(bool,X42)) | s('type.sum.sum'(bool,X41),X40) = s('type.sum.sum'(bool,X41),'const.sum.INR_1'(s(X41,sK13(X40,X41,bool)))) | s('type.sum.sum'(bool,X41),X40) = s('type.sum.sum'(bool,X41),'const.sum.INL_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK12(X40,X41,bool)) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f22313,f19171])).
fof(f21380,plain,(
  ( ! [X639,X644,X640,X642,X638,X641,X643] : (s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X638,X639),'const.sum.++_3'(s(fun(X640,X638),X641),s(fun(X642,X639),X643),s('type.sum.sum'(X640,X642),'const.sum.INL_1'(s(X640,X644)))))))) ) | $spl282),
  inference(cnf_transformation,[],[f21380_D])).
fof(f21380_D,plain,(
  ( ! [X639,X644,X640,X642,X638,X641,X643] : (s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X638,X639),'const.sum.++_3'(s(fun(X640,X638),X641),s(fun(X642,X639),X643),s('type.sum.sum'(X640,X642),'const.sum.INL_1'(s(X640,X644)))))))) ) <=> ~$spl282),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl282])])).
fof(f24589,plain,(
  ( ! [X158,X156,X161,X159,X157,X155,X162,X160] : (s(bool,'T_0') = s(bool,X162) | s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X155,X156),'const.sum.++_3'(s(fun(X157,X155),X158),s(fun(X159,X156),X160),s('type.sum.sum'(X157,X159),'const.sum.INL_1'(s(X157,X161)))))))) )),
  inference(subsumption_resolution,[],[f24578,f123])).
fof(f24578,plain,(
  ( ! [X158,X156,X161,X159,X157,X155,X162,X160] : (p(s(bool,X162)) | s(bool,'T_0') = s(bool,X162) | s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X155,X156),'const.sum.++_3'(s(fun(X157,X155),X158),s(fun(X159,X156),X160),s('type.sum.sum'(X157,X159),'const.sum.INL_1'(s(X157,X161)))))))) )),
  inference(superposition,[],[f9289,f19171])).
fof(f9289,plain,(
  ( ! [X6,X4,X8,X7,X5,X3,X9] : (p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X3,X4),'const.sum.++_3'(s(fun(X5,X3),X6),s(fun(X8,X4),X9),s('type.sum.sum'(X5,X8),'const.sum.INL_1'(s(X5,X7))))))))) )),
  inference(superposition,[],[f134,f131])).
fof(f24588,plain,(
  ( ! [X154,X152,X151,X149,X147,X153,X150,X148] : (s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X147,X148),'const.sum.++_3'(s(fun(X149,X147),X150),s(fun(X151,X148),X152),s('type.sum.sum'(X149,X151),'const.sum.INL_1'(s(X149,X153))))))) | s(bool,'T_0') = s(bool,X154)) )),
  inference(subsumption_resolution,[],[f24577,f123])).
fof(f24577,plain,(
  ( ! [X154,X152,X151,X149,X147,X153,X150,X148] : (p(s(bool,X154)) | s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X147,X148),'const.sum.++_3'(s(fun(X149,X147),X150),s(fun(X151,X148),X152),s('type.sum.sum'(X149,X151),'const.sum.INL_1'(s(X149,X153))))))) | s(bool,'T_0') = s(bool,X154)) )),
  inference(superposition,[],[f9289,f19171])).
fof(f21368,plain,(
  ( ! [X546,X542,X545,X547,X543,X548,X544] : (s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X542,X543),'const.sum.++_3'(s(fun(X544,X542),X545),s(fun(X546,X543),X547),s('type.sum.sum'(X544,X546),'const.sum.INR_1'(s(X546,X548)))))))) ) | $spl276),
  inference(cnf_transformation,[],[f21368_D])).
fof(f21368_D,plain,(
  ( ! [X546,X542,X545,X547,X543,X548,X544] : (s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X542,X543),'const.sum.++_3'(s(fun(X544,X542),X545),s(fun(X546,X543),X547),s('type.sum.sum'(X544,X546),'const.sum.INR_1'(s(X546,X548)))))))) ) <=> ~$spl276),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl276])])).
fof(f24485,plain,(
  ( ! [X134,X140,X138,X136,X135,X141,X139,X137] : (s(bool,'T_0') = s(bool,X141) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X134,X135),'const.sum.++_3'(s(fun(X136,X134),X137),s(fun(X138,X135),X139),s('type.sum.sum'(X136,X138),'const.sum.INR_1'(s(X138,X140)))))))) )),
  inference(subsumption_resolution,[],[f24471,f123])).
fof(f24471,plain,(
  ( ! [X134,X140,X138,X136,X135,X141,X139,X137] : (p(s(bool,X141)) | s(bool,'T_0') = s(bool,X141) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X134,X135),'const.sum.++_3'(s(fun(X136,X134),X137),s(fun(X138,X135),X139),s('type.sum.sum'(X136,X138),'const.sum.INR_1'(s(X138,X140)))))))) )),
  inference(superposition,[],[f9190,f19171])).
fof(f9190,plain,(
  ( ! [X6,X4,X8,X7,X5,X3,X9] : (p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X3,X4),'const.sum.++_3'(s(fun(X8,X3),X9),s(fun(X5,X4),X6),s('type.sum.sum'(X8,X5),'const.sum.INR_1'(s(X5,X7))))))))) )),
  inference(superposition,[],[f129,f132])).
fof(f24484,plain,(
  ( ! [X127,X132,X130,X128,X126,X133,X131,X129] : (s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X126,X127),'const.sum.++_3'(s(fun(X128,X126),X129),s(fun(X130,X127),X131),s('type.sum.sum'(X128,X130),'const.sum.INR_1'(s(X130,X132))))))) | s(bool,'T_0') = s(bool,X133)) )),
  inference(subsumption_resolution,[],[f24470,f123])).
fof(f24470,plain,(
  ( ! [X127,X132,X130,X128,X126,X133,X131,X129] : (p(s(bool,X133)) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X126,X127),'const.sum.++_3'(s(fun(X128,X126),X129),s(fun(X130,X127),X131),s('type.sum.sum'(X128,X130),'const.sum.INR_1'(s(X130,X132))))))) | s(bool,'T_0') = s(bool,X133)) )),
  inference(superposition,[],[f9190,f19171])).
fof(f21378,plain,(
  ( ! [X609,X608,X610] : (s('type.sum.sum'(X608,X609),X610) = s('type.sum.sum'(X608,X609),'const.sum.INR_1'(s(X609,sK13(X610,X609,X608)))) | s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X608,X609),X610)))) ) | $spl280),
  inference(cnf_transformation,[],[f21378_D])).
fof(f21378_D,plain,(
  ( ! [X609,X608,X610] : (s('type.sum.sum'(X608,X609),X610) = s('type.sum.sum'(X608,X609),'const.sum.INR_1'(s(X609,sK13(X610,X609,X608)))) | s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X608,X609),X610)))) ) <=> ~$spl280),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl280])])).
fof(f24380,plain,(
  ( ! [X94,X92,X93,X91] : (s('type.sum.sum'(X91,X92),X93) = s('type.sum.sum'(X91,X92),'const.sum.INR_1'(s(X92,sK13(X93,X92,X91)))) | s(bool,'T_0') = s(bool,X94) | s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X91,X92),X93)))) )),
  inference(subsumption_resolution,[],[f24364,f123])).
fof(f24364,plain,(
  ( ! [X94,X92,X93,X91] : (p(s(bool,X94)) | s('type.sum.sum'(X91,X92),X93) = s('type.sum.sum'(X91,X92),'const.sum.INR_1'(s(X92,sK13(X93,X92,X91)))) | s(bool,'T_0') = s(bool,X94) | s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X91,X92),X93)))) )),
  inference(superposition,[],[f9288,f19171])).
fof(f24379,plain,(
  ( ! [X90,X88,X87,X89] : (s('type.sum.sum'(X87,X88),X89) = s('type.sum.sum'(X87,X88),'const.sum.INR_1'(s(X88,sK13(X89,X88,X87)))) | s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X87,X88),X89))) | s(bool,'T_0') = s(bool,X90)) )),
  inference(subsumption_resolution,[],[f24363,f123])).
fof(f24363,plain,(
  ( ! [X90,X88,X87,X89] : (p(s(bool,X90)) | s('type.sum.sum'(X87,X88),X89) = s('type.sum.sum'(X87,X88),'const.sum.INR_1'(s(X88,sK13(X89,X88,X87)))) | s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X87,X88),X89))) | s(bool,'T_0') = s(bool,X90)) )),
  inference(superposition,[],[f9288,f19171])).
fof(f24317,plain,(
  ( ! [X105,X104] : (s(bool,X105) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,X105) | s(bool,'T_0') = s(bool,X104)) ) | $spl10),
  inference(subsumption_resolution,[],[f24294,f123])).
fof(f24294,plain,(
  ( ! [X105,X104] : (p(s(bool,X104)) | s(bool,X105) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,X105) | s(bool,'T_0') = s(bool,X104)) ) | $spl10),
  inference(superposition,[],[f22799,f19305])).
fof(f19305,plain,(
  ( ! [X106] : (s(bool,X106) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'T_0') = s(bool,X106)) ) | $spl10),
  inference(subsumption_resolution,[],[f19259,f150])).
fof(f19259,plain,(
  ( ! [X106] : (p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,X106) | s(bool,X106) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11)))) ) | $spl10),
  inference(superposition,[],[f19124,f176])).
fof(f21366,plain,(
  ( ! [X528,X526,X527] : (s('type.sum.sum'(X526,X527),X528) = s('type.sum.sum'(X526,X527),'const.sum.INL_1'(s(X526,sK12(X528,X527,X526)))) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X526,X527),X528)))) ) | $spl274),
  inference(cnf_transformation,[],[f21366_D])).
fof(f21366_D,plain,(
  ( ! [X528,X526,X527] : (s('type.sum.sum'(X526,X527),X528) = s('type.sum.sum'(X526,X527),'const.sum.INL_1'(s(X526,sK12(X528,X527,X526)))) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X526,X527),X528)))) ) <=> ~$spl274),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl274])])).
fof(f24148,plain,(
  ( ! [X85,X83,X84,X82] : (s('type.sum.sum'(X82,X83),X84) = s('type.sum.sum'(X82,X83),'const.sum.INL_1'(s(X82,sK12(X84,X83,X82)))) | s(bool,'T_0') = s(bool,X85) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X82,X83),X84)))) )),
  inference(subsumption_resolution,[],[f24129,f123])).
fof(f24129,plain,(
  ( ! [X85,X83,X84,X82] : (p(s(bool,X85)) | s('type.sum.sum'(X82,X83),X84) = s('type.sum.sum'(X82,X83),'const.sum.INL_1'(s(X82,sK12(X84,X83,X82)))) | s(bool,'T_0') = s(bool,X85) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X82,X83),X84)))) )),
  inference(superposition,[],[f9189,f19171])).
fof(f24147,plain,(
  ( ! [X80,X78,X81,X79] : (s('type.sum.sum'(X78,X79),X80) = s('type.sum.sum'(X78,X79),'const.sum.INL_1'(s(X78,sK12(X80,X79,X78)))) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X78,X79),X80))) | s(bool,'T_0') = s(bool,X81)) )),
  inference(subsumption_resolution,[],[f24128,f123])).
fof(f24128,plain,(
  ( ! [X80,X78,X81,X79] : (p(s(bool,X81)) | s('type.sum.sum'(X78,X79),X80) = s('type.sum.sum'(X78,X79),'const.sum.INL_1'(s(X78,sK12(X80,X79,X78)))) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X78,X79),X80))) | s(bool,'T_0') = s(bool,X81)) )),
  inference(superposition,[],[f9189,f19171])).
fof(f23989,plain,(
  ( ! [X105] : (s(bool,X105) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') != s(bool,X105)) ) | $spl292),
  inference(cnf_transformation,[],[f23989_D])).
fof(f23989_D,plain,(
  ( ! [X105] : (s(bool,X105) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') != s(bool,X105)) ) <=> ~$spl292),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl292])])).
fof(f23988,plain,(
  ( ! [X105,X104] : (s(bool,X105) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') != s(bool,X105) | s(bool,'T_0') = s(bool,X104)) ) | $spl10),
  inference(subsumption_resolution,[],[f23965,f123])).
fof(f23965,plain,(
  ( ! [X105,X104] : (p(s(bool,X104)) | s(bool,X105) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') != s(bool,X105) | s(bool,'T_0') = s(bool,X104)) ) | $spl10),
  inference(superposition,[],[f21924,f19305])).
fof(f8635,plain,(
  ( ! [X222,X231,X229,X225,X233,X223,X230,X228,X226,X224,X234,X232] : (s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X222,X228),X229))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X222,X228),X230))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X222,fun(X222,bool)),X223),s(fun(X222,X224),X225),s(fun(X224,X222),X226))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X228,fun(X228,bool)),X231),s(fun(X228,X232),X233),s(fun(X232,X228),X234)))) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X222,fun(X222,bool)),X223),s(fun(X228,fun(X228,bool)),X231),s('type.sum.sum'(X222,X228),X229),s('type.sum.sum'(X222,X228),X230))))) ) | $spl56),
  inference(cnf_transformation,[],[f8635_D])).
fof(f8635_D,plain,(
  ( ! [X222,X231,X229,X225,X233,X223,X230,X228,X226,X224,X234,X232] : (s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X222,X228),X229))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X222,X228),X230))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X222,fun(X222,bool)),X223),s(fun(X222,X224),X225),s(fun(X224,X222),X226))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X228,fun(X228,bool)),X231),s(fun(X228,X232),X233),s(fun(X232,X228),X234)))) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X222,fun(X222,bool)),X223),s(fun(X228,fun(X228,bool)),X231),s('type.sum.sum'(X222,X228),X229),s('type.sum.sum'(X222,X228),X230))))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f23380,plain,(
  ( ! [X177,X167,X165,X175,X173,X171,X169,X176,X166,X174,X172,X170,X168] : (s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X165,X171),X172))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X165,X171),X173))) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X165,fun(X165,bool)),X166),s(fun(X171,fun(X171,bool)),X174),s('type.sum.sum'(X165,X171),X172),s('type.sum.sum'(X165,X171),X173)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X171,fun(X171,bool)),X174),s(fun(X171,X175),X176),s(fun(X175,X171),X177)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X165,fun(X165,bool)),X166),s(fun(X165,X167),X168),s(fun(X167,X165),X169))) | s(bool,'F_0') = s(bool,X170)) )),
  inference(subsumption_resolution,[],[f23172,f119])).
fof(f23172,plain,(
  ( ! [X177,X167,X165,X175,X173,X171,X169,X176,X166,X174,X172,X170,X168] : (~p(s(bool,X170)) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X165,X171),X172))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X165,X171),X173))) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X165,fun(X165,bool)),X166),s(fun(X171,fun(X171,bool)),X174),s('type.sum.sum'(X165,X171),X172),s('type.sum.sum'(X165,X171),X173)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X171,fun(X171,bool)),X174),s(fun(X171,X175),X176),s(fun(X175,X171),X177)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X165,fun(X165,bool)),X166),s(fun(X165,X167),X168),s(fun(X167,X165),X169))) | s(bool,'F_0') = s(bool,X170)) )),
  inference(superposition,[],[f130,f5187])).
fof(f23379,plain,(
  ( ! [X158,X156,X154,X152,X163,X161,X159,X157,X155,X153,X164,X162,X160] : (s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X152,X158),X159))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X152,X158),X160))) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X152,fun(X152,bool)),X153),s(fun(X158,fun(X158,bool)),X161),s('type.sum.sum'(X152,X158),X159),s('type.sum.sum'(X152,X158),X160)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X158,fun(X158,bool)),X161),s(fun(X158,X162),X163),s(fun(X162,X158),X164)))) | s(bool,'F_0') = s(bool,X157) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X152,fun(X152,bool)),X153),s(fun(X152,X154),X155),s(fun(X154,X152),X156)))) )),
  inference(subsumption_resolution,[],[f23171,f119])).
fof(f23171,plain,(
  ( ! [X158,X156,X154,X152,X163,X161,X159,X157,X155,X153,X164,X162,X160] : (~p(s(bool,X157)) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X152,X158),X159))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X152,X158),X160))) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X152,fun(X152,bool)),X153),s(fun(X158,fun(X158,bool)),X161),s('type.sum.sum'(X152,X158),X159),s('type.sum.sum'(X152,X158),X160)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X158,fun(X158,bool)),X161),s(fun(X158,X162),X163),s(fun(X162,X158),X164)))) | s(bool,'F_0') = s(bool,X157) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X152,fun(X152,bool)),X153),s(fun(X152,X154),X155),s(fun(X154,X152),X156)))) )),
  inference(superposition,[],[f130,f5187])).
fof(f8637,plain,(
  ( ! [X243,X241,X239,X237,X235,X246,X244,X242,X238,X236,X247,X245] : (s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X241,X235),X242))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X241,X235),X243))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X235,fun(X235,bool)),X236),s(fun(X235,X237),X238),s(fun(X237,X235),X239))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X241,fun(X241,bool)),X244),s(fun(X241,X245),X246),s(fun(X245,X241),X247)))) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X241,fun(X241,bool)),X244),s(fun(X235,fun(X235,bool)),X236),s('type.sum.sum'(X241,X235),X242),s('type.sum.sum'(X241,X235),X243))))) ) | $spl58),
  inference(cnf_transformation,[],[f8637_D])).
fof(f8637_D,plain,(
  ( ! [X243,X241,X239,X237,X235,X246,X244,X242,X238,X236,X247,X245] : (s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X241,X235),X242))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X241,X235),X243))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X235,fun(X235,bool)),X236),s(fun(X235,X237),X238),s(fun(X237,X235),X239))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X241,fun(X241,bool)),X244),s(fun(X241,X245),X246),s(fun(X245,X241),X247)))) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X241,fun(X241,bool)),X244),s(fun(X235,fun(X235,bool)),X236),s('type.sum.sum'(X241,X235),X242),s('type.sum.sum'(X241,X235),X243))))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f23370,plain,(
  ( ! [X177,X167,X165,X175,X173,X171,X169,X176,X166,X174,X172,X170,X168] : (s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X171,X165),X172))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X171,X165),X173))) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X171,fun(X171,bool)),X174),s(fun(X165,fun(X165,bool)),X166),s('type.sum.sum'(X171,X165),X172),s('type.sum.sum'(X171,X165),X173)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X171,fun(X171,bool)),X174),s(fun(X171,X175),X176),s(fun(X175,X171),X177)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X165,fun(X165,bool)),X166),s(fun(X165,X167),X168),s(fun(X167,X165),X169))) | s(bool,'F_0') = s(bool,X170)) )),
  inference(subsumption_resolution,[],[f23152,f119])).
fof(f23152,plain,(
  ( ! [X177,X167,X165,X175,X173,X171,X169,X176,X166,X174,X172,X170,X168] : (~p(s(bool,X170)) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X171,X165),X172))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X171,X165),X173))) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X171,fun(X171,bool)),X174),s(fun(X165,fun(X165,bool)),X166),s('type.sum.sum'(X171,X165),X172),s('type.sum.sum'(X171,X165),X173)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X171,fun(X171,bool)),X174),s(fun(X171,X175),X176),s(fun(X175,X171),X177)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X165,fun(X165,bool)),X166),s(fun(X165,X167),X168),s(fun(X167,X165),X169))) | s(bool,'F_0') = s(bool,X170)) )),
  inference(superposition,[],[f130,f5187])).
fof(f23369,plain,(
  ( ! [X158,X156,X154,X152,X163,X161,X159,X157,X155,X153,X164,X162,X160] : (s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X158,X152),X159))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X158,X152),X160))) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X158,fun(X158,bool)),X161),s(fun(X152,fun(X152,bool)),X153),s('type.sum.sum'(X158,X152),X159),s('type.sum.sum'(X158,X152),X160)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X158,fun(X158,bool)),X161),s(fun(X158,X162),X163),s(fun(X162,X158),X164)))) | s(bool,'F_0') = s(bool,X157) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X152,fun(X152,bool)),X153),s(fun(X152,X154),X155),s(fun(X154,X152),X156)))) )),
  inference(subsumption_resolution,[],[f23151,f119])).
fof(f23151,plain,(
  ( ! [X158,X156,X154,X152,X163,X161,X159,X157,X155,X153,X164,X162,X160] : (~p(s(bool,X157)) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X158,X152),X159))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X158,X152),X160))) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X158,fun(X158,bool)),X161),s(fun(X152,fun(X152,bool)),X153),s('type.sum.sum'(X158,X152),X159),s('type.sum.sum'(X158,X152),X160)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X158,fun(X158,bool)),X161),s(fun(X158,X162),X163),s(fun(X162,X158),X164)))) | s(bool,'F_0') = s(bool,X157) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X152,fun(X152,bool)),X153),s(fun(X152,X154),X155),s(fun(X154,X152),X156)))) )),
  inference(superposition,[],[f130,f5187])).
fof(f8653,plain,(
  ( ! [X325,X327,X321,X329,X318,X324,X326,X320,X322,X328,X317,X319] : (s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X317,X319),X321))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X317,X319),X322))) | s(bool,'F_0') = s(bool,'const.quotient_sum.+++_4'(s(fun(X317,fun(X317,bool)),X318),s(fun(X319,fun(X319,bool)),X320),s('type.sum.sum'(X317,X319),X321),s('type.sum.sum'(X317,X319),X322))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X317,fun(X317,bool)),X318),s(fun(X317,X327),X328),s(fun(X327,X317),X329)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X319,fun(X319,bool)),X320),s(fun(X319,X324),X325),s(fun(X324,X319),X326))))) ) | $spl74),
  inference(cnf_transformation,[],[f8653_D])).
fof(f8653_D,plain,(
  ( ! [X325,X327,X321,X329,X318,X324,X326,X320,X322,X328,X317,X319] : (s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X317,X319),X321))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X317,X319),X322))) | s(bool,'F_0') = s(bool,'const.quotient_sum.+++_4'(s(fun(X317,fun(X317,bool)),X318),s(fun(X319,fun(X319,bool)),X320),s('type.sum.sum'(X317,X319),X321),s('type.sum.sum'(X317,X319),X322))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X317,fun(X317,bool)),X318),s(fun(X317,X327),X328),s(fun(X327,X317),X329)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X319,fun(X319,bool)),X320),s(fun(X319,X324),X325),s(fun(X324,X319),X326))))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f23360,plain,(
  ( ! [X158,X167,X165,X163,X161,X169,X159,X157,X166,X164,X162,X160,X168] : (s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X157,X159),X161))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X157,X159),X162))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X159,fun(X159,bool)),X160),s(fun(X159,X164),X165),s(fun(X164,X159),X166)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X157,fun(X157,bool)),X158),s(fun(X157,X167),X168),s(fun(X167,X157),X169)))) | s(bool,'F_0') = s(bool,'const.quotient_sum.+++_4'(s(fun(X157,fun(X157,bool)),X158),s(fun(X159,fun(X159,bool)),X160),s('type.sum.sum'(X157,X159),X161),s('type.sum.sum'(X157,X159),X162))) | s(bool,'F_0') = s(bool,X163)) )),
  inference(subsumption_resolution,[],[f23132,f119])).
fof(f23132,plain,(
  ( ! [X158,X167,X165,X163,X161,X169,X159,X157,X166,X164,X162,X160,X168] : (~p(s(bool,X163)) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X157,X159),X161))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X157,X159),X162))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X159,fun(X159,bool)),X160),s(fun(X159,X164),X165),s(fun(X164,X159),X166)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X157,fun(X157,bool)),X158),s(fun(X157,X167),X168),s(fun(X167,X157),X169)))) | s(bool,'F_0') = s(bool,'const.quotient_sum.+++_4'(s(fun(X157,fun(X157,bool)),X158),s(fun(X159,fun(X159,bool)),X160),s('type.sum.sum'(X157,X159),X161),s('type.sum.sum'(X157,X159),X162))) | s(bool,'F_0') = s(bool,X163)) )),
  inference(superposition,[],[f130,f5187])).
fof(f23359,plain,(
  ( ! [X146,X144,X156,X154,X152,X151,X149,X147,X145,X155,X153,X150,X148] : (s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X144,X146),X148))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X144,X146),X149))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X146,fun(X146,bool)),X147),s(fun(X146,X151),X152),s(fun(X151,X146),X153)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X144,fun(X144,bool)),X145),s(fun(X144,X154),X155),s(fun(X154,X144),X156)))) | s(bool,'F_0') = s(bool,X150) | s(bool,'F_0') = s(bool,'const.quotient_sum.+++_4'(s(fun(X144,fun(X144,bool)),X145),s(fun(X146,fun(X146,bool)),X147),s('type.sum.sum'(X144,X146),X148),s('type.sum.sum'(X144,X146),X149)))) )),
  inference(subsumption_resolution,[],[f23131,f119])).
fof(f23131,plain,(
  ( ! [X146,X144,X156,X154,X152,X151,X149,X147,X145,X155,X153,X150,X148] : (~p(s(bool,X150)) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X144,X146),X148))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X144,X146),X149))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X146,fun(X146,bool)),X147),s(fun(X146,X151),X152),s(fun(X151,X146),X153)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X144,fun(X144,bool)),X145),s(fun(X144,X154),X155),s(fun(X154,X144),X156)))) | s(bool,'F_0') = s(bool,X150) | s(bool,'F_0') = s(bool,'const.quotient_sum.+++_4'(s(fun(X144,fun(X144,bool)),X145),s(fun(X146,fun(X146,bool)),X147),s('type.sum.sum'(X144,X146),X148),s('type.sum.sum'(X144,X146),X149)))) )),
  inference(superposition,[],[f130,f5187])).
fof(f22866,plain,(
  ( ! [X64,X62] : (~p(s(bool,X62)) | s(bool,'F_0') = s(bool,X64) | s(bool,X62) = s(bool,X64)) ) | $spl290),
  inference(cnf_transformation,[],[f22866_D])).
fof(f22866_D,plain,(
  ( ! [X64,X62] : (~p(s(bool,X62)) | s(bool,'F_0') = s(bool,X64) | s(bool,X62) = s(bool,X64)) ) <=> ~$spl290),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl290])])).
fof(f22867,plain,(
  ( ! [X66,X67,X65] : (~p(s(bool,X65)) | s(bool,'F_0') != s(bool,X66) | s(bool,X65) = s(bool,X67) | s(bool,'F_0') = s(bool,X67)) )),
  inference(subsumption_resolution,[],[f22836,f118])).
fof(f22836,plain,(
  ( ! [X66,X67,X65] : (~p(s(bool,X65)) | s(bool,'F_0') != s(bool,X66) | p(s(bool,X66)) | s(bool,X65) = s(bool,X67) | s(bool,'F_0') = s(bool,X67)) )),
  inference(superposition,[],[f22072,f5187])).
fof(f22865,plain,(
  ( ! [X64,X62,X63] : (~p(s(bool,X62)) | s(bool,'F_0') != s(bool,X63) | s(bool,X62) = s(bool,X64) | s(bool,'F_0') = s(bool,X64)) )),
  inference(subsumption_resolution,[],[f22835,f118])).
fof(f22835,plain,(
  ( ! [X64,X62,X63] : (~p(s(bool,X62)) | s(bool,'F_0') != s(bool,X63) | p(s(bool,X63)) | s(bool,X62) = s(bool,X64) | s(bool,'F_0') = s(bool,X64)) )),
  inference(superposition,[],[f22072,f5187])).
fof(f11978,plain,(
  ( ! [X114,X115] : (p(s(bool,X115)) | s(bool,'F_0') = s(bool,X114) | s(bool,X114) != s(bool,X115)) ) | $spl160),
  inference(cnf_transformation,[],[f11978_D])).
fof(f11978_D,plain,(
  ( ! [X114,X115] : (p(s(bool,X115)) | s(bool,'F_0') = s(bool,X114) | s(bool,X114) != s(bool,X115)) ) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f22860,plain,(
  ( ! [X33,X31,X32] : (s(bool,X31) != s(bool,X33) | p(s(bool,X33)) | s(bool,'F_0') = s(bool,X31) | s(bool,'F_0') = s(bool,X32)) )),
  inference(subsumption_resolution,[],[f22823,f119])).
fof(f22823,plain,(
  ( ! [X33,X31,X32] : (~p(s(bool,X32)) | s(bool,X31) != s(bool,X33) | p(s(bool,X33)) | s(bool,'F_0') = s(bool,X31) | s(bool,'F_0') = s(bool,X32)) )),
  inference(superposition,[],[f22072,f5187])).
fof(f22859,plain,(
  ( ! [X30,X28,X29] : (s(bool,X28) != s(bool,X30) | p(s(bool,X30)) | s(bool,'F_0') = s(bool,X29) | s(bool,'F_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f22822,f119])).
fof(f22822,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,X29)) | s(bool,X28) != s(bool,X30) | p(s(bool,X30)) | s(bool,'F_0') = s(bool,X29) | s(bool,'F_0') = s(bool,X28)) )),
  inference(superposition,[],[f22072,f5187])).
fof(f8631,plain,(
  ( ! [X208,X198,X206,X204,X202,X200,X209,X199,X207,X205,X201] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X200,X204),X205))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X198,X206),'const.sum.++_3'(s(fun(X200,X198),X202),s(fun(X204,X206),X207),s('type.sum.sum'(X200,X204),X205))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X198,fun(X198,bool)),X199),s(fun(X198,X200),X201),s(fun(X200,X198),X202))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X206,fun(X206,bool)),X208),s(fun(X206,X204),X209),s(fun(X204,X206),X207))))) ) | $spl52),
  inference(cnf_transformation,[],[f8631_D])).
fof(f8631_D,plain,(
  ( ! [X208,X198,X206,X204,X202,X200,X209,X199,X207,X205,X201] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X200,X204),X205))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X198,X206),'const.sum.++_3'(s(fun(X200,X198),X202),s(fun(X204,X206),X207),s('type.sum.sum'(X200,X204),X205))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X198,fun(X198,bool)),X199),s(fun(X198,X200),X201),s(fun(X200,X198),X202))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X206,fun(X206,bool)),X208),s(fun(X206,X204),X209),s(fun(X204,X206),X207))))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f22776,plain,(
  ( ! [X158,X156,X154,X152,X161,X151,X159,X157,X155,X153,X162,X160] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X153,X157),X158))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X151,X159),'const.sum.++_3'(s(fun(X153,X151),X155),s(fun(X157,X159),X160),s('type.sum.sum'(X153,X157),X158))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X159,fun(X159,bool)),X161),s(fun(X159,X157),X162),s(fun(X157,X159),X160)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X151,fun(X151,bool)),X152),s(fun(X151,X153),X154),s(fun(X153,X151),X155))) | s(bool,'F_0') = s(bool,X156)) )),
  inference(subsumption_resolution,[],[f22687,f119])).
fof(f22687,plain,(
  ( ! [X158,X156,X154,X152,X161,X151,X159,X157,X155,X153,X162,X160] : (~p(s(bool,X156)) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X153,X157),X158))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X151,X159),'const.sum.++_3'(s(fun(X153,X151),X155),s(fun(X157,X159),X160),s('type.sum.sum'(X153,X157),X158))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X159,fun(X159,bool)),X161),s(fun(X159,X157),X162),s(fun(X157,X159),X160)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X151,fun(X151,bool)),X152),s(fun(X151,X153),X154),s(fun(X153,X151),X155))) | s(bool,'F_0') = s(bool,X156)) )),
  inference(superposition,[],[f127,f5187])).
fof(f22775,plain,(
  ( ! [X146,X144,X142,X140,X149,X147,X145,X143,X141,X139,X150,X148] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X141,X145),X146))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X139,X147),'const.sum.++_3'(s(fun(X141,X139),X143),s(fun(X145,X147),X148),s('type.sum.sum'(X141,X145),X146))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X147,fun(X147,bool)),X149),s(fun(X147,X145),X150),s(fun(X145,X147),X148)))) | s(bool,'F_0') = s(bool,X144) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X139,fun(X139,bool)),X140),s(fun(X139,X141),X142),s(fun(X141,X139),X143)))) )),
  inference(subsumption_resolution,[],[f22686,f119])).
fof(f22686,plain,(
  ( ! [X146,X144,X142,X140,X149,X147,X145,X143,X141,X139,X150,X148] : (~p(s(bool,X144)) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X141,X145),X146))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X139,X147),'const.sum.++_3'(s(fun(X141,X139),X143),s(fun(X145,X147),X148),s('type.sum.sum'(X141,X145),X146))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X147,fun(X147,bool)),X149),s(fun(X147,X145),X150),s(fun(X145,X147),X148)))) | s(bool,'F_0') = s(bool,X144) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X139,fun(X139,bool)),X140),s(fun(X139,X141),X142),s(fun(X141,X139),X143)))) )),
  inference(superposition,[],[f127,f5187])).
fof(f8633,plain,(
  ( ! [X212,X218,X216,X210,X220,X213,X219,X217,X211,X221,X214] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X216,X212),X217))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X218,X210),'const.sum.++_3'(s(fun(X216,X218),X219),s(fun(X212,X210),X214),s('type.sum.sum'(X216,X212),X217))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X210,fun(X210,bool)),X211),s(fun(X210,X212),X213),s(fun(X212,X210),X214))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X218,fun(X218,bool)),X220),s(fun(X218,X216),X221),s(fun(X216,X218),X219))))) ) | $spl54),
  inference(cnf_transformation,[],[f8633_D])).
fof(f8633_D,plain,(
  ( ! [X212,X218,X216,X210,X220,X213,X219,X217,X211,X221,X214] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X216,X212),X217))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X218,X210),'const.sum.++_3'(s(fun(X216,X218),X219),s(fun(X212,X210),X214),s('type.sum.sum'(X216,X212),X217))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X210,fun(X210,bool)),X211),s(fun(X210,X212),X213),s(fun(X212,X210),X214))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X218,fun(X218,bool)),X220),s(fun(X218,X216),X221),s(fun(X216,X218),X219))))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f22766,plain,(
  ( ! [X158,X156,X154,X152,X161,X151,X159,X157,X155,X153,X162,X160] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X157,X153),X158))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X159,X151),'const.sum.++_3'(s(fun(X157,X159),X160),s(fun(X153,X151),X155),s('type.sum.sum'(X157,X153),X158))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X159,fun(X159,bool)),X161),s(fun(X159,X157),X162),s(fun(X157,X159),X160)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X151,fun(X151,bool)),X152),s(fun(X151,X153),X154),s(fun(X153,X151),X155))) | s(bool,'F_0') = s(bool,X156)) )),
  inference(subsumption_resolution,[],[f22667,f119])).
fof(f22667,plain,(
  ( ! [X158,X156,X154,X152,X161,X151,X159,X157,X155,X153,X162,X160] : (~p(s(bool,X156)) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X157,X153),X158))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X159,X151),'const.sum.++_3'(s(fun(X157,X159),X160),s(fun(X153,X151),X155),s('type.sum.sum'(X157,X153),X158))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X159,fun(X159,bool)),X161),s(fun(X159,X157),X162),s(fun(X157,X159),X160)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X151,fun(X151,bool)),X152),s(fun(X151,X153),X154),s(fun(X153,X151),X155))) | s(bool,'F_0') = s(bool,X156)) )),
  inference(superposition,[],[f127,f5187])).
fof(f22765,plain,(
  ( ! [X146,X144,X142,X140,X149,X147,X145,X143,X141,X139,X150,X148] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X145,X141),X146))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X147,X139),'const.sum.++_3'(s(fun(X145,X147),X148),s(fun(X141,X139),X143),s('type.sum.sum'(X145,X141),X146))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X147,fun(X147,bool)),X149),s(fun(X147,X145),X150),s(fun(X145,X147),X148)))) | s(bool,'F_0') = s(bool,X144) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X139,fun(X139,bool)),X140),s(fun(X139,X141),X142),s(fun(X141,X139),X143)))) )),
  inference(subsumption_resolution,[],[f22666,f119])).
fof(f22666,plain,(
  ( ! [X146,X144,X142,X140,X149,X147,X145,X143,X141,X139,X150,X148] : (~p(s(bool,X144)) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X145,X141),X146))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X147,X139),'const.sum.++_3'(s(fun(X145,X147),X148),s(fun(X141,X139),X143),s('type.sum.sum'(X145,X141),X146))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X147,fun(X147,bool)),X149),s(fun(X147,X145),X150),s(fun(X145,X147),X148)))) | s(bool,'F_0') = s(bool,X144) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X139,fun(X139,bool)),X140),s(fun(X139,X141),X142),s(fun(X141,X139),X143)))) )),
  inference(superposition,[],[f127,f5187])).
fof(f22378,plain,(
  ( ! [X39] : (s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(bool,X39),'const.sum.INL_1'(s(bool,'F_0')))))) ) | $spl288),
  inference(cnf_transformation,[],[f22378_D])).
fof(f22378_D,plain,(
  ( ! [X39] : (s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(bool,X39),'const.sum.INL_1'(s(bool,'F_0')))))) ) <=> ~$spl288),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl288])])).
fof(f22377,plain,(
  ( ! [X39,X38] : (s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(bool,X39),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X38)) )),
  inference(subsumption_resolution,[],[f22254,f123])).
fof(f22254,plain,(
  ( ! [X39,X38] : (s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(bool,X39),'const.sum.INL_1'(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X38) | p(s(bool,X38))) )),
  inference(superposition,[],[f5507,f119])).
fof(f167,plain,(
  ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11)))) | $spl1),
  inference(cnf_transformation,[],[f167_D])).
fof(f167_D,plain,(
  ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f22077,plain,(
  ( ! [X46] : (~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11)))) | s(bool,'T_0') = s(bool,X46)) ) | $spl10),
  inference(subsumption_resolution,[],[f22022,f126])).
fof(f22022,plain,(
  ( ! [X46] : (s(bool,'F_0') != s(bool,X46) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11)))) | s(bool,'T_0') = s(bool,X46)) ) | $spl10),
  inference(superposition,[],[f118,f19305])).
fof(f22049,plain,(
  ( ! [X28] : (~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11)))) | s(bool,'T_0') = s(bool,X28)) ) | $spl10),
  inference(subsumption_resolution,[],[f21953,f126])).
fof(f21953,plain,(
  ( ! [X28] : (~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11)))) | s(bool,'F_0') != s(bool,X28) | s(bool,'T_0') = s(bool,X28)) ) | $spl10),
  inference(superposition,[],[f118,f19305])).
fof(f21890,plain,(
  ~p(s(bool,'F_0')) | $spl287),
  inference(cnf_transformation,[],[f21890_D])).
fof(f21890_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl287),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl287])])).
fof(f138,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f98])).
fof(f98,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f97])).
fof(f97,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X22] : ((p(s(bool,X22)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X22)))),
  file('thm.quotient_sum.ISR_RSP','thm.bool.IMP_CLAUSES|split|4')).
fof(f12004,plain,(
  ( ! [X31] : (p(s(bool,X31))) ) | $spl180),
  inference(cnf_transformation,[],[f12004_D])).
fof(f12004_D,plain,(
  ( ! [X31] : (p(s(bool,X31))) ) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f21891,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(subsumption_resolution,[],[f145,f141])).
fof(f141,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f100])).
fof(f100,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))))),
  inference(flattening,[],[f99])).
fof(f99,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f85])).
fof(f85,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X22] : ((p(s(bool,X22)) => p(s(bool,X22))) <=> p(s(bool,'T_0')))),
  file('thm.quotient_sum.ISR_RSP','thm.bool.IMP_CLAUSES|split|3')).
fof(f145,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f101])).
fof(f101,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f86])).
fof(f86,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X22] : ((p(s(bool,'F_0')) => p(s(bool,X22))) <=> p(s(bool,'T_0')))),
  file('thm.quotient_sum.ISR_RSP','thm.bool.IMP_CLAUSES|split|2')).
fof(f149,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f88])).
fof(f88,plain,(
  ! [X0] : (~p(s(bool,'F_0')) | p(s(bool,X0)))),
  inference(ennf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0] : (p(s(bool,'F_0')) => p(s(bool,X0)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X22] : (p(s(bool,'F_0')) => p(s(bool,X22)))),
  file('thm.quotient_sum.ISR_RSP','thm.bool.FALSITY')).
fof(f21869,plain,(
  $false | $spl266),
  inference(subsumption_resolution,[],[f21868,f17369])).
fof(f21868,plain,(
  ( ! [X21] : (s(bool,'T_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X21)))) ) | $spl266),
  inference(subsumption_resolution,[],[f21567,f155])).
fof(f21567,plain,(
  ( ! [X21] : (~p(s(bool,'T_0')) | s(bool,'T_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X21)))) ) | $spl266),
  inference(backward_demodulation,[],[f17369,f2598])).
fof(f21867,plain,(
  $false | $spl266),
  inference(subsumption_resolution,[],[f21866,f156])).
fof(f156,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.quotient_sum.ISR_RSP',aHL_TRUTH)).
fof(f21866,plain,(
  ~p(s(bool,'T_0')) | $spl266),
  inference(forward_demodulation,[],[f21774,f17369])).
fof(f21774,plain,(
  ( ! [X6] : (~p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X6))))) ) | $spl266),
  inference(trivial_inequality_removal,[],[f21566])).
fof(f21566,plain,(
  ( ! [X6] : (s(bool,'T_0') != s(bool,'T_0') | ~p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X6))))) ) | $spl266),
  inference(backward_demodulation,[],[f17369,f2215])).
fof(f21813,plain,(
  $false | $spl266),
  inference(subsumption_resolution,[],[f21812,f156])).
fof(f21812,plain,(
  ~p(s(bool,'T_0')) | $spl266),
  inference(forward_demodulation,[],[f21811,f17369])).
fof(f21811,plain,(
  ( ! [X118,X122,X120] : (~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X118,X120),X122))))) ) | $spl266),
  inference(subsumption_resolution,[],[f21810,f141])).
fof(f21810,plain,(
  ( ! [X118,X122,X120] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X118,X120),X122))))) ) | $spl266),
  inference(forward_demodulation,[],[f21809,f17369])).
fof(f21809,plain,(
  ( ! [X123,X121,X118,X124,X122,X120,X117] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X117,fun(X117,bool)),X123),s(fun(X117,X120),X124),s(fun(X120,X117),X121)))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X118,X120),X122))))) ) | $spl266),
  inference(subsumption_resolution,[],[f21808,f141])).
fof(f21808,plain,(
  ( ! [X123,X121,X118,X124,X122,X120,X117] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X117,fun(X117,bool)),X123),s(fun(X117,X120),X124),s(fun(X120,X117),X121)))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X118,X120),X122))))) ) | $spl266),
  inference(forward_demodulation,[],[f21807,f17369])).
fof(f21807,plain,(
  ( ! [X125,X123,X121,X118,X116,X126,X124,X122,X120,X119,X117] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X116,fun(X116,bool)),X125),s(fun(X116,X118),X126),s(fun(X118,X116),X119)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X117,fun(X117,bool)),X123),s(fun(X117,X120),X124),s(fun(X120,X117),X121)))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X118,X120),X122))))) ) | $spl266),
  inference(subsumption_resolution,[],[f21537,f17369])).
fof(f21537,plain,(
  ( ! [X125,X123,X121,X118,X116,X126,X124,X122,X120,X119,X117] : (s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X116,X117),'const.sum.++_3'(s(fun(X118,X116),X119),s(fun(X120,X117),X121),s('type.sum.sum'(X118,X120),X122))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X116,fun(X116,bool)),X125),s(fun(X116,X118),X126),s(fun(X118,X116),X119)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X117,fun(X117,bool)),X123),s(fun(X117,X120),X124),s(fun(X120,X117),X121)))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X118,X120),X122))))) ) | $spl266),
  inference(backward_demodulation,[],[f17369,f9511])).
fof(f21806,plain,(
  $false | $spl266),
  inference(subsumption_resolution,[],[f21805,f156])).
fof(f21805,plain,(
  ~p(s(bool,'T_0')) | $spl266),
  inference(forward_demodulation,[],[f21804,f17369])).
fof(f21804,plain,(
  ( ! [X109,X112,X110,X106,X113] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X106,fun(X106,bool)),X112),s(fun(X106,X109),X113),s(fun(X109,X106),X110))))) ) | $spl266),
  inference(subsumption_resolution,[],[f21803,f141])).
fof(f21803,plain,(
  ( ! [X109,X112,X110,X106,X113] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X106,fun(X106,bool)),X112),s(fun(X106,X109),X113),s(fun(X109,X106),X110))))) ) | $spl266),
  inference(forward_demodulation,[],[f21802,f17369])).
fof(f21802,plain,(
  ( ! [X111,X109,X107,X105,X112,X110,X108,X106,X113] : (~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X105,X106),'const.sum.++_3'(s(fun(X107,X105),X108),s(fun(X109,X106),X110),s('type.sum.sum'(X107,X109),X111)))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X106,fun(X106,bool)),X112),s(fun(X106,X109),X113),s(fun(X109,X106),X110))))) ) | $spl266),
  inference(subsumption_resolution,[],[f21801,f141])).
fof(f21801,plain,(
  ( ! [X111,X109,X107,X105,X112,X110,X108,X106,X113] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X105,X106),'const.sum.++_3'(s(fun(X107,X105),X108),s(fun(X109,X106),X110),s('type.sum.sum'(X107,X109),X111)))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X106,fun(X106,bool)),X112),s(fun(X106,X109),X113),s(fun(X109,X106),X110))))) ) | $spl266),
  inference(forward_demodulation,[],[f21800,f17369])).
fof(f21800,plain,(
  ( ! [X111,X109,X107,X105,X114,X112,X110,X108,X106,X115,X113] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X105,fun(X105,bool)),X114),s(fun(X105,X107),X115),s(fun(X107,X105),X108)))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X105,X106),'const.sum.++_3'(s(fun(X107,X105),X108),s(fun(X109,X106),X110),s('type.sum.sum'(X107,X109),X111)))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X106,fun(X106,bool)),X112),s(fun(X106,X109),X113),s(fun(X109,X106),X110))))) ) | $spl266),
  inference(subsumption_resolution,[],[f21536,f17369])).
fof(f21536,plain,(
  ( ! [X111,X109,X107,X105,X114,X112,X110,X108,X106,X115,X113] : (s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X107,X109),X111))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X105,fun(X105,bool)),X114),s(fun(X105,X107),X115),s(fun(X107,X105),X108)))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X105,X106),'const.sum.++_3'(s(fun(X107,X105),X108),s(fun(X109,X106),X110),s('type.sum.sum'(X107,X109),X111)))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X106,fun(X106,bool)),X112),s(fun(X106,X109),X113),s(fun(X109,X106),X110))))) ) | $spl266),
  inference(backward_demodulation,[],[f17369,f9510])).
fof(f21798,plain,(
  $false | $spl266),
  inference(subsumption_resolution,[],[f21777,f17369])).
fof(f21777,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X1)))) ) | $spl266),
  inference(duplicate_literal_removal,[],[f21528])).
fof(f21528,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X1))) | s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X1)))) ) | $spl266),
  inference(backward_demodulation,[],[f17369,f4326])).
fof(f4326,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X1))) | s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X1)))) )),
  inference(subsumption_resolution,[],[f4299,f2351])).
fof(f4299,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X1))) | s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,sK13(X1,sK0,sK0)))))) | s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X1)))) )),
  inference(superposition,[],[f4134,f1692])).
fof(f4134,plain,(
  ( ! [X9] : (s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X9))) | s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X9)))) )),
  inference(resolution,[],[f4105,f118])).
fof(f21795,plain,(
  $false | $spl266),
  inference(subsumption_resolution,[],[f21778,f17369])).
fof(f21778,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X3)))) ) | $spl266),
  inference(duplicate_literal_removal,[],[f21526])).
fof(f21526,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X3))) | s(bool,'T_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X3)))) ) | $spl266),
  inference(backward_demodulation,[],[f17369,f3609])).
fof(f3609,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X3))) | s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X3)))) )),
  inference(subsumption_resolution,[],[f3584,f1592])).
fof(f3584,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X3))) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,sK12(X3,sK0,sK0)))))) | s(bool,'T_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X3)))) )),
  inference(superposition,[],[f3541,f2467])).
fof(f2467,plain,(
  ( ! [X2] : (s('type.sum.sum'(sK0,sK0),X2) = s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,sK12(X2,sK0,sK0)))) | s(bool,'T_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X2)))) )),
  inference(superposition,[],[f2457,f111])).
fof(f2457,plain,(
  ( ! [X50] : (s(bool,'T_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X50)))))) )),
  inference(subsumption_resolution,[],[f2424,f150])).
fof(f2424,plain,(
  ( ! [X50] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X50)))))) )),
  inference(superposition,[],[f122,f2351])).
fof(f21794,plain,(
  $false | $spl266),
  inference(subsumption_resolution,[],[f21779,f156])).
fof(f21779,plain,(
  ~p(s(bool,'T_0')) | $spl266),
  inference(duplicate_literal_removal,[],[f21519])).
fof(f21519,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'T_0')) | $spl266),
  inference(backward_demodulation,[],[f17369,f154])).
fof(f154,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f106])).
fof(f106,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f78])).
fof(f78,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f15])).
fof(f15,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.quotient_sum.ISR_RSP','thm.bool.NOT_CLAUSES|split|2')).
fof(f21793,plain,(
  $false | $spl266),
  inference(subsumption_resolution,[],[f21781,f156])).
fof(f21781,plain,(
  ~p(s(bool,'T_0')) | $spl266),
  inference(duplicate_literal_removal,[],[f21517])).
fof(f21517,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'T_0')) | $spl266),
  inference(backward_demodulation,[],[f17369,f152])).
fof(f152,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f105])).
fof(f105,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f77])).
fof(f77,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f14])).
fof(f14,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.quotient_sum.ISR_RSP','thm.bool.NOT_CLAUSES|split|1')).
fof(f21792,plain,(
  $false | $spl266),
  inference(subsumption_resolution,[],[f21515,f156])).
fof(f21515,plain,(
  ~p(s(bool,'T_0')) | $spl266),
  inference(backward_demodulation,[],[f17369,f150])).
fof(f21788,plain,(
  $false | $spl266),
  inference(subsumption_resolution,[],[f21787,f17369])).
fof(f21787,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0)) ) | $spl266),
  inference(subsumption_resolution,[],[f21509,f124])).
fof(f124,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f95])).
fof(f21509,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl266),
  inference(backward_demodulation,[],[f17369,f120])).
fof(f120,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f93])).
fof(f21786,plain,(
  $false | $spl266),
  inference(subsumption_resolution,[],[f21785,f17369])).
fof(f21785,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0)) ) | $spl266),
  inference(subsumption_resolution,[],[f21507,f124])).
fof(f21507,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl266),
  inference(backward_demodulation,[],[f17369,f118])).
fof(f21751,plain,(
  $false | $spl266),
  inference(trivial_inequality_removal,[],[f21716])).
fof(f21716,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl266),
  inference(superposition,[],[f1128,f17369])).
fof(f21750,plain,(
  $false | $spl266),
  inference(trivial_inequality_removal,[],[f21721])).
fof(f21721,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl266),
  inference(superposition,[],[f5633,f17369])).
fof(f21748,plain,(
  $false | $spl266),
  inference(trivial_inequality_removal,[],[f21724])).
fof(f21724,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl266),
  inference(superposition,[],[f1592,f17369])).
fof(f21747,plain,(
  $false | $spl266),
  inference(trivial_inequality_removal,[],[f21726])).
fof(f21726,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl266),
  inference(superposition,[],[f1593,f17369])).
fof(f1593,plain,(
  ( ! [X26] : (s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK2),'const.sum.INL_1'(s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,X26)))))))) )),
  inference(subsumption_resolution,[],[f1556,f128])).
fof(f1556,plain,(
  ( ! [X26] : (p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,X26)))))) | s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK2),'const.sum.INL_1'(s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,X26)))))))) )),
  inference(superposition,[],[f122,f1442])).
fof(f21745,plain,(
  $false | ($spl8 | $spl266)),
  inference(trivial_inequality_removal,[],[f21728])).
fof(f21728,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl8 | $spl266)),
  inference(superposition,[],[f12281,f17369])).
fof(f12281,plain,(
  s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | $spl8),
  inference(backward_demodulation,[],[f173,f110])).
fof(f110,plain,(
  s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11)))),
  inference(cnf_transformation,[],[f89])).
fof(f21744,plain,(
  $false | $spl266),
  inference(trivial_inequality_removal,[],[f21730])).
fof(f21730,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl266),
  inference(superposition,[],[f5815,f17369])).
fof(f21743,plain,(
  $false | $spl266),
  inference(trivial_inequality_removal,[],[f21732])).
fof(f21732,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl266),
  inference(superposition,[],[f2457,f17369])).
fof(f21423,plain,(
  ( ! [X668,X669] : (s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X668))) | s(bool,'T_0') = s(bool,X669) | s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X668)))) )),
  inference(subsumption_resolution,[],[f20693,f123])).
fof(f20693,plain,(
  ( ! [X668,X669] : (p(s(bool,X669)) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X668))) | s(bool,'T_0') = s(bool,X669) | s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X668)))) )),
  inference(superposition,[],[f3517,f19171])).
fof(f21420,plain,(
  ( ! [X639,X644,X640,X642,X638,X645,X641,X643] : (s(bool,'T_0') = s(bool,X645) | s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X638,X639),'const.sum.++_3'(s(fun(X640,X638),X641),s(fun(X642,X639),X643),s('type.sum.sum'(X640,X642),'const.sum.INL_1'(s(X640,X644)))))))) )),
  inference(subsumption_resolution,[],[f20685,f123])).
fof(f20685,plain,(
  ( ! [X639,X644,X640,X642,X638,X645,X641,X643] : (p(s(bool,X645)) | s(bool,'T_0') = s(bool,X645) | s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X638,X639),'const.sum.++_3'(s(fun(X640,X638),X641),s(fun(X642,X639),X643),s('type.sum.sum'(X640,X642),'const.sum.INL_1'(s(X640,X644)))))))) )),
  inference(superposition,[],[f9289,f19171])).
fof(f21419,plain,(
  ( ! [X609,X611,X608,X610] : (s('type.sum.sum'(X608,X609),X610) = s('type.sum.sum'(X608,X609),'const.sum.INR_1'(s(X609,sK13(X610,X609,X608)))) | s(bool,'T_0') = s(bool,X611) | s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X608,X609),X610)))) )),
  inference(subsumption_resolution,[],[f20682,f123])).
fof(f20682,plain,(
  ( ! [X609,X611,X608,X610] : (p(s(bool,X611)) | s('type.sum.sum'(X608,X609),X610) = s('type.sum.sum'(X608,X609),'const.sum.INR_1'(s(X609,sK13(X610,X609,X608)))) | s(bool,'T_0') = s(bool,X611) | s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X608,X609),X610)))) )),
  inference(superposition,[],[f9288,f19171])).
fof(f21414,plain,(
  ( ! [X580,X579] : (s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X579))) | s(bool,'T_0') = s(bool,X580) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X579)))) )),
  inference(subsumption_resolution,[],[f20667,f123])).
fof(f20667,plain,(
  ( ! [X580,X579] : (p(s(bool,X580)) | s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X579))) | s(bool,'T_0') = s(bool,X580) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X579)))) )),
  inference(superposition,[],[f4105,f19171])).
fof(f21411,plain,(
  ( ! [X546,X542,X549,X545,X547,X543,X548,X544] : (s(bool,'T_0') = s(bool,X549) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X542,X543),'const.sum.++_3'(s(fun(X544,X542),X545),s(fun(X546,X543),X547),s('type.sum.sum'(X544,X546),'const.sum.INR_1'(s(X546,X548)))))))) )),
  inference(subsumption_resolution,[],[f20656,f123])).
fof(f20656,plain,(
  ( ! [X546,X542,X549,X545,X547,X543,X548,X544] : (p(s(bool,X549)) | s(bool,'T_0') = s(bool,X549) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X542,X543),'const.sum.++_3'(s(fun(X544,X542),X545),s(fun(X546,X543),X547),s('type.sum.sum'(X544,X546),'const.sum.INR_1'(s(X546,X548)))))))) )),
  inference(superposition,[],[f9190,f19171])).
fof(f21410,plain,(
  ( ! [X528,X526,X529,X527] : (s('type.sum.sum'(X526,X527),X528) = s('type.sum.sum'(X526,X527),'const.sum.INL_1'(s(X526,sK12(X528,X527,X526)))) | s(bool,'T_0') = s(bool,X529) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X526,X527),X528)))) )),
  inference(subsumption_resolution,[],[f20654,f123])).
fof(f20654,plain,(
  ( ! [X528,X526,X529,X527] : (p(s(bool,X529)) | s('type.sum.sum'(X526,X527),X528) = s('type.sum.sum'(X526,X527),'const.sum.INL_1'(s(X526,sK12(X528,X527,X526)))) | s(bool,'T_0') = s(bool,X529) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X526,X527),X528)))) )),
  inference(superposition,[],[f9189,f19171])).
fof(f21383,plain,(
  ( ! [X668,X669] : (s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X668))) | s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X668))) | s(bool,'T_0') = s(bool,X669)) )),
  inference(subsumption_resolution,[],[f20501,f123])).
fof(f20501,plain,(
  ( ! [X668,X669] : (p(s(bool,X669)) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X668))) | s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X668))) | s(bool,'T_0') = s(bool,X669)) )),
  inference(superposition,[],[f3517,f19171])).
fof(f21379,plain,(
  ( ! [X639,X644,X640,X642,X638,X645,X641,X643] : (s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X638,X639),'const.sum.++_3'(s(fun(X640,X638),X641),s(fun(X642,X639),X643),s('type.sum.sum'(X640,X642),'const.sum.INL_1'(s(X640,X644))))))) | s(bool,'T_0') = s(bool,X645)) )),
  inference(subsumption_resolution,[],[f20493,f123])).
fof(f20493,plain,(
  ( ! [X639,X644,X640,X642,X638,X645,X641,X643] : (p(s(bool,X645)) | s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X638,X639),'const.sum.++_3'(s(fun(X640,X638),X641),s(fun(X642,X639),X643),s('type.sum.sum'(X640,X642),'const.sum.INL_1'(s(X640,X644))))))) | s(bool,'T_0') = s(bool,X645)) )),
  inference(superposition,[],[f9289,f19171])).
fof(f21377,plain,(
  ( ! [X609,X611,X608,X610] : (s('type.sum.sum'(X608,X609),X610) = s('type.sum.sum'(X608,X609),'const.sum.INR_1'(s(X609,sK13(X610,X609,X608)))) | s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X608,X609),X610))) | s(bool,'T_0') = s(bool,X611)) )),
  inference(subsumption_resolution,[],[f20490,f123])).
fof(f20490,plain,(
  ( ! [X609,X611,X608,X610] : (p(s(bool,X611)) | s('type.sum.sum'(X608,X609),X610) = s('type.sum.sum'(X608,X609),'const.sum.INR_1'(s(X609,sK13(X610,X609,X608)))) | s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X608,X609),X610))) | s(bool,'T_0') = s(bool,X611)) )),
  inference(superposition,[],[f9288,f19171])).
fof(f21371,plain,(
  ( ! [X580,X579] : (s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X579))) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X579))) | s(bool,'T_0') = s(bool,X580)) )),
  inference(subsumption_resolution,[],[f20475,f123])).
fof(f20475,plain,(
  ( ! [X580,X579] : (p(s(bool,X580)) | s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X579))) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X579))) | s(bool,'T_0') = s(bool,X580)) )),
  inference(superposition,[],[f4105,f19171])).
fof(f21367,plain,(
  ( ! [X546,X542,X549,X545,X547,X543,X548,X544] : (s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X542,X543),'const.sum.++_3'(s(fun(X544,X542),X545),s(fun(X546,X543),X547),s('type.sum.sum'(X544,X546),'const.sum.INR_1'(s(X546,X548))))))) | s(bool,'T_0') = s(bool,X549)) )),
  inference(subsumption_resolution,[],[f20464,f123])).
fof(f20464,plain,(
  ( ! [X546,X542,X549,X545,X547,X543,X548,X544] : (p(s(bool,X549)) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X542,X543),'const.sum.++_3'(s(fun(X544,X542),X545),s(fun(X546,X543),X547),s('type.sum.sum'(X544,X546),'const.sum.INR_1'(s(X546,X548))))))) | s(bool,'T_0') = s(bool,X549)) )),
  inference(superposition,[],[f9190,f19171])).
fof(f21365,plain,(
  ( ! [X528,X526,X529,X527] : (s('type.sum.sum'(X526,X527),X528) = s('type.sum.sum'(X526,X527),'const.sum.INL_1'(s(X526,sK12(X528,X527,X526)))) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X526,X527),X528))) | s(bool,'T_0') = s(bool,X529)) )),
  inference(subsumption_resolution,[],[f20462,f123])).
fof(f20462,plain,(
  ( ! [X528,X526,X529,X527] : (p(s(bool,X529)) | s('type.sum.sum'(X526,X527),X528) = s('type.sum.sum'(X526,X527),'const.sum.INL_1'(s(X526,sK12(X528,X527,X526)))) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X526,X527),X528))) | s(bool,'T_0') = s(bool,X529)) )),
  inference(superposition,[],[f9189,f19171])).
fof(f18199,plain,(
  ( ! [X177,X185,X175,X182,X180,X178,X176,X184,X174,X183,X181,X179] : (s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X174,X180),'const.sum.++_3'(s(fun(X176,X174),X178),s(fun(X181,X180),X182),s('type.sum.sum'(X176,X181),X183))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X180,fun(X180,bool)),X184),s(fun(X180,X181),X185),s(fun(X181,X180),X182)))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X176,X181),X183)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X174,fun(X174,bool)),X175),s(fun(X174,X176),X177),s(fun(X176,X174),X178))) | s(bool,'F_0') = s(bool,X179)) )),
  inference(subsumption_resolution,[],[f18180,f119])).
fof(f18180,plain,(
  ( ! [X177,X185,X175,X182,X180,X178,X176,X184,X174,X183,X181,X179] : (~p(s(bool,X179)) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X174,X180),'const.sum.++_3'(s(fun(X176,X174),X178),s(fun(X181,X180),X182),s('type.sum.sum'(X176,X181),X183))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X180,fun(X180,bool)),X184),s(fun(X180,X181),X185),s(fun(X181,X180),X182)))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X176,X181),X183)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X174,fun(X174,bool)),X175),s(fun(X174,X176),X177),s(fun(X176,X174),X178))) | s(bool,'F_0') = s(bool,X179)) )),
  inference(superposition,[],[f9521,f5187])).
fof(f18197,plain,(
  ( ! [X167,X165,X163,X173,X171,X169,X166,X164,X162,X172,X170,X168] : (s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X162,X168),'const.sum.++_3'(s(fun(X164,X162),X166),s(fun(X169,X168),X170),s('type.sum.sum'(X164,X169),X171))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X168,fun(X168,bool)),X172),s(fun(X168,X169),X173),s(fun(X169,X168),X170)))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X164,X169),X171)))) | s(bool,'F_0') = s(bool,X167) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X162,fun(X162,bool)),X163),s(fun(X162,X164),X165),s(fun(X164,X162),X166)))) )),
  inference(subsumption_resolution,[],[f18179,f119])).
fof(f18179,plain,(
  ( ! [X167,X165,X163,X173,X171,X169,X166,X164,X162,X172,X170,X168] : (~p(s(bool,X167)) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X162,X168),'const.sum.++_3'(s(fun(X164,X162),X166),s(fun(X169,X168),X170),s('type.sum.sum'(X164,X169),X171))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X168,fun(X168,bool)),X172),s(fun(X168,X169),X173),s(fun(X169,X168),X170)))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X164,X169),X171)))) | s(bool,'F_0') = s(bool,X167) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X162,fun(X162,bool)),X163),s(fun(X162,X164),X165),s(fun(X164,X162),X166)))) )),
  inference(superposition,[],[f9521,f5187])).
fof(f18146,plain,(
  ( ! [X177,X185,X175,X182,X180,X178,X176,X184,X174,X183,X181,X179] : (s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X174,X180),'const.sum.++_3'(s(fun(X176,X174),X178),s(fun(X181,X180),X182),s('type.sum.sum'(X176,X181),X183))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X180,fun(X180,bool)),X184),s(fun(X180,X181),X185),s(fun(X181,X180),X182)))) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X176,X181),X183)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X174,fun(X174,bool)),X175),s(fun(X174,X176),X177),s(fun(X176,X174),X178))) | s(bool,'F_0') = s(bool,X179)) )),
  inference(subsumption_resolution,[],[f18129,f119])).
fof(f18129,plain,(
  ( ! [X177,X185,X175,X182,X180,X178,X176,X184,X174,X183,X181,X179] : (~p(s(bool,X179)) | s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X174,X180),'const.sum.++_3'(s(fun(X176,X174),X178),s(fun(X181,X180),X182),s('type.sum.sum'(X176,X181),X183))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X180,fun(X180,bool)),X184),s(fun(X180,X181),X185),s(fun(X181,X180),X182)))) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X176,X181),X183)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X174,fun(X174,bool)),X175),s(fun(X174,X176),X177),s(fun(X176,X174),X178))) | s(bool,'F_0') = s(bool,X179)) )),
  inference(superposition,[],[f9519,f5187])).
fof(f18144,plain,(
  ( ! [X167,X165,X163,X173,X171,X169,X166,X164,X162,X172,X170,X168] : (s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X162,X168),'const.sum.++_3'(s(fun(X164,X162),X166),s(fun(X169,X168),X170),s('type.sum.sum'(X164,X169),X171))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X168,fun(X168,bool)),X172),s(fun(X168,X169),X173),s(fun(X169,X168),X170)))) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X164,X169),X171)))) | s(bool,'F_0') = s(bool,X167) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X162,fun(X162,bool)),X163),s(fun(X162,X164),X165),s(fun(X164,X162),X166)))) )),
  inference(subsumption_resolution,[],[f18128,f119])).
fof(f18128,plain,(
  ( ! [X167,X165,X163,X173,X171,X169,X166,X164,X162,X172,X170,X168] : (~p(s(bool,X167)) | s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X162,X168),'const.sum.++_3'(s(fun(X164,X162),X166),s(fun(X169,X168),X170),s('type.sum.sum'(X164,X169),X171))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X168,fun(X168,bool)),X172),s(fun(X168,X169),X173),s(fun(X169,X168),X170)))) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X164,X169),X171)))) | s(bool,'F_0') = s(bool,X167) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X162,fun(X162,bool)),X163),s(fun(X162,X164),X165),s(fun(X164,X162),X166)))) )),
  inference(superposition,[],[f9519,f5187])).
fof(f11885,plain,(
  ( ! [X26,X25] : (s(bool,X26) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X25))))) | s(bool,'T_0') = s(bool,X26)) ) | $spl94),
  inference(cnf_transformation,[],[f11885_D])).
fof(f11885_D,plain,(
  ( ! [X26,X25] : (s(bool,X26) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X25))))) | s(bool,'T_0') = s(bool,X26)) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f18101,plain,(
  ( ! [X90,X88,X89] : (s(bool,'T_0') = s(bool,X90) | s(bool,X90) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X88))))) | s(bool,'F_0') = s(bool,X89)) ) | $spl92),
  inference(subsumption_resolution,[],[f18059,f126])).
fof(f18059,plain,(
  ( ! [X90,X88,X89] : (s(bool,'T_0') != s(bool,X89) | s(bool,'T_0') = s(bool,X90) | s(bool,X90) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X88))))) | s(bool,'F_0') = s(bool,X89)) ) | $spl92),
  inference(superposition,[],[f15157,f5092])).
fof(f18007,plain,(
  ( ! [X177,X185,X175,X182,X180,X178,X176,X184,X174,X183,X181,X179] : (p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X174,X180),'const.sum.++_3'(s(fun(X176,X174),X178),s(fun(X181,X180),X182),s('type.sum.sum'(X176,X181),X183)))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X180,fun(X180,bool)),X184),s(fun(X180,X181),X185),s(fun(X181,X180),X182)))) | s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X176,X181),X183))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X174,fun(X174,bool)),X175),s(fun(X174,X176),X177),s(fun(X176,X174),X178))) | s(bool,'F_0') = s(bool,X179)) )),
  inference(subsumption_resolution,[],[f17990,f119])).
fof(f17990,plain,(
  ( ! [X177,X185,X175,X182,X180,X178,X176,X184,X174,X183,X181,X179] : (~p(s(bool,X179)) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X174,X180),'const.sum.++_3'(s(fun(X176,X174),X178),s(fun(X181,X180),X182),s('type.sum.sum'(X176,X181),X183)))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X180,fun(X180,bool)),X184),s(fun(X180,X181),X185),s(fun(X181,X180),X182)))) | s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X176,X181),X183))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X174,fun(X174,bool)),X175),s(fun(X174,X176),X177),s(fun(X176,X174),X178))) | s(bool,'F_0') = s(bool,X179)) )),
  inference(superposition,[],[f9518,f5187])).
fof(f18005,plain,(
  ( ! [X167,X165,X163,X173,X171,X169,X166,X164,X162,X172,X170,X168] : (p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X162,X168),'const.sum.++_3'(s(fun(X164,X162),X166),s(fun(X169,X168),X170),s('type.sum.sum'(X164,X169),X171)))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X168,fun(X168,bool)),X172),s(fun(X168,X169),X173),s(fun(X169,X168),X170)))) | s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X164,X169),X171))) | s(bool,'F_0') = s(bool,X167) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X162,fun(X162,bool)),X163),s(fun(X162,X164),X165),s(fun(X164,X162),X166)))) )),
  inference(subsumption_resolution,[],[f17989,f119])).
fof(f17989,plain,(
  ( ! [X167,X165,X163,X173,X171,X169,X166,X164,X162,X172,X170,X168] : (~p(s(bool,X167)) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X162,X168),'const.sum.++_3'(s(fun(X164,X162),X166),s(fun(X169,X168),X170),s('type.sum.sum'(X164,X169),X171)))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X168,fun(X168,bool)),X172),s(fun(X168,X169),X173),s(fun(X169,X168),X170)))) | s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X164,X169),X171))) | s(bool,'F_0') = s(bool,X167) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X162,fun(X162,bool)),X163),s(fun(X162,X164),X165),s(fun(X164,X162),X166)))) )),
  inference(superposition,[],[f9518,f5187])).
fof(f17934,plain,(
  ( ! [X26,X24,X25] : (s(bool,X26) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X25))))) | s(bool,'F_0') = s(bool,X24) | s(bool,'T_0') = s(bool,X26)) ) | $spl160),
  inference(duplicate_literal_removal,[],[f17837])).
fof(f17837,plain,(
  ( ! [X26,X24,X25] : (s(bool,X26) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X25))))) | s(bool,'F_0') = s(bool,X24) | s(bool,'T_0') = s(bool,X26) | s(bool,'F_0') = s(bool,X24)) ) | $spl160),
  inference(superposition,[],[f15050,f5092])).
fof(f17415,plain,(
  $false | ($spl8 | $spl96)),
  inference(subsumption_resolution,[],[f17363,f12281])).
fof(f17363,plain,(
  s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | $spl96),
  inference(equality_resolution,[],[f11886])).
fof(f17395,plain,(
  $false | ($spl10 | $spl96)),
  inference(subsumption_resolution,[],[f17394,f2457])).
fof(f17394,plain,(
  ( ! [X90] : (s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X90)))))) ) | ($spl10 | $spl96)),
  inference(subsumption_resolution,[],[f17342,f176])).
fof(f17342,plain,(
  ( ! [X90] : (s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X90)))))) ) | $spl96),
  inference(superposition,[],[f11886,f2351])).
fof(f17391,plain,(
  $false | ($spl10 | $spl96)),
  inference(subsumption_resolution,[],[f17390,f5815])).
fof(f17390,plain,(
  ( ! [X61,X59,X60] : (s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X59,X60),'const.sum.INR_1'(s(X60,X61)))))) ) | ($spl10 | $spl96)),
  inference(subsumption_resolution,[],[f17337,f176])).
fof(f17337,plain,(
  ( ! [X61,X59,X60] : (s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X59,X60),'const.sum.INR_1'(s(X60,X61)))))) ) | $spl96),
  inference(superposition,[],[f11886,f5444])).
fof(f17389,plain,(
  $false | ($spl8 | $spl10 | $spl96)),
  inference(subsumption_resolution,[],[f17364,f12281])).
fof(f17364,plain,(
  s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | ($spl10 | $spl96)),
  inference(trivial_inequality_removal,[],[f17336])).
fof(f17336,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | ($spl10 | $spl96)),
  inference(superposition,[],[f11886,f176])).
fof(f17388,plain,(
  $false | ($spl10 | $spl96)),
  inference(subsumption_resolution,[],[f17387,f1593])).
fof(f17387,plain,(
  ( ! [X57] : (s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK2),'const.sum.INL_1'(s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,X57)))))))) ) | ($spl10 | $spl96)),
  inference(subsumption_resolution,[],[f17335,f176])).
fof(f17335,plain,(
  ( ! [X57] : (s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK2),'const.sum.INL_1'(s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,X57)))))))) ) | $spl96),
  inference(superposition,[],[f11886,f1599])).
fof(f1599,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK2),'const.sum.INL_1'(s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,X0)))))))) )),
  inference(backward_demodulation,[],[f1580,f1442])).
fof(f17386,plain,(
  $false | ($spl10 | $spl96)),
  inference(subsumption_resolution,[],[f17385,f1592])).
fof(f17385,plain,(
  ( ! [X52] : (s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,X52)))))) ) | ($spl10 | $spl96)),
  inference(subsumption_resolution,[],[f17331,f176])).
fof(f17331,plain,(
  ( ! [X52] : (s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,X52)))))) ) | $spl96),
  inference(superposition,[],[f11886,f1580])).
fof(f17383,plain,(
  $false | ($spl10 | $spl96)),
  inference(subsumption_resolution,[],[f17382,f5633])).
fof(f17382,plain,(
  ( ! [X37,X35,X36] : (s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X35,X36),'const.sum.INL_1'(s(X35,X37)))))) ) | ($spl10 | $spl96)),
  inference(subsumption_resolution,[],[f17328,f176])).
fof(f17328,plain,(
  ( ! [X37,X35,X36] : (s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X35,X36),'const.sum.INL_1'(s(X35,X37)))))) ) | $spl96),
  inference(superposition,[],[f11886,f5433])).
fof(f17372,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,X3) | s(bool,'T_0') = s(bool,X2)) ) | ($spl10 | $spl96)),
  inference(subsumption_resolution,[],[f17371,f118])).
fof(f17371,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,X3) | s(bool,'T_0') = s(bool,X2) | p(s(bool,X3))) ) | ($spl10 | $spl96)),
  inference(forward_demodulation,[],[f17370,f176])).
fof(f17370,plain,(
  ( ! [X2,X3] : (s(bool,X3) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'T_0') = s(bool,X2) | p(s(bool,X3))) ) | $spl96),
  inference(subsumption_resolution,[],[f17308,f123])).
fof(f17308,plain,(
  ( ! [X2,X3] : (s(bool,X3) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'T_0') = s(bool,X2) | p(s(bool,X2)) | p(s(bool,X3))) ) | $spl96),
  inference(superposition,[],[f11886,f112])).
fof(f17368,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | ($spl10 | $spl96)),
  inference(subsumption_resolution,[],[f17367,f118])).
fof(f17367,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,'T_0') = s(bool,X0) | p(s(bool,X1))) ) | ($spl10 | $spl96)),
  inference(forward_demodulation,[],[f17366,f176])).
fof(f17366,plain,(
  ( ! [X0,X1] : (s(bool,X1) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'T_0') = s(bool,X0) | p(s(bool,X1))) ) | $spl96),
  inference(subsumption_resolution,[],[f17307,f123])).
fof(f17307,plain,(
  ( ! [X0,X1] : (s(bool,X1) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X0))) ) | $spl96),
  inference(superposition,[],[f11886,f112])).
fof(f17212,plain,(
  ( ! [X177,X185,X175,X182,X180,X178,X176,X184,X174,X183,X181,X179] : (s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X174,X180),'const.sum.++_3'(s(fun(X176,X174),X178),s(fun(X181,X180),X182),s('type.sum.sum'(X176,X181),X183))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X180,fun(X180,bool)),X184),s(fun(X180,X181),X185),s(fun(X181,X180),X182)))) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X176,X181),X183)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X174,fun(X174,bool)),X175),s(fun(X174,X176),X177),s(fun(X176,X174),X178))) | s(bool,'F_0') = s(bool,X179)) )),
  inference(subsumption_resolution,[],[f17195,f119])).
fof(f17195,plain,(
  ( ! [X177,X185,X175,X182,X180,X178,X176,X184,X174,X183,X181,X179] : (~p(s(bool,X179)) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X174,X180),'const.sum.++_3'(s(fun(X176,X174),X178),s(fun(X181,X180),X182),s('type.sum.sum'(X176,X181),X183))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X180,fun(X180,bool)),X184),s(fun(X180,X181),X185),s(fun(X181,X180),X182)))) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X176,X181),X183)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X174,fun(X174,bool)),X175),s(fun(X174,X176),X177),s(fun(X176,X174),X178))) | s(bool,'F_0') = s(bool,X179)) )),
  inference(superposition,[],[f9513,f5187])).
fof(f17210,plain,(
  ( ! [X167,X165,X163,X173,X171,X169,X166,X164,X162,X172,X170,X168] : (s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X162,X168),'const.sum.++_3'(s(fun(X164,X162),X166),s(fun(X169,X168),X170),s('type.sum.sum'(X164,X169),X171))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X168,fun(X168,bool)),X172),s(fun(X168,X169),X173),s(fun(X169,X168),X170)))) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X164,X169),X171)))) | s(bool,'F_0') = s(bool,X167) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X162,fun(X162,bool)),X163),s(fun(X162,X164),X165),s(fun(X164,X162),X166)))) )),
  inference(subsumption_resolution,[],[f17194,f119])).
fof(f17194,plain,(
  ( ! [X167,X165,X163,X173,X171,X169,X166,X164,X162,X172,X170,X168] : (~p(s(bool,X167)) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X162,X168),'const.sum.++_3'(s(fun(X164,X162),X166),s(fun(X169,X168),X170),s('type.sum.sum'(X164,X169),X171))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X168,fun(X168,bool)),X172),s(fun(X168,X169),X173),s(fun(X169,X168),X170)))) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X164,X169),X171)))) | s(bool,'F_0') = s(bool,X167) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X162,fun(X162,bool)),X163),s(fun(X162,X164),X165),s(fun(X164,X162),X166)))) )),
  inference(superposition,[],[f9513,f5187])).
fof(f17166,plain,(
  ( ! [X58] : (s(bool,X58) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'T_0') = s(bool,X58)) ) | $spl262),
  inference(cnf_transformation,[],[f17166_D])).
fof(f17166_D,plain,(
  ( ! [X58] : (s(bool,X58) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'T_0') = s(bool,X58)) ) <=> ~$spl262),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl262])])).
fof(f17165,plain,(
  ( ! [X57,X58] : (s(bool,X58) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'T_0') = s(bool,X58) | s(bool,'F_0') = s(bool,X57)) ) | ($spl10 | $spl92)),
  inference(subsumption_resolution,[],[f17002,f119])).
fof(f17002,plain,(
  ( ! [X57,X58] : (~p(s(bool,X57)) | s(bool,X58) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'T_0') = s(bool,X58) | s(bool,'F_0') = s(bool,X57)) ) | ($spl10 | $spl92)),
  inference(superposition,[],[f11882,f14796])).
fof(f14796,plain,(
  ( ! [X6] : (s(bool,X6) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'F_0') = s(bool,X6)) ) | $spl10),
  inference(trivial_inequality_removal,[],[f14793])).
fof(f14793,plain,(
  ( ! [X6] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,X6) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'F_0') = s(bool,X6)) ) | $spl10),
  inference(superposition,[],[f13912,f5187])).
fof(f13912,plain,(
  s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | $spl10),
  inference(forward_demodulation,[],[f110,f176])).
fof(f16780,plain,(
  ( ! [X177,X185,X175,X182,X180,X178,X176,X184,X174,X183,X181,X179] : (~p(s(bool,X179)) | s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X174,X180),'const.sum.++_3'(s(fun(X176,X174),X178),s(fun(X181,X180),X182),s('type.sum.sum'(X176,X181),X183))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X180,fun(X180,bool)),X184),s(fun(X180,X181),X185),s(fun(X181,X180),X182)))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X176,X181),X183)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X174,fun(X174,bool)),X175),s(fun(X174,X176),X177),s(fun(X176,X174),X178)))) )),
  inference(subsumption_resolution,[],[f16762,f118])).
fof(f16762,plain,(
  ( ! [X177,X185,X175,X182,X180,X178,X176,X184,X174,X183,X181,X179] : (~p(s(bool,X179)) | s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X174,X180),'const.sum.++_3'(s(fun(X176,X174),X178),s(fun(X181,X180),X182),s('type.sum.sum'(X176,X181),X183))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X180,fun(X180,bool)),X184),s(fun(X180,X181),X185),s(fun(X181,X180),X182)))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X176,X181),X183)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X174,fun(X174,bool)),X175),s(fun(X174,X176),X177),s(fun(X176,X174),X178))) | s(bool,'F_0') = s(bool,X179)) )),
  inference(superposition,[],[f9511,f5187])).
fof(f16778,plain,(
  ( ! [X167,X165,X163,X173,X171,X169,X166,X164,X162,X172,X170,X168] : (~p(s(bool,X167)) | s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X162,X168),'const.sum.++_3'(s(fun(X164,X162),X166),s(fun(X169,X168),X170),s('type.sum.sum'(X164,X169),X171))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X168,fun(X168,bool)),X172),s(fun(X168,X169),X173),s(fun(X169,X168),X170)))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X164,X169),X171)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X162,fun(X162,bool)),X163),s(fun(X162,X164),X165),s(fun(X164,X162),X166)))) )),
  inference(subsumption_resolution,[],[f16761,f118])).
fof(f16761,plain,(
  ( ! [X167,X165,X163,X173,X171,X169,X166,X164,X162,X172,X170,X168] : (~p(s(bool,X167)) | s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X162,X168),'const.sum.++_3'(s(fun(X164,X162),X166),s(fun(X169,X168),X170),s('type.sum.sum'(X164,X169),X171))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X168,fun(X168,bool)),X172),s(fun(X168,X169),X173),s(fun(X169,X168),X170)))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X164,X169),X171)))) | s(bool,'F_0') = s(bool,X167) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X162,fun(X162,bool)),X163),s(fun(X162,X164),X165),s(fun(X164,X162),X166)))) )),
  inference(superposition,[],[f9511,f5187])).
fof(f16732,plain,(
  $false | $spl164),
  inference(subsumption_resolution,[],[f16662,f11980])).
fof(f16662,plain,(
  p(s(bool,'F_0')) | $spl164),
  inference(resolution,[],[f11980,f153])).
fof(f153,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f106])).
fof(f16731,plain,(
  $false | $spl164),
  inference(subsumption_resolution,[],[f16661,f11980])).
fof(f16661,plain,(
  p(s(bool,'F_0')) | $spl164),
  inference(resolution,[],[f11980,f151])).
fof(f151,plain,(
  p(s(bool,'F_0')) | p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f105])).
fof(f16730,plain,(
  $false | $spl164),
  inference(subsumption_resolution,[],[f16660,f11980])).
fof(f16660,plain,(
  ( ! [X13] : (p(s(bool,X13))) ) | $spl164),
  inference(resolution,[],[f11980,f146])).
fof(f146,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f104])).
fof(f104,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,'T_0'))))),
  inference(flattening,[],[f103])).
fof(f103,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f87])).
fof(f87,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X22] : ((p(s(bool,X22)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  file('thm.quotient_sum.ISR_RSP','thm.bool.IMP_CLAUSES|split|1')).
fof(f16729,plain,(
  $false | $spl164),
  inference(subsumption_resolution,[],[f16658,f11980])).
fof(f16658,plain,(
  p(s(bool,'F_0')) | $spl164),
  inference(resolution,[],[f11980,f143])).
fof(f143,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f102])).
fof(f16728,plain,(
  $false | $spl164),
  inference(subsumption_resolution,[],[f16656,f11980])).
fof(f16656,plain,(
  ( ! [X10] : (p(s(bool,X10))) ) | $spl164),
  inference(resolution,[],[f11980,f140])).
fof(f140,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f100])).
fof(f16663,plain,(
  $false | $spl164),
  inference(resolution,[],[f11980,f155])).
fof(f16664,plain,(
  $false | $spl164),
  inference(resolution,[],[f11980,f156])).
fof(f16669,plain,(
  $false | $spl164),
  inference(resolution,[],[f11980,f9190])).
fof(f16673,plain,(
  $false | $spl164),
  inference(resolution,[],[f11980,f9289])).
fof(f16640,plain,(
  ( ! [X177,X185,X175,X182,X180,X178,X176,X184,X174,X183,X181,X179] : (~p(s(bool,X179)) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X174,X180),'const.sum.++_3'(s(fun(X176,X174),X178),s(fun(X181,X180),X182),s('type.sum.sum'(X176,X181),X183)))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X180,fun(X180,bool)),X184),s(fun(X180,X181),X185),s(fun(X181,X180),X182)))) | s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X176,X181),X183))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X174,fun(X174,bool)),X175),s(fun(X174,X176),X177),s(fun(X176,X174),X178)))) )),
  inference(subsumption_resolution,[],[f16622,f118])).
fof(f16622,plain,(
  ( ! [X177,X185,X175,X182,X180,X178,X176,X184,X174,X183,X181,X179] : (~p(s(bool,X179)) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X174,X180),'const.sum.++_3'(s(fun(X176,X174),X178),s(fun(X181,X180),X182),s('type.sum.sum'(X176,X181),X183)))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X180,fun(X180,bool)),X184),s(fun(X180,X181),X185),s(fun(X181,X180),X182)))) | s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X176,X181),X183))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X174,fun(X174,bool)),X175),s(fun(X174,X176),X177),s(fun(X176,X174),X178))) | s(bool,'F_0') = s(bool,X179)) )),
  inference(superposition,[],[f9510,f5187])).
fof(f16638,plain,(
  ( ! [X167,X165,X163,X173,X171,X169,X166,X164,X162,X172,X170,X168] : (~p(s(bool,X167)) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X162,X168),'const.sum.++_3'(s(fun(X164,X162),X166),s(fun(X169,X168),X170),s('type.sum.sum'(X164,X169),X171)))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X168,fun(X168,bool)),X172),s(fun(X168,X169),X173),s(fun(X169,X168),X170)))) | s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X164,X169),X171))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X162,fun(X162,bool)),X163),s(fun(X162,X164),X165),s(fun(X164,X162),X166)))) )),
  inference(subsumption_resolution,[],[f16621,f118])).
fof(f16621,plain,(
  ( ! [X167,X165,X163,X173,X171,X169,X166,X164,X162,X172,X170,X168] : (~p(s(bool,X167)) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X162,X168),'const.sum.++_3'(s(fun(X164,X162),X166),s(fun(X169,X168),X170),s('type.sum.sum'(X164,X169),X171)))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X168,fun(X168,bool)),X172),s(fun(X168,X169),X173),s(fun(X169,X168),X170)))) | s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X164,X169),X171))) | s(bool,'F_0') = s(bool,X167) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X162,fun(X162,bool)),X163),s(fun(X162,X164),X165),s(fun(X164,X162),X166)))) )),
  inference(superposition,[],[f9510,f5187])).
fof(f11888,plain,(
  ( ! [X44] : (s(bool,X44) != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'T_0') = s(bool,X44)) ) | $spl100),
  inference(cnf_transformation,[],[f11888_D])).
fof(f11888_D,plain,(
  ( ! [X44] : (s(bool,X44) != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'T_0') = s(bool,X44)) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f16593,plain,(
  ( ! [X78,X77] : (s(bool,X78) != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'T_0') = s(bool,X78) | s(bool,'F_0') = s(bool,X77)) ) | ($spl92 | $spl234)),
  inference(subsumption_resolution,[],[f16438,f119])).
fof(f16438,plain,(
  ( ! [X78,X77] : (~p(s(bool,X77)) | s(bool,X78) != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'T_0') = s(bool,X78) | s(bool,'F_0') = s(bool,X77)) ) | ($spl92 | $spl234)),
  inference(superposition,[],[f11882,f12640])).
fof(f11887,plain,(
  ( ! [X42] : (s(bool,X42) != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'T_0') = s(bool,X42)) ) | $spl98),
  inference(cnf_transformation,[],[f11887_D])).
fof(f11887_D,plain,(
  ( ! [X42] : (s(bool,X42) != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'T_0') = s(bool,X42)) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f16226,plain,(
  ( ! [X78,X77] : (s(bool,X78) != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'T_0') = s(bool,X78) | s(bool,'F_0') = s(bool,X77)) ) | ($spl92 | $spl232)),
  inference(subsumption_resolution,[],[f16075,f119])).
fof(f16075,plain,(
  ( ! [X78,X77] : (~p(s(bool,X77)) | s(bool,X78) != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'T_0') = s(bool,X78) | s(bool,'F_0') = s(bool,X77)) ) | ($spl92 | $spl232)),
  inference(superposition,[],[f11882,f12639])).
fof(f15870,plain,(
  ( ! [X177,X185,X175,X182,X180,X178,X176,X184,X174,X183,X181,X179] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X180,fun(X180,bool)),X181),s(fun(X180,X182),X183),s(fun(X182,X180),X184)))) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X176,X182),X185))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X174,X180),'const.sum.++_3'(s(fun(X176,X174),X178),s(fun(X182,X180),X184),s('type.sum.sum'(X176,X182),X185)))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X174,fun(X174,bool)),X175),s(fun(X174,X176),X177),s(fun(X176,X174),X178))) | s(bool,'F_0') = s(bool,X179)) )),
  inference(subsumption_resolution,[],[f15854,f119])).
fof(f15854,plain,(
  ( ! [X177,X185,X175,X182,X180,X178,X176,X184,X174,X183,X181,X179] : (~p(s(bool,X179)) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X180,fun(X180,bool)),X181),s(fun(X180,X182),X183),s(fun(X182,X180),X184)))) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X176,X182),X185))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X174,X180),'const.sum.++_3'(s(fun(X176,X174),X178),s(fun(X182,X180),X184),s('type.sum.sum'(X176,X182),X185)))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X174,fun(X174,bool)),X175),s(fun(X174,X176),X177),s(fun(X176,X174),X178))) | s(bool,'F_0') = s(bool,X179)) )),
  inference(superposition,[],[f9449,f5187])).
fof(f15868,plain,(
  ( ! [X167,X165,X163,X173,X171,X169,X166,X164,X162,X172,X170,X168] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X168,fun(X168,bool)),X169),s(fun(X168,X170),X171),s(fun(X170,X168),X172)))) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X164,X170),X173))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X162,X168),'const.sum.++_3'(s(fun(X164,X162),X166),s(fun(X170,X168),X172),s('type.sum.sum'(X164,X170),X173)))))) | s(bool,'F_0') = s(bool,X167) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X162,fun(X162,bool)),X163),s(fun(X162,X164),X165),s(fun(X164,X162),X166)))) )),
  inference(subsumption_resolution,[],[f15853,f119])).
fof(f15853,plain,(
  ( ! [X167,X165,X163,X173,X171,X169,X166,X164,X162,X172,X170,X168] : (~p(s(bool,X167)) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X168,fun(X168,bool)),X169),s(fun(X168,X170),X171),s(fun(X170,X168),X172)))) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X164,X170),X173))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X162,X168),'const.sum.++_3'(s(fun(X164,X162),X166),s(fun(X170,X168),X172),s('type.sum.sum'(X164,X170),X173)))))) | s(bool,'F_0') = s(bool,X167) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X162,fun(X162,bool)),X163),s(fun(X162,X164),X165),s(fun(X164,X162),X166)))) )),
  inference(superposition,[],[f9449,f5187])).
fof(f15826,plain,(
  $false | ($spl8 | $spl230)),
  inference(subsumption_resolution,[],[f15825,f1128])).
fof(f15825,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl8 | $spl230)),
  inference(forward_demodulation,[],[f15749,f173])).
fof(f15749,plain,(
  s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | $spl230),
  inference(equality_resolution,[],[f12224])).
fof(f15802,plain,(
  $false | ($spl8 | $spl32 | $spl230)),
  inference(subsumption_resolution,[],[f15801,f1128])).
fof(f15801,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl8 | $spl32 | $spl230)),
  inference(forward_demodulation,[],[f15800,f771])).
fof(f15800,plain,(
  s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) | ($spl8 | $spl32 | $spl230)),
  inference(subsumption_resolution,[],[f15733,f173])).
fof(f15733,plain,(
  s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) | ($spl32 | $spl230)),
  inference(superposition,[],[f12224,f771])).
fof(f15799,plain,(
  $false | ($spl8 | $spl26 | $spl32 | $spl230)),
  inference(subsumption_resolution,[],[f15798,f1128])).
fof(f15798,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl8 | $spl26 | $spl32 | $spl230)),
  inference(forward_demodulation,[],[f15797,f1035])).
fof(f15797,plain,(
  s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) | ($spl8 | $spl26 | $spl32 | $spl230)),
  inference(subsumption_resolution,[],[f15732,f173])).
fof(f15732,plain,(
  s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) | ($spl26 | $spl32 | $spl230)),
  inference(superposition,[],[f12224,f1035])).
fof(f15796,plain,(
  $false | ($spl8 | $spl230)),
  inference(subsumption_resolution,[],[f15795,f1128])).
fof(f15795,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl8 | $spl230)),
  inference(forward_demodulation,[],[f15794,f4197])).
fof(f4197,plain,(
  ( ! [X35] : (s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,X35)))))) )),
  inference(trivial_inequality_removal,[],[f4196])).
fof(f4196,plain,(
  ( ! [X35] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,X35)))))) )),
  inference(superposition,[],[f4131,f126])).
fof(f4131,plain,(
  ( ! [X4] : (s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,X4)))))) )),
  inference(resolution,[],[f4105,f128])).
fof(f15794,plain,(
  ( ! [X85] : (s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,X85)))))) ) | ($spl8 | $spl230)),
  inference(subsumption_resolution,[],[f15730,f173])).
fof(f15730,plain,(
  ( ! [X85] : (s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,X85)))))) ) | $spl230),
  inference(superposition,[],[f12224,f4197])).
fof(f15793,plain,(
  $false | ($spl8 | $spl230)),
  inference(subsumption_resolution,[],[f15792,f1128])).
fof(f15792,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl8 | $spl230)),
  inference(forward_demodulation,[],[f15791,f9229])).
fof(f15791,plain,(
  ( ! [X83,X84,X82] : (s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X82,X83),'const.sum.INL_1'(s(X82,X84)))))) ) | ($spl8 | $spl230)),
  inference(subsumption_resolution,[],[f15729,f173])).
fof(f15729,plain,(
  ( ! [X83,X84,X82] : (s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X82,X83),'const.sum.INL_1'(s(X82,X84)))))) ) | $spl230),
  inference(superposition,[],[f12224,f9229])).
fof(f15788,plain,(
  $false | ($spl8 | $spl230)),
  inference(subsumption_resolution,[],[f15787,f1128])).
fof(f15787,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl8 | $spl230)),
  inference(forward_demodulation,[],[f15750,f173])).
fof(f15750,plain,(
  s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | ($spl8 | $spl230)),
  inference(trivial_inequality_removal,[],[f15724])).
fof(f15724,plain,(
  s(bool,'T_0') != s(bool,'T_0') | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | ($spl8 | $spl230)),
  inference(superposition,[],[f12224,f173])).
fof(f15786,plain,(
  $false | ($spl8 | $spl230)),
  inference(subsumption_resolution,[],[f15785,f1128])).
fof(f15785,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl8 | $spl230)),
  inference(forward_demodulation,[],[f15784,f3224])).
fof(f3224,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK2),'const.sum.INR_1'(s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,X1)))))))) )),
  inference(backward_demodulation,[],[f3205,f1443])).
fof(f15784,plain,(
  ( ! [X53] : (s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK2),'const.sum.INR_1'(s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,X53)))))))) ) | ($spl8 | $spl230)),
  inference(subsumption_resolution,[],[f15722,f173])).
fof(f15722,plain,(
  ( ! [X53] : (s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK2),'const.sum.INR_1'(s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,X53)))))))) ) | $spl230),
  inference(superposition,[],[f12224,f3224])).
fof(f15783,plain,(
  $false | ($spl8 | $spl230)),
  inference(subsumption_resolution,[],[f15782,f1128])).
fof(f15782,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl8 | $spl230)),
  inference(forward_demodulation,[],[f15781,f3205])).
fof(f15781,plain,(
  ( ! [X52] : (s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X52)))))) ) | ($spl8 | $spl230)),
  inference(subsumption_resolution,[],[f15721,f173])).
fof(f15721,plain,(
  ( ! [X52] : (s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X52)))))) ) | $spl230),
  inference(superposition,[],[f12224,f3205])).
fof(f15780,plain,(
  $false | ($spl8 | $spl230)),
  inference(subsumption_resolution,[],[f15779,f1128])).
fof(f15779,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl8 | $spl230)),
  inference(forward_demodulation,[],[f15778,f9130])).
fof(f15778,plain,(
  ( ! [X47,X48,X46] : (s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X46,X47),'const.sum.INR_1'(s(X47,X48)))))) ) | ($spl8 | $spl230)),
  inference(subsumption_resolution,[],[f15718,f173])).
fof(f15718,plain,(
  ( ! [X47,X48,X46] : (s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X46,X47),'const.sum.INR_1'(s(X47,X48)))))) ) | $spl230),
  inference(superposition,[],[f12224,f9130])).
fof(f15776,plain,(
  $false | ($spl8 | $spl230)),
  inference(subsumption_resolution,[],[f15775,f1128])).
fof(f15775,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl8 | $spl230)),
  inference(forward_demodulation,[],[f15774,f409])).
fof(f409,plain,(
  s(bool,'T_0') = s(bool,'const.quotient_sum.+++_4'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(sK3,fun(sK3,bool)),sK7),s('type.sum.sum'(sK2,sK3),sK10),s('type.sum.sum'(sK2,sK3),sK11)))),
  inference(resolution,[],[f109,f123])).
fof(f109,plain,(
  p(s(bool,'const.quotient_sum.+++_4'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(sK3,fun(sK3,bool)),sK7),s('type.sum.sum'(sK2,sK3),sK10),s('type.sum.sum'(sK2,sK3),sK11))))),
  inference(cnf_transformation,[],[f89])).
fof(f15774,plain,(
  s(bool,'F_0') = s(bool,'const.quotient_sum.+++_4'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(sK3,fun(sK3,bool)),sK7),s('type.sum.sum'(sK2,sK3),sK10),s('type.sum.sum'(sK2,sK3),sK11))) | ($spl8 | $spl230)),
  inference(subsumption_resolution,[],[f15715,f173])).
fof(f15715,plain,(
  s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'F_0') = s(bool,'const.quotient_sum.+++_4'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(sK3,fun(sK3,bool)),sK7),s('type.sum.sum'(sK2,sK3),sK10),s('type.sum.sum'(sK2,sK3),sK11))) | $spl230),
  inference(superposition,[],[f12224,f409])).
fof(f15773,plain,(
  $false | ($spl8 | $spl230)),
  inference(subsumption_resolution,[],[f15772,f1128])).
fof(f15772,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl8 | $spl230)),
  inference(forward_demodulation,[],[f15771,f308])).
fof(f308,plain,(
  s(bool,'T_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK3,fun(sK3,bool)),sK7),s(fun(sK3,sK1),sK8),s(fun(sK1,sK3),sK9)))),
  inference(resolution,[],[f108,f123])).
fof(f15771,plain,(
  s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK3,fun(sK3,bool)),sK7),s(fun(sK3,sK1),sK8),s(fun(sK1,sK3),sK9))) | ($spl8 | $spl230)),
  inference(subsumption_resolution,[],[f15714,f173])).
fof(f15714,plain,(
  s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK3,fun(sK3,bool)),sK7),s(fun(sK3,sK1),sK8),s(fun(sK1,sK3),sK9))) | $spl230),
  inference(superposition,[],[f12224,f308])).
fof(f15770,plain,(
  $false | ($spl8 | $spl230)),
  inference(subsumption_resolution,[],[f15769,f1128])).
fof(f15769,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl8 | $spl230)),
  inference(forward_demodulation,[],[f15768,f187])).
fof(f15768,plain,(
  s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(sK2,sK0),sK5),s(fun(sK0,sK2),sK6))) | ($spl8 | $spl230)),
  inference(subsumption_resolution,[],[f15713,f173])).
fof(f15713,plain,(
  s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(sK2,sK0),sK5),s(fun(sK0,sK2),sK6))) | $spl230),
  inference(superposition,[],[f12224,f187])).
fof(f15760,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X7) | s(bool,'F_0') = s(bool,X6)) ) | ($spl8 | $spl230)),
  inference(subsumption_resolution,[],[f15759,f122])).
fof(f15759,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X7) | s(bool,'F_0') = s(bool,X6) | ~p(s(bool,X7))) ) | ($spl8 | $spl230)),
  inference(forward_demodulation,[],[f15758,f173])).
fof(f15758,plain,(
  ( ! [X6,X7] : (s(bool,X7) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'F_0') = s(bool,X6) | ~p(s(bool,X7))) ) | $spl230),
  inference(subsumption_resolution,[],[f15701,f119])).
fof(f15701,plain,(
  ( ! [X6,X7] : (s(bool,X7) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'F_0') = s(bool,X6) | ~p(s(bool,X7)) | ~p(s(bool,X6))) ) | $spl230),
  inference(superposition,[],[f12224,f115])).
fof(f15757,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X5) | s(bool,'F_0') = s(bool,X4)) ) | ($spl8 | $spl230)),
  inference(subsumption_resolution,[],[f15756,f122])).
fof(f15756,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X5) | s(bool,'F_0') = s(bool,X4) | ~p(s(bool,X5))) ) | ($spl8 | $spl230)),
  inference(forward_demodulation,[],[f15755,f173])).
fof(f15755,plain,(
  ( ! [X4,X5] : (s(bool,X5) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'F_0') = s(bool,X4) | ~p(s(bool,X5))) ) | $spl230),
  inference(subsumption_resolution,[],[f15700,f119])).
fof(f15700,plain,(
  ( ! [X4,X5] : (s(bool,X5) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'F_0') = s(bool,X4) | ~p(s(bool,X4)) | ~p(s(bool,X5))) ) | $spl230),
  inference(superposition,[],[f12224,f115])).
fof(f15658,plain,(
  $false | ($spl8 | $spl32 | $spl228)),
  inference(subsumption_resolution,[],[f15657,f1128])).
fof(f15657,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl8 | $spl32 | $spl228)),
  inference(forward_demodulation,[],[f15656,f771])).
fof(f15656,plain,(
  s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) | ($spl8 | $spl32 | $spl228)),
  inference(subsumption_resolution,[],[f15404,f12281])).
fof(f15404,plain,(
  s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) | ($spl32 | $spl228)),
  inference(superposition,[],[f771,f12223])).
fof(f15655,plain,(
  $false | ($spl8 | $spl26 | $spl32 | $spl228)),
  inference(subsumption_resolution,[],[f15654,f1128])).
fof(f15654,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl8 | $spl26 | $spl32 | $spl228)),
  inference(forward_demodulation,[],[f15653,f1035])).
fof(f15653,plain,(
  s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) | ($spl8 | $spl26 | $spl32 | $spl228)),
  inference(subsumption_resolution,[],[f15403,f12281])).
fof(f15403,plain,(
  s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) | ($spl26 | $spl32 | $spl228)),
  inference(superposition,[],[f1035,f12223])).
fof(f15652,plain,(
  $false | ($spl8 | $spl228)),
  inference(subsumption_resolution,[],[f15651,f1128])).
fof(f15651,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl8 | $spl228)),
  inference(forward_demodulation,[],[f15650,f4197])).
fof(f15650,plain,(
  ( ! [X328] : (s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,X328)))))) ) | ($spl8 | $spl228)),
  inference(subsumption_resolution,[],[f15402,f12281])).
fof(f15402,plain,(
  ( ! [X328] : (s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,X328)))))) ) | $spl228),
  inference(superposition,[],[f4197,f12223])).
fof(f15645,plain,(
  $false | ($spl8 | $spl228)),
  inference(subsumption_resolution,[],[f15644,f1128])).
fof(f15644,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl8 | $spl228)),
  inference(forward_demodulation,[],[f15643,f9229])).
fof(f15643,plain,(
  ( ! [X313,X315,X314] : (s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X313,X314),'const.sum.INL_1'(s(X313,X315)))))) ) | ($spl8 | $spl228)),
  inference(subsumption_resolution,[],[f15393,f12281])).
fof(f15393,plain,(
  ( ! [X313,X315,X314] : (s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X313,X314),'const.sum.INL_1'(s(X313,X315)))))) ) | $spl228),
  inference(superposition,[],[f9229,f12223])).
fof(f15636,plain,(
  $false | ($spl8 | $spl228)),
  inference(subsumption_resolution,[],[f15635,f1128])).
fof(f15635,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl8 | $spl228)),
  inference(forward_demodulation,[],[f15634,f3224])).
fof(f15634,plain,(
  ( ! [X276] : (s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK2),'const.sum.INR_1'(s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,X276)))))))) ) | ($spl8 | $spl228)),
  inference(subsumption_resolution,[],[f15384,f12281])).
fof(f15384,plain,(
  ( ! [X276] : (s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK2),'const.sum.INR_1'(s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,X276)))))))) ) | $spl228),
  inference(superposition,[],[f3224,f12223])).
fof(f15633,plain,(
  $false | ($spl8 | $spl228)),
  inference(subsumption_resolution,[],[f15632,f1128])).
fof(f15632,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl8 | $spl228)),
  inference(forward_demodulation,[],[f15631,f3205])).
fof(f15631,plain,(
  ( ! [X275] : (s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X275)))))) ) | ($spl8 | $spl228)),
  inference(subsumption_resolution,[],[f15383,f12281])).
fof(f15383,plain,(
  ( ! [X275] : (s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X275)))))) ) | $spl228),
  inference(superposition,[],[f3205,f12223])).
fof(f15624,plain,(
  $false | ($spl8 | $spl228)),
  inference(subsumption_resolution,[],[f15623,f1128])).
fof(f15623,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl8 | $spl228)),
  inference(forward_demodulation,[],[f15622,f9130])).
fof(f15622,plain,(
  ( ! [X265,X264,X266] : (s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X264,X265),'const.sum.INR_1'(s(X265,X266)))))) ) | ($spl8 | $spl228)),
  inference(subsumption_resolution,[],[f15375,f12281])).
fof(f15375,plain,(
  ( ! [X265,X264,X266] : (s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X264,X265),'const.sum.INR_1'(s(X265,X266)))))) ) | $spl228),
  inference(superposition,[],[f9130,f12223])).
fof(f15615,plain,(
  $false | ($spl8 | $spl228)),
  inference(subsumption_resolution,[],[f15614,f1128])).
fof(f15614,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl8 | $spl228)),
  inference(forward_demodulation,[],[f15613,f409])).
fof(f15613,plain,(
  s(bool,'F_0') = s(bool,'const.quotient_sum.+++_4'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(sK3,fun(sK3,bool)),sK7),s('type.sum.sum'(sK2,sK3),sK10),s('type.sum.sum'(sK2,sK3),sK11))) | ($spl8 | $spl228)),
  inference(subsumption_resolution,[],[f15368,f12281])).
fof(f15368,plain,(
  s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,'const.quotient_sum.+++_4'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(sK3,fun(sK3,bool)),sK7),s('type.sum.sum'(sK2,sK3),sK10),s('type.sum.sum'(sK2,sK3),sK11))) | $spl228),
  inference(superposition,[],[f409,f12223])).
fof(f15595,plain,(
  $false | ($spl8 | $spl228)),
  inference(subsumption_resolution,[],[f15594,f1128])).
fof(f15594,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl8 | $spl228)),
  inference(forward_demodulation,[],[f15593,f187])).
fof(f15593,plain,(
  s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(sK2,sK0),sK5),s(fun(sK0,sK2),sK6))) | ($spl8 | $spl228)),
  inference(subsumption_resolution,[],[f15348,f12281])).
fof(f15348,plain,(
  s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(sK2,sK0),sK5),s(fun(sK0,sK2),sK6))) | $spl228),
  inference(superposition,[],[f187,f12223])).
fof(f15592,plain,(
  $false | ($spl8 | $spl228)),
  inference(subsumption_resolution,[],[f15591,f1128])).
fof(f15591,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl8 | $spl228)),
  inference(forward_demodulation,[],[f15590,f308])).
fof(f15590,plain,(
  s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK3,fun(sK3,bool)),sK7),s(fun(sK3,sK1),sK8),s(fun(sK1,sK3),sK9))) | ($spl8 | $spl228)),
  inference(subsumption_resolution,[],[f15347,f12281])).
fof(f15347,plain,(
  s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK3,fun(sK3,bool)),sK7),s(fun(sK3,sK1),sK8),s(fun(sK1,sK3),sK9))) | $spl228),
  inference(superposition,[],[f308,f12223])).
fof(f15541,plain,(
  $false | ($spl8 | $spl32 | $spl228)),
  inference(subsumption_resolution,[],[f15540,f1128])).
fof(f15540,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl8 | $spl32 | $spl228)),
  inference(forward_demodulation,[],[f15539,f771])).
fof(f15539,plain,(
  s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) | ($spl8 | $spl32 | $spl228)),
  inference(subsumption_resolution,[],[f15287,f12281])).
fof(f15287,plain,(
  s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) | ($spl32 | $spl228)),
  inference(superposition,[],[f12223,f771])).
fof(f15538,plain,(
  $false | ($spl8 | $spl26 | $spl32 | $spl228)),
  inference(subsumption_resolution,[],[f15537,f1128])).
fof(f15537,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl8 | $spl26 | $spl32 | $spl228)),
  inference(forward_demodulation,[],[f15536,f1035])).
fof(f15536,plain,(
  s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) | ($spl8 | $spl26 | $spl32 | $spl228)),
  inference(subsumption_resolution,[],[f15286,f12281])).
fof(f15286,plain,(
  s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) | ($spl26 | $spl32 | $spl228)),
  inference(superposition,[],[f12223,f1035])).
fof(f15535,plain,(
  $false | ($spl8 | $spl228)),
  inference(subsumption_resolution,[],[f15534,f1128])).
fof(f15534,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl8 | $spl228)),
  inference(forward_demodulation,[],[f15533,f4197])).
fof(f15533,plain,(
  ( ! [X87] : (s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,X87)))))) ) | ($spl8 | $spl228)),
  inference(subsumption_resolution,[],[f15284,f12281])).
fof(f15284,plain,(
  ( ! [X87] : (s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,X87)))))) ) | $spl228),
  inference(superposition,[],[f12223,f4197])).
fof(f15532,plain,(
  $false | ($spl8 | $spl228)),
  inference(subsumption_resolution,[],[f15531,f1128])).
fof(f15531,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl8 | $spl228)),
  inference(forward_demodulation,[],[f15530,f9229])).
fof(f15530,plain,(
  ( ! [X85,X86,X84] : (s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X84,X85),'const.sum.INL_1'(s(X84,X86)))))) ) | ($spl8 | $spl228)),
  inference(subsumption_resolution,[],[f15283,f12281])).
fof(f15283,plain,(
  ( ! [X85,X86,X84] : (s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X84,X85),'const.sum.INL_1'(s(X84,X86)))))) ) | $spl228),
  inference(superposition,[],[f12223,f9229])).
fof(f15527,plain,(
  $false | ($spl8 | $spl228)),
  inference(subsumption_resolution,[],[f15526,f1128])).
fof(f15526,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl8 | $spl228)),
  inference(forward_demodulation,[],[f15525,f3224])).
fof(f15525,plain,(
  ( ! [X54] : (s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK2),'const.sum.INR_1'(s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,X54)))))))) ) | ($spl8 | $spl228)),
  inference(subsumption_resolution,[],[f15278,f12281])).
fof(f15278,plain,(
  ( ! [X54] : (s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK2),'const.sum.INR_1'(s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,X54)))))))) ) | $spl228),
  inference(superposition,[],[f12223,f3224])).
fof(f15524,plain,(
  $false | ($spl8 | $spl228)),
  inference(subsumption_resolution,[],[f15523,f1128])).
fof(f15523,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl8 | $spl228)),
  inference(forward_demodulation,[],[f15522,f3205])).
fof(f15522,plain,(
  ( ! [X53] : (s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X53)))))) ) | ($spl8 | $spl228)),
  inference(subsumption_resolution,[],[f15277,f12281])).
fof(f15277,plain,(
  ( ! [X53] : (s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X53)))))) ) | $spl228),
  inference(superposition,[],[f12223,f3205])).
fof(f15521,plain,(
  $false | ($spl8 | $spl228)),
  inference(subsumption_resolution,[],[f15520,f1128])).
fof(f15520,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl8 | $spl228)),
  inference(forward_demodulation,[],[f15519,f9130])).
fof(f15519,plain,(
  ( ! [X47,X48,X49] : (s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X47,X48),'const.sum.INR_1'(s(X48,X49)))))) ) | ($spl8 | $spl228)),
  inference(subsumption_resolution,[],[f15274,f12281])).
fof(f15274,plain,(
  ( ! [X47,X48,X49] : (s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X47,X48),'const.sum.INR_1'(s(X48,X49)))))) ) | $spl228),
  inference(superposition,[],[f12223,f9130])).
fof(f15517,plain,(
  $false | ($spl8 | $spl228)),
  inference(subsumption_resolution,[],[f15516,f1128])).
fof(f15516,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl8 | $spl228)),
  inference(forward_demodulation,[],[f15515,f409])).
fof(f15515,plain,(
  s(bool,'F_0') = s(bool,'const.quotient_sum.+++_4'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(sK3,fun(sK3,bool)),sK7),s('type.sum.sum'(sK2,sK3),sK10),s('type.sum.sum'(sK2,sK3),sK11))) | ($spl8 | $spl228)),
  inference(subsumption_resolution,[],[f15271,f12281])).
fof(f15271,plain,(
  s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,'const.quotient_sum.+++_4'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(sK3,fun(sK3,bool)),sK7),s('type.sum.sum'(sK2,sK3),sK10),s('type.sum.sum'(sK2,sK3),sK11))) | $spl228),
  inference(superposition,[],[f12223,f409])).
fof(f15514,plain,(
  $false | ($spl8 | $spl228)),
  inference(subsumption_resolution,[],[f15513,f1128])).
fof(f15513,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl8 | $spl228)),
  inference(forward_demodulation,[],[f15512,f308])).
fof(f15512,plain,(
  s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK3,fun(sK3,bool)),sK7),s(fun(sK3,sK1),sK8),s(fun(sK1,sK3),sK9))) | ($spl8 | $spl228)),
  inference(subsumption_resolution,[],[f15270,f12281])).
fof(f15270,plain,(
  s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK3,fun(sK3,bool)),sK7),s(fun(sK3,sK1),sK8),s(fun(sK1,sK3),sK9))) | $spl228),
  inference(superposition,[],[f12223,f308])).
fof(f15511,plain,(
  $false | ($spl8 | $spl228)),
  inference(subsumption_resolution,[],[f15510,f1128])).
fof(f15510,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl8 | $spl228)),
  inference(forward_demodulation,[],[f15509,f187])).
fof(f15509,plain,(
  s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(sK2,sK0),sK5),s(fun(sK0,sK2),sK6))) | ($spl8 | $spl228)),
  inference(subsumption_resolution,[],[f15269,f12281])).
fof(f15269,plain,(
  s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(sK2,sK0),sK5),s(fun(sK0,sK2),sK6))) | $spl228),
  inference(superposition,[],[f12223,f187])).
fof(f15501,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') = s(bool,X7) | s(bool,'F_0') = s(bool,X6)) ) | ($spl10 | $spl228)),
  inference(subsumption_resolution,[],[f15500,f119])).
fof(f15500,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') = s(bool,X7) | s(bool,'F_0') = s(bool,X6) | ~p(s(bool,X7))) ) | ($spl10 | $spl228)),
  inference(forward_demodulation,[],[f15499,f176])).
fof(f15499,plain,(
  ( ! [X6,X7] : (s(bool,X7) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,X6) | ~p(s(bool,X7))) ) | $spl228),
  inference(subsumption_resolution,[],[f15257,f119])).
fof(f15257,plain,(
  ( ! [X6,X7] : (s(bool,X7) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,X6) | ~p(s(bool,X7)) | ~p(s(bool,X6))) ) | $spl228),
  inference(superposition,[],[f12223,f115])).
fof(f15498,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') = s(bool,X5) | s(bool,'F_0') = s(bool,X4)) ) | ($spl10 | $spl228)),
  inference(subsumption_resolution,[],[f15497,f119])).
fof(f15497,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') = s(bool,X5) | s(bool,'F_0') = s(bool,X4) | ~p(s(bool,X5))) ) | ($spl10 | $spl228)),
  inference(forward_demodulation,[],[f15496,f176])).
fof(f15496,plain,(
  ( ! [X4,X5] : (s(bool,X5) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,X4) | ~p(s(bool,X5))) ) | $spl228),
  inference(subsumption_resolution,[],[f15256,f119])).
fof(f15256,plain,(
  ( ! [X4,X5] : (s(bool,X5) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,X4) | ~p(s(bool,X4)) | ~p(s(bool,X5))) ) | $spl228),
  inference(superposition,[],[f12223,f115])).
fof(f15246,plain,(
  ( ! [X83,X84,X82] : (s(bool,X84) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X82))))) | s(bool,'T_0') = s(bool,X84) | s(bool,'F_0') = s(bool,X83)) ) | $spl92),
  inference(subsumption_resolution,[],[f15212,f119])).
fof(f15212,plain,(
  ( ! [X83,X84,X82] : (~p(s(bool,X83)) | s(bool,X84) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X82))))) | s(bool,'T_0') = s(bool,X84) | s(bool,'F_0') = s(bool,X83)) ) | $spl92),
  inference(superposition,[],[f11882,f5092])).
fof(f15238,plain,(
  ( ! [X35,X33] : (s(bool,X33) != s(bool,X35) | s(bool,'F_0') = s(bool,X33) | s(bool,'T_0') = s(bool,X35)) ) | $spl254),
  inference(cnf_transformation,[],[f15238_D])).
fof(f15238_D,plain,(
  ( ! [X35,X33] : (s(bool,X33) != s(bool,X35) | s(bool,'F_0') = s(bool,X33) | s(bool,'T_0') = s(bool,X35)) ) <=> ~$spl254),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl254])])).
fof(f15239,plain,(
  ( ! [X37,X38,X36] : (s(bool,X36) != s(bool,X38) | s(bool,'T_0') = s(bool,X38) | s(bool,'F_0') = s(bool,X36) | s(bool,'F_0') = s(bool,X37)) ) | $spl92),
  inference(subsumption_resolution,[],[f15196,f119])).
fof(f15196,plain,(
  ( ! [X37,X38,X36] : (~p(s(bool,X37)) | s(bool,X36) != s(bool,X38) | s(bool,'T_0') = s(bool,X38) | s(bool,'F_0') = s(bool,X36) | s(bool,'F_0') = s(bool,X37)) ) | $spl92),
  inference(superposition,[],[f11882,f5187])).
fof(f15237,plain,(
  ( ! [X35,X33,X34] : (s(bool,X33) != s(bool,X35) | s(bool,'T_0') = s(bool,X35) | s(bool,'F_0') = s(bool,X34) | s(bool,'F_0') = s(bool,X33)) ) | $spl92),
  inference(subsumption_resolution,[],[f15195,f119])).
fof(f15195,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X34)) | s(bool,X33) != s(bool,X35) | s(bool,'T_0') = s(bool,X35) | s(bool,'F_0') = s(bool,X34) | s(bool,'F_0') = s(bool,X33)) ) | $spl92),
  inference(superposition,[],[f11882,f5187])).
fof(f176,plain,(
  s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | $spl10),
  inference(cnf_transformation,[],[f176_D])).
fof(f176_D,plain,(
  s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f12224,plain,(
  ( ! [X7] : (s(bool,X7) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'F_0') = s(bool,X7)) ) | $spl230),
  inference(cnf_transformation,[],[f12224_D])).
fof(f12224_D,plain,(
  ( ! [X7] : (s(bool,X7) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'F_0') = s(bool,X7)) ) <=> ~$spl230),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl230])])).
fof(f14849,plain,(
  ( ! [X8] : (s(bool,X8) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,X8)) )),
  inference(superposition,[],[f110,f5187])).
fof(f14848,plain,(
  ( ! [X7] : (s(bool,X7) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'F_0') = s(bool,X7) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11)))) )),
  inference(superposition,[],[f110,f5187])).
fof(f174,plain,(
  s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | $spl9),
  inference(cnf_transformation,[],[f174_D])).
fof(f174_D,plain,(
  s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f14844,plain,(
  s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11)))),
  inference(superposition,[],[f110,f126])).
fof(f251,plain,(
  p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10)))) | $spl14),
  inference(cnf_transformation,[],[f251_D])).
fof(f251_D,plain,(
  p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10)))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f14802,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))))) ) | $spl10),
  inference(subsumption_resolution,[],[f14783,f118])).
fof(f14783,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10)))) | p(s(bool,X1))) ) | $spl10),
  inference(superposition,[],[f13912,f112])).
fof(f14801,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))))) ) | $spl10),
  inference(subsumption_resolution,[],[f14782,f118])).
fof(f14782,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,X0)) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))))) ) | $spl10),
  inference(superposition,[],[f13912,f112])).
fof(f14781,plain,(
  ( ! [X8] : (s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,X8)) ) | $spl1),
  inference(subsumption_resolution,[],[f14772,f119])).
fof(f14772,plain,(
  ( ! [X8] : (~p(s(bool,X8)) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,X8)) ) | $spl1),
  inference(superposition,[],[f167,f5187])).
fof(f14780,plain,(
  ( ! [X7] : (s(bool,'F_0') = s(bool,X7) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11)))) ) | $spl1),
  inference(subsumption_resolution,[],[f14771,f119])).
fof(f14771,plain,(
  ( ! [X7] : (~p(s(bool,X7)) | s(bool,'F_0') = s(bool,X7) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11)))) ) | $spl1),
  inference(superposition,[],[f167,f5187])).
fof(f14747,plain,(
  ( ! [X8] : (s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,X8)) ) | $spl8),
  inference(subsumption_resolution,[],[f14740,f126])).
fof(f14740,plain,(
  ( ! [X8] : (s(bool,'T_0') != s(bool,X8) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,X8)) ) | $spl8),
  inference(superposition,[],[f12281,f5187])).
fof(f14746,plain,(
  ( ! [X7] : (s(bool,'F_0') = s(bool,X7) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11)))) ) | $spl8),
  inference(subsumption_resolution,[],[f14739,f126])).
fof(f14739,plain,(
  ( ! [X7] : (s(bool,'T_0') != s(bool,X7) | s(bool,'F_0') = s(bool,X7) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11)))) ) | $spl8),
  inference(superposition,[],[f12281,f5187])).
fof(f14745,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))))) ) | $spl8),
  inference(subsumption_resolution,[],[f14730,f122])).
fof(f14730,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,X3)) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))))) ) | $spl8),
  inference(superposition,[],[f12281,f115])).
fof(f14744,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))))) ) | $spl8),
  inference(subsumption_resolution,[],[f14729,f122])).
fof(f14729,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11)))) | ~p(s(bool,X2))) ) | $spl8),
  inference(superposition,[],[f12281,f115])).
fof(f168,plain,(
  p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11)))) | $spl0),
  inference(cnf_transformation,[],[f168_D])).
fof(f168_D,plain,(
  p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f14726,plain,(
  s('type.sum.sum'(sK2,sK3),sK11) = s('type.sum.sum'(sK2,sK3),'const.sum.INL_1'(s(sK2,sK12('const.sum.INL_1'(s(sK2,sK12(sK11,sK3,sK2))),sK3,sK2)))) | $spl252),
  inference(cnf_transformation,[],[f14726_D])).
fof(f14726_D,plain,(
  s('type.sum.sum'(sK2,sK3),sK11) = s('type.sum.sum'(sK2,sK3),'const.sum.INL_1'(s(sK2,sK12('const.sum.INL_1'(s(sK2,sK12(sK11,sK3,sK2))),sK3,sK2)))) <=> ~$spl252),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl252])])).
fof(f14724,plain,(
  s('type.sum.sum'(sK2,sK3),sK11) = s('type.sum.sum'(sK2,sK3),'const.sum.INL_1'(s(sK2,sK12('const.sum.INL_1'(s(sK2,sK12(sK11,sK3,sK2))),sK3,sK2)))) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11)))) | $spl237),
  inference(forward_demodulation,[],[f14698,f14654])).
fof(f14654,plain,(
  s('type.sum.sum'(sK2,sK3),sK11) = s('type.sum.sum'(sK2,sK3),'const.sum.INL_1'(s(sK2,sK12(sK11,sK3,sK2)))) | $spl237),
  inference(trivial_inequality_removal,[],[f14653])).
fof(f14653,plain,(
  s('type.sum.sum'(sK2,sK3),sK11) != s('type.sum.sum'(sK2,sK3),sK11) | s('type.sum.sum'(sK2,sK3),sK11) = s('type.sum.sum'(sK2,sK3),'const.sum.INL_1'(s(sK2,sK12(sK11,sK3,sK2)))) | $spl237),
  inference(superposition,[],[f12641,f111])).
fof(f12641,plain,(
  s('type.sum.sum'(sK2,sK3),sK11) != s('type.sum.sum'(sK2,sK3),'const.sum.INR_1'(s(sK3,sK13(sK11,sK3,sK2)))) | $spl237),
  inference(cnf_transformation,[],[f12641_D])).
fof(f12641_D,plain,(
  s('type.sum.sum'(sK2,sK3),sK11) != s('type.sum.sum'(sK2,sK3),'const.sum.INR_1'(s(sK3,sK13(sK11,sK3,sK2)))) <=> ~$spl237),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl237])])).
fof(f14698,plain,(
  p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11)))) | s('type.sum.sum'(sK2,sK3),'const.sum.INL_1'(s(sK2,sK12(sK11,sK3,sK2)))) = s('type.sum.sum'(sK2,sK3),'const.sum.INL_1'(s(sK2,sK12('const.sum.INL_1'(s(sK2,sK12(sK11,sK3,sK2))),sK3,sK2)))) | $spl237),
  inference(superposition,[],[f9189,f14654])).
fof(f14712,plain,(
  $false | ($spl12 | $spl237)),
  inference(subsumption_resolution,[],[f14711,f1128])).
fof(f14711,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl12 | $spl237)),
  inference(forward_demodulation,[],[f14691,f178])).
fof(f14691,plain,(
  s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | $spl237),
  inference(superposition,[],[f5433,f14654])).
fof(f14710,plain,(
  $false | ($spl12 | $spl237)),
  inference(subsumption_resolution,[],[f14690,f178])).
fof(f14690,plain,(
  s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | $spl237),
  inference(superposition,[],[f5633,f14654])).
fof(f14652,plain,(
  ( ! [X177,X185,X175,X182,X180,X178,X176,X184,X174,X183,X181,X179] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X180,fun(X180,bool)),X181),s(fun(X180,X182),X183),s(fun(X182,X180),X184)))) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X176,X182),X185))) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X174,X180),'const.sum.++_3'(s(fun(X176,X174),X178),s(fun(X182,X180),X184),s('type.sum.sum'(X176,X182),X185)))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X174,fun(X174,bool)),X175),s(fun(X174,X176),X177),s(fun(X176,X174),X178))) | s(bool,'F_0') = s(bool,X179)) )),
  inference(subsumption_resolution,[],[f14638,f119])).
fof(f14638,plain,(
  ( ! [X177,X185,X175,X182,X180,X178,X176,X184,X174,X183,X181,X179] : (~p(s(bool,X179)) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X180,fun(X180,bool)),X181),s(fun(X180,X182),X183),s(fun(X182,X180),X184)))) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X176,X182),X185))) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X174,X180),'const.sum.++_3'(s(fun(X176,X174),X178),s(fun(X182,X180),X184),s('type.sum.sum'(X176,X182),X185)))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X174,fun(X174,bool)),X175),s(fun(X174,X176),X177),s(fun(X176,X174),X178))) | s(bool,'F_0') = s(bool,X179)) )),
  inference(superposition,[],[f9447,f5187])).
fof(f14650,plain,(
  ( ! [X167,X165,X163,X173,X171,X169,X166,X164,X162,X172,X170,X168] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X168,fun(X168,bool)),X169),s(fun(X168,X170),X171),s(fun(X170,X168),X172)))) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X164,X170),X173))) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X162,X168),'const.sum.++_3'(s(fun(X164,X162),X166),s(fun(X170,X168),X172),s('type.sum.sum'(X164,X170),X173)))))) | s(bool,'F_0') = s(bool,X167) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X162,fun(X162,bool)),X163),s(fun(X162,X164),X165),s(fun(X164,X162),X166)))) )),
  inference(subsumption_resolution,[],[f14637,f119])).
fof(f14637,plain,(
  ( ! [X167,X165,X163,X173,X171,X169,X166,X164,X162,X172,X170,X168] : (~p(s(bool,X167)) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X168,fun(X168,bool)),X169),s(fun(X168,X170),X171),s(fun(X170,X168),X172)))) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X164,X170),X173))) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X162,X168),'const.sum.++_3'(s(fun(X164,X162),X166),s(fun(X170,X168),X172),s('type.sum.sum'(X164,X170),X173)))))) | s(bool,'F_0') = s(bool,X167) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X162,fun(X162,bool)),X163),s(fun(X162,X164),X165),s(fun(X164,X162),X166)))) )),
  inference(superposition,[],[f9447,f5187])).
fof(f14603,plain,(
  $false | ($spl32 | $spl236)),
  inference(subsumption_resolution,[],[f14602,f1128])).
fof(f14602,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl32 | $spl236)),
  inference(forward_demodulation,[],[f14582,f771])).
fof(f14582,plain,(
  s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) | $spl236),
  inference(superposition,[],[f5444,f12642])).
fof(f14601,plain,(
  $false | ($spl32 | $spl236)),
  inference(subsumption_resolution,[],[f14581,f771])).
fof(f14581,plain,(
  s(bool,'T_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) | $spl236),
  inference(superposition,[],[f5815,f12642])).
fof(f14567,plain,(
  ( ! [X177,X185,X175,X182,X180,X178,X176,X184,X174,X183,X181,X179] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X180,X176),X181))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X182,X174),'const.sum.++_3'(s(fun(X180,X182),X183),s(fun(X176,X174),X178),s('type.sum.sum'(X180,X176),X181))))) | s(bool,'T_0') != s(bool,'const.quotient.QUOTIENT_3'(s(fun(X182,fun(X182,bool)),X184),s(fun(X182,X180),X185),s(fun(X180,X182),X183))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X174,fun(X174,bool)),X175),s(fun(X174,X176),X177),s(fun(X176,X174),X178))) | s(bool,'F_0') = s(bool,X179)) )),
  inference(subsumption_resolution,[],[f14553,f119])).
fof(f14553,plain,(
  ( ! [X177,X185,X175,X182,X180,X178,X176,X184,X174,X183,X181,X179] : (~p(s(bool,X179)) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X180,X176),X181))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X182,X174),'const.sum.++_3'(s(fun(X180,X182),X183),s(fun(X176,X174),X178),s('type.sum.sum'(X180,X176),X181))))) | s(bool,'T_0') != s(bool,'const.quotient.QUOTIENT_3'(s(fun(X182,fun(X182,bool)),X184),s(fun(X182,X180),X185),s(fun(X180,X182),X183))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X174,fun(X174,bool)),X175),s(fun(X174,X176),X177),s(fun(X176,X174),X178))) | s(bool,'F_0') = s(bool,X179)) )),
  inference(superposition,[],[f9426,f5187])).
fof(f14565,plain,(
  ( ! [X167,X165,X163,X173,X171,X169,X166,X164,X162,X172,X170,X168] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X168,X164),X169))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X170,X162),'const.sum.++_3'(s(fun(X168,X170),X171),s(fun(X164,X162),X166),s('type.sum.sum'(X168,X164),X169))))) | s(bool,'T_0') != s(bool,'const.quotient.QUOTIENT_3'(s(fun(X170,fun(X170,bool)),X172),s(fun(X170,X168),X173),s(fun(X168,X170),X171))) | s(bool,'F_0') = s(bool,X167) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X162,fun(X162,bool)),X163),s(fun(X162,X164),X165),s(fun(X164,X162),X166)))) )),
  inference(subsumption_resolution,[],[f14552,f119])).
fof(f14552,plain,(
  ( ! [X167,X165,X163,X173,X171,X169,X166,X164,X162,X172,X170,X168] : (~p(s(bool,X167)) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X168,X164),X169))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X170,X162),'const.sum.++_3'(s(fun(X168,X170),X171),s(fun(X164,X162),X166),s('type.sum.sum'(X168,X164),X169))))) | s(bool,'T_0') != s(bool,'const.quotient.QUOTIENT_3'(s(fun(X170,fun(X170,bool)),X172),s(fun(X170,X168),X173),s(fun(X168,X170),X171))) | s(bool,'F_0') = s(bool,X167) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X162,fun(X162,bool)),X163),s(fun(X162,X164),X165),s(fun(X164,X162),X166)))) )),
  inference(superposition,[],[f9426,f5187])).
fof(f14521,plain,(
  $false | ($spl26 | $spl32 | $spl34)),
  inference(subsumption_resolution,[],[f14520,f1128])).
fof(f14520,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl26 | $spl32 | $spl34)),
  inference(forward_demodulation,[],[f773,f1035])).
fof(f14073,plain,(
  $false | ($spl26 | $spl32 | $spl225)),
  inference(subsumption_resolution,[],[f12174,f13927])).
fof(f13927,plain,(
  p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10)))) | ($spl26 | $spl32)),
  inference(trivial_inequality_removal,[],[f1064])).
fof(f1064,plain,(
  s(bool,'T_0') != s(bool,'T_0') | p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10)))) | ($spl26 | $spl32)),
  inference(superposition,[],[f122,f1035])).
fof(f12174,plain,(
  ~p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10)))) | $spl225),
  inference(cnf_transformation,[],[f12174_D])).
fof(f12174_D,plain,(
  ~p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10)))) <=> ~$spl225),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl225])])).
fof(f171,plain,(
  s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | $spl6),
  inference(cnf_transformation,[],[f171_D])).
fof(f171_D,plain,(
  s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f14072,plain,(
  ( ! [X8] : (s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'F_0') = s(bool,X8)) ) | $spl12),
  inference(subsumption_resolution,[],[f14067,f126])).
fof(f14067,plain,(
  ( ! [X8] : (s(bool,'T_0') != s(bool,X8) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'F_0') = s(bool,X8)) ) | $spl12),
  inference(superposition,[],[f12180,f5187])).
fof(f12180,plain,(
  s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | $spl12),
  inference(forward_demodulation,[],[f110,f178])).
fof(f14071,plain,(
  ( ! [X7] : (s(bool,'F_0') = s(bool,X7) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10)))) ) | $spl12),
  inference(subsumption_resolution,[],[f14066,f126])).
fof(f14066,plain,(
  ( ! [X7] : (s(bool,'T_0') != s(bool,X7) | s(bool,'F_0') = s(bool,X7) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10)))) ) | $spl12),
  inference(superposition,[],[f12180,f5187])).
fof(f14018,plain,(
  ( ! [X8] : (s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'F_0') = s(bool,X8)) ) | $spl15),
  inference(subsumption_resolution,[],[f14009,f119])).
fof(f14009,plain,(
  ( ! [X8] : (~p(s(bool,X8)) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'F_0') = s(bool,X8)) ) | $spl15),
  inference(superposition,[],[f250,f5187])).
fof(f14017,plain,(
  ( ! [X7] : (s(bool,'F_0') = s(bool,X7) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10)))) ) | $spl15),
  inference(subsumption_resolution,[],[f14008,f119])).
fof(f14008,plain,(
  ( ! [X7] : (~p(s(bool,X7)) | s(bool,'F_0') = s(bool,X7) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10)))) ) | $spl15),
  inference(superposition,[],[f250,f5187])).
fof(f13926,plain,(
  $false | ($spl29 | $spl33)),
  inference(subsumption_resolution,[],[f13131,f770])).
fof(f770,plain,(
  s(bool,'T_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) | $spl33),
  inference(cnf_transformation,[],[f770_D])).
fof(f770_D,plain,(
  s(bool,'T_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f13131,plain,(
  s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) | $spl29),
  inference(trivial_inequality_removal,[],[f13124])).
fof(f13124,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) | $spl29),
  inference(superposition,[],[f766,f126])).
fof(f766,plain,(
  s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) | $spl29),
  inference(cnf_transformation,[],[f766_D])).
fof(f766_D,plain,(
  s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f13135,plain,(
  ~p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11)))) | $spl241),
  inference(cnf_transformation,[],[f13135_D])).
fof(f13135_D,plain,(
  ~p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11)))) <=> ~$spl241),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl241])])).
fof(f13656,plain,(
  ( ! [X20] : (s(bool,'T_0') != s(bool,X20) | ~p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))))) ) | $spl33),
  inference(subsumption_resolution,[],[f13641,f122])).
fof(f13641,plain,(
  ( ! [X20] : (s(bool,'T_0') != s(bool,X20) | ~p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11)))) | ~p(s(bool,X20))) ) | $spl33),
  inference(superposition,[],[f770,f115])).
fof(f13658,plain,(
  ( ! [X21] : (s(bool,'T_0') != s(bool,X21) | ~p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))))) ) | $spl33),
  inference(subsumption_resolution,[],[f13642,f122])).
fof(f13642,plain,(
  ( ! [X21] : (s(bool,'T_0') != s(bool,X21) | ~p(s(bool,X21)) | ~p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))))) ) | $spl33),
  inference(superposition,[],[f770,f115])).
fof(f767,plain,(
  s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) | $spl28),
  inference(cnf_transformation,[],[f767_D])).
fof(f767_D,plain,(
  s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f13925,plain,(
  ( ! [X25] : (s(bool,'F_0') = s(bool,X25) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11)))) ) | $spl33),
  inference(subsumption_resolution,[],[f13651,f126])).
fof(f13651,plain,(
  ( ! [X25] : (s(bool,'T_0') != s(bool,X25) | s(bool,'F_0') = s(bool,X25) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11)))) ) | $spl33),
  inference(superposition,[],[f770,f5187])).
fof(f13924,plain,(
  ( ! [X26] : (s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,X26)) ) | $spl33),
  inference(subsumption_resolution,[],[f13652,f126])).
fof(f13652,plain,(
  ( ! [X26] : (s(bool,'T_0') != s(bool,X26) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,X26)) ) | $spl33),
  inference(superposition,[],[f770,f5187])).
fof(f13921,plain,(
  $false | ($spl0 | $spl10)),
  inference(subsumption_resolution,[],[f13920,f150])).
fof(f13920,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl10)),
  inference(forward_demodulation,[],[f168,f176])).
fof(f13919,plain,(
  $false | ($spl6 | $spl10)),
  inference(subsumption_resolution,[],[f436,f176])).
fof(f436,plain,(
  s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | $spl6),
  inference(backward_demodulation,[],[f171,f110])).
fof(f13918,plain,(
  $false | ($spl6 | $spl31)),
  inference(subsumption_resolution,[],[f10633,f768])).
fof(f10633,plain,(
  s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) | $spl6),
  inference(superposition,[],[f9229,f10620])).
fof(f10620,plain,(
  s('type.sum.sum'(sK2,sK3),sK10) = s('type.sum.sum'(sK2,sK3),'const.sum.INL_1'(s(sK2,sK12(sK10,sK3,sK2)))) | $spl6),
  inference(subsumption_resolution,[],[f10590,f150])).
fof(f10590,plain,(
  p(s(bool,'F_0')) | s('type.sum.sum'(sK2,sK3),sK10) = s('type.sum.sum'(sK2,sK3),'const.sum.INL_1'(s(sK2,sK12(sK10,sK3,sK2)))) | $spl6),
  inference(superposition,[],[f9189,f171])).
fof(f13917,plain,(
  $false | ($spl6 | $spl10)),
  inference(subsumption_resolution,[],[f13916,f150])).
fof(f13916,plain,(
  p(s(bool,'F_0')) | ($spl6 | $spl10)),
  inference(forward_demodulation,[],[f521,f176])).
fof(f521,plain,(
  p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11)))) | $spl6),
  inference(trivial_inequality_removal,[],[f515])).
fof(f515,plain,(
  s(bool,'F_0') != s(bool,'F_0') | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11)))) | $spl6),
  inference(superposition,[],[f436,f121])).
fof(f13915,plain,(
  $false | ($spl6 | $spl10)),
  inference(subsumption_resolution,[],[f13914,f150])).
fof(f13914,plain,(
  p(s(bool,'F_0')) | ($spl6 | $spl10)),
  inference(forward_demodulation,[],[f522,f176])).
fof(f522,plain,(
  p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11)))) | $spl6),
  inference(trivial_inequality_removal,[],[f514])).
fof(f514,plain,(
  s(bool,'F_0') != s(bool,'F_0') | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11)))) | $spl6),
  inference(superposition,[],[f436,f119])).
fof(f13913,plain,(
  $false | ($spl10 | $spl12)),
  inference(subsumption_resolution,[],[f12181,f1128])).
fof(f12181,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl10 | $spl12)),
  inference(forward_demodulation,[],[f176,f178])).
fof(f13894,plain,(
  $false | ($spl1 | $spl33)),
  inference(subsumption_resolution,[],[f13873,f770])).
fof(f13873,plain,(
  s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) | $spl1),
  inference(superposition,[],[f9229,f12225])).
fof(f12225,plain,(
  s('type.sum.sum'(sK2,sK3),sK11) = s('type.sum.sum'(sK2,sK3),'const.sum.INL_1'(s(sK2,sK12(sK11,sK3,sK2)))) | $spl1),
  inference(resolution,[],[f167,f9189])).
fof(f13870,plain,(
  ( ! [X177,X185,X175,X182,X180,X178,X176,X184,X174,X183,X181,X179] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X180,X176),X181))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X182,X174),'const.sum.++_3'(s(fun(X180,X182),X183),s(fun(X176,X174),X178),s('type.sum.sum'(X180,X176),X181))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X182,fun(X182,bool)),X184),s(fun(X182,X180),X185),s(fun(X180,X182),X183))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X174,fun(X174,bool)),X175),s(fun(X174,X176),X177),s(fun(X176,X174),X178))) | s(bool,'F_0') = s(bool,X179)) )),
  inference(subsumption_resolution,[],[f13857,f119])).
fof(f13857,plain,(
  ( ! [X177,X185,X175,X182,X180,X178,X176,X184,X174,X183,X181,X179] : (~p(s(bool,X179)) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X180,X176),X181))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X182,X174),'const.sum.++_3'(s(fun(X180,X182),X183),s(fun(X176,X174),X178),s('type.sum.sum'(X180,X176),X181))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X182,fun(X182,bool)),X184),s(fun(X182,X180),X185),s(fun(X180,X182),X183))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X174,fun(X174,bool)),X175),s(fun(X174,X176),X177),s(fun(X176,X174),X178))) | s(bool,'F_0') = s(bool,X179)) )),
  inference(superposition,[],[f9424,f5187])).
fof(f13869,plain,(
  ( ! [X167,X165,X163,X173,X171,X169,X166,X164,X162,X172,X170,X168] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X168,X164),X169))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X170,X162),'const.sum.++_3'(s(fun(X168,X170),X171),s(fun(X164,X162),X166),s('type.sum.sum'(X168,X164),X169))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X170,fun(X170,bool)),X172),s(fun(X170,X168),X173),s(fun(X168,X170),X171))) | s(bool,'F_0') = s(bool,X167) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X162,fun(X162,bool)),X163),s(fun(X162,X164),X165),s(fun(X164,X162),X166)))) )),
  inference(subsumption_resolution,[],[f13856,f119])).
fof(f13856,plain,(
  ( ! [X167,X165,X163,X173,X171,X169,X166,X164,X162,X172,X170,X168] : (~p(s(bool,X167)) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X168,X164),X169))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X170,X162),'const.sum.++_3'(s(fun(X168,X170),X171),s(fun(X164,X162),X166),s('type.sum.sum'(X168,X164),X169))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X170,fun(X170,bool)),X172),s(fun(X170,X168),X173),s(fun(X168,X170),X171))) | s(bool,'F_0') = s(bool,X167) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X162,fun(X162,bool)),X163),s(fun(X162,X164),X165),s(fun(X164,X162),X166)))) )),
  inference(superposition,[],[f9424,f5187])).
fof(f13825,plain,(
  ( ! [X177,X185,X175,X182,X180,X178,X176,X184,X174,X183,X181,X179] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X176,X180),X181))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X174,X182),'const.sum.++_3'(s(fun(X176,X174),X178),s(fun(X180,X182),X183),s('type.sum.sum'(X176,X180),X181))))) | s(bool,'T_0') != s(bool,'const.quotient.QUOTIENT_3'(s(fun(X182,fun(X182,bool)),X184),s(fun(X182,X180),X185),s(fun(X180,X182),X183))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X174,fun(X174,bool)),X175),s(fun(X174,X176),X177),s(fun(X176,X174),X178))) | s(bool,'F_0') = s(bool,X179)) )),
  inference(subsumption_resolution,[],[f13811,f119])).
fof(f13811,plain,(
  ( ! [X177,X185,X175,X182,X180,X178,X176,X184,X174,X183,X181,X179] : (~p(s(bool,X179)) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X176,X180),X181))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X174,X182),'const.sum.++_3'(s(fun(X176,X174),X178),s(fun(X180,X182),X183),s('type.sum.sum'(X176,X180),X181))))) | s(bool,'T_0') != s(bool,'const.quotient.QUOTIENT_3'(s(fun(X182,fun(X182,bool)),X184),s(fun(X182,X180),X185),s(fun(X180,X182),X183))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X174,fun(X174,bool)),X175),s(fun(X174,X176),X177),s(fun(X176,X174),X178))) | s(bool,'F_0') = s(bool,X179)) )),
  inference(superposition,[],[f9416,f5187])).
fof(f13823,plain,(
  ( ! [X167,X165,X163,X173,X171,X169,X166,X164,X162,X172,X170,X168] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X164,X168),X169))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X162,X170),'const.sum.++_3'(s(fun(X164,X162),X166),s(fun(X168,X170),X171),s('type.sum.sum'(X164,X168),X169))))) | s(bool,'T_0') != s(bool,'const.quotient.QUOTIENT_3'(s(fun(X170,fun(X170,bool)),X172),s(fun(X170,X168),X173),s(fun(X168,X170),X171))) | s(bool,'F_0') = s(bool,X167) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X162,fun(X162,bool)),X163),s(fun(X162,X164),X165),s(fun(X164,X162),X166)))) )),
  inference(subsumption_resolution,[],[f13810,f119])).
fof(f13810,plain,(
  ( ! [X167,X165,X163,X173,X171,X169,X166,X164,X162,X172,X170,X168] : (~p(s(bool,X167)) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X164,X168),X169))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X162,X170),'const.sum.++_3'(s(fun(X164,X162),X166),s(fun(X168,X170),X171),s('type.sum.sum'(X164,X168),X169))))) | s(bool,'T_0') != s(bool,'const.quotient.QUOTIENT_3'(s(fun(X170,fun(X170,bool)),X172),s(fun(X170,X168),X173),s(fun(X168,X170),X171))) | s(bool,'F_0') = s(bool,X167) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X162,fun(X162,bool)),X163),s(fun(X162,X164),X165),s(fun(X164,X162),X166)))) )),
  inference(superposition,[],[f9416,f5187])).
fof(f13784,plain,(
  ( ! [X177,X185,X175,X182,X180,X178,X176,X184,X174,X183,X181,X179] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X176,X180),X181))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X174,X182),'const.sum.++_3'(s(fun(X176,X174),X178),s(fun(X180,X182),X183),s('type.sum.sum'(X176,X180),X181))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X182,fun(X182,bool)),X184),s(fun(X182,X180),X185),s(fun(X180,X182),X183))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X174,fun(X174,bool)),X175),s(fun(X174,X176),X177),s(fun(X176,X174),X178))) | s(bool,'F_0') = s(bool,X179)) )),
  inference(subsumption_resolution,[],[f13770,f119])).
fof(f13770,plain,(
  ( ! [X177,X185,X175,X182,X180,X178,X176,X184,X174,X183,X181,X179] : (~p(s(bool,X179)) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X176,X180),X181))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X174,X182),'const.sum.++_3'(s(fun(X176,X174),X178),s(fun(X180,X182),X183),s('type.sum.sum'(X176,X180),X181))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X182,fun(X182,bool)),X184),s(fun(X182,X180),X185),s(fun(X180,X182),X183))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X174,fun(X174,bool)),X175),s(fun(X174,X176),X177),s(fun(X176,X174),X178))) | s(bool,'F_0') = s(bool,X179)) )),
  inference(superposition,[],[f9414,f5187])).
fof(f13782,plain,(
  ( ! [X167,X165,X163,X173,X171,X169,X166,X164,X162,X172,X170,X168] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X164,X168),X169))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X162,X170),'const.sum.++_3'(s(fun(X164,X162),X166),s(fun(X168,X170),X171),s('type.sum.sum'(X164,X168),X169))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X170,fun(X170,bool)),X172),s(fun(X170,X168),X173),s(fun(X168,X170),X171))) | s(bool,'F_0') = s(bool,X167) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X162,fun(X162,bool)),X163),s(fun(X162,X164),X165),s(fun(X164,X162),X166)))) )),
  inference(subsumption_resolution,[],[f13769,f119])).
fof(f13769,plain,(
  ( ! [X167,X165,X163,X173,X171,X169,X166,X164,X162,X172,X170,X168] : (~p(s(bool,X167)) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X164,X168),X169))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X162,X170),'const.sum.++_3'(s(fun(X164,X162),X166),s(fun(X168,X170),X171),s('type.sum.sum'(X164,X168),X169))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X170,fun(X170,bool)),X172),s(fun(X170,X168),X173),s(fun(X168,X170),X171))) | s(bool,'F_0') = s(bool,X167) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X162,fun(X162,bool)),X163),s(fun(X162,X164),X165),s(fun(X164,X162),X166)))) )),
  inference(superposition,[],[f9414,f5187])).
fof(f13742,plain,(
  ( ! [X111,X116,X114,X112,X110,X117,X115,X113] : (~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X110,fun(X110,bool)),X111),s(fun(X110,fun(X110,bool)),X111),s('type.sum.sum'(X110,X110),X116),s('type.sum.sum'(X110,X110),X117)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X110,X110),X116))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X110,X110),X117))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X110,fun(X110,bool)),X111),s(fun(X110,X112),X113),s(fun(X112,X110),X114))) | s(bool,'F_0') = s(bool,X115)) )),
  inference(subsumption_resolution,[],[f13728,f119])).
fof(f13728,plain,(
  ( ! [X111,X116,X114,X112,X110,X117,X115,X113] : (~p(s(bool,X115)) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X110,fun(X110,bool)),X111),s(fun(X110,fun(X110,bool)),X111),s('type.sum.sum'(X110,X110),X116),s('type.sum.sum'(X110,X110),X117)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X110,X110),X116))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X110,X110),X117))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X110,fun(X110,bool)),X111),s(fun(X110,X112),X113),s(fun(X112,X110),X114))) | s(bool,'F_0') = s(bool,X115)) )),
  inference(superposition,[],[f9892,f5187])).
fof(f13740,plain,(
  ( ! [X103,X109,X107,X105,X102,X108,X106,X104] : (~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X102,fun(X102,bool)),X103),s(fun(X102,fun(X102,bool)),X103),s('type.sum.sum'(X102,X102),X108),s('type.sum.sum'(X102,X102),X109)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X102,X102),X108))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X102,X102),X109))) | s(bool,'F_0') = s(bool,X107) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X102,fun(X102,bool)),X103),s(fun(X102,X104),X105),s(fun(X104,X102),X106)))) )),
  inference(subsumption_resolution,[],[f13727,f119])).
fof(f13727,plain,(
  ( ! [X103,X109,X107,X105,X102,X108,X106,X104] : (~p(s(bool,X107)) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X102,fun(X102,bool)),X103),s(fun(X102,fun(X102,bool)),X103),s('type.sum.sum'(X102,X102),X108),s('type.sum.sum'(X102,X102),X109)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X102,X102),X108))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X102,X102),X109))) | s(bool,'F_0') = s(bool,X107) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X102,fun(X102,bool)),X103),s(fun(X102,X104),X105),s(fun(X104,X102),X106)))) )),
  inference(superposition,[],[f9892,f5187])).
fof(f13701,plain,(
  ( ! [X111,X116,X114,X112,X110,X117,X115,X113] : (~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X110,fun(X110,bool)),X111),s(fun(sK3,fun(sK3,bool)),sK7),s('type.sum.sum'(X110,sK3),X116),s('type.sum.sum'(X110,sK3),X117)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X110,sK3),X116))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X110,sK3),X117))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X110,fun(X110,bool)),X111),s(fun(X110,X112),X113),s(fun(X112,X110),X114))) | s(bool,'F_0') = s(bool,X115)) )),
  inference(subsumption_resolution,[],[f13688,f119])).
fof(f13688,plain,(
  ( ! [X111,X116,X114,X112,X110,X117,X115,X113] : (~p(s(bool,X115)) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X110,fun(X110,bool)),X111),s(fun(sK3,fun(sK3,bool)),sK7),s('type.sum.sum'(X110,sK3),X116),s('type.sum.sum'(X110,sK3),X117)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X110,sK3),X116))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X110,sK3),X117))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X110,fun(X110,bool)),X111),s(fun(X110,X112),X113),s(fun(X112,X110),X114))) | s(bool,'F_0') = s(bool,X115)) )),
  inference(superposition,[],[f301,f5187])).
fof(f13700,plain,(
  ( ! [X103,X109,X107,X105,X102,X108,X106,X104] : (~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X102,fun(X102,bool)),X103),s(fun(sK3,fun(sK3,bool)),sK7),s('type.sum.sum'(X102,sK3),X108),s('type.sum.sum'(X102,sK3),X109)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X102,sK3),X108))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X102,sK3),X109))) | s(bool,'F_0') = s(bool,X107) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X102,fun(X102,bool)),X103),s(fun(X102,X104),X105),s(fun(X104,X102),X106)))) )),
  inference(subsumption_resolution,[],[f13687,f119])).
fof(f13687,plain,(
  ( ! [X103,X109,X107,X105,X102,X108,X106,X104] : (~p(s(bool,X107)) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X102,fun(X102,bool)),X103),s(fun(sK3,fun(sK3,bool)),sK7),s('type.sum.sum'(X102,sK3),X108),s('type.sum.sum'(X102,sK3),X109)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X102,sK3),X108))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X102,sK3),X109))) | s(bool,'F_0') = s(bool,X107) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X102,fun(X102,bool)),X103),s(fun(X102,X104),X105),s(fun(X104,X102),X106)))) )),
  inference(superposition,[],[f301,f5187])).
fof(f13482,plain,(
  $false | ($spl26 | $spl28 | $spl30)),
  inference(subsumption_resolution,[],[f13414,f1128])).
fof(f13414,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl26 | $spl28 | $spl30)),
  inference(backward_demodulation,[],[f769,f798])).
fof(f798,plain,(
  s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) | ($spl26 | $spl28)),
  inference(backward_demodulation,[],[f767,f429])).
fof(f773,plain,(
  s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) | $spl34),
  inference(cnf_transformation,[],[f773_D])).
fof(f773_D,plain,(
  s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f13183,plain,(
  ( ! [X26] : (s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'F_0') = s(bool,X26)) ) | ($spl26 | $spl28)),
  inference(subsumption_resolution,[],[f13174,f119])).
fof(f13174,plain,(
  ( ! [X26] : (~p(s(bool,X26)) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'F_0') = s(bool,X26)) ) | ($spl26 | $spl28)),
  inference(superposition,[],[f13139,f5187])).
fof(f13139,plain,(
  ~p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10)))) | ($spl26 | $spl28)),
  inference(trivial_inequality_removal,[],[f12699])).
fof(f12699,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ~p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10)))) | ($spl26 | $spl28)),
  inference(superposition,[],[f118,f798])).
fof(f13182,plain,(
  ( ! [X25] : (s(bool,'F_0') = s(bool,X25) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10)))) ) | ($spl26 | $spl28)),
  inference(subsumption_resolution,[],[f13173,f119])).
fof(f13173,plain,(
  ( ! [X25] : (~p(s(bool,X25)) | s(bool,'F_0') = s(bool,X25) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10)))) ) | ($spl26 | $spl28)),
  inference(superposition,[],[f13139,f5187])).
fof(f13136,plain,(
  p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11)))) | $spl240),
  inference(cnf_transformation,[],[f13136_D])).
fof(f13136_D,plain,(
  p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11)))) <=> ~$spl240),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl240])])).
fof(f13137,plain,(
  ( ! [X19] : (s(bool,'F_0') != s(bool,X19) | p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))))) ) | $spl29),
  inference(subsumption_resolution,[],[f13116,f118])).
fof(f13116,plain,(
  ( ! [X19] : (s(bool,'F_0') != s(bool,X19) | p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11)))) | p(s(bool,X19))) ) | $spl29),
  inference(superposition,[],[f766,f112])).
fof(f13134,plain,(
  ( ! [X18] : (s(bool,'F_0') != s(bool,X18) | p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))))) ) | $spl29),
  inference(subsumption_resolution,[],[f13115,f118])).
fof(f13115,plain,(
  ( ! [X18] : (s(bool,'F_0') != s(bool,X18) | p(s(bool,X18)) | p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))))) ) | $spl29),
  inference(superposition,[],[f766,f112])).
fof(f13077,plain,(
  $false | ($spl26 | $spl28 | $spl224)),
  inference(subsumption_resolution,[],[f13076,f150])).
fof(f13076,plain,(
  p(s(bool,'F_0')) | ($spl26 | $spl28 | $spl224)),
  inference(forward_demodulation,[],[f12175,f798])).
fof(f13075,plain,(
  ( ! [X111,X116,X114,X112,X110,X117,X115,X113] : (~p(s(bool,'const.quotient_sum.+++_4'(s(fun(sK3,fun(sK3,bool)),sK7),s(fun(X110,fun(X110,bool)),X111),s('type.sum.sum'(sK3,X110),X116),s('type.sum.sum'(sK3,X110),X117)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK3,X110),X116))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK3,X110),X117))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X110,fun(X110,bool)),X111),s(fun(X110,X112),X113),s(fun(X112,X110),X114))) | s(bool,'F_0') = s(bool,X115)) )),
  inference(subsumption_resolution,[],[f13062,f119])).
fof(f13062,plain,(
  ( ! [X111,X116,X114,X112,X110,X117,X115,X113] : (~p(s(bool,X115)) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(sK3,fun(sK3,bool)),sK7),s(fun(X110,fun(X110,bool)),X111),s('type.sum.sum'(sK3,X110),X116),s('type.sum.sum'(sK3,X110),X117)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK3,X110),X116))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK3,X110),X117))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X110,fun(X110,bool)),X111),s(fun(X110,X112),X113),s(fun(X112,X110),X114))) | s(bool,'F_0') = s(bool,X115)) )),
  inference(superposition,[],[f300,f5187])).
fof(f13074,plain,(
  ( ! [X103,X109,X107,X105,X102,X108,X106,X104] : (~p(s(bool,'const.quotient_sum.+++_4'(s(fun(sK3,fun(sK3,bool)),sK7),s(fun(X102,fun(X102,bool)),X103),s('type.sum.sum'(sK3,X102),X108),s('type.sum.sum'(sK3,X102),X109)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK3,X102),X108))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK3,X102),X109))) | s(bool,'F_0') = s(bool,X107) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X102,fun(X102,bool)),X103),s(fun(X102,X104),X105),s(fun(X104,X102),X106)))) )),
  inference(subsumption_resolution,[],[f13061,f119])).
fof(f13061,plain,(
  ( ! [X103,X109,X107,X105,X102,X108,X106,X104] : (~p(s(bool,X107)) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(sK3,fun(sK3,bool)),sK7),s(fun(X102,fun(X102,bool)),X103),s('type.sum.sum'(sK3,X102),X108),s('type.sum.sum'(sK3,X102),X109)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK3,X102),X108))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK3,X102),X109))) | s(bool,'F_0') = s(bool,X107) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X102,fun(X102,bool)),X103),s(fun(X102,X104),X105),s(fun(X104,X102),X106)))) )),
  inference(superposition,[],[f300,f5187])).
fof(f13030,plain,(
  ( ! [X111,X116,X114,X112,X110,X117,X115,X113] : (~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X110,fun(X110,bool)),X111),s(fun(sK2,fun(sK2,bool)),sK4),s('type.sum.sum'(X110,sK2),X116),s('type.sum.sum'(X110,sK2),X117)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X110,sK2),X116))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X110,sK2),X117))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X110,fun(X110,bool)),X111),s(fun(X110,X112),X113),s(fun(X112,X110),X114))) | s(bool,'F_0') = s(bool,X115)) )),
  inference(subsumption_resolution,[],[f13017,f119])).
fof(f13017,plain,(
  ( ! [X111,X116,X114,X112,X110,X117,X115,X113] : (~p(s(bool,X115)) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X110,fun(X110,bool)),X111),s(fun(sK2,fun(sK2,bool)),sK4),s('type.sum.sum'(X110,sK2),X116),s('type.sum.sum'(X110,sK2),X117)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X110,sK2),X116))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X110,sK2),X117))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X110,fun(X110,bool)),X111),s(fun(X110,X112),X113),s(fun(X112,X110),X114))) | s(bool,'F_0') = s(bool,X115)) )),
  inference(superposition,[],[f180,f5187])).
fof(f13029,plain,(
  ( ! [X103,X109,X107,X105,X102,X108,X106,X104] : (~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X102,fun(X102,bool)),X103),s(fun(sK2,fun(sK2,bool)),sK4),s('type.sum.sum'(X102,sK2),X108),s('type.sum.sum'(X102,sK2),X109)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X102,sK2),X108))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X102,sK2),X109))) | s(bool,'F_0') = s(bool,X107) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X102,fun(X102,bool)),X103),s(fun(X102,X104),X105),s(fun(X104,X102),X106)))) )),
  inference(subsumption_resolution,[],[f13016,f119])).
fof(f13016,plain,(
  ( ! [X103,X109,X107,X105,X102,X108,X106,X104] : (~p(s(bool,X107)) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X102,fun(X102,bool)),X103),s(fun(sK2,fun(sK2,bool)),sK4),s('type.sum.sum'(X102,sK2),X108),s('type.sum.sum'(X102,sK2),X109)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X102,sK2),X108))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X102,sK2),X109))) | s(bool,'F_0') = s(bool,X107) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X102,fun(X102,bool)),X103),s(fun(X102,X104),X105),s(fun(X104,X102),X106)))) )),
  inference(superposition,[],[f180,f5187])).
fof(f12985,plain,(
  ( ! [X111,X116,X114,X112,X110,X117,X115,X113] : (~p(s(bool,'const.quotient_sum.+++_4'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(X110,fun(X110,bool)),X111),s('type.sum.sum'(sK2,X110),X116),s('type.sum.sum'(sK2,X110),X117)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,X110),X116))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,X110),X117))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X110,fun(X110,bool)),X111),s(fun(X110,X112),X113),s(fun(X112,X110),X114))) | s(bool,'F_0') = s(bool,X115)) )),
  inference(subsumption_resolution,[],[f12972,f119])).
fof(f12972,plain,(
  ( ! [X111,X116,X114,X112,X110,X117,X115,X113] : (~p(s(bool,X115)) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(X110,fun(X110,bool)),X111),s('type.sum.sum'(sK2,X110),X116),s('type.sum.sum'(sK2,X110),X117)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,X110),X116))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,X110),X117))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X110,fun(X110,bool)),X111),s(fun(X110,X112),X113),s(fun(X112,X110),X114))) | s(bool,'F_0') = s(bool,X115)) )),
  inference(superposition,[],[f179,f5187])).
fof(f12984,plain,(
  ( ! [X103,X109,X107,X105,X102,X108,X106,X104] : (~p(s(bool,'const.quotient_sum.+++_4'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(X102,fun(X102,bool)),X103),s('type.sum.sum'(sK2,X102),X108),s('type.sum.sum'(sK2,X102),X109)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,X102),X108))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,X102),X109))) | s(bool,'F_0') = s(bool,X107) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X102,fun(X102,bool)),X103),s(fun(X102,X104),X105),s(fun(X104,X102),X106)))) )),
  inference(subsumption_resolution,[],[f12971,f119])).
fof(f12971,plain,(
  ( ! [X103,X109,X107,X105,X102,X108,X106,X104] : (~p(s(bool,X107)) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(X102,fun(X102,bool)),X103),s('type.sum.sum'(sK2,X102),X108),s('type.sum.sum'(sK2,X102),X109)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,X102),X108))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,X102),X109))) | s(bool,'F_0') = s(bool,X107) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X102,fun(X102,bool)),X103),s(fun(X102,X104),X105),s(fun(X104,X102),X106)))) )),
  inference(superposition,[],[f179,f5187])).
fof(f12945,plain,(
  ( ! [X94,X99,X97,X95,X100,X98,X96] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X96,X96),X100))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X94,X94),'const.sum.++_3'(s(fun(X96,X94),X98),s(fun(X96,X94),X98),s('type.sum.sum'(X96,X96),X100))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X94,fun(X94,bool)),X95),s(fun(X94,X96),X97),s(fun(X96,X94),X98))) | s(bool,'F_0') = s(bool,X99)) )),
  inference(subsumption_resolution,[],[f12931,f119])).
fof(f12931,plain,(
  ( ! [X94,X99,X97,X95,X100,X98,X96] : (~p(s(bool,X99)) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X96,X96),X100))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X94,X94),'const.sum.++_3'(s(fun(X96,X94),X98),s(fun(X96,X94),X98),s('type.sum.sum'(X96,X96),X100))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X94,fun(X94,bool)),X95),s(fun(X94,X96),X97),s(fun(X96,X94),X98))) | s(bool,'F_0') = s(bool,X99)) )),
  inference(superposition,[],[f9432,f5187])).
fof(f12943,plain,(
  ( ! [X92,X90,X88,X87,X93,X91,X89] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X89,X89),X93))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X87,X87),'const.sum.++_3'(s(fun(X89,X87),X91),s(fun(X89,X87),X91),s('type.sum.sum'(X89,X89),X93))))) | s(bool,'F_0') = s(bool,X92) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X87,fun(X87,bool)),X88),s(fun(X87,X89),X90),s(fun(X89,X87),X91)))) )),
  inference(subsumption_resolution,[],[f12930,f119])).
fof(f12930,plain,(
  ( ! [X92,X90,X88,X87,X93,X91,X89] : (~p(s(bool,X92)) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X89,X89),X93))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X87,X87),'const.sum.++_3'(s(fun(X89,X87),X91),s(fun(X89,X87),X91),s('type.sum.sum'(X89,X89),X93))))) | s(bool,'F_0') = s(bool,X92) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X87,fun(X87,bool)),X88),s(fun(X87,X89),X90),s(fun(X89,X87),X91)))) )),
  inference(superposition,[],[f9432,f5187])).
fof(f12904,plain,(
  ( ! [X94,X99,X97,X95,X100,X98,X96] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X96,sK1),X100))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X94,sK3),'const.sum.++_3'(s(fun(X96,X94),X98),s(fun(sK1,sK3),sK9),s('type.sum.sum'(X96,sK1),X100))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X94,fun(X94,bool)),X95),s(fun(X94,X96),X97),s(fun(X96,X94),X98))) | s(bool,'F_0') = s(bool,X99)) )),
  inference(subsumption_resolution,[],[f12891,f119])).
fof(f12891,plain,(
  ( ! [X94,X99,X97,X95,X100,X98,X96] : (~p(s(bool,X99)) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X96,sK1),X100))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X94,sK3),'const.sum.++_3'(s(fun(X96,X94),X98),s(fun(sK1,sK3),sK9),s('type.sum.sum'(X96,sK1),X100))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X94,fun(X94,bool)),X95),s(fun(X94,X96),X97),s(fun(X96,X94),X98))) | s(bool,'F_0') = s(bool,X99)) )),
  inference(superposition,[],[f303,f5187])).
fof(f12903,plain,(
  ( ! [X92,X90,X88,X87,X93,X91,X89] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X89,sK1),X93))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X87,sK3),'const.sum.++_3'(s(fun(X89,X87),X91),s(fun(sK1,sK3),sK9),s('type.sum.sum'(X89,sK1),X93))))) | s(bool,'F_0') = s(bool,X92) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X87,fun(X87,bool)),X88),s(fun(X87,X89),X90),s(fun(X89,X87),X91)))) )),
  inference(subsumption_resolution,[],[f12890,f119])).
fof(f12890,plain,(
  ( ! [X92,X90,X88,X87,X93,X91,X89] : (~p(s(bool,X92)) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X89,sK1),X93))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X87,sK3),'const.sum.++_3'(s(fun(X89,X87),X91),s(fun(sK1,sK3),sK9),s('type.sum.sum'(X89,sK1),X93))))) | s(bool,'F_0') = s(bool,X92) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X87,fun(X87,bool)),X88),s(fun(X87,X89),X90),s(fun(X89,X87),X91)))) )),
  inference(superposition,[],[f303,f5187])).
fof(f12639,plain,(
  ( ! [X23] : (s(bool,X23) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,X23)) ) | $spl232),
  inference(cnf_transformation,[],[f12639_D])).
fof(f12639_D,plain,(
  ( ! [X23] : (s(bool,X23) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,X23)) ) <=> ~$spl232),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl232])])).
fof(f768,plain,(
  s(bool,'T_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) | $spl31),
  inference(cnf_transformation,[],[f768_D])).
fof(f768_D,plain,(
  s(bool,'T_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) <=> ~$spl31),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl31])])).
fof(f12620,plain,(
  ( ! [X32] : (s(bool,'T_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,X32) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,X32)) ) | $spl26),
  inference(superposition,[],[f6189,f429])).
fof(f6189,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') != s(bool,X11) | s(bool,X10) = s(bool,X11) | s(bool,'F_0') = s(bool,X10)) )),
  inference(resolution,[],[f5397,f122])).
fof(f5397,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | s(bool,'F_0') = s(bool,X13) | s(bool,X12) = s(bool,X13)) )),
  inference(subsumption_resolution,[],[f5191,f129])).
fof(f5191,plain,(
  ( ! [X12,X13,X11] : (s(bool,X12) = s(bool,X13) | s(bool,'F_0') = s(bool,X13) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X11)))))) | ~p(s(bool,X12))) )),
  inference(superposition,[],[f5092,f115])).
fof(f12640,plain,(
  ( ! [X25] : (s(bool,X25) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'F_0') = s(bool,X25)) ) | $spl234),
  inference(cnf_transformation,[],[f12640_D])).
fof(f12640_D,plain,(
  ( ! [X25] : (s(bool,X25) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'F_0') = s(bool,X25)) ) <=> ~$spl234),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl234])])).
fof(f12618,plain,(
  ( ! [X30] : (s(bool,X30) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,X30)) ) | $spl26),
  inference(superposition,[],[f5187,f429])).
fof(f12617,plain,(
  ( ! [X29] : (s(bool,X29) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'F_0') = s(bool,X29) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11)))) ) | $spl26),
  inference(superposition,[],[f5187,f429])).
fof(f12616,plain,(
  ( ! [X28] : (s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,X28) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,X28)) ) | $spl26),
  inference(superposition,[],[f5187,f429])).
fof(f12615,plain,(
  ( ! [X27] : (s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,X27) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,X27)) ) | $spl26),
  inference(superposition,[],[f5187,f429])).
fof(f771,plain,(
  s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) | $spl32),
  inference(cnf_transformation,[],[f771_D])).
fof(f771_D,plain,(
  s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f12607,plain,(
  s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) | $spl26),
  inference(superposition,[],[f126,f429])).
fof(f769,plain,(
  s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) | $spl30),
  inference(cnf_transformation,[],[f769_D])).
fof(f769_D,plain,(
  s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f12606,plain,(
  s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) | $spl26),
  inference(superposition,[],[f126,f429])).
fof(f12642,plain,(
  s('type.sum.sum'(sK2,sK3),sK11) = s('type.sum.sum'(sK2,sK3),'const.sum.INR_1'(s(sK3,sK13(sK11,sK3,sK2)))) | $spl236),
  inference(cnf_transformation,[],[f12642_D])).
fof(f12642_D,plain,(
  s('type.sum.sum'(sK2,sK3),sK11) = s('type.sum.sum'(sK2,sK3),'const.sum.INR_1'(s(sK3,sK13(sK11,sK3,sK2)))) <=> ~$spl236),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl236])])).
fof(f12175,plain,(
  p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10)))) | $spl224),
  inference(cnf_transformation,[],[f12175_D])).
fof(f12175_D,plain,(
  p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10)))) <=> ~$spl224),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl224])])).
fof(f12579,plain,(
  p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10)))) | s('type.sum.sum'(sK2,sK3),sK11) = s('type.sum.sum'(sK2,sK3),'const.sum.INR_1'(s(sK3,sK13(sK11,sK3,sK2)))) | $spl26),
  inference(superposition,[],[f9288,f429])).
fof(f12578,plain,(
  ( ! [X26] : (s(bool,X26) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,X26)) ) | $spl26),
  inference(superposition,[],[f429,f5187])).
fof(f12577,plain,(
  ( ! [X25] : (s(bool,X25) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'F_0') = s(bool,X25) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11)))) ) | $spl26),
  inference(superposition,[],[f429,f5187])).
fof(f12576,plain,(
  ( ! [X24] : (s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,X24) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,X24)) ) | $spl26),
  inference(superposition,[],[f429,f5187])).
fof(f12575,plain,(
  ( ! [X23] : (s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,X23) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,X23)) ) | $spl26),
  inference(superposition,[],[f429,f5187])).
fof(f12574,plain,(
  s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) | $spl26),
  inference(superposition,[],[f429,f126])).
fof(f12573,plain,(
  s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) | $spl26),
  inference(superposition,[],[f429,f126])).
fof(f12552,plain,(
  $false | ($spl4 | $spl8)),
  inference(trivial_inequality_removal,[],[f12524])).
fof(f12524,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl4 | $spl8)),
  inference(superposition,[],[f12488,f187])).
fof(f12488,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2)) ) | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f12487,f124])).
fof(f12487,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | ~p(s(bool,X2))) ) | ($spl4 | $spl8)),
  inference(forward_demodulation,[],[f170,f173])).
fof(f12551,plain,(
  $false | ($spl4 | $spl8)),
  inference(trivial_inequality_removal,[],[f12525])).
fof(f12525,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl4 | $spl8)),
  inference(superposition,[],[f12488,f308])).
fof(f12550,plain,(
  $false | ($spl4 | $spl8)),
  inference(trivial_inequality_removal,[],[f12526])).
fof(f12526,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl4 | $spl8)),
  inference(superposition,[],[f12488,f409])).
fof(f12549,plain,(
  $false | ($spl4 | $spl8)),
  inference(trivial_inequality_removal,[],[f12529])).
fof(f12529,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl4 | $spl8)),
  inference(superposition,[],[f12488,f9130])).
fof(f12548,plain,(
  $false | ($spl4 | $spl8)),
  inference(trivial_inequality_removal,[],[f12532])).
fof(f12532,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl4 | $spl8)),
  inference(superposition,[],[f12488,f3205])).
fof(f12547,plain,(
  $false | ($spl4 | $spl8)),
  inference(trivial_inequality_removal,[],[f12533])).
fof(f12533,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl4 | $spl8)),
  inference(superposition,[],[f12488,f3224])).
fof(f12546,plain,(
  $false | ($spl4 | $spl8)),
  inference(trivial_inequality_removal,[],[f12535])).
fof(f12535,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl4 | $spl8)),
  inference(superposition,[],[f12488,f173])).
fof(f12545,plain,(
  $false | ($spl4 | $spl8)),
  inference(trivial_inequality_removal,[],[f12540])).
fof(f12540,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl4 | $spl8)),
  inference(superposition,[],[f12488,f9229])).
fof(f12544,plain,(
  $false | ($spl4 | $spl8)),
  inference(trivial_inequality_removal,[],[f12541])).
fof(f12541,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl4 | $spl8)),
  inference(superposition,[],[f12488,f4197])).
fof(f12543,plain,(
  $false | ($spl4 | $spl8)),
  inference(equality_resolution,[],[f12488])).
fof(f12395,plain,(
  ( ! [X8] : (s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,X8)) ) | $spl8),
  inference(subsumption_resolution,[],[f12390,f126])).
fof(f12390,plain,(
  ( ! [X8] : (s(bool,'T_0') != s(bool,X8) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,X8)) ) | $spl8),
  inference(superposition,[],[f12281,f5187])).
fof(f12394,plain,(
  ( ! [X7] : (s(bool,'F_0') = s(bool,X7) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11)))) ) | $spl8),
  inference(subsumption_resolution,[],[f12389,f126])).
fof(f12389,plain,(
  ( ! [X7] : (s(bool,'T_0') != s(bool,X7) | s(bool,'F_0') = s(bool,X7) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11)))) ) | $spl8),
  inference(superposition,[],[f12281,f5187])).
fof(f12279,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))))) ) | $spl7),
  inference(subsumption_resolution,[],[f12260,f118])).
fof(f12260,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10)))) | p(s(bool,X1))) ) | $spl7),
  inference(superposition,[],[f172,f112])).
fof(f12278,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))))) ) | $spl7),
  inference(subsumption_resolution,[],[f12259,f118])).
fof(f12259,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,X0)) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))))) ) | $spl7),
  inference(superposition,[],[f172,f112])).
fof(f12258,plain,(
  ( ! [X8] : (s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,X8)) ) | $spl1),
  inference(subsumption_resolution,[],[f12249,f119])).
fof(f12249,plain,(
  ( ! [X8] : (~p(s(bool,X8)) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,X8)) ) | $spl1),
  inference(superposition,[],[f167,f5187])).
fof(f12257,plain,(
  ( ! [X7] : (s(bool,'F_0') = s(bool,X7) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11)))) ) | $spl1),
  inference(subsumption_resolution,[],[f12248,f119])).
fof(f12248,plain,(
  ( ! [X7] : (~p(s(bool,X7)) | s(bool,'F_0') = s(bool,X7) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11)))) ) | $spl1),
  inference(superposition,[],[f167,f5187])).
fof(f12222,plain,(
  ( ! [X8] : (s(bool,X8) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,X8)) )),
  inference(superposition,[],[f110,f5187])).
fof(f12221,plain,(
  ( ! [X7] : (s(bool,X7) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'F_0') = s(bool,X7) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11)))) )),
  inference(superposition,[],[f110,f5187])).
fof(f12223,plain,(
  ( ! [X5] : (s(bool,X5) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,X5)) ) | $spl228),
  inference(cnf_transformation,[],[f12223_D])).
fof(f12223_D,plain,(
  ( ! [X5] : (s(bool,X5) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,X5)) ) <=> ~$spl228),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl228])])).
fof(f172,plain,(
  s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | $spl7),
  inference(cnf_transformation,[],[f172_D])).
fof(f172_D,plain,(
  s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f12220,plain,(
  ( ! [X6] : (s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,X6) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,X6)) )),
  inference(superposition,[],[f110,f5187])).
fof(f12219,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,X5) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,X5)) )),
  inference(superposition,[],[f110,f5187])).
fof(f178,plain,(
  s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | $spl12),
  inference(cnf_transformation,[],[f178_D])).
fof(f178_D,plain,(
  s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f12218,plain,(
  s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11)))),
  inference(superposition,[],[f110,f126])).
fof(f12217,plain,(
  s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11)))),
  inference(superposition,[],[f110,f126])).
fof(f12216,plain,(
  s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))))),
  inference(superposition,[],[f110,f125])).
fof(f12215,plain,(
  s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))))),
  inference(superposition,[],[f110,f123])).
fof(f12214,plain,(
  s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))))),
  inference(superposition,[],[f110,f121])).
fof(f12213,plain,(
  s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))))),
  inference(superposition,[],[f110,f119])).
fof(f170,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | ~p(s(bool,X2))) ) | $spl4),
  inference(cnf_transformation,[],[f170_D])).
fof(f170_D,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | ~p(s(bool,X2))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f12212,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | ~p(s(bool,X3)) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))))) )),
  inference(superposition,[],[f110,f115])).
fof(f12211,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11)))) | ~p(s(bool,X2))) )),
  inference(superposition,[],[f110,f115])).
fof(f169,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | p(s(bool,X0))) ) | $spl2),
  inference(cnf_transformation,[],[f169_D])).
fof(f169_D,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | p(s(bool,X0))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f12210,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11)))) | p(s(bool,X1))) )),
  inference(superposition,[],[f110,f112])).
fof(f12209,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | p(s(bool,X0)) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))))) )),
  inference(superposition,[],[f110,f112])).
fof(f12183,plain,(
  $false | ($spl1 | $spl12)),
  inference(subsumption_resolution,[],[f12182,f156])).
fof(f12182,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl12)),
  inference(forward_demodulation,[],[f167,f178])).
fof(f426,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(sK2,X3),X4),s(fun(X3,sK2),X5))))) ) | $spl22),
  inference(cnf_transformation,[],[f426_D])).
fof(f426_D,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(sK2,X3),X4),s(fun(X3,sK2),X5))))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f427,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK3,fun(sK3,bool)),sK7),s(fun(sK3,X0),X1),s(fun(X0,sK3),X2))))) ) | $spl24),
  inference(cnf_transformation,[],[f427_D])).
fof(f427_D,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK3,fun(sK3,bool)),sK7),s(fun(sK3,X0),X1),s(fun(X0,sK3),X2))))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f429,plain,(
  s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) | $spl26),
  inference(cnf_transformation,[],[f429_D])).
fof(f429_D,plain,(
  s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f1384,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK3,fun(sK3,bool)),sK7),s(fun(sK3,X0),X1),s(fun(X0,sK3),X2)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(sK2,X3),X4),s(fun(X3,sK2),X5))))) )),
  inference(subsumption_resolution,[],[f1335,f141])).
fof(f1335,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'T_0')) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK3,fun(sK3,bool)),sK7),s(fun(sK3,X0),X1),s(fun(X0,sK3),X2)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(sK2,X3),X4),s(fun(X3,sK2),X5))))) )),
  inference(superposition,[],[f130,f409])).
fof(f10208,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK3,fun(sK3,bool)),sK7),s(fun(sK3,X0),X1),s(fun(X0,sK3),X2)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(sK2,X3),X4),s(fun(X3,sK2),X5))))) )),
  inference(subsumption_resolution,[],[f9979,f141])).
fof(f9979,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'T_0')) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK3,fun(sK3,bool)),sK7),s(fun(sK3,X0),X1),s(fun(X0,sK3),X2)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(sK2,X3),X4),s(fun(X3,sK2),X5))))) )),
  inference(superposition,[],[f130,f409])).
fof(f12177,plain,(
  s('type.sum.sum'(sK2,sK3),sK10) = s('type.sum.sum'(sK2,sK3),'const.sum.INR_1'(s(sK3,sK13('const.sum.INL_1'(s(sK2,sK12(sK10,sK3,sK2))),sK3,sK2)))) | $spl226),
  inference(cnf_transformation,[],[f12177_D])).
fof(f12177_D,plain,(
  s('type.sum.sum'(sK2,sK3),sK10) = s('type.sum.sum'(sK2,sK3),'const.sum.INR_1'(s(sK3,sK13('const.sum.INL_1'(s(sK2,sK12(sK10,sK3,sK2))),sK3,sK2)))) <=> ~$spl226),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl226])])).
fof(f12173,plain,(
  s('type.sum.sum'(sK2,sK3),sK10) = s('type.sum.sum'(sK2,sK3),'const.sum.INR_1'(s(sK3,sK13('const.sum.INL_1'(s(sK2,sK12(sK10,sK3,sK2))),sK3,sK2)))) | p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10)))) | $spl6),
  inference(forward_demodulation,[],[f10961,f10620])).
fof(f10961,plain,(
  p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10)))) | s('type.sum.sum'(sK2,sK3),'const.sum.INL_1'(s(sK2,sK12(sK10,sK3,sK2)))) = s('type.sum.sum'(sK2,sK3),'const.sum.INR_1'(s(sK3,sK13('const.sum.INL_1'(s(sK2,sK12(sK10,sK3,sK2))),sK3,sK2)))) | $spl6),
  inference(superposition,[],[f9288,f10620])).
fof(f12102,plain,(
  ( ! [X30] : (s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X30)))) ) | $spl222),
  inference(cnf_transformation,[],[f12102_D])).
fof(f12102_D,plain,(
  ( ! [X30] : (s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X30)))) ) <=> ~$spl222),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl222])])).
fof(f255,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl19),
  inference(cnf_transformation,[],[f255_D])).
fof(f255_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f12105,plain,(
  ( ! [X47] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X47)))) )),
  inference(subsumption_resolution,[],[f2001,f1886])).
fof(f2001,plain,(
  ( ! [X47] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X47)))) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X47)))) )),
  inference(superposition,[],[f122,f1934])).
fof(f1934,plain,(
  ( ! [X5] : (s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X5))) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X5)))) )),
  inference(resolution,[],[f1886,f119])).
fof(f12104,plain,(
  ( ! [X51] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X51)))) )),
  inference(subsumption_resolution,[],[f2005,f1886])).
fof(f2005,plain,(
  ( ! [X51] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X51)))) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X51)))) )),
  inference(superposition,[],[f124,f1934])).
fof(f2062,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X0)))) )),
  inference(duplicate_literal_removal,[],[f2046])).
fof(f2046,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X0))) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X0)))) )),
  inference(superposition,[],[f1936,f1934])).
fof(f1936,plain,(
  ( ! [X7] : (s(bool,'T_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X7))) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X7)))) )),
  inference(resolution,[],[f1886,f122])).
fof(f12103,plain,(
  ( ! [X29] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X29)))) )),
  inference(subsumption_resolution,[],[f2053,f1886])).
fof(f2053,plain,(
  ( ! [X29] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X29))) | p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X29))))) )),
  inference(superposition,[],[f1936,f119])).
fof(f12101,plain,(
  ( ! [X30] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X30)))) )),
  inference(subsumption_resolution,[],[f2054,f1886])).
fof(f2054,plain,(
  ( ! [X30] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X30))) | p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X30))))) )),
  inference(superposition,[],[f1936,f121])).
fof(f12100,plain,(
  ( ! [X34] : (s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X34))) | s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X34)))) ) | $spl220),
  inference(cnf_transformation,[],[f12100_D])).
fof(f12100_D,plain,(
  ( ! [X34] : (s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X34))) | s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X34)))) ) <=> ~$spl220),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl220])])).
fof(f2058,plain,(
  ( ! [X34] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X34))) | s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X34)))) )),
  inference(superposition,[],[f1936,f126])).
fof(f12097,plain,(
  ( ! [X19] : (s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X19)))) ) | $spl218),
  inference(cnf_transformation,[],[f12097_D])).
fof(f12097_D,plain,(
  ( ! [X19] : (s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X19)))) ) <=> ~$spl218),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl218])])).
fof(f12098,plain,(
  ( ! [X18] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X18)))) )),
  inference(subsumption_resolution,[],[f2281,f2288])).
fof(f2281,plain,(
  ( ! [X18] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X18))) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X18))))) )),
  inference(superposition,[],[f2232,f119])).
fof(f12096,plain,(
  ( ! [X19] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X19)))) )),
  inference(subsumption_resolution,[],[f2282,f2288])).
fof(f2282,plain,(
  ( ! [X19] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X19))) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X19))))) )),
  inference(superposition,[],[f2232,f121])).
fof(f12069,plain,(
  ( ! [X34] : (s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X34))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X34)))) ) | $spl206),
  inference(cnf_transformation,[],[f12069_D])).
fof(f12069_D,plain,(
  ( ! [X34] : (s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X34))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X34)))) ) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f2286,plain,(
  ( ! [X23] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X23))) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X23)))) )),
  inference(superposition,[],[f2232,f126])).
fof(f12054,plain,(
  ( ! [X5] : (s(bool,'T_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X5)))) ) | $spl202),
  inference(cnf_transformation,[],[f12054_D])).
fof(f12054_D,plain,(
  ( ! [X5] : (s(bool,'T_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X5)))) ) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f12087,plain,(
  ( ! [X18] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X18)))) )),
  inference(subsumption_resolution,[],[f2591,f2598])).
fof(f2591,plain,(
  ( ! [X18] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X18))) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X18))))) )),
  inference(superposition,[],[f2234,f119])).
fof(f12086,plain,(
  ( ! [X19] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X19)))) )),
  inference(subsumption_resolution,[],[f2592,f2598])).
fof(f2592,plain,(
  ( ! [X19] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X19))) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X19))))) )),
  inference(superposition,[],[f2234,f121])).
fof(f12085,plain,(
  ( ! [X23] : (s(bool,'T_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X23))) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X23)))) ) | $spl216),
  inference(cnf_transformation,[],[f12085_D])).
fof(f12085_D,plain,(
  ( ! [X23] : (s(bool,'T_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X23))) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X23)))) ) <=> ~$spl216),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl216])])).
fof(f2596,plain,(
  ( ! [X23] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X23))) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X23)))) )),
  inference(superposition,[],[f2234,f126])).
fof(f12084,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X1))) )),
  inference(subsumption_resolution,[],[f12083,f122])).
fof(f12083,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X1))) )),
  inference(forward_demodulation,[],[f2742,f4197])).
fof(f2742,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X1) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,X0))))) | p(s(bool,X1))) )),
  inference(superposition,[],[f2230,f1580])).
fof(f2230,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X1))) | s(bool,X2) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X1))) | p(s(bool,X2))) )),
  inference(resolution,[],[f2215,f112])).
fof(f12082,plain,(
  ( ! [X26,X27] : (s(bool,X27) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X26))) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X26)))) | p(s(bool,X27))) ) | $spl214),
  inference(cnf_transformation,[],[f12082_D])).
fof(f12082_D,plain,(
  ( ! [X26,X27] : (s(bool,X27) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X26))) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X26)))) | p(s(bool,X27))) ) <=> ~$spl214),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl214])])).
fof(f2748,plain,(
  ( ! [X24,X25] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X25) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X24))) | p(s(bool,X25)) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X24))))) )),
  inference(superposition,[],[f2230,f119])).
fof(f2749,plain,(
  ( ! [X26,X27] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X27) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X26))) | p(s(bool,X27)) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X26))))) )),
  inference(superposition,[],[f2230,f121])).
fof(f12035,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X0))) | s(bool,X1) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X0)))) ) | $spl192),
  inference(cnf_transformation,[],[f12035_D])).
fof(f12035_D,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X0))) | s(bool,X1) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X0)))) ) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f2753,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X35) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X34))) | p(s(bool,X35)) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X34)))) )),
  inference(superposition,[],[f2230,f126])).
fof(f12081,plain,(
  ( ! [X0,X1] : (s(bool,X1) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X0))) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X0)))) ) | $spl212),
  inference(cnf_transformation,[],[f12081_D])).
fof(f12081_D,plain,(
  ( ! [X0,X1] : (s(bool,X1) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X0))) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X0)))) ) <=> ~$spl212),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl212])])).
fof(f12080,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X1) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X0))) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X0)))) )),
  inference(subsumption_resolution,[],[f3009,f5397])).
fof(f3009,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X1) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X0))) | p(s(bool,X1)) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X0)))) )),
  inference(superposition,[],[f2606,f1934])).
fof(f2606,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X1))) | s(bool,X2) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X1))) | p(s(bool,X2))) )),
  inference(resolution,[],[f2598,f112])).
fof(f12079,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X3))) )),
  inference(subsumption_resolution,[],[f12078,f122])).
fof(f12078,plain,(
  ( ! [X3] : (s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X3))) )),
  inference(forward_demodulation,[],[f3010,f3205])).
fof(f3010,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X3) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X2))))) | p(s(bool,X3))) )),
  inference(superposition,[],[f2606,f2351])).
fof(f12077,plain,(
  ( ! [X41,X40] : (s(bool,X41) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X40))) | p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X40)))) | p(s(bool,X41))) ) | $spl210),
  inference(cnf_transformation,[],[f12077_D])).
fof(f12077_D,plain,(
  ( ! [X41,X40] : (s(bool,X41) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X40))) | p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X40)))) | p(s(bool,X41))) ) <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f3017,plain,(
  ( ! [X39,X38] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X39) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X38))) | p(s(bool,X39)) | p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X38))))) )),
  inference(superposition,[],[f2606,f119])).
fof(f3018,plain,(
  ( ! [X41,X40] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X41) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X40))) | p(s(bool,X41)) | p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X40))))) )),
  inference(superposition,[],[f2606,f121])).
fof(f12076,plain,(
  ( ! [X48,X49] : (s(bool,X49) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X48))) | s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X48))) | p(s(bool,X49))) ) | $spl208),
  inference(cnf_transformation,[],[f12076_D])).
fof(f12076_D,plain,(
  ( ! [X48,X49] : (s(bool,X49) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X48))) | s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X48))) | p(s(bool,X49))) ) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f3022,plain,(
  ( ! [X48,X49] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X49) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X48))) | p(s(bool,X49)) | s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X48)))) )),
  inference(superposition,[],[f2606,f126])).
fof(f11983,plain,(
  ( ! [X280] : (s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X280)))) ) | $spl170),
  inference(cnf_transformation,[],[f11983_D])).
fof(f11983_D,plain,(
  ( ! [X280] : (s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X280)))) ) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f12071,plain,(
  ( ! [X32] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X32)))) )),
  inference(subsumption_resolution,[],[f3600,f3517])).
fof(f3600,plain,(
  ( ! [X32] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X32))) | ~p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X32))))) )),
  inference(superposition,[],[f3541,f123])).
fof(f12070,plain,(
  ( ! [X33] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X33)))) )),
  inference(subsumption_resolution,[],[f3601,f3517])).
fof(f3601,plain,(
  ( ! [X33] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X33))) | ~p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X33))))) )),
  inference(superposition,[],[f3541,f125])).
fof(f3602,plain,(
  ( ! [X34] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X34))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X34)))) )),
  inference(superposition,[],[f3541,f126])).
fof(f3693,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X9)))) )),
  inference(duplicate_literal_removal,[],[f3643])).
fof(f3643,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X9))) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X9)))) )),
  inference(superposition,[],[f3541,f3543])).
fof(f3543,plain,(
  ( ! [X11] : (s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X11))) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X11)))) )),
  inference(resolution,[],[f3517,f123])).
fof(f12068,plain,(
  ( ! [X50] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X50)))) )),
  inference(subsumption_resolution,[],[f3658,f3517])).
fof(f3658,plain,(
  ( ! [X50] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X50)))) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X50)))) )),
  inference(superposition,[],[f118,f3543])).
fof(f12067,plain,(
  ( ! [X54] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X54)))) )),
  inference(subsumption_resolution,[],[f3662,f3517])).
fof(f3662,plain,(
  ( ! [X54] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X54)))) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X54)))) )),
  inference(superposition,[],[f120,f3543])).
fof(f12046,plain,(
  ( ! [X23] : (s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X23))) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X23)))) ) | $spl196),
  inference(cnf_transformation,[],[f12046_D])).
fof(f12046_D,plain,(
  ( ! [X23] : (s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X23))) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X23)))) ) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f3719,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X1))) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X1)))) )),
  inference(superposition,[],[f3609,f1934])).
fof(f11985,plain,(
  ( ! [X218] : (s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X218)))) ) | $spl174),
  inference(cnf_transformation,[],[f11985_D])).
fof(f11985_D,plain,(
  ( ! [X218] : (s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X218)))) ) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f12065,plain,(
  ( ! [X31] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X31)))) )),
  inference(subsumption_resolution,[],[f3727,f118])).
fof(f3727,plain,(
  ( ! [X31] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X31))) | p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X31))))) )),
  inference(superposition,[],[f3609,f119])).
fof(f12064,plain,(
  ( ! [X32] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X32)))) )),
  inference(subsumption_resolution,[],[f3728,f118])).
fof(f3728,plain,(
  ( ! [X32] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X32))) | p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X32))))) )),
  inference(superposition,[],[f3609,f121])).
fof(f12063,plain,(
  ( ! [X36] : (s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X36)))) ) | $spl204),
  inference(cnf_transformation,[],[f12063_D])).
fof(f12063_D,plain,(
  ( ! [X36] : (s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X36)))) ) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f12062,plain,(
  ( ! [X36] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X36)))) )),
  inference(subsumption_resolution,[],[f3732,f126])).
fof(f3732,plain,(
  ( ! [X36] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X36))) | s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X36)))) )),
  inference(superposition,[],[f3609,f126])).
fof(f4330,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X4)))) )),
  inference(forward_demodulation,[],[f4329,f4197])).
fof(f4329,plain,(
  ( ! [X4] : (s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,sK12(X4,sK0,sK0)))))) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X4)))) )),
  inference(subsumption_resolution,[],[f4302,f126])).
fof(f4302,plain,(
  ( ! [X4] : (s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X4))) | s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,sK12(X4,sK0,sK0)))))) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X4)))) )),
  inference(superposition,[],[f4134,f3227])).
fof(f4332,plain,(
  ( ! [X5] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X5)))) )),
  inference(forward_demodulation,[],[f4331,f4197])).
fof(f4331,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,sK12(X5,sK0,sK0)))))) | s(bool,'T_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X5)))) )),
  inference(subsumption_resolution,[],[f4303,f1936])).
fof(f4303,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X5))) | s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,sK12(X5,sK0,sK0)))))) | s(bool,'T_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X5)))) )),
  inference(superposition,[],[f4134,f2467])).
fof(f12053,plain,(
  ( ! [X6] : (s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X6)))) ) | $spl200),
  inference(cnf_transformation,[],[f12053_D])).
fof(f12053_D,plain,(
  ( ! [X6] : (s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X6)))) ) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f4334,plain,(
  ( ! [X6] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X6)))) )),
  inference(forward_demodulation,[],[f4333,f4197])).
fof(f4333,plain,(
  ( ! [X6] : (s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X6))) | s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,sK12(X6,sK0,sK0))))))) )),
  inference(subsumption_resolution,[],[f4304,f4134])).
fof(f4304,plain,(
  ( ! [X6] : (s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X6))) | s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,sK12(X6,sK0,sK0)))))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X6)))) )),
  inference(superposition,[],[f4134,f2393])).
fof(f2393,plain,(
  ( ! [X2] : (s('type.sum.sum'(sK0,sK0),X2) = s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,sK12(X2,sK0,sK0)))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X2)))) )),
  inference(superposition,[],[f2351,f111])).
fof(f12052,plain,(
  ( ! [X13] : (s('type.sum.sum'(sK0,sK0),X13) = s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,sK13(X13,sK0,sK0))))) ) | $spl198),
  inference(cnf_transformation,[],[f12052_D])).
fof(f12052_D,plain,(
  ( ! [X13] : (s('type.sum.sum'(sK0,sK0),X13) = s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,sK13(X13,sK0,sK0))))) ) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f4337,plain,(
  ( ! [X13] : (s(bool,'T_0') != s(bool,'F_0') | s('type.sum.sum'(sK0,sK0),X13) = s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,sK13(X13,sK0,sK0))))) )),
  inference(forward_demodulation,[],[f4336,f4197])).
fof(f4336,plain,(
  ( ! [X13] : (s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,sK12(X13,sK0,sK0)))))) | s('type.sum.sum'(sK0,sK0),X13) = s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,sK13(X13,sK0,sK0))))) )),
  inference(subsumption_resolution,[],[f4306,f1600])).
fof(f4306,plain,(
  ( ! [X13] : (s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X13))) | s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,sK12(X13,sK0,sK0)))))) | s('type.sum.sum'(sK0,sK0),X13) = s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,sK13(X13,sK0,sK0))))) )),
  inference(superposition,[],[f4134,f111])).
fof(f12048,plain,(
  ( ! [X21] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X21)))) )),
  inference(subsumption_resolution,[],[f4318,f4105])).
fof(f4318,plain,(
  ( ! [X21] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X21))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X21))))) )),
  inference(superposition,[],[f4134,f123])).
fof(f12047,plain,(
  ( ! [X22] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X22)))) )),
  inference(subsumption_resolution,[],[f4319,f4105])).
fof(f4319,plain,(
  ( ! [X22] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X22))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X22))))) )),
  inference(superposition,[],[f4134,f125])).
fof(f4320,plain,(
  ( ! [X23] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X23))) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X23)))) )),
  inference(superposition,[],[f4134,f126])).
fof(f11984,plain,(
  ( ! [X220] : (s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X220)))) ) | $spl172),
  inference(cnf_transformation,[],[f11984_D])).
fof(f11984_D,plain,(
  ( ! [X220] : (s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X220)))) ) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f4380,plain,(
  ( ! [X21] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X21)))) )),
  inference(subsumption_resolution,[],[f4365,f122])).
fof(f4365,plain,(
  ( ! [X21] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X21))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X21))))) )),
  inference(superposition,[],[f4326,f123])).
fof(f4381,plain,(
  ( ! [X22] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X22)))) )),
  inference(subsumption_resolution,[],[f4366,f122])).
fof(f4366,plain,(
  ( ! [X22] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X22))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X22))))) )),
  inference(superposition,[],[f4326,f125])).
fof(f12042,plain,(
  ( ! [X23] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X23)))) )),
  inference(subsumption_resolution,[],[f4367,f4326])).
fof(f4367,plain,(
  ( ! [X23] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X23))) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X23)))) )),
  inference(superposition,[],[f4326,f126])).
fof(f12041,plain,(
  ( ! [X6] : (p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X6))))) ) | $spl194),
  inference(cnf_transformation,[],[f12041_D])).
fof(f12041_D,plain,(
  ( ! [X6] : (p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X6))))) ) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f4489,plain,(
  ( ! [X6] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X6))))) )),
  inference(forward_demodulation,[],[f4170,f4197])).
fof(f4170,plain,(
  ( ! [X6] : (p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X6)))) | s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,sK12(X6,sK0,sK0))))))) )),
  inference(subsumption_resolution,[],[f4148,f4105])).
fof(f4148,plain,(
  ( ! [X6] : (p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X6)))) | s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,sK12(X6,sK0,sK0)))))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X6)))) )),
  inference(superposition,[],[f4105,f2393])).
fof(f12040,plain,(
  ( ! [X15] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X15))) )),
  inference(subsumption_resolution,[],[f12039,f122])).
fof(f12039,plain,(
  ( ! [X15] : (s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X15))) )),
  inference(forward_demodulation,[],[f4531,f4197])).
fof(f4531,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X15)) | s(bool,X15) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,sK12('const.sum.INL_1'(s(sK0,X14)),sK0,sK0))))))) )),
  inference(forward_demodulation,[],[f4497,f4197])).
fof(f4497,plain,(
  ( ! [X14,X15] : (s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,X14))))) | p(s(bool,X15)) | s(bool,X15) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,sK12('const.sum.INL_1'(s(sK0,X14)),sK0,sK0))))))) )),
  inference(superposition,[],[f4128,f4385])).
fof(f4385,plain,(
  ( ! [X4] : (s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,X4))) = s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,sK12('const.sum.INL_1'(s(sK0,X4)),sK0,sK0))))) )),
  inference(resolution,[],[f4171,f128])).
fof(f4171,plain,(
  ( ! [X7] : (p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X7)))) | s('type.sum.sum'(sK0,sK0),X7) = s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,sK12(X7,sK0,sK0))))) )),
  inference(subsumption_resolution,[],[f4149,f2351])).
fof(f4149,plain,(
  ( ! [X7] : (p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X7)))) | s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,sK13(X7,sK0,sK0)))))) | s('type.sum.sum'(sK0,sK0),X7) = s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,sK12(X7,sK0,sK0))))) )),
  inference(superposition,[],[f4105,f111])).
fof(f4128,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X0))) | p(s(bool,X1)) | s(bool,X1) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X0)))) )),
  inference(resolution,[],[f4105,f2755])).
fof(f2755,plain,(
  ( ! [X30,X31] : (~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X30)))) | p(s(bool,X31)) | s(bool,X31) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X30)))) )),
  inference(trivial_inequality_removal,[],[f2751])).
fof(f2751,plain,(
  ( ! [X30,X31] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X31) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X30))) | p(s(bool,X31)) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X30))))) )),
  inference(superposition,[],[f2230,f125])).
fof(f12038,plain,(
  ( ! [X35] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X35))) )),
  inference(subsumption_resolution,[],[f12037,f122])).
fof(f12037,plain,(
  ( ! [X35] : (s(bool,'T_0') = s(bool,X35) | s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X35))) )),
  inference(forward_demodulation,[],[f12036,f4197])).
fof(f12036,plain,(
  ( ! [X30,X35,X31,X34] : (s(bool,X35) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,happ(s(fun(X30,sK0),X31),s(X30,X34))))))) | s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X35))) )),
  inference(forward_demodulation,[],[f4534,f131])).
fof(f4534,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X35)) | s(bool,X35) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.++_3'(s(fun(X30,sK0),X31),s(fun(X32,sK0),X33),s('type.sum.sum'(X30,X32),'const.sum.INL_1'(s(X30,X34)))))))) )),
  inference(forward_demodulation,[],[f4500,f4197])).
fof(f4500,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,happ(s(fun(X30,sK0),X31),s(X30,X34))))))) | p(s(bool,X35)) | s(bool,X35) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.++_3'(s(fun(X30,sK0),X31),s(fun(X32,sK0),X33),s('type.sum.sum'(X30,X32),'const.sum.INL_1'(s(X30,X34)))))))) )),
  inference(superposition,[],[f4128,f131])).
fof(f4502,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X1)) | s(bool,X1) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X0))) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X0)))) )),
  inference(superposition,[],[f4128,f3543])).
fof(f12034,plain,(
  ( ! [X7] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X7))) )),
  inference(subsumption_resolution,[],[f12033,f122])).
fof(f12033,plain,(
  ( ! [X7] : (s(bool,'T_0') = s(bool,X7) | s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X7))) )),
  inference(forward_demodulation,[],[f4505,f4197])).
fof(f4505,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X7)) | s(bool,X7) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,X6)))))) )),
  inference(superposition,[],[f4128,f4197])).
fof(f12019,plain,(
  ( ! [X409,X411] : (p(s(bool,X411)) | s(bool,X411) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X409)))) ) | $spl184),
  inference(cnf_transformation,[],[f12019_D])).
fof(f12019_D,plain,(
  ( ! [X409,X411] : (p(s(bool,X411)) | s(bool,X411) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X409)))) ) <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f12032,plain,(
  ( ! [X47,X46] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X47)) | s(bool,X47) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X46)))) )),
  inference(subsumption_resolution,[],[f4514,f112])).
fof(f4514,plain,(
  ( ! [X47,X46] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X47)) | s(bool,X47) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X46))) | ~p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X46))))) )),
  inference(superposition,[],[f4128,f123])).
fof(f12031,plain,(
  ( ! [X48,X49] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X49)) | s(bool,X49) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X48)))) )),
  inference(subsumption_resolution,[],[f4515,f112])).
fof(f4515,plain,(
  ( ! [X48,X49] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X49)) | s(bool,X49) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X48))) | ~p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X48))))) )),
  inference(superposition,[],[f4128,f125])).
fof(f11982,plain,(
  ( ! [X285,X284] : (s(bool,X285) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X284))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X284)))) ) | $spl168),
  inference(cnf_transformation,[],[f11982_D])).
fof(f11982_D,plain,(
  ( ! [X285,X284] : (s(bool,X285) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X284))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X284)))) ) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f12030,plain,(
  ( ! [X50,X51] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X51) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X50))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X50)))) )),
  inference(subsumption_resolution,[],[f4516,f5397])).
fof(f4516,plain,(
  ( ! [X50,X51] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X51)) | s(bool,X51) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X50))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X50)))) )),
  inference(superposition,[],[f4128,f126])).
fof(f12029,plain,(
  ( ! [X15] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X15))) )),
  inference(subsumption_resolution,[],[f12028,f118])).
fof(f12028,plain,(
  ( ! [X15] : (s(bool,'F_0') = s(bool,X15) | s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X15))) )),
  inference(forward_demodulation,[],[f4582,f1580])).
fof(f4582,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X15) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,sK12('const.sum.INL_1'(s(sK0,X14)),sK0,sK0)))))) | ~p(s(bool,X15))) )),
  inference(forward_demodulation,[],[f4548,f4197])).
fof(f4548,plain,(
  ( ! [X14,X15] : (s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,X14))))) | s(bool,X15) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,sK12('const.sum.INL_1'(s(sK0,X14)),sK0,sK0)))))) | ~p(s(bool,X15))) )),
  inference(superposition,[],[f4132,f4385])).
fof(f12027,plain,(
  ( ! [X35] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X35))) )),
  inference(subsumption_resolution,[],[f12026,f118])).
fof(f12026,plain,(
  ( ! [X35] : (s(bool,'F_0') = s(bool,X35) | s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X35))) )),
  inference(forward_demodulation,[],[f12025,f1580])).
fof(f12025,plain,(
  ( ! [X30,X35,X31,X34] : (s(bool,X35) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,happ(s(fun(X30,sK0),X31),s(X30,X34))))))) | s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X35))) )),
  inference(forward_demodulation,[],[f4585,f131])).
fof(f4585,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X35) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.++_3'(s(fun(X30,sK0),X31),s(fun(X32,sK0),X33),s('type.sum.sum'(X30,X32),'const.sum.INL_1'(s(X30,X34))))))) | ~p(s(bool,X35))) )),
  inference(forward_demodulation,[],[f4551,f4197])).
fof(f4551,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,happ(s(fun(X30,sK0),X31),s(X30,X34))))))) | s(bool,X35) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.++_3'(s(fun(X30,sK0),X31),s(fun(X32,sK0),X33),s('type.sum.sum'(X30,X32),'const.sum.INL_1'(s(X30,X34))))))) | ~p(s(bool,X35))) )),
  inference(superposition,[],[f4132,f131])).
fof(f12024,plain,(
  ( ! [X0,X1] : (s(bool,X1) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X0))) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X0))) | ~p(s(bool,X1))) ) | $spl190),
  inference(cnf_transformation,[],[f12024_D])).
fof(f12024_D,plain,(
  ( ! [X0,X1] : (s(bool,X1) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X0))) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X0))) | ~p(s(bool,X1))) ) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f4553,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X1) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X0))) | ~p(s(bool,X1)) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X0)))) )),
  inference(superposition,[],[f4132,f3543])).
fof(f12023,plain,(
  ( ! [X7] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X7))) )),
  inference(subsumption_resolution,[],[f12022,f118])).
fof(f12022,plain,(
  ( ! [X7] : (s(bool,'F_0') = s(bool,X7) | s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X7))) )),
  inference(forward_demodulation,[],[f4556,f1580])).
fof(f4556,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X7) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,X6))))) | ~p(s(bool,X7))) )),
  inference(superposition,[],[f4132,f4197])).
fof(f12021,plain,(
  ( ! [X48,X49] : (s(bool,X49) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X48))) | ~p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X48)))) | ~p(s(bool,X49))) ) | $spl188),
  inference(cnf_transformation,[],[f12021_D])).
fof(f12021_D,plain,(
  ( ! [X48,X49] : (s(bool,X49) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X48))) | ~p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X48)))) | ~p(s(bool,X49))) ) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f4565,plain,(
  ( ! [X47,X46] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X47) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X46))) | ~p(s(bool,X47)) | ~p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X46))))) )),
  inference(superposition,[],[f4132,f123])).
fof(f4566,plain,(
  ( ! [X48,X49] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X49) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X48))) | ~p(s(bool,X49)) | ~p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X48))))) )),
  inference(superposition,[],[f4132,f125])).
fof(f12020,plain,(
  ( ! [X50,X51] : (s(bool,X51) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X50))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X50))) | ~p(s(bool,X51))) ) | $spl186),
  inference(cnf_transformation,[],[f12020_D])).
fof(f12020_D,plain,(
  ( ! [X50,X51] : (s(bool,X51) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X50))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X50))) | ~p(s(bool,X51))) ) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f4567,plain,(
  ( ! [X50,X51] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X51) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X50))) | ~p(s(bool,X51)) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X50)))) )),
  inference(superposition,[],[f4132,f126])).
fof(f5406,plain,(
  ( ! [X24] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X24))) )),
  inference(forward_demodulation,[],[f5405,f3205])).
fof(f5405,plain,(
  ( ! [X24,X25] : (s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X25))))) | ~p(s(bool,X24))) )),
  inference(subsumption_resolution,[],[f5208,f118])).
fof(f5208,plain,(
  ( ! [X24,X25] : (s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X25))))) | ~p(s(bool,X24)) | s(bool,'F_0') = s(bool,X24)) )),
  inference(superposition,[],[f118,f5092])).
fof(f5408,plain,(
  ( ! [X32] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X32))) )),
  inference(forward_demodulation,[],[f5407,f3205])).
fof(f5407,plain,(
  ( ! [X33,X32] : (s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X33))))) | ~p(s(bool,X32))) )),
  inference(subsumption_resolution,[],[f5212,f118])).
fof(f5212,plain,(
  ( ! [X33,X32] : (s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X33))))) | ~p(s(bool,X32)) | s(bool,'F_0') = s(bool,X32)) )),
  inference(superposition,[],[f120,f5092])).
fof(f5435,plain,(
  ( ! [X324] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X324)))) )),
  inference(forward_demodulation,[],[f5434,f3205])).
fof(f5434,plain,(
  ( ! [X325,X324] : (s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X325))))) | s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X324)))) )),
  inference(subsumption_resolution,[],[f5273,f4134])).
fof(f5273,plain,(
  ( ! [X325,X324] : (s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X325))))) | s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X324))) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X324)))) )),
  inference(superposition,[],[f4134,f5092])).
fof(f5437,plain,(
  ( ! [X328] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X328)))) )),
  inference(forward_demodulation,[],[f5436,f3205])).
fof(f5436,plain,(
  ( ! [X329,X328] : (s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X329))))) | s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X328)))) )),
  inference(subsumption_resolution,[],[f5275,f4326])).
fof(f5275,plain,(
  ( ! [X329,X328] : (s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X329))))) | s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X328))) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X328)))) )),
  inference(superposition,[],[f4326,f5092])).
fof(f5446,plain,(
  ( ! [X401] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X401)))) )),
  inference(forward_demodulation,[],[f5445,f3205])).
fof(f5445,plain,(
  ( ! [X401,X402] : (s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X402))))) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X401)))) )),
  inference(subsumption_resolution,[],[f5295,f3541])).
fof(f5295,plain,(
  ( ! [X401,X402] : (s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X402))))) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X401))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X401)))) )),
  inference(superposition,[],[f3541,f5092])).
fof(f5448,plain,(
  ( ! [X409,X411] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X411)) | s(bool,X411) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X409)))) )),
  inference(forward_demodulation,[],[f5447,f3205])).
fof(f5447,plain,(
  ( ! [X410,X409,X411] : (s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X410))))) | p(s(bool,X411)) | s(bool,X411) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X409)))) )),
  inference(subsumption_resolution,[],[f5299,f4128])).
fof(f5299,plain,(
  ( ! [X410,X409,X411] : (s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X410))))) | p(s(bool,X411)) | s(bool,X411) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X409))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X409)))) )),
  inference(superposition,[],[f4128,f5092])).
fof(f11981,plain,(
  ( ! [X287,X286] : (s(bool,X287) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X286))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X286)))) ) | $spl166),
  inference(cnf_transformation,[],[f11981_D])).
fof(f11981_D,plain,(
  ( ! [X287,X286] : (s(bool,X287) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X286))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X286)))) ) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f5450,plain,(
  ( ! [X412,X414] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X414) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X412))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X412)))) )),
  inference(forward_demodulation,[],[f5449,f3205])).
fof(f5449,plain,(
  ( ! [X412,X414,X413] : (s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X413))))) | s(bool,X414) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X412))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X412)))) )),
  inference(subsumption_resolution,[],[f5300,f2352])).
fof(f2352,plain,(
  ( ! [X2,X1] : (p(s(bool,X2)) | s(bool,X2) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X1))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X1)))) )),
  inference(resolution,[],[f2288,f112])).
fof(f5300,plain,(
  ( ! [X412,X414,X413] : (s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X413))))) | s(bool,X414) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X412))) | ~p(s(bool,X414)) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X412)))) )),
  inference(superposition,[],[f4132,f5092])).
fof(f5465,plain,(
  ( ! [X53] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X53))) )),
  inference(forward_demodulation,[],[f5464,f3205])).
fof(f5464,plain,(
  ( ! [X52,X53] : (~p(s(bool,X53)) | s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X52)))))) )),
  inference(subsumption_resolution,[],[f5334,f118])).
fof(f5334,plain,(
  ( ! [X52,X53] : (~p(s(bool,X53)) | s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X52))))) | s(bool,'F_0') = s(bool,X53)) )),
  inference(superposition,[],[f118,f5092])).
fof(f5467,plain,(
  ( ! [X61] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X61))) )),
  inference(forward_demodulation,[],[f5466,f3205])).
fof(f5466,plain,(
  ( ! [X61,X60] : (~p(s(bool,X61)) | s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X60)))))) )),
  inference(subsumption_resolution,[],[f5338,f118])).
fof(f5338,plain,(
  ( ! [X61,X60] : (~p(s(bool,X61)) | s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X60))))) | s(bool,'F_0') = s(bool,X61)) )),
  inference(superposition,[],[f120,f5092])).
fof(f12018,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X1)))))) ) | $spl182),
  inference(cnf_transformation,[],[f12018_D])).
fof(f12018_D,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X1)))))) ) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f5477,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X0) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X1)))))) )),
  inference(forward_demodulation,[],[f5372,f3205])).
fof(f5372,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X1))))) | s(bool,X0) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X1)))))) )),
  inference(equality_factoring,[],[f5092])).
fof(f12015,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X3))) )),
  inference(subsumption_resolution,[],[f12014,f122])).
fof(f12014,plain,(
  ( ! [X3] : (s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X3))) )),
  inference(forward_demodulation,[],[f5549,f4197])).
fof(f5549,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X3) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,X2))))) | p(s(bool,X3))) )),
  inference(superposition,[],[f2230,f5433])).
fof(f12003,plain,(
  ( ! [X31] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X31))) )),
  inference(subsumption_resolution,[],[f12002,f122])).
fof(f12002,plain,(
  ( ! [X31] : (s(bool,'T_0') = s(bool,X31) | s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X31))) )),
  inference(forward_demodulation,[],[f5732,f3205])).
fof(f5732,plain,(
  ( ! [X30,X31] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X31) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X30))))) | p(s(bool,X31))) )),
  inference(superposition,[],[f2606,f5444])).
fof(f11954,plain,(
  ( ! [X47,X45,X43,X48] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,X45),X48))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,X43),'const.sum.++_3'(s(fun(sK0,sK2),sK6),s(fun(X45,X43),X47),s('type.sum.sum'(sK0,X45),X48)))))) ) | $spl134),
  inference(cnf_transformation,[],[f11954_D])).
fof(f11954_D,plain,(
  ( ! [X47,X45,X43,X48] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,X45),X48))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,X43),'const.sum.++_3'(s(fun(sK0,sK2),sK6),s(fun(X45,X43),X47),s('type.sum.sum'(sK0,X45),X48)))))) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f11989,plain,(
  ( ! [X218,X216,X219,X214] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,X216),X219))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,X214),'const.sum.++_3'(s(fun(sK0,sK2),sK6),s(fun(X216,X214),X218),s('type.sum.sum'(sK0,X216),X219)))))) )),
  inference(subsumption_resolution,[],[f6126,f181])).
fof(f6126,plain,(
  ( ! [X218,X216,X215,X219,X217,X214] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X214,fun(X214,bool)),X215),s(fun(X214,X216),X217),s(fun(X216,X214),X218)))) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,X216),X219))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,X214),'const.sum.++_3'(s(fun(sK0,sK2),sK6),s(fun(X216,X214),X218),s('type.sum.sum'(sK0,X216),X219)))))) )),
  inference(superposition,[],[f122,f1409])).
fof(f1409,plain,(
  ( ! [X14,X19,X17,X15,X18,X16] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,X14),X15))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,X16),'const.sum.++_3'(s(fun(sK0,sK2),sK6),s(fun(X14,X16),X17),s('type.sum.sum'(sK0,X14),X15))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X16,fun(X16,bool)),X18),s(fun(X16,X14),X19),s(fun(X14,X16),X17)))) )),
  inference(resolution,[],[f181,f119])).
fof(f11988,plain,(
  ( ! [X243,X242,X240,X238] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,X240),X243))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,X238),'const.sum.++_3'(s(fun(sK0,sK2),sK6),s(fun(X240,X238),X242),s('type.sum.sum'(sK0,X240),X243)))))) )),
  inference(subsumption_resolution,[],[f6130,f181])).
fof(f6130,plain,(
  ( ! [X243,X241,X239,X242,X240,X238] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X238,fun(X238,bool)),X239),s(fun(X238,X240),X241),s(fun(X240,X238),X242)))) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,X240),X243))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,X238),'const.sum.++_3'(s(fun(sK0,sK2),sK6),s(fun(X240,X238),X242),s('type.sum.sum'(sK0,X240),X243)))))) )),
  inference(superposition,[],[f124,f1409])).
fof(f6559,plain,(
  ( ! [X16] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X16))) ) | $spl12),
  inference(forward_demodulation,[],[f6558,f178])).
fof(f6558,plain,(
  ( ! [X16] : (s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | ~p(s(bool,X16))) ) | $spl12),
  inference(subsumption_resolution,[],[f6349,f118])).
fof(f6349,plain,(
  ( ! [X16] : (s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | ~p(s(bool,X16)) | s(bool,'F_0') = s(bool,X16)) ) | $spl12),
  inference(superposition,[],[f118,f6282])).
fof(f6282,plain,(
  ( ! [X73] : (s(bool,X73) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,X73)) ) | $spl12),
  inference(subsumption_resolution,[],[f6245,f155])).
fof(f6245,plain,(
  ( ! [X73] : (~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,X73) | s(bool,X73) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11)))) ) | $spl12),
  inference(superposition,[],[f5397,f178])).
fof(f6562,plain,(
  ( ! [X20] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X20))) ) | $spl12),
  inference(forward_demodulation,[],[f6561,f178])).
fof(f6561,plain,(
  ( ! [X20] : (s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | ~p(s(bool,X20))) ) | $spl12),
  inference(subsumption_resolution,[],[f6353,f118])).
fof(f6353,plain,(
  ( ! [X20] : (s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | ~p(s(bool,X20)) | s(bool,'F_0') = s(bool,X20)) ) | $spl12),
  inference(superposition,[],[f120,f6282])).
fof(f6618,plain,(
  ( ! [X216] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X216)))) ) | $spl12),
  inference(forward_demodulation,[],[f6617,f178])).
fof(f6617,plain,(
  ( ! [X216] : (s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X216)))) ) | $spl12),
  inference(subsumption_resolution,[],[f6418,f4134])).
fof(f6418,plain,(
  ( ! [X216] : (s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X216))) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X216)))) ) | $spl12),
  inference(superposition,[],[f4134,f6282])).
fof(f6621,plain,(
  ( ! [X218] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X218)))) ) | $spl12),
  inference(forward_demodulation,[],[f6620,f178])).
fof(f6620,plain,(
  ( ! [X218] : (s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X218)))) ) | $spl12),
  inference(subsumption_resolution,[],[f6420,f4326])).
fof(f6420,plain,(
  ( ! [X218] : (s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X218))) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X218)))) ) | $spl12),
  inference(superposition,[],[f4326,f6282])).
fof(f6633,plain,(
  ( ! [X277] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X277)))) ) | $spl12),
  inference(forward_demodulation,[],[f6632,f178])).
fof(f6632,plain,(
  ( ! [X277] : (s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X277)))) ) | $spl12),
  inference(subsumption_resolution,[],[f6446,f3541])).
fof(f6446,plain,(
  ( ! [X277] : (s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X277))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X277)))) ) | $spl12),
  inference(superposition,[],[f3541,f6282])).
fof(f6637,plain,(
  ( ! [X282,X281] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X282) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X281))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X281)))) ) | $spl12),
  inference(forward_demodulation,[],[f6636,f178])).
fof(f6636,plain,(
  ( ! [X282,X281] : (s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,X282) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X281))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X281)))) ) | $spl12),
  inference(subsumption_resolution,[],[f6450,f5397])).
fof(f6450,plain,(
  ( ! [X282,X281] : (s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | p(s(bool,X282)) | s(bool,X282) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X281))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X281)))) ) | $spl12),
  inference(superposition,[],[f4128,f6282])).
fof(f6639,plain,(
  ( ! [X283,X284] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X284) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X283))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X283)))) ) | $spl12),
  inference(forward_demodulation,[],[f6638,f178])).
fof(f6638,plain,(
  ( ! [X283,X284] : (s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,X284) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X283))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X283)))) ) | $spl12),
  inference(subsumption_resolution,[],[f6451,f2352])).
fof(f6451,plain,(
  ( ! [X283,X284] : (s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,X284) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X283))) | ~p(s(bool,X284)) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X283)))) ) | $spl12),
  inference(superposition,[],[f4132,f6282])).
fof(f6652,plain,(
  ( ! [X18] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X18))) ) | $spl12),
  inference(forward_demodulation,[],[f6651,f178])).
fof(f6651,plain,(
  ( ! [X18] : (~p(s(bool,X18)) | s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11)))) ) | $spl12),
  inference(subsumption_resolution,[],[f6476,f118])).
fof(f6476,plain,(
  ( ! [X18] : (~p(s(bool,X18)) | s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,X18)) ) | $spl12),
  inference(superposition,[],[f118,f6282])).
fof(f6655,plain,(
  ( ! [X22] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X22))) ) | $spl12),
  inference(forward_demodulation,[],[f6654,f178])).
fof(f6654,plain,(
  ( ! [X22] : (~p(s(bool,X22)) | s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11)))) ) | $spl12),
  inference(subsumption_resolution,[],[f6480,f118])).
fof(f6480,plain,(
  ( ! [X22] : (~p(s(bool,X22)) | s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,X22)) ) | $spl12),
  inference(superposition,[],[f120,f6282])).
fof(f11987,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11)))) ) | $spl178),
  inference(cnf_transformation,[],[f11987_D])).
fof(f11987_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11)))) ) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f6670,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X0) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11)))) ) | $spl12),
  inference(forward_demodulation,[],[f6517,f178])).
fof(f6517,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,X0) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11)))) ) | $spl12),
  inference(equality_factoring,[],[f6282])).
fof(f6945,plain,(
  ( ! [X16] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X16))) ) | $spl32),
  inference(forward_demodulation,[],[f6944,f771])).
fof(f6944,plain,(
  ( ! [X16] : (s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) | ~p(s(bool,X16))) ) | $spl32),
  inference(subsumption_resolution,[],[f6730,f118])).
fof(f6730,plain,(
  ( ! [X16] : (s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) | ~p(s(bool,X16)) | s(bool,'F_0') = s(bool,X16)) ) | $spl32),
  inference(superposition,[],[f118,f6288])).
fof(f6288,plain,(
  ( ! [X123] : (s(bool,X123) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,X123)) ) | $spl32),
  inference(subsumption_resolution,[],[f6256,f155])).
fof(f6256,plain,(
  ( ! [X123] : (~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,X123) | s(bool,X123) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11)))) ) | $spl32),
  inference(superposition,[],[f5397,f771])).
fof(f6948,plain,(
  ( ! [X20] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X20))) ) | $spl32),
  inference(forward_demodulation,[],[f6947,f771])).
fof(f6947,plain,(
  ( ! [X20] : (s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) | ~p(s(bool,X20))) ) | $spl32),
  inference(subsumption_resolution,[],[f6734,f118])).
fof(f6734,plain,(
  ( ! [X20] : (s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) | ~p(s(bool,X20)) | s(bool,'F_0') = s(bool,X20)) ) | $spl32),
  inference(superposition,[],[f120,f6288])).
fof(f7005,plain,(
  ( ! [X217] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X217)))) ) | $spl32),
  inference(forward_demodulation,[],[f7004,f771])).
fof(f7004,plain,(
  ( ! [X217] : (s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X217)))) ) | $spl32),
  inference(subsumption_resolution,[],[f6800,f4134])).
fof(f6800,plain,(
  ( ! [X217] : (s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X217))) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X217)))) ) | $spl32),
  inference(superposition,[],[f4134,f6288])).
fof(f7008,plain,(
  ( ! [X219] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X219)))) ) | $spl32),
  inference(forward_demodulation,[],[f7007,f771])).
fof(f7007,plain,(
  ( ! [X219] : (s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X219)))) ) | $spl32),
  inference(subsumption_resolution,[],[f6802,f4326])).
fof(f6802,plain,(
  ( ! [X219] : (s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X219))) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X219)))) ) | $spl32),
  inference(superposition,[],[f4326,f6288])).
fof(f7020,plain,(
  ( ! [X278] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X278)))) ) | $spl32),
  inference(forward_demodulation,[],[f7019,f771])).
fof(f7019,plain,(
  ( ! [X278] : (s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X278)))) ) | $spl32),
  inference(subsumption_resolution,[],[f6826,f3541])).
fof(f6826,plain,(
  ( ! [X278] : (s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X278))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X278)))) ) | $spl32),
  inference(superposition,[],[f3541,f6288])).
fof(f7024,plain,(
  ( ! [X282,X283] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X283) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X282))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X282)))) ) | $spl32),
  inference(forward_demodulation,[],[f7023,f771])).
fof(f7023,plain,(
  ( ! [X282,X283] : (s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,X283) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X282))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X282)))) ) | $spl32),
  inference(subsumption_resolution,[],[f6830,f5397])).
fof(f6830,plain,(
  ( ! [X282,X283] : (s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) | p(s(bool,X283)) | s(bool,X283) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X282))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X282)))) ) | $spl32),
  inference(superposition,[],[f4128,f6288])).
fof(f7026,plain,(
  ( ! [X285,X284] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X285) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X284))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X284)))) ) | $spl32),
  inference(forward_demodulation,[],[f7025,f771])).
fof(f7025,plain,(
  ( ! [X285,X284] : (s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,X285) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X284))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X284)))) ) | $spl32),
  inference(subsumption_resolution,[],[f6831,f2352])).
fof(f6831,plain,(
  ( ! [X285,X284] : (s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,X285) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X284))) | ~p(s(bool,X285)) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X284)))) ) | $spl32),
  inference(superposition,[],[f4132,f6288])).
fof(f7039,plain,(
  ( ! [X38] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X38))) ) | $spl32),
  inference(forward_demodulation,[],[f7038,f771])).
fof(f7038,plain,(
  ( ! [X38] : (~p(s(bool,X38)) | s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11)))) ) | $spl32),
  inference(subsumption_resolution,[],[f6858,f118])).
fof(f6858,plain,(
  ( ! [X38] : (~p(s(bool,X38)) | s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,X38)) ) | $spl32),
  inference(superposition,[],[f118,f6288])).
fof(f7042,plain,(
  ( ! [X42] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X42))) ) | $spl32),
  inference(forward_demodulation,[],[f7041,f771])).
fof(f7041,plain,(
  ( ! [X42] : (~p(s(bool,X42)) | s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11)))) ) | $spl32),
  inference(subsumption_resolution,[],[f6862,f118])).
fof(f6862,plain,(
  ( ! [X42] : (~p(s(bool,X42)) | s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,X42)) ) | $spl32),
  inference(superposition,[],[f120,f6288])).
fof(f11986,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11)))) ) | $spl176),
  inference(cnf_transformation,[],[f11986_D])).
fof(f11986_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11)))) ) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f7057,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X0) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11)))) ) | $spl32),
  inference(forward_demodulation,[],[f6900,f771])).
fof(f6900,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,X0) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11)))) ) | $spl32),
  inference(equality_factoring,[],[f6288])).
fof(f7341,plain,(
  ( ! [X16] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X16))) ) | ($spl26 | $spl32)),
  inference(forward_demodulation,[],[f7340,f1035])).
fof(f7340,plain,(
  ( ! [X16] : (s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) | ~p(s(bool,X16))) ) | ($spl26 | $spl32)),
  inference(subsumption_resolution,[],[f7120,f118])).
fof(f7120,plain,(
  ( ! [X16] : (s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) | ~p(s(bool,X16)) | s(bool,'F_0') = s(bool,X16)) ) | ($spl26 | $spl32)),
  inference(superposition,[],[f118,f6289])).
fof(f6289,plain,(
  ( ! [X124] : (s(bool,X124) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'F_0') = s(bool,X124)) ) | ($spl26 | $spl32)),
  inference(subsumption_resolution,[],[f6257,f155])).
fof(f6257,plain,(
  ( ! [X124] : (~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,X124) | s(bool,X124) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10)))) ) | ($spl26 | $spl32)),
  inference(superposition,[],[f5397,f1035])).
fof(f7344,plain,(
  ( ! [X20] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X20))) ) | ($spl26 | $spl32)),
  inference(forward_demodulation,[],[f7343,f1035])).
fof(f7343,plain,(
  ( ! [X20] : (s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) | ~p(s(bool,X20))) ) | ($spl26 | $spl32)),
  inference(subsumption_resolution,[],[f7124,f118])).
fof(f7124,plain,(
  ( ! [X20] : (s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) | ~p(s(bool,X20)) | s(bool,'F_0') = s(bool,X20)) ) | ($spl26 | $spl32)),
  inference(superposition,[],[f120,f6289])).
fof(f7402,plain,(
  ( ! [X218] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X218)))) ) | ($spl26 | $spl32)),
  inference(forward_demodulation,[],[f7401,f1035])).
fof(f7401,plain,(
  ( ! [X218] : (s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X218)))) ) | ($spl26 | $spl32)),
  inference(subsumption_resolution,[],[f7191,f4134])).
fof(f7191,plain,(
  ( ! [X218] : (s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X218))) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X218)))) ) | ($spl26 | $spl32)),
  inference(superposition,[],[f4134,f6289])).
fof(f7405,plain,(
  ( ! [X220] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X220)))) ) | ($spl26 | $spl32)),
  inference(forward_demodulation,[],[f7404,f1035])).
fof(f7404,plain,(
  ( ! [X220] : (s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X220)))) ) | ($spl26 | $spl32)),
  inference(subsumption_resolution,[],[f7193,f4326])).
fof(f7193,plain,(
  ( ! [X220] : (s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X220))) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X220)))) ) | ($spl26 | $spl32)),
  inference(superposition,[],[f4326,f6289])).
fof(f7417,plain,(
  ( ! [X280] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X280)))) ) | ($spl26 | $spl32)),
  inference(forward_demodulation,[],[f7416,f1035])).
fof(f7416,plain,(
  ( ! [X280] : (s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X280)))) ) | ($spl26 | $spl32)),
  inference(subsumption_resolution,[],[f7219,f3541])).
fof(f7219,plain,(
  ( ! [X280] : (s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X280))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X280)))) ) | ($spl26 | $spl32)),
  inference(superposition,[],[f3541,f6289])).
fof(f7421,plain,(
  ( ! [X285,X284] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X285) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X284))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X284)))) ) | ($spl26 | $spl32)),
  inference(forward_demodulation,[],[f7420,f1035])).
fof(f7420,plain,(
  ( ! [X285,X284] : (s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,X285) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X284))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X284)))) ) | ($spl26 | $spl32)),
  inference(subsumption_resolution,[],[f7223,f5397])).
fof(f7223,plain,(
  ( ! [X285,X284] : (s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) | p(s(bool,X285)) | s(bool,X285) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X284))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X284)))) ) | ($spl26 | $spl32)),
  inference(superposition,[],[f4128,f6289])).
fof(f7423,plain,(
  ( ! [X287,X286] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X287) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X286))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X286)))) ) | ($spl26 | $spl32)),
  inference(forward_demodulation,[],[f7422,f1035])).
fof(f7422,plain,(
  ( ! [X287,X286] : (s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,X287) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X286))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X286)))) ) | ($spl26 | $spl32)),
  inference(subsumption_resolution,[],[f7224,f2352])).
fof(f7224,plain,(
  ( ! [X287,X286] : (s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,X287) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X286))) | ~p(s(bool,X287)) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X286)))) ) | ($spl26 | $spl32)),
  inference(superposition,[],[f4132,f6289])).
fof(f7436,plain,(
  ( ! [X38] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X38))) ) | ($spl26 | $spl32)),
  inference(forward_demodulation,[],[f7435,f1035])).
fof(f7435,plain,(
  ( ! [X38] : (~p(s(bool,X38)) | s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10)))) ) | ($spl26 | $spl32)),
  inference(subsumption_resolution,[],[f7251,f118])).
fof(f7251,plain,(
  ( ! [X38] : (~p(s(bool,X38)) | s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'F_0') = s(bool,X38)) ) | ($spl26 | $spl32)),
  inference(superposition,[],[f118,f6289])).
fof(f7439,plain,(
  ( ! [X42] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X42))) ) | ($spl26 | $spl32)),
  inference(forward_demodulation,[],[f7438,f1035])).
fof(f7438,plain,(
  ( ! [X42] : (~p(s(bool,X42)) | s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10)))) ) | ($spl26 | $spl32)),
  inference(subsumption_resolution,[],[f7255,f118])).
fof(f7255,plain,(
  ( ! [X42] : (~p(s(bool,X42)) | s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'F_0') = s(bool,X42)) ) | ($spl26 | $spl32)),
  inference(superposition,[],[f120,f6289])).
fof(f11979,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10)))) ) | $spl162),
  inference(cnf_transformation,[],[f11979_D])).
fof(f11979_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10)))) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f7454,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X0) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10)))) ) | ($spl26 | $spl32)),
  inference(forward_demodulation,[],[f7293,f1035])).
fof(f7293,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,X0) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10)))) ) | ($spl26 | $spl32)),
  inference(equality_factoring,[],[f6289])).
fof(f11968,plain,(
  ( ! [X250] : (s(bool,X250) = s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(sK2,sK0),sK5),s(fun(sK0,sK2),sK6))) | s(bool,'F_0') = s(bool,X250)) ) | $spl158),
  inference(cnf_transformation,[],[f11968_D])).
fof(f11968_D,plain,(
  ( ! [X250] : (s(bool,X250) = s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(sK2,sK0),sK5),s(fun(sK0,sK2),sK6))) | s(bool,'F_0') = s(bool,X250)) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f254,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl18),
  inference(cnf_transformation,[],[f254_D])).
fof(f254_D,plain,(
  s(bool,'T_0') = s(bool,'F_0') <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f7599,plain,(
  ( ! [X46] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X46) = s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(sK2,sK0),sK5),s(fun(sK0,sK2),sK6))) | s(bool,'F_0') = s(bool,X46)) )),
  inference(superposition,[],[f5187,f187])).
fof(f11967,plain,(
  ( ! [X251] : (s(bool,X251) = s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK3,fun(sK3,bool)),sK7),s(fun(sK3,sK1),sK8),s(fun(sK1,sK3),sK9))) | s(bool,'F_0') = s(bool,X251)) ) | $spl156),
  inference(cnf_transformation,[],[f11967_D])).
fof(f11967_D,plain,(
  ( ! [X251] : (s(bool,X251) = s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK3,fun(sK3,bool)),sK7),s(fun(sK3,sK1),sK8),s(fun(sK1,sK3),sK9))) | s(bool,'F_0') = s(bool,X251)) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f7600,plain,(
  ( ! [X47] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X47) = s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK3,fun(sK3,bool)),sK7),s(fun(sK3,sK1),sK8),s(fun(sK1,sK3),sK9))) | s(bool,'F_0') = s(bool,X47)) )),
  inference(superposition,[],[f5187,f308])).
fof(f11966,plain,(
  ( ! [X265] : (s(bool,X265) = s(bool,'const.quotient_sum.+++_4'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(sK3,fun(sK3,bool)),sK7),s('type.sum.sum'(sK2,sK3),sK10),s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,X265)) ) | $spl154),
  inference(cnf_transformation,[],[f11966_D])).
fof(f11966_D,plain,(
  ( ! [X265] : (s(bool,X265) = s(bool,'const.quotient_sum.+++_4'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(sK3,fun(sK3,bool)),sK7),s('type.sum.sum'(sK2,sK3),sK10),s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,X265)) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f7601,plain,(
  ( ! [X53] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X53) = s(bool,'const.quotient_sum.+++_4'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(sK3,fun(sK3,bool)),sK7),s('type.sum.sum'(sK2,sK3),sK10),s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,X53)) )),
  inference(superposition,[],[f5187,f409])).
fof(f11962,plain,(
  ( ! [X343,X342] : (s(bool,X343) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK2),'const.sum.INR_1'(s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,X342))))))) | s(bool,'F_0') = s(bool,X343)) ) | $spl146),
  inference(cnf_transformation,[],[f11962_D])).
fof(f11962_D,plain,(
  ( ! [X343,X342] : (s(bool,X343) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK2),'const.sum.INR_1'(s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,X342))))))) | s(bool,'F_0') = s(bool,X343)) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f7609,plain,(
  ( ! [X83,X84] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X84) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK2),'const.sum.INR_1'(s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,X83))))))) | s(bool,'F_0') = s(bool,X84)) )),
  inference(superposition,[],[f5187,f3224])).
fof(f8622,plain,(
  ( ! [X80,X82] : (s(bool,X80) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X82))) | s(bool,'F_0') = s(bool,X80) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X82)))) ) | $spl48),
  inference(cnf_transformation,[],[f8622_D])).
fof(f8622_D,plain,(
  ( ! [X80,X82] : (s(bool,X80) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X82))) | s(bool,'F_0') = s(bool,X80) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X82)))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f7619,plain,(
  ( ! [X132,X131] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X132) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X131))) | s(bool,'F_0') = s(bool,X132) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X131)))) )),
  inference(superposition,[],[f5187,f3543])).
fof(f11957,plain,(
  ( ! [X434,X433] : (s(bool,X434) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,X433))))) | s(bool,'F_0') = s(bool,X434)) ) | $spl138),
  inference(cnf_transformation,[],[f11957_D])).
fof(f11957_D,plain,(
  ( ! [X434,X433] : (s(bool,X434) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,X433))))) | s(bool,'F_0') = s(bool,X434)) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f7620,plain,(
  ( ! [X136,X135] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X136) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,X135))))) | s(bool,'F_0') = s(bool,X136)) )),
  inference(superposition,[],[f5187,f4197])).
fof(f7653,plain,(
  ( ! [X46] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X46) = s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(sK2,sK0),sK5),s(fun(sK0,sK2),sK6))) | s(bool,'F_0') = s(bool,X46)) )),
  inference(superposition,[],[f5187,f187])).
fof(f7654,plain,(
  ( ! [X47] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X47) = s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK3,fun(sK3,bool)),sK7),s(fun(sK3,sK1),sK8),s(fun(sK1,sK3),sK9))) | s(bool,'F_0') = s(bool,X47)) )),
  inference(superposition,[],[f5187,f308])).
fof(f7655,plain,(
  ( ! [X53] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X53) = s(bool,'const.quotient_sum.+++_4'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(sK3,fun(sK3,bool)),sK7),s('type.sum.sum'(sK2,sK3),sK10),s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,X53)) )),
  inference(superposition,[],[f5187,f409])).
fof(f7663,plain,(
  ( ! [X83,X84] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X84) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK2),'const.sum.INR_1'(s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,X83))))))) | s(bool,'F_0') = s(bool,X84)) )),
  inference(superposition,[],[f5187,f3224])).
fof(f7673,plain,(
  ( ! [X132,X131] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X132) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X131))) | s(bool,'F_0') = s(bool,X132) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X131)))) )),
  inference(superposition,[],[f5187,f3543])).
fof(f7674,plain,(
  ( ! [X136,X135] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X136) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,X135))))) | s(bool,'F_0') = s(bool,X136)) )),
  inference(superposition,[],[f5187,f4197])).
fof(f7721,plain,(
  ( ! [X111,X110] : (s(bool,X110) != s(bool,X111) | p(s(bool,X111)) | s(bool,'F_0') = s(bool,X110) | s(bool,'T_0') = s(bool,'F_0')) )),
  inference(superposition,[],[f122,f5187])).
fof(f7723,plain,(
  ( ! [X114,X115] : (s(bool,X114) != s(bool,X115) | p(s(bool,X115)) | s(bool,'F_0') = s(bool,X114) | s(bool,'T_0') = s(bool,'F_0')) )),
  inference(superposition,[],[f124,f5187])).
fof(f7886,plain,(
  ( ! [X111,X110] : (s(bool,X110) != s(bool,X111) | p(s(bool,X111)) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X110)) )),
  inference(superposition,[],[f122,f5187])).
fof(f7888,plain,(
  ( ! [X114,X115] : (s(bool,X114) != s(bool,X115) | p(s(bool,X115)) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X114)) )),
  inference(superposition,[],[f124,f5187])).
fof(f11923,plain,(
  ( ! [X23,X22] : (s(bool,X22) = s(bool,X23) | s(bool,'F_0') = s(bool,X23)) ) | $spl116),
  inference(cnf_transformation,[],[f11923_D])).
fof(f11923_D,plain,(
  ( ! [X23,X22] : (s(bool,X22) = s(bool,X23) | s(bool,'F_0') = s(bool,X23)) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f11977,plain,(
  ( ! [X87,X86] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X86) = s(bool,X87) | s(bool,'F_0') = s(bool,X87)) )),
  inference(subsumption_resolution,[],[f8035,f5397])).
fof(f8035,plain,(
  ( ! [X87,X86] : (p(s(bool,X86)) | s(bool,'T_0') != s(bool,'F_0') | s(bool,X86) = s(bool,X87) | s(bool,'F_0') = s(bool,X87)) )),
  inference(superposition,[],[f122,f5187])).
fof(f11976,plain,(
  ( ! [X94,X95] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X94) = s(bool,X95) | s(bool,'F_0') = s(bool,X95)) )),
  inference(subsumption_resolution,[],[f8039,f5397])).
fof(f8039,plain,(
  ( ! [X94,X95] : (p(s(bool,X94)) | s(bool,'T_0') != s(bool,'F_0') | s(bool,X94) = s(bool,X95) | s(bool,'F_0') = s(bool,X95)) )),
  inference(superposition,[],[f124,f5187])).
fof(f11975,plain,(
  ( ! [X39,X38] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X38) = s(bool,X39) | s(bool,'F_0') = s(bool,X39)) )),
  inference(subsumption_resolution,[],[f8081,f5397])).
fof(f8081,plain,(
  ( ! [X39,X38] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X38)) | s(bool,X38) = s(bool,X39) | s(bool,'F_0') = s(bool,X39)) )),
  inference(superposition,[],[f122,f5187])).
fof(f11974,plain,(
  ( ! [X47,X46] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X46) = s(bool,X47) | s(bool,'F_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f8085,f5397])).
fof(f8085,plain,(
  ( ! [X47,X46] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X46)) | s(bool,X46) = s(bool,X47) | s(bool,'F_0') = s(bool,X47)) )),
  inference(superposition,[],[f124,f5187])).
fof(f8137,plain,(
  ( ! [X250] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X250) = s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(sK2,sK0),sK5),s(fun(sK0,sK2),sK6))) | s(bool,'F_0') = s(bool,X250)) )),
  inference(superposition,[],[f187,f5187])).
fof(f8138,plain,(
  ( ! [X251] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X251) = s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK3,fun(sK3,bool)),sK7),s(fun(sK3,sK1),sK8),s(fun(sK1,sK3),sK9))) | s(bool,'F_0') = s(bool,X251)) )),
  inference(superposition,[],[f308,f5187])).
fof(f8140,plain,(
  ( ! [X265] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X265) = s(bool,'const.quotient_sum.+++_4'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(sK3,fun(sK3,bool)),sK7),s('type.sum.sum'(sK2,sK3),sK10),s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,X265)) )),
  inference(superposition,[],[f409,f5187])).
fof(f11920,plain,(
  ( ! [X68,X66,X69,X67] : (s(bool,X69) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X66,X67),'const.sum.INL_1'(s(X66,X68))))) | s(bool,'F_0') = s(bool,X69)) ) | $spl112),
  inference(cnf_transformation,[],[f11920_D])).
fof(f11920_D,plain,(
  ( ! [X68,X66,X69,X67] : (s(bool,X69) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X66,X67),'const.sum.INL_1'(s(X66,X68))))) | s(bool,'F_0') = s(bool,X69)) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f8147,plain,(
  ( ! [X294,X296,X295,X297] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X297) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X294,X295),'const.sum.INL_1'(s(X294,X296))))) | s(bool,'F_0') = s(bool,X297)) )),
  inference(superposition,[],[f5633,f5187])).
fof(f11919,plain,(
  ( ! [X87,X86] : (s(bool,X87) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,X86))))) | s(bool,'F_0') = s(bool,X87)) ) | $spl110),
  inference(cnf_transformation,[],[f11919_D])).
fof(f11919_D,plain,(
  ( ! [X87,X86] : (s(bool,X87) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,X86))))) | s(bool,'F_0') = s(bool,X87)) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f8148,plain,(
  ( ! [X302,X301] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X302) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,X301))))) | s(bool,'F_0') = s(bool,X302)) )),
  inference(superposition,[],[f1592,f5187])).
fof(f11965,plain,(
  ( ! [X304,X306,X305] : (s(bool,X305) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X304))) | s(bool,'F_0') = s(bool,X306) | s(bool,X306) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X304))) | p(s(bool,X305))) ) | $spl152),
  inference(cnf_transformation,[],[f11965_D])).
fof(f11965_D,plain,(
  ( ! [X304,X306,X305] : (s(bool,X305) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X304))) | s(bool,'F_0') = s(bool,X306) | s(bool,X306) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X304))) | p(s(bool,X305))) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f8149,plain,(
  ( ! [X304,X306,X305] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X305) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X304))) | p(s(bool,X305)) | s(bool,X306) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X304))) | s(bool,'F_0') = s(bool,X306)) )),
  inference(superposition,[],[f2230,f5187])).
fof(f11964,plain,(
  ( ! [X308,X307] : (s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X307))) | s(bool,'F_0') = s(bool,X308) | s(bool,X308) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X307)))) ) | $spl150),
  inference(cnf_transformation,[],[f11964_D])).
fof(f11964_D,plain,(
  ( ! [X308,X307] : (s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X307))) | s(bool,'F_0') = s(bool,X308) | s(bool,X308) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X307)))) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f8150,plain,(
  ( ! [X308,X307] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X307))) | s(bool,X308) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X307))) | s(bool,'F_0') = s(bool,X308)) )),
  inference(superposition,[],[f2232,f5187])).
fof(f11963,plain,(
  ( ! [X310,X309] : (s(bool,'T_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X309))) | s(bool,'F_0') = s(bool,X310) | s(bool,X310) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X309)))) ) | $spl148),
  inference(cnf_transformation,[],[f11963_D])).
fof(f11963_D,plain,(
  ( ! [X310,X309] : (s(bool,'T_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X309))) | s(bool,'F_0') = s(bool,X310) | s(bool,X310) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X309)))) ) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f8151,plain,(
  ( ! [X310,X309] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X309))) | s(bool,X310) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X309))) | s(bool,'F_0') = s(bool,X310)) )),
  inference(superposition,[],[f2234,f5187])).
fof(f11918,plain,(
  ( ! [X93] : (s(bool,X93) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'F_0') = s(bool,X93)) ) | $spl108),
  inference(cnf_transformation,[],[f11918_D])).
fof(f11918_D,plain,(
  ( ! [X93] : (s(bool,X93) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'F_0') = s(bool,X93)) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f8161,plain,(
  ( ! [X331] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X331) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'F_0') = s(bool,X331)) ) | $spl9),
  inference(superposition,[],[f174,f5187])).
fof(f11917,plain,(
  ( ! [X99,X100] : (s(bool,X100) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK2),'const.sum.INL_1'(s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,X99))))))) | s(bool,'F_0') = s(bool,X100)) ) | $spl106),
  inference(cnf_transformation,[],[f11917_D])).
fof(f11917_D,plain,(
  ( ! [X99,X100] : (s(bool,X100) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK2),'const.sum.INL_1'(s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,X99))))))) | s(bool,'F_0') = s(bool,X100)) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f8166,plain,(
  ( ! [X339,X340] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X340) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK2),'const.sum.INL_1'(s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,X339))))))) | s(bool,'F_0') = s(bool,X340)) )),
  inference(superposition,[],[f1593,f5187])).
fof(f8167,plain,(
  ( ! [X343,X342] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X343) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK2),'const.sum.INR_1'(s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,X342))))))) | s(bool,'F_0') = s(bool,X343)) )),
  inference(superposition,[],[f3224,f5187])).
fof(f11916,plain,(
  ( ! [X103,X101,X102,X104] : (s(bool,X104) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X101,X102),'const.sum.INR_1'(s(X102,X103))))) | s(bool,'F_0') = s(bool,X104)) ) | $spl104),
  inference(cnf_transformation,[],[f11916_D])).
fof(f11916_D,plain,(
  ( ! [X103,X101,X102,X104] : (s(bool,X104) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X101,X102),'const.sum.INR_1'(s(X102,X103))))) | s(bool,'F_0') = s(bool,X104)) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f8174,plain,(
  ( ! [X385,X387,X388,X386] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X388) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X385,X386),'const.sum.INR_1'(s(X386,X387))))) | s(bool,'F_0') = s(bool,X388)) )),
  inference(superposition,[],[f5815,f5187])).
fof(f11961,plain,(
  ( ! [X401,X402] : (s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X401))) | s(bool,'F_0') = s(bool,X402) | s(bool,X402) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X401)))) ) | $spl144),
  inference(cnf_transformation,[],[f11961_D])).
fof(f11961_D,plain,(
  ( ! [X401,X402] : (s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X401))) | s(bool,'F_0') = s(bool,X402) | s(bool,X402) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X401)))) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f8180,plain,(
  ( ! [X401,X402] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X401))) | s(bool,X402) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X401))) | s(bool,'F_0') = s(bool,X402)) )),
  inference(superposition,[],[f1936,f5187])).
fof(f11960,plain,(
  ( ! [X405,X407,X406] : (s(bool,X406) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X405))) | s(bool,'F_0') = s(bool,X407) | s(bool,X407) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X405))) | p(s(bool,X406))) ) | $spl142),
  inference(cnf_transformation,[],[f11960_D])).
fof(f11960_D,plain,(
  ( ! [X405,X407,X406] : (s(bool,X406) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X405))) | s(bool,'F_0') = s(bool,X407) | s(bool,X407) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X405))) | p(s(bool,X406))) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f8182,plain,(
  ( ! [X405,X407,X406] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X406) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X405))) | p(s(bool,X406)) | s(bool,X407) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X405))) | s(bool,'F_0') = s(bool,X407)) )),
  inference(superposition,[],[f2606,f5187])).
fof(f8186,plain,(
  ( ! [X416,X415] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X415))) | s(bool,X416) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X415))) | s(bool,'F_0') = s(bool,X416)) )),
  inference(superposition,[],[f3543,f5187])).
fof(f11959,plain,(
  ( ! [X420,X419] : (s(bool,X420) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X419))) | s(bool,'F_0') = s(bool,X420)) ) | $spl140),
  inference(cnf_transformation,[],[f11959_D])).
fof(f11959_D,plain,(
  ( ! [X420,X419] : (s(bool,X420) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X419))) | s(bool,'F_0') = s(bool,X420)) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f11973,plain,(
  ( ! [X420,X419] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X420) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X419))) | s(bool,'F_0') = s(bool,X420)) )),
  inference(subsumption_resolution,[],[f8188,f5187])).
fof(f8188,plain,(
  ( ! [X420,X419] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X419))) | s(bool,X420) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X419))) | s(bool,'F_0') = s(bool,X420)) )),
  inference(superposition,[],[f3609,f5187])).
fof(f11915,plain,(
  ( ! [X144,X143] : (s(bool,X144) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X143))))) | s(bool,'F_0') = s(bool,X144)) ) | $spl102),
  inference(cnf_transformation,[],[f11915_D])).
fof(f11915_D,plain,(
  ( ! [X144,X143] : (s(bool,X144) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X143))))) | s(bool,'F_0') = s(bool,X144)) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f8192,plain,(
  ( ! [X430,X431] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X431) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X430))))) | s(bool,'F_0') = s(bool,X431)) )),
  inference(superposition,[],[f2457,f5187])).
fof(f8193,plain,(
  ( ! [X434,X433] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X434) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,X433))))) | s(bool,'F_0') = s(bool,X434)) )),
  inference(superposition,[],[f4197,f5187])).
fof(f11972,plain,(
  ( ! [X87,X86] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X86) = s(bool,X87) | s(bool,'F_0') = s(bool,X87)) )),
  inference(subsumption_resolution,[],[f8234,f5397])).
fof(f8234,plain,(
  ( ! [X87,X86] : (p(s(bool,X86)) | s(bool,'T_0') != s(bool,'F_0') | s(bool,X86) = s(bool,X87) | s(bool,'F_0') = s(bool,X87)) )),
  inference(superposition,[],[f122,f5187])).
fof(f11971,plain,(
  ( ! [X94,X95] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X94) = s(bool,X95) | s(bool,'F_0') = s(bool,X95)) )),
  inference(subsumption_resolution,[],[f8238,f5397])).
fof(f8238,plain,(
  ( ! [X94,X95] : (p(s(bool,X94)) | s(bool,'T_0') != s(bool,'F_0') | s(bool,X94) = s(bool,X95) | s(bool,'F_0') = s(bool,X95)) )),
  inference(superposition,[],[f124,f5187])).
fof(f11970,plain,(
  ( ! [X39,X38] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X38) = s(bool,X39) | s(bool,'F_0') = s(bool,X39)) )),
  inference(subsumption_resolution,[],[f8280,f5397])).
fof(f8280,plain,(
  ( ! [X39,X38] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X38)) | s(bool,X38) = s(bool,X39) | s(bool,'F_0') = s(bool,X39)) )),
  inference(superposition,[],[f122,f5187])).
fof(f11969,plain,(
  ( ! [X47,X46] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X46) = s(bool,X47) | s(bool,'F_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f8284,f5397])).
fof(f8284,plain,(
  ( ! [X47,X46] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X46)) | s(bool,X46) = s(bool,X47) | s(bool,'F_0') = s(bool,X47)) )),
  inference(superposition,[],[f124,f5187])).
fof(f8336,plain,(
  ( ! [X250] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X250) = s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(sK2,sK0),sK5),s(fun(sK0,sK2),sK6))) | s(bool,'F_0') = s(bool,X250)) )),
  inference(superposition,[],[f187,f5187])).
fof(f8337,plain,(
  ( ! [X251] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X251) = s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK3,fun(sK3,bool)),sK7),s(fun(sK3,sK1),sK8),s(fun(sK1,sK3),sK9))) | s(bool,'F_0') = s(bool,X251)) )),
  inference(superposition,[],[f308,f5187])).
fof(f8339,plain,(
  ( ! [X265] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X265) = s(bool,'const.quotient_sum.+++_4'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(sK3,fun(sK3,bool)),sK7),s('type.sum.sum'(sK2,sK3),sK10),s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,X265)) )),
  inference(superposition,[],[f409,f5187])).
fof(f8346,plain,(
  ( ! [X294,X296,X295,X297] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X297) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X294,X295),'const.sum.INL_1'(s(X294,X296))))) | s(bool,'F_0') = s(bool,X297)) )),
  inference(superposition,[],[f5633,f5187])).
fof(f8347,plain,(
  ( ! [X302,X301] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X302) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,X301))))) | s(bool,'F_0') = s(bool,X302)) )),
  inference(superposition,[],[f1592,f5187])).
fof(f8348,plain,(
  ( ! [X304,X306,X305] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X305) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X304))) | p(s(bool,X305)) | s(bool,X306) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X304))) | s(bool,'F_0') = s(bool,X306)) )),
  inference(superposition,[],[f2230,f5187])).
fof(f8349,plain,(
  ( ! [X308,X307] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X307))) | s(bool,X308) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X307))) | s(bool,'F_0') = s(bool,X308)) )),
  inference(superposition,[],[f2232,f5187])).
fof(f8350,plain,(
  ( ! [X310,X309] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X309))) | s(bool,X310) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X309))) | s(bool,'F_0') = s(bool,X310)) )),
  inference(superposition,[],[f2234,f5187])).
fof(f8360,plain,(
  ( ! [X331] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X331) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'F_0') = s(bool,X331)) ) | $spl9),
  inference(superposition,[],[f174,f5187])).
fof(f8365,plain,(
  ( ! [X339,X340] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X340) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK2),'const.sum.INL_1'(s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,X339))))))) | s(bool,'F_0') = s(bool,X340)) )),
  inference(superposition,[],[f1593,f5187])).
fof(f8366,plain,(
  ( ! [X343,X342] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X343) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK2),'const.sum.INR_1'(s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,X342))))))) | s(bool,'F_0') = s(bool,X343)) )),
  inference(superposition,[],[f3224,f5187])).
fof(f8373,plain,(
  ( ! [X385,X387,X388,X386] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X388) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X385,X386),'const.sum.INR_1'(s(X386,X387))))) | s(bool,'F_0') = s(bool,X388)) )),
  inference(superposition,[],[f5815,f5187])).
fof(f8379,plain,(
  ( ! [X401,X402] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X401))) | s(bool,X402) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X401))) | s(bool,'F_0') = s(bool,X402)) )),
  inference(superposition,[],[f1936,f5187])).
fof(f8381,plain,(
  ( ! [X405,X407,X406] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X406) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X405))) | p(s(bool,X406)) | s(bool,X407) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X405))) | s(bool,'F_0') = s(bool,X407)) )),
  inference(superposition,[],[f2606,f5187])).
fof(f8385,plain,(
  ( ! [X416,X415] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X415))) | s(bool,X416) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X415))) | s(bool,'F_0') = s(bool,X416)) )),
  inference(superposition,[],[f3543,f5187])).
fof(f11958,plain,(
  ( ! [X420,X419] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X420) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X419))) | s(bool,'F_0') = s(bool,X420)) )),
  inference(subsumption_resolution,[],[f8387,f5187])).
fof(f8387,plain,(
  ( ! [X420,X419] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X419))) | s(bool,X420) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X419))) | s(bool,'F_0') = s(bool,X420)) )),
  inference(superposition,[],[f3609,f5187])).
fof(f8391,plain,(
  ( ! [X430,X431] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X431) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X430))))) | s(bool,'F_0') = s(bool,X431)) )),
  inference(superposition,[],[f2457,f5187])).
fof(f8392,plain,(
  ( ! [X434,X433] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X434) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,X433))))) | s(bool,'F_0') = s(bool,X434)) )),
  inference(superposition,[],[f4197,f5187])).
fof(f11956,plain,(
  ( ! [X6,X4,X2,X8,X7] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,X4),X7))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,X2),'const.sum.++_3'(s(fun(sK0,sK2),sK6),s(fun(X4,X2),X6),s('type.sum.sum'(sK0,X4),X7))))) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,X4),X8))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,X2),'const.sum.++_3'(s(fun(sK0,sK2),sK6),s(fun(X4,X2),X6),s('type.sum.sum'(sK0,X4),X8)))))) ) | $spl136),
  inference(cnf_transformation,[],[f11956_D])).
fof(f11956_D,plain,(
  ( ! [X6,X4,X2,X8,X7] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,X4),X7))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,X2),'const.sum.++_3'(s(fun(sK0,sK2),sK6),s(fun(X4,X2),X6),s('type.sum.sum'(sK0,X4),X7))))) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,X4),X8))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,X2),'const.sum.++_3'(s(fun(sK0,sK2),sK6),s(fun(X4,X2),X6),s('type.sum.sum'(sK0,X4),X8)))))) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f8802,plain,(
  ( ! [X6,X4,X2,X8,X7] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,X4),X7))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,X2),'const.sum.++_3'(s(fun(sK0,sK2),sK6),s(fun(X4,X2),X6),s('type.sum.sum'(sK0,X4),X7))))) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,X4),X8))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,X2),'const.sum.++_3'(s(fun(sK0,sK2),sK6),s(fun(X4,X2),X6),s('type.sum.sum'(sK0,X4),X8)))))) )),
  inference(superposition,[],[f1411,f1409])).
fof(f1411,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (s(bool,'T_0') != s(bool,'const.quotient.QUOTIENT_3'(s(fun(X28,fun(X28,bool)),X30),s(fun(X28,X26),X31),s(fun(X26,X28),X29))) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,X26),X27))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,X28),'const.sum.++_3'(s(fun(sK0,sK2),sK6),s(fun(X26,X28),X29),s('type.sum.sum'(sK0,X26),X27)))))) )),
  inference(resolution,[],[f181,f122])).
fof(f11955,plain,(
  ( ! [X39,X37,X41,X42] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,X39),X42))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,X37),'const.sum.++_3'(s(fun(sK0,sK2),sK6),s(fun(X39,X37),X41),s('type.sum.sum'(sK0,X39),X42)))))) )),
  inference(subsumption_resolution,[],[f8807,f181])).
fof(f8807,plain,(
  ( ! [X39,X37,X41,X38,X42,X40] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,X39),X42))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,X37),'const.sum.++_3'(s(fun(sK0,sK2),sK6),s(fun(X39,X37),X41),s('type.sum.sum'(sK0,X39),X42))))) | p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X37,fun(X37,bool)),X38),s(fun(X37,X39),X40),s(fun(X39,X37),X41))))) )),
  inference(superposition,[],[f1411,f119])).
fof(f11953,plain,(
  ( ! [X47,X45,X43,X48] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,X45),X48))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,X43),'const.sum.++_3'(s(fun(sK0,sK2),sK6),s(fun(X45,X43),X47),s('type.sum.sum'(sK0,X45),X48)))))) )),
  inference(subsumption_resolution,[],[f8808,f181])).
fof(f8808,plain,(
  ( ! [X47,X45,X43,X48,X46,X44] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,X45),X48))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,X43),'const.sum.++_3'(s(fun(sK0,sK2),sK6),s(fun(X45,X43),X47),s('type.sum.sum'(sK0,X45),X48))))) | p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X43,fun(X43,bool)),X44),s(fun(X43,X45),X46),s(fun(X45,X43),X47))))) )),
  inference(superposition,[],[f1411,f121])).
fof(f11952,plain,(
  ( ! [X70,X68,X72,X71,X69,X67] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,X69),X72))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,X67),'const.sum.++_3'(s(fun(sK0,sK2),sK6),s(fun(X69,X67),X71),s('type.sum.sum'(sK0,X69),X72))))) | s(bool,'T_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X67,fun(X67,bool)),X68),s(fun(X67,X69),X70),s(fun(X69,X67),X71)))) ) | $spl132),
  inference(cnf_transformation,[],[f11952_D])).
fof(f11952_D,plain,(
  ( ! [X70,X68,X72,X71,X69,X67] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,X69),X72))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,X67),'const.sum.++_3'(s(fun(sK0,sK2),sK6),s(fun(X69,X67),X71),s('type.sum.sum'(sK0,X69),X72))))) | s(bool,'T_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X67,fun(X67,bool)),X68),s(fun(X67,X69),X70),s(fun(X69,X67),X71)))) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f8812,plain,(
  ( ! [X70,X68,X72,X71,X69,X67] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,X69),X72))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,X67),'const.sum.++_3'(s(fun(sK0,sK2),sK6),s(fun(X69,X67),X71),s('type.sum.sum'(sK0,X69),X72))))) | s(bool,'T_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X67,fun(X67,bool)),X68),s(fun(X67,X69),X70),s(fun(X69,X67),X71)))) )),
  inference(superposition,[],[f1411,f126])).
fof(f11951,plain,(
  ( ! [X92,X90,X88,X87,X93,X91,X89] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,X89),X92))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,X87),'const.sum.++_3'(s(fun(sK0,sK2),sK6),s(fun(X89,X87),X91),s('type.sum.sum'(sK0,X89),X92))))) | s(bool,'F_0') = s(bool,X93) | s(bool,X93) = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X87,fun(X87,bool)),X88),s(fun(X87,X89),X90),s(fun(X89,X87),X91)))) ) | $spl130),
  inference(cnf_transformation,[],[f11951_D])).
fof(f11951_D,plain,(
  ( ! [X92,X90,X88,X87,X93,X91,X89] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,X89),X92))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,X87),'const.sum.++_3'(s(fun(sK0,sK2),sK6),s(fun(X89,X87),X91),s('type.sum.sum'(sK0,X89),X92))))) | s(bool,'F_0') = s(bool,X93) | s(bool,X93) = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X87,fun(X87,bool)),X88),s(fun(X87,X89),X90),s(fun(X89,X87),X91)))) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f8814,plain,(
  ( ! [X80,X85,X83,X81,X86,X84,X82] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,X82),X85))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,X80),'const.sum.++_3'(s(fun(sK0,sK2),sK6),s(fun(X82,X80),X84),s('type.sum.sum'(sK0,X82),X85))))) | s(bool,X86) = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X80,fun(X80,bool)),X81),s(fun(X80,X82),X83),s(fun(X82,X80),X84))) | s(bool,'F_0') = s(bool,X86)) )),
  inference(superposition,[],[f1411,f5187])).
fof(f8815,plain,(
  ( ! [X92,X90,X88,X87,X93,X91,X89] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,X89),X92))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,X87),'const.sum.++_3'(s(fun(sK0,sK2),sK6),s(fun(X89,X87),X91),s('type.sum.sum'(sK0,X89),X92))))) | s(bool,X93) = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X87,fun(X87,bool)),X88),s(fun(X87,X89),X90),s(fun(X89,X87),X91))) | s(bool,'F_0') = s(bool,X93)) )),
  inference(superposition,[],[f1411,f5187])).
fof(f11940,plain,(
  ( ! [X123,X121,X122,X120] : (s(bool,X123) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X120,X121),'const.sum.INR_1'(s(X121,X122))))) | s(bool,'F_0') = s(bool,X123)) ) | $spl128),
  inference(cnf_transformation,[],[f11940_D])).
fof(f11940_D,plain,(
  ( ! [X123,X121,X122,X120] : (s(bool,X123) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X120,X121),'const.sum.INR_1'(s(X121,X122))))) | s(bool,'F_0') = s(bool,X123)) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f9677,plain,(
  ( ! [X47,X45,X48,X46] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X48) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X45,X46),'const.sum.INR_1'(s(X46,X47))))) | s(bool,'F_0') = s(bool,X48)) )),
  inference(superposition,[],[f9130,f5187])).
fof(f9678,plain,(
  ( ! [X52,X50,X51,X49] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X52) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X49,X50),'const.sum.INR_1'(s(X50,X51))))) | s(bool,'F_0') = s(bool,X52)) )),
  inference(superposition,[],[f9130,f5187])).
fof(f9717,plain,(
  ( ! [X118,X116,X119,X117] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X119) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X116,X117),'const.sum.INR_1'(s(X117,X118))))) | s(bool,'F_0') = s(bool,X119)) )),
  inference(superposition,[],[f5187,f9130])).
fof(f9718,plain,(
  ( ! [X123,X121,X122,X120] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X123) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X120,X121),'const.sum.INR_1'(s(X121,X122))))) | s(bool,'F_0') = s(bool,X123)) )),
  inference(superposition,[],[f5187,f9130])).
fof(f11930,plain,(
  ( ! [X144,X142,X143,X141] : (s(bool,X144) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X141,X142),'const.sum.INL_1'(s(X141,X143))))) | s(bool,'F_0') = s(bool,X144)) ) | $spl126),
  inference(cnf_transformation,[],[f11930_D])).
fof(f11930_D,plain,(
  ( ! [X144,X142,X143,X141] : (s(bool,X144) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X141,X142),'const.sum.INL_1'(s(X141,X143))))) | s(bool,'F_0') = s(bool,X144)) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f10362,plain,(
  ( ! [X68,X66,X69,X67] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X69) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X66,X67),'const.sum.INL_1'(s(X66,X68))))) | s(bool,'F_0') = s(bool,X69)) )),
  inference(superposition,[],[f9229,f5187])).
fof(f10363,plain,(
  ( ! [X70,X72,X71,X73] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X73) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X70,X71),'const.sum.INL_1'(s(X70,X72))))) | s(bool,'F_0') = s(bool,X73)) )),
  inference(superposition,[],[f9229,f5187])).
fof(f10403,plain,(
  ( ! [X140,X138,X139,X137] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X140) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X137,X138),'const.sum.INL_1'(s(X137,X139))))) | s(bool,'F_0') = s(bool,X140)) )),
  inference(superposition,[],[f5187,f9229])).
fof(f10404,plain,(
  ( ! [X144,X142,X143,X141] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X144) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X141,X142),'const.sum.INL_1'(s(X141,X143))))) | s(bool,'F_0') = s(bool,X144)) )),
  inference(superposition,[],[f5187,f9229])).
fof(f11929,plain,(
  ( ! [X2] : (~p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X2)))) | s('type.sum.sum'(sK0,sK0),X2) = s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,sK12(X2,sK0,sK0))))) ) | $spl124),
  inference(cnf_transformation,[],[f11929_D])).
fof(f11929_D,plain,(
  ( ! [X2] : (~p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X2)))) | s('type.sum.sum'(sK0,sK0),X2) = s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,sK12(X2,sK0,sK0))))) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f10791,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,'F_0') | ~p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X2)))) | s('type.sum.sum'(sK0,sK0),X2) = s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,sK12(X2,sK0,sK0))))) )),
  inference(forward_demodulation,[],[f10760,f3205])).
fof(f10760,plain,(
  ( ! [X2] : (~p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X2)))) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,sK13(X2,sK0,sK0)))))) | s('type.sum.sum'(sK0,sK0),X2) = s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,sK12(X2,sK0,sK0))))) )),
  inference(superposition,[],[f1886,f111])).
fof(f11928,plain,(
  ( ! [X8] : (~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X8)))) | s('type.sum.sum'(sK0,sK0),X8) = s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,sK13(X8,sK0,sK0))))) ) | $spl122),
  inference(cnf_transformation,[],[f11928_D])).
fof(f11928_D,plain,(
  ( ! [X8] : (~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X8)))) | s('type.sum.sum'(sK0,sK0),X8) = s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,sK13(X8,sK0,sK0))))) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f11054,plain,(
  ( ! [X8] : (s(bool,'T_0') = s(bool,'F_0') | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X8)))) | s('type.sum.sum'(sK0,sK0),X8) = s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,sK13(X8,sK0,sK0))))) )),
  inference(forward_demodulation,[],[f11024,f4197])).
fof(f11024,plain,(
  ( ! [X8] : (~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X8)))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,sK12(X8,sK0,sK0)))))) | s('type.sum.sum'(sK0,sK0),X8) = s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,sK13(X8,sK0,sK0))))) )),
  inference(superposition,[],[f2288,f111])).
fof(f11927,plain,(
  ( ! [X8] : (p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X8)))) | s('type.sum.sum'(sK0,sK0),X8) = s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,sK13(X8,sK0,sK0))))) ) | $spl120),
  inference(cnf_transformation,[],[f11927_D])).
fof(f11927_D,plain,(
  ( ! [X8] : (p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X8)))) | s('type.sum.sum'(sK0,sK0),X8) = s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,sK13(X8,sK0,sK0))))) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f11192,plain,(
  ( ! [X8] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X8)))) | s('type.sum.sum'(sK0,sK0),X8) = s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,sK13(X8,sK0,sK0))))) )),
  inference(forward_demodulation,[],[f11158,f4197])).
fof(f11158,plain,(
  ( ! [X8] : (p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X8)))) | s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,sK12(X8,sK0,sK0)))))) | s('type.sum.sum'(sK0,sK0),X8) = s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,sK13(X8,sK0,sK0))))) )),
  inference(superposition,[],[f4105,f111])).
fof(f11926,plain,(
  ( ! [X33,X31,X32] : (s(bool,X31) != s(bool,X32) | s(bool,'F_0') = s(bool,X31) | s(bool,'F_0') = s(bool,X33) | s(bool,X32) = s(bool,X33)) ) | $spl118),
  inference(cnf_transformation,[],[f11926_D])).
fof(f11926_D,plain,(
  ( ! [X33,X31,X32] : (s(bool,X31) != s(bool,X32) | s(bool,'F_0') = s(bool,X31) | s(bool,'F_0') = s(bool,X33) | s(bool,X32) = s(bool,X33)) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f11336,plain,(
  ( ! [X30,X28,X29] : (s(bool,X28) != s(bool,X29) | s(bool,X29) = s(bool,X30) | s(bool,'F_0') = s(bool,X30) | s(bool,'F_0') = s(bool,X28) | s(bool,'T_0') = s(bool,'F_0')) )),
  inference(superposition,[],[f6189,f5187])).
fof(f11337,plain,(
  ( ! [X33,X31,X32] : (s(bool,X31) != s(bool,X32) | s(bool,X32) = s(bool,X33) | s(bool,'F_0') = s(bool,X33) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X31)) )),
  inference(superposition,[],[f6189,f5187])).
fof(f11925,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X12) = s(bool,X13) | s(bool,'F_0') = s(bool,X13)) )),
  inference(subsumption_resolution,[],[f11342,f5397])).
fof(f11342,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X12) = s(bool,X13) | s(bool,'F_0') = s(bool,X13) | p(s(bool,X12))) )),
  inference(superposition,[],[f6189,f119])).
fof(f11924,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X14) = s(bool,X15) | s(bool,'F_0') = s(bool,X15)) )),
  inference(subsumption_resolution,[],[f11343,f5397])).
fof(f11343,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X14) = s(bool,X15) | s(bool,'F_0') = s(bool,X15) | p(s(bool,X14))) )),
  inference(superposition,[],[f6189,f121])).
fof(f11922,plain,(
  ( ! [X23,X22] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X22) = s(bool,X23) | s(bool,'F_0') = s(bool,X23)) )),
  inference(subsumption_resolution,[],[f11347,f6189])).
fof(f11347,plain,(
  ( ! [X23,X22] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X22) = s(bool,X23) | s(bool,'F_0') = s(bool,X23) | s(bool,'T_0') = s(bool,X22)) )),
  inference(superposition,[],[f6189,f126])).
fof(f11921,plain,(
  ( ! [X30,X31,X32] : (s(bool,X30) = s(bool,X31) | s(bool,'F_0') = s(bool,X32) | s(bool,X30) = s(bool,X32) | s(bool,'F_0') = s(bool,X31)) ) | $spl114),
  inference(cnf_transformation,[],[f11921_D])).
fof(f11921_D,plain,(
  ( ! [X30,X31,X32] : (s(bool,X30) = s(bool,X31) | s(bool,'F_0') = s(bool,X32) | s(bool,X30) = s(bool,X32) | s(bool,'F_0') = s(bool,X31)) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f11349,plain,(
  ( ! [X28,X29,X27] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X27) = s(bool,X28) | s(bool,'F_0') = s(bool,X28) | s(bool,X27) = s(bool,X29) | s(bool,'F_0') = s(bool,X29)) )),
  inference(superposition,[],[f6189,f5187])).
fof(f11350,plain,(
  ( ! [X30,X31,X32] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X30) = s(bool,X31) | s(bool,'F_0') = s(bool,X31) | s(bool,X30) = s(bool,X32) | s(bool,'F_0') = s(bool,X32)) )),
  inference(superposition,[],[f6189,f5187])).
fof(f11367,plain,(
  ( ! [X68,X66,X69,X67] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X69) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X66,X67),'const.sum.INL_1'(s(X66,X68))))) | s(bool,'F_0') = s(bool,X69)) )),
  inference(superposition,[],[f6189,f5433])).
fof(f11370,plain,(
  ( ! [X87,X86] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X87) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INL_1'(s(sK0,X86))))) | s(bool,'F_0') = s(bool,X87)) )),
  inference(superposition,[],[f6189,f1580])).
fof(f11373,plain,(
  ( ! [X93] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X93) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'F_0') = s(bool,X93)) ) | $spl6),
  inference(superposition,[],[f6189,f171])).
fof(f11377,plain,(
  ( ! [X99,X100] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X100) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK2),'const.sum.INL_1'(s(sK2,happ(s(fun(sK0,sK2),sK6),s(sK0,X99))))))) | s(bool,'F_0') = s(bool,X100)) )),
  inference(superposition,[],[f6189,f1599])).
fof(f11378,plain,(
  ( ! [X103,X101,X102,X104] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X104) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X101,X102),'const.sum.INR_1'(s(X102,X103))))) | s(bool,'F_0') = s(bool,X104)) )),
  inference(superposition,[],[f6189,f5444])).
fof(f11387,plain,(
  ( ! [X144,X143] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X144) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X143))))) | s(bool,'F_0') = s(bool,X144)) )),
  inference(superposition,[],[f6189,f2351])).
fof(f11874,plain,(
  ( ! [X43,X44] : (s(bool,X44) != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'F_0') = s(bool,X43) | s(bool,'T_0') = s(bool,X44)) ) | ($spl19 | $spl26 | $spl32)),
  inference(duplicate_literal_removal,[],[f11789])).
fof(f11789,plain,(
  ( ! [X43,X44] : (s(bool,X44) != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'F_0') = s(bool,X43) | s(bool,'T_0') = s(bool,X44) | s(bool,'F_0') = s(bool,X43)) ) | ($spl19 | $spl26 | $spl32)),
  inference(superposition,[],[f10463,f6289])).
fof(f10463,plain,(
  ( ! [X10,X11] : (s(bool,X10) != s(bool,X11) | s(bool,'F_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X11)) ) | $spl19),
  inference(resolution,[],[f8625,f123])).
fof(f8625,plain,(
  ( ! [X111,X110] : (p(s(bool,X111)) | s(bool,X110) != s(bool,X111) | s(bool,'F_0') = s(bool,X110)) ) | $spl19),
  inference(subsumption_resolution,[],[f7721,f255])).
fof(f11875,plain,(
  ( ! [X41,X42] : (s(bool,X42) != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,X41) | s(bool,'T_0') = s(bool,X42)) ) | ($spl19 | $spl32)),
  inference(duplicate_literal_removal,[],[f11788])).
fof(f11788,plain,(
  ( ! [X41,X42] : (s(bool,X42) != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,X41) | s(bool,'T_0') = s(bool,X42) | s(bool,'F_0') = s(bool,X41)) ) | ($spl19 | $spl32)),
  inference(superposition,[],[f10463,f6288])).
fof(f11886,plain,(
  ( ! [X40] : (s(bool,X40) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'T_0') = s(bool,X40)) ) | $spl96),
  inference(cnf_transformation,[],[f11886_D])).
fof(f11886_D,plain,(
  ( ! [X40] : (s(bool,X40) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'T_0') = s(bool,X40)) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f11876,plain,(
  ( ! [X39,X40] : (s(bool,X40) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,X39) | s(bool,'T_0') = s(bool,X40)) ) | ($spl12 | $spl19)),
  inference(duplicate_literal_removal,[],[f11787])).
fof(f11787,plain,(
  ( ! [X39,X40] : (s(bool,X40) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))) | s(bool,'F_0') = s(bool,X39) | s(bool,'T_0') = s(bool,X40) | s(bool,'F_0') = s(bool,X39)) ) | ($spl12 | $spl19)),
  inference(superposition,[],[f10463,f6282])).
fof(f11879,plain,(
  ( ! [X26,X24,X25] : (s(bool,X26) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X25))))) | s(bool,'F_0') = s(bool,X24) | s(bool,'T_0') = s(bool,X26)) ) | $spl19),
  inference(duplicate_literal_removal,[],[f11782])).
fof(f11782,plain,(
  ( ! [X26,X24,X25] : (s(bool,X26) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),'const.sum.INR_1'(s(sK0,X25))))) | s(bool,'F_0') = s(bool,X24) | s(bool,'T_0') = s(bool,X26) | s(bool,'F_0') = s(bool,X24)) ) | $spl19),
  inference(superposition,[],[f10463,f5092])).
fof(f11882,plain,(
  ( ! [X8,X7] : (~p(s(bool,X7)) | s(bool,X7) != s(bool,X8) | s(bool,'T_0') = s(bool,X8)) ) | $spl92),
  inference(cnf_transformation,[],[f11882_D])).
fof(f11882_D,plain,(
  ( ! [X8,X7] : (~p(s(bool,X7)) | s(bool,X7) != s(bool,X8) | s(bool,'T_0') = s(bool,X8)) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f11883,plain,(
  ( ! [X10,X11,X9] : (s(bool,X10) != s(bool,X11) | s(bool,'F_0') = s(bool,X9) | s(bool,'T_0') = s(bool,X11) | ~p(s(bool,X10))) ) | $spl19),
  inference(subsumption_resolution,[],[f11775,f119])).
fof(f11775,plain,(
  ( ! [X10,X11,X9] : (s(bool,X10) != s(bool,X11) | s(bool,'F_0') = s(bool,X9) | s(bool,'T_0') = s(bool,X11) | ~p(s(bool,X10)) | ~p(s(bool,X9))) ) | $spl19),
  inference(superposition,[],[f10463,f115])).
fof(f11881,plain,(
  ( ! [X6,X8,X7] : (s(bool,X7) != s(bool,X8) | s(bool,'F_0') = s(bool,X6) | s(bool,'T_0') = s(bool,X8) | ~p(s(bool,X7))) ) | $spl19),
  inference(subsumption_resolution,[],[f11774,f119])).
fof(f11774,plain,(
  ( ! [X6,X8,X7] : (s(bool,X7) != s(bool,X8) | s(bool,'F_0') = s(bool,X6) | s(bool,'T_0') = s(bool,X8) | ~p(s(bool,X6)) | ~p(s(bool,X7))) ) | $spl19),
  inference(superposition,[],[f10463,f115])).
fof(f11673,plain,(
  ( ! [X94,X99,X97,X95,X100,X98,X96] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK1,X96),X100))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK3,X94),'const.sum.++_3'(s(fun(sK1,sK3),sK9),s(fun(X96,X94),X98),s('type.sum.sum'(sK1,X96),X100))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X94,fun(X94,bool)),X95),s(fun(X94,X96),X97),s(fun(X96,X94),X98))) | s(bool,'F_0') = s(bool,X99)) )),
  inference(subsumption_resolution,[],[f11657,f119])).
fof(f11657,plain,(
  ( ! [X94,X99,X97,X95,X100,X98,X96] : (~p(s(bool,X99)) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK1,X96),X100))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK3,X94),'const.sum.++_3'(s(fun(sK1,sK3),sK9),s(fun(X96,X94),X98),s('type.sum.sum'(sK1,X96),X100))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X94,fun(X94,bool)),X95),s(fun(X94,X96),X97),s(fun(X96,X94),X98))) | s(bool,'F_0') = s(bool,X99)) )),
  inference(superposition,[],[f302,f5187])).
fof(f11672,plain,(
  ( ! [X92,X90,X88,X87,X93,X91,X89] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK1,X89),X93))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK3,X87),'const.sum.++_3'(s(fun(sK1,sK3),sK9),s(fun(X89,X87),X91),s('type.sum.sum'(sK1,X89),X93))))) | s(bool,'F_0') = s(bool,X92) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X87,fun(X87,bool)),X88),s(fun(X87,X89),X90),s(fun(X89,X87),X91)))) )),
  inference(subsumption_resolution,[],[f11656,f119])).
fof(f11656,plain,(
  ( ! [X92,X90,X88,X87,X93,X91,X89] : (~p(s(bool,X92)) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK1,X89),X93))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK3,X87),'const.sum.++_3'(s(fun(sK1,sK3),sK9),s(fun(X89,X87),X91),s('type.sum.sum'(sK1,X89),X93))))) | s(bool,'F_0') = s(bool,X92) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X87,fun(X87,bool)),X88),s(fun(X87,X89),X90),s(fun(X89,X87),X91)))) )),
  inference(superposition,[],[f302,f5187])).
fof(f11623,plain,(
  ( ! [X94,X99,X97,X95,X100,X98,X96] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X96,sK0),X100))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X94,sK2),'const.sum.++_3'(s(fun(X96,X94),X98),s(fun(sK0,sK2),sK6),s('type.sum.sum'(X96,sK0),X100))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X94,fun(X94,bool)),X95),s(fun(X94,X96),X97),s(fun(X96,X94),X98))) | s(bool,'F_0') = s(bool,X99)) )),
  inference(subsumption_resolution,[],[f11607,f119])).
fof(f11607,plain,(
  ( ! [X94,X99,X97,X95,X100,X98,X96] : (~p(s(bool,X99)) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X96,sK0),X100))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X94,sK2),'const.sum.++_3'(s(fun(X96,X94),X98),s(fun(sK0,sK2),sK6),s('type.sum.sum'(X96,sK0),X100))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X94,fun(X94,bool)),X95),s(fun(X94,X96),X97),s(fun(X96,X94),X98))) | s(bool,'F_0') = s(bool,X99)) )),
  inference(superposition,[],[f182,f5187])).
fof(f11622,plain,(
  ( ! [X92,X90,X88,X87,X93,X91,X89] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X89,sK0),X93))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X87,sK2),'const.sum.++_3'(s(fun(X89,X87),X91),s(fun(sK0,sK2),sK6),s('type.sum.sum'(X89,sK0),X93))))) | s(bool,'F_0') = s(bool,X92) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X87,fun(X87,bool)),X88),s(fun(X87,X89),X90),s(fun(X89,X87),X91)))) )),
  inference(subsumption_resolution,[],[f11606,f119])).
fof(f11606,plain,(
  ( ! [X92,X90,X88,X87,X93,X91,X89] : (~p(s(bool,X92)) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X89,sK0),X93))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X87,sK2),'const.sum.++_3'(s(fun(X89,X87),X91),s(fun(sK0,sK2),sK6),s('type.sum.sum'(X89,sK0),X93))))) | s(bool,'F_0') = s(bool,X92) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X87,fun(X87,bool)),X88),s(fun(X87,X89),X90),s(fun(X89,X87),X91)))) )),
  inference(superposition,[],[f182,f5187])).
fof(f11573,plain,(
  ( ! [X94,X99,X97,X95,X100,X98,X96] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,X96),X100))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,X94),'const.sum.++_3'(s(fun(sK0,sK2),sK6),s(fun(X96,X94),X98),s('type.sum.sum'(sK0,X96),X100))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X94,fun(X94,bool)),X95),s(fun(X94,X96),X97),s(fun(X96,X94),X98))) | s(bool,'F_0') = s(bool,X99)) )),
  inference(subsumption_resolution,[],[f11556,f119])).
fof(f11556,plain,(
  ( ! [X94,X99,X97,X95,X100,X98,X96] : (~p(s(bool,X99)) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,X96),X100))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,X94),'const.sum.++_3'(s(fun(sK0,sK2),sK6),s(fun(X96,X94),X98),s('type.sum.sum'(sK0,X96),X100))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X94,fun(X94,bool)),X95),s(fun(X94,X96),X97),s(fun(X96,X94),X98))) | s(bool,'F_0') = s(bool,X99)) )),
  inference(superposition,[],[f181,f5187])).
fof(f11571,plain,(
  ( ! [X92,X90,X88,X87,X93,X91,X89] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,X89),X93))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,X87),'const.sum.++_3'(s(fun(sK0,sK2),sK6),s(fun(X89,X87),X91),s('type.sum.sum'(sK0,X89),X93))))) | s(bool,'F_0') = s(bool,X92) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X87,fun(X87,bool)),X88),s(fun(X87,X89),X90),s(fun(X89,X87),X91)))) )),
  inference(subsumption_resolution,[],[f11555,f119])).
fof(f11555,plain,(
  ( ! [X92,X90,X88,X87,X93,X91,X89] : (~p(s(bool,X92)) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,X89),X93))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,X87),'const.sum.++_3'(s(fun(sK0,sK2),sK6),s(fun(X89,X87),X91),s('type.sum.sum'(sK0,X89),X93))))) | s(bool,'F_0') = s(bool,X92) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X87,fun(X87,bool)),X88),s(fun(X87,X89),X90),s(fun(X89,X87),X91)))) )),
  inference(superposition,[],[f181,f5187])).
fof(f11136,plain,(
  ( ! [X37,X38] : (s(bool,'T_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X37))) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X37))) | s(bool,'F_0') = s(bool,X38)) )),
  inference(subsumption_resolution,[],[f11114,f119])).
fof(f11114,plain,(
  ( ! [X37,X38] : (~p(s(bool,X38)) | s(bool,'T_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X37))) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X37))) | s(bool,'F_0') = s(bool,X38)) )),
  inference(superposition,[],[f2598,f5187])).
fof(f11134,plain,(
  ( ! [X35,X36] : (s(bool,'T_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X35))) | s(bool,'F_0') = s(bool,X36) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X35)))) )),
  inference(subsumption_resolution,[],[f11113,f119])).
fof(f11113,plain,(
  ( ! [X35,X36] : (~p(s(bool,X36)) | s(bool,'T_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X35))) | s(bool,'F_0') = s(bool,X36) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X35)))) )),
  inference(superposition,[],[f2598,f5187])).
fof(f11069,plain,(
  ( ! [X37,X38] : (s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X37))) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X37))) | s(bool,'F_0') = s(bool,X38)) )),
  inference(subsumption_resolution,[],[f11047,f119])).
fof(f11047,plain,(
  ( ! [X37,X38] : (~p(s(bool,X38)) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X37))) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X37))) | s(bool,'F_0') = s(bool,X38)) )),
  inference(superposition,[],[f2288,f5187])).
fof(f11068,plain,(
  ( ! [X35,X36] : (s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X35))) | s(bool,'F_0') = s(bool,X36) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X35)))) )),
  inference(subsumption_resolution,[],[f11046,f119])).
fof(f11046,plain,(
  ( ! [X35,X36] : (~p(s(bool,X36)) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X35))) | s(bool,'F_0') = s(bool,X36) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X35)))) )),
  inference(superposition,[],[f2288,f5187])).
fof(f10872,plain,(
  ( ! [X47,X46] : (s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X46))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X46))) | s(bool,'F_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f10850,f119])).
fof(f10850,plain,(
  ( ! [X47,X46] : (~p(s(bool,X47)) | s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X46))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X46))) | s(bool,'F_0') = s(bool,X47)) )),
  inference(superposition,[],[f2215,f5187])).
fof(f10870,plain,(
  ( ! [X45,X44] : (s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X44))) | s(bool,'F_0') = s(bool,X45) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X44)))) )),
  inference(subsumption_resolution,[],[f10849,f119])).
fof(f10849,plain,(
  ( ! [X45,X44] : (~p(s(bool,X45)) | s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X44))) | s(bool,'F_0') = s(bool,X45) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X44)))) )),
  inference(superposition,[],[f2215,f5187])).
fof(f10806,plain,(
  ( ! [X47,X46] : (s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X46))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X46))) | s(bool,'F_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f10784,f119])).
fof(f10784,plain,(
  ( ! [X47,X46] : (~p(s(bool,X47)) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X46))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X46))) | s(bool,'F_0') = s(bool,X47)) )),
  inference(superposition,[],[f1886,f5187])).
fof(f10804,plain,(
  ( ! [X45,X44] : (s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X44))) | s(bool,'F_0') = s(bool,X45) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X44)))) )),
  inference(subsumption_resolution,[],[f10783,f119])).
fof(f10783,plain,(
  ( ! [X45,X44] : (~p(s(bool,X45)) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X44))) | s(bool,'F_0') = s(bool,X45) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X44)))) )),
  inference(superposition,[],[f1886,f5187])).
fof(f10246,plain,(
  ( ! [X191,X198,X196,X194,X192,X202,X200,X190,X199,X197,X195,X193,X201] : (s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X190,X196),X197))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X190,X196),X198))) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X190,fun(X190,bool)),X191),s(fun(X196,fun(X196,bool)),X199),s('type.sum.sum'(X190,X196),X197),s('type.sum.sum'(X190,X196),X198)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X196,fun(X196,bool)),X199),s(fun(X196,X200),X201),s(fun(X200,X196),X202)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X190,fun(X190,bool)),X191),s(fun(X190,X192),X193),s(fun(X192,X190),X194))) | s(bool,'F_0') = s(bool,X195)) )),
  inference(subsumption_resolution,[],[f10034,f119])).
fof(f10034,plain,(
  ( ! [X191,X198,X196,X194,X192,X202,X200,X190,X199,X197,X195,X193,X201] : (~p(s(bool,X195)) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X190,X196),X197))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X190,X196),X198))) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X190,fun(X190,bool)),X191),s(fun(X196,fun(X196,bool)),X199),s('type.sum.sum'(X190,X196),X197),s('type.sum.sum'(X190,X196),X198)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X196,fun(X196,bool)),X199),s(fun(X196,X200),X201),s(fun(X200,X196),X202)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X190,fun(X190,bool)),X191),s(fun(X190,X192),X193),s(fun(X192,X190),X194))) | s(bool,'F_0') = s(bool,X195)) )),
  inference(superposition,[],[f130,f5187])).
fof(f10245,plain,(
  ( ! [X177,X189,X187,X185,X182,X180,X178,X188,X186,X184,X183,X181,X179] : (s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X177,X183),X184))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X177,X183),X185))) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X177,fun(X177,bool)),X178),s(fun(X183,fun(X183,bool)),X186),s('type.sum.sum'(X177,X183),X184),s('type.sum.sum'(X177,X183),X185)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X183,fun(X183,bool)),X186),s(fun(X183,X187),X188),s(fun(X187,X183),X189)))) | s(bool,'F_0') = s(bool,X182) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X177,fun(X177,bool)),X178),s(fun(X177,X179),X180),s(fun(X179,X177),X181)))) )),
  inference(subsumption_resolution,[],[f10033,f119])).
fof(f10033,plain,(
  ( ! [X177,X189,X187,X185,X182,X180,X178,X188,X186,X184,X183,X181,X179] : (~p(s(bool,X182)) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X177,X183),X184))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X177,X183),X185))) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X177,fun(X177,bool)),X178),s(fun(X183,fun(X183,bool)),X186),s('type.sum.sum'(X177,X183),X184),s('type.sum.sum'(X177,X183),X185)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X183,fun(X183,bool)),X186),s(fun(X183,X187),X188),s(fun(X187,X183),X189)))) | s(bool,'F_0') = s(bool,X182) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X177,fun(X177,bool)),X178),s(fun(X177,X179),X180),s(fun(X179,X177),X181)))) )),
  inference(superposition,[],[f130,f5187])).
fof(f10231,plain,(
  ( ! [X191,X198,X196,X194,X192,X202,X200,X190,X199,X197,X195,X193,X201] : (s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X196,X190),X197))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X196,X190),X198))) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X196,fun(X196,bool)),X199),s(fun(X190,fun(X190,bool)),X191),s('type.sum.sum'(X196,X190),X197),s('type.sum.sum'(X196,X190),X198)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X196,fun(X196,bool)),X199),s(fun(X196,X200),X201),s(fun(X200,X196),X202)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X190,fun(X190,bool)),X191),s(fun(X190,X192),X193),s(fun(X192,X190),X194))) | s(bool,'F_0') = s(bool,X195)) )),
  inference(subsumption_resolution,[],[f10014,f119])).
fof(f10014,plain,(
  ( ! [X191,X198,X196,X194,X192,X202,X200,X190,X199,X197,X195,X193,X201] : (~p(s(bool,X195)) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X196,X190),X197))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X196,X190),X198))) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X196,fun(X196,bool)),X199),s(fun(X190,fun(X190,bool)),X191),s('type.sum.sum'(X196,X190),X197),s('type.sum.sum'(X196,X190),X198)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X196,fun(X196,bool)),X199),s(fun(X196,X200),X201),s(fun(X200,X196),X202)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X190,fun(X190,bool)),X191),s(fun(X190,X192),X193),s(fun(X192,X190),X194))) | s(bool,'F_0') = s(bool,X195)) )),
  inference(superposition,[],[f130,f5187])).
fof(f10230,plain,(
  ( ! [X177,X189,X187,X185,X182,X180,X178,X188,X186,X184,X183,X181,X179] : (s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X183,X177),X184))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X183,X177),X185))) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X183,fun(X183,bool)),X186),s(fun(X177,fun(X177,bool)),X178),s('type.sum.sum'(X183,X177),X184),s('type.sum.sum'(X183,X177),X185)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X183,fun(X183,bool)),X186),s(fun(X183,X187),X188),s(fun(X187,X183),X189)))) | s(bool,'F_0') = s(bool,X182) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X177,fun(X177,bool)),X178),s(fun(X177,X179),X180),s(fun(X179,X177),X181)))) )),
  inference(subsumption_resolution,[],[f10013,f119])).
fof(f10013,plain,(
  ( ! [X177,X189,X187,X185,X182,X180,X178,X188,X186,X184,X183,X181,X179] : (~p(s(bool,X182)) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X183,X177),X184))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X183,X177),X185))) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X183,fun(X183,bool)),X186),s(fun(X177,fun(X177,bool)),X178),s('type.sum.sum'(X183,X177),X184),s('type.sum.sum'(X183,X177),X185)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X183,fun(X183,bool)),X186),s(fun(X183,X187),X188),s(fun(X187,X183),X189)))) | s(bool,'F_0') = s(bool,X182) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X177,fun(X177,bool)),X178),s(fun(X177,X179),X180),s(fun(X179,X177),X181)))) )),
  inference(superposition,[],[f130,f5187])).
fof(f10216,plain,(
  ( ! [X191,X189,X187,X185,X194,X192,X182,X190,X188,X186,X184,X193,X183] : (s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X182,X184),X186))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X182,X184),X187))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X184,fun(X184,bool)),X185),s(fun(X184,X189),X190),s(fun(X189,X184),X191)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X182,fun(X182,bool)),X183),s(fun(X182,X192),X193),s(fun(X192,X182),X194)))) | s(bool,'F_0') = s(bool,'const.quotient_sum.+++_4'(s(fun(X182,fun(X182,bool)),X183),s(fun(X184,fun(X184,bool)),X185),s('type.sum.sum'(X182,X184),X186),s('type.sum.sum'(X182,X184),X187))) | s(bool,'F_0') = s(bool,X188)) )),
  inference(subsumption_resolution,[],[f9994,f119])).
fof(f9994,plain,(
  ( ! [X191,X189,X187,X185,X194,X192,X182,X190,X188,X186,X184,X193,X183] : (~p(s(bool,X188)) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X182,X184),X186))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X182,X184),X187))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X184,fun(X184,bool)),X185),s(fun(X184,X189),X190),s(fun(X189,X184),X191)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X182,fun(X182,bool)),X183),s(fun(X182,X192),X193),s(fun(X192,X182),X194)))) | s(bool,'F_0') = s(bool,'const.quotient_sum.+++_4'(s(fun(X182,fun(X182,bool)),X183),s(fun(X184,fun(X184,bool)),X185),s('type.sum.sum'(X182,X184),X186),s('type.sum.sum'(X182,X184),X187))) | s(bool,'F_0') = s(bool,X188)) )),
  inference(superposition,[],[f130,f5187])).
fof(f10215,plain,(
  ( ! [X177,X175,X173,X171,X169,X180,X178,X176,X174,X172,X170,X181,X179] : (s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X169,X171),X173))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X169,X171),X174))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X171,fun(X171,bool)),X172),s(fun(X171,X176),X177),s(fun(X176,X171),X178)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X169,fun(X169,bool)),X170),s(fun(X169,X179),X180),s(fun(X179,X169),X181)))) | s(bool,'F_0') = s(bool,X175) | s(bool,'F_0') = s(bool,'const.quotient_sum.+++_4'(s(fun(X169,fun(X169,bool)),X170),s(fun(X171,fun(X171,bool)),X172),s('type.sum.sum'(X169,X171),X173),s('type.sum.sum'(X169,X171),X174)))) )),
  inference(subsumption_resolution,[],[f9993,f119])).
fof(f9993,plain,(
  ( ! [X177,X175,X173,X171,X169,X180,X178,X176,X174,X172,X170,X181,X179] : (~p(s(bool,X175)) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X169,X171),X173))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X169,X171),X174))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X171,fun(X171,bool)),X172),s(fun(X171,X176),X177),s(fun(X176,X171),X178)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X169,fun(X169,bool)),X170),s(fun(X169,X179),X180),s(fun(X179,X169),X181)))) | s(bool,'F_0') = s(bool,X175) | s(bool,'F_0') = s(bool,'const.quotient_sum.+++_4'(s(fun(X169,fun(X169,bool)),X170),s(fun(X171,fun(X171,bool)),X172),s('type.sum.sum'(X169,X171),X173),s('type.sum.sum'(X169,X171),X174)))) )),
  inference(superposition,[],[f130,f5187])).
fof(f9587,plain,(
  ( ! [X177,X185,X175,X182,X180,X178,X176,X184,X174,X183,X181,X179] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X176,X180),X181))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X174,X182),'const.sum.++_3'(s(fun(X176,X174),X178),s(fun(X180,X182),X183),s('type.sum.sum'(X176,X180),X181))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X182,fun(X182,bool)),X184),s(fun(X182,X180),X185),s(fun(X180,X182),X183)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X174,fun(X174,bool)),X175),s(fun(X174,X176),X177),s(fun(X176,X174),X178))) | s(bool,'F_0') = s(bool,X179)) )),
  inference(subsumption_resolution,[],[f9497,f119])).
fof(f9497,plain,(
  ( ! [X177,X185,X175,X182,X180,X178,X176,X184,X174,X183,X181,X179] : (~p(s(bool,X179)) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X176,X180),X181))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X174,X182),'const.sum.++_3'(s(fun(X176,X174),X178),s(fun(X180,X182),X183),s('type.sum.sum'(X176,X180),X181))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X182,fun(X182,bool)),X184),s(fun(X182,X180),X185),s(fun(X180,X182),X183)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X174,fun(X174,bool)),X175),s(fun(X174,X176),X177),s(fun(X176,X174),X178))) | s(bool,'F_0') = s(bool,X179)) )),
  inference(superposition,[],[f127,f5187])).
fof(f9586,plain,(
  ( ! [X167,X165,X163,X173,X171,X169,X166,X164,X162,X172,X170,X168] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X164,X168),X169))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X162,X170),'const.sum.++_3'(s(fun(X164,X162),X166),s(fun(X168,X170),X171),s('type.sum.sum'(X164,X168),X169))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X170,fun(X170,bool)),X172),s(fun(X170,X168),X173),s(fun(X168,X170),X171)))) | s(bool,'F_0') = s(bool,X167) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X162,fun(X162,bool)),X163),s(fun(X162,X164),X165),s(fun(X164,X162),X166)))) )),
  inference(subsumption_resolution,[],[f9496,f119])).
fof(f9496,plain,(
  ( ! [X167,X165,X163,X173,X171,X169,X166,X164,X162,X172,X170,X168] : (~p(s(bool,X167)) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X164,X168),X169))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X162,X170),'const.sum.++_3'(s(fun(X164,X162),X166),s(fun(X168,X170),X171),s('type.sum.sum'(X164,X168),X169))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X170,fun(X170,bool)),X172),s(fun(X170,X168),X173),s(fun(X168,X170),X171)))) | s(bool,'F_0') = s(bool,X167) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X162,fun(X162,bool)),X163),s(fun(X162,X164),X165),s(fun(X164,X162),X166)))) )),
  inference(superposition,[],[f127,f5187])).
fof(f9573,plain,(
  ( ! [X177,X185,X175,X182,X180,X178,X176,X184,X174,X183,X181,X179] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X180,X176),X181))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X182,X174),'const.sum.++_3'(s(fun(X180,X182),X183),s(fun(X176,X174),X178),s('type.sum.sum'(X180,X176),X181))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X182,fun(X182,bool)),X184),s(fun(X182,X180),X185),s(fun(X180,X182),X183)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X174,fun(X174,bool)),X175),s(fun(X174,X176),X177),s(fun(X176,X174),X178))) | s(bool,'F_0') = s(bool,X179)) )),
  inference(subsumption_resolution,[],[f9477,f119])).
fof(f9477,plain,(
  ( ! [X177,X185,X175,X182,X180,X178,X176,X184,X174,X183,X181,X179] : (~p(s(bool,X179)) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X180,X176),X181))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X182,X174),'const.sum.++_3'(s(fun(X180,X182),X183),s(fun(X176,X174),X178),s('type.sum.sum'(X180,X176),X181))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X182,fun(X182,bool)),X184),s(fun(X182,X180),X185),s(fun(X180,X182),X183)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X174,fun(X174,bool)),X175),s(fun(X174,X176),X177),s(fun(X176,X174),X178))) | s(bool,'F_0') = s(bool,X179)) )),
  inference(superposition,[],[f127,f5187])).
fof(f9572,plain,(
  ( ! [X167,X165,X163,X173,X171,X169,X166,X164,X162,X172,X170,X168] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X168,X164),X169))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X170,X162),'const.sum.++_3'(s(fun(X168,X170),X171),s(fun(X164,X162),X166),s('type.sum.sum'(X168,X164),X169))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X170,fun(X170,bool)),X172),s(fun(X170,X168),X173),s(fun(X168,X170),X171)))) | s(bool,'F_0') = s(bool,X167) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X162,fun(X162,bool)),X163),s(fun(X162,X164),X165),s(fun(X164,X162),X166)))) )),
  inference(subsumption_resolution,[],[f9476,f119])).
fof(f9476,plain,(
  ( ! [X167,X165,X163,X173,X171,X169,X166,X164,X162,X172,X170,X168] : (~p(s(bool,X167)) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X168,X164),X169))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X170,X162),'const.sum.++_3'(s(fun(X168,X170),X171),s(fun(X164,X162),X166),s('type.sum.sum'(X168,X164),X169))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X170,fun(X170,bool)),X172),s(fun(X170,X168),X173),s(fun(X168,X170),X171)))) | s(bool,'F_0') = s(bool,X167) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X162,fun(X162,bool)),X163),s(fun(X162,X164),X165),s(fun(X164,X162),X166)))) )),
  inference(superposition,[],[f127,f5187])).
fof(f9099,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f8874,f150])).
fof(f8874,plain,(
  p(s(bool,'F_0')) | $spl46),
  inference(backward_demodulation,[],[f8621,f134])).
fof(f9097,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f9092,f8621])).
fof(f9092,plain,(
  ( ! [X9] : (s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X9)))) ) | $spl46),
  inference(trivial_inequality_removal,[],[f8863])).
fof(f8863,plain,(
  ( ! [X9] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X9)))) ) | $spl46),
  inference(backward_demodulation,[],[f8621,f4134])).
fof(f9096,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f9095,f8621])).
fof(f9095,plain,(
  ( ! [X35] : (s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X35)))) ) | $spl46),
  inference(subsumption_resolution,[],[f8861,f150])).
fof(f8861,plain,(
  ( ! [X35] : (p(s(bool,'F_0')) | s(bool,'F_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X35)))) ) | $spl46),
  inference(backward_demodulation,[],[f8621,f4105])).
fof(f9093,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f8842,f150])).
fof(f8842,plain,(
  p(s(bool,'F_0')) | $spl46),
  inference(backward_demodulation,[],[f8621,f129])).
fof(f8661,plain,(
  ( ! [X341,X340] : (s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK3,sK3),X340))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK3,sK3),X341))) | s(bool,'F_0') = s(bool,'const.quotient_sum.+++_4'(s(fun(sK3,fun(sK3,bool)),sK7),s(fun(sK3,fun(sK3,bool)),sK7),s('type.sum.sum'(sK3,sK3),X340),s('type.sum.sum'(sK3,sK3),X341)))) ) | $spl82),
  inference(cnf_transformation,[],[f8661_D])).
fof(f8661_D,plain,(
  ( ! [X341,X340] : (s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK3,sK3),X340))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK3,sK3),X341))) | s(bool,'F_0') = s(bool,'const.quotient_sum.+++_4'(s(fun(sK3,fun(sK3,bool)),sK7),s(fun(sK3,fun(sK3,bool)),sK7),s('type.sum.sum'(sK3,sK3),X340),s('type.sum.sum'(sK3,sK3),X341)))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f8696,plain,(
  ( ! [X341,X340,X342] : (s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK3,sK3),X340))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK3,sK3),X341))) | s(bool,'F_0') = s(bool,'const.quotient_sum.+++_4'(s(fun(sK3,fun(sK3,bool)),sK7),s(fun(sK3,fun(sK3,bool)),sK7),s('type.sum.sum'(sK3,sK3),X340),s('type.sum.sum'(sK3,sK3),X341))) | s(bool,'F_0') = s(bool,X342)) )),
  inference(subsumption_resolution,[],[f7938,f119])).
fof(f7938,plain,(
  ( ! [X341,X340,X342] : (~p(s(bool,X342)) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK3,sK3),X340))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK3,sK3),X341))) | s(bool,'F_0') = s(bool,'const.quotient_sum.+++_4'(s(fun(sK3,fun(sK3,bool)),sK7),s(fun(sK3,fun(sK3,bool)),sK7),s('type.sum.sum'(sK3,sK3),X340),s('type.sum.sum'(sK3,sK3),X341))) | s(bool,'F_0') = s(bool,X342)) )),
  inference(superposition,[],[f2094,f5187])).
fof(f2094,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.quotient_sum.+++_4'(s(fun(sK3,fun(sK3,bool)),sK7),s(fun(sK3,fun(sK3,bool)),sK7),s('type.sum.sum'(sK3,sK3),X2),s('type.sum.sum'(sK3,sK3),X3)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK3,sK3),X2))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK3,sK3),X3)))) )),
  inference(subsumption_resolution,[],[f2078,f141])).
fof(f2078,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(sK3,fun(sK3,bool)),sK7),s(fun(sK3,fun(sK3,bool)),sK7),s('type.sum.sum'(sK3,sK3),X2),s('type.sum.sum'(sK3,sK3),X3)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK3,sK3),X2))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK3,sK3),X3)))) )),
  inference(superposition,[],[f300,f308])).
fof(f8659,plain,(
  ( ! [X337,X338] : (s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK3,sK2),X337))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK3,sK2),X338))) | s(bool,'F_0') = s(bool,'const.quotient_sum.+++_4'(s(fun(sK3,fun(sK3,bool)),sK7),s(fun(sK2,fun(sK2,bool)),sK4),s('type.sum.sum'(sK3,sK2),X337),s('type.sum.sum'(sK3,sK2),X338)))) ) | $spl80),
  inference(cnf_transformation,[],[f8659_D])).
fof(f8659_D,plain,(
  ( ! [X337,X338] : (s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK3,sK2),X337))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK3,sK2),X338))) | s(bool,'F_0') = s(bool,'const.quotient_sum.+++_4'(s(fun(sK3,fun(sK3,bool)),sK7),s(fun(sK2,fun(sK2,bool)),sK4),s('type.sum.sum'(sK3,sK2),X337),s('type.sum.sum'(sK3,sK2),X338)))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f8695,plain,(
  ( ! [X337,X339,X338] : (s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK3,sK2),X337))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK3,sK2),X338))) | s(bool,'F_0') = s(bool,'const.quotient_sum.+++_4'(s(fun(sK3,fun(sK3,bool)),sK7),s(fun(sK2,fun(sK2,bool)),sK4),s('type.sum.sum'(sK3,sK2),X337),s('type.sum.sum'(sK3,sK2),X338))) | s(bool,'F_0') = s(bool,X339)) )),
  inference(subsumption_resolution,[],[f7937,f119])).
fof(f7937,plain,(
  ( ! [X337,X339,X338] : (~p(s(bool,X339)) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK3,sK2),X337))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK3,sK2),X338))) | s(bool,'F_0') = s(bool,'const.quotient_sum.+++_4'(s(fun(sK3,fun(sK3,bool)),sK7),s(fun(sK2,fun(sK2,bool)),sK4),s('type.sum.sum'(sK3,sK2),X337),s('type.sum.sum'(sK3,sK2),X338))) | s(bool,'F_0') = s(bool,X339)) )),
  inference(superposition,[],[f1927,f5187])).
fof(f1927,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.quotient_sum.+++_4'(s(fun(sK3,fun(sK3,bool)),sK7),s(fun(sK2,fun(sK2,bool)),sK4),s('type.sum.sum'(sK3,sK2),X2),s('type.sum.sum'(sK3,sK2),X3)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK3,sK2),X2))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK3,sK2),X3)))) )),
  inference(subsumption_resolution,[],[f1911,f141])).
fof(f1911,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(sK3,fun(sK3,bool)),sK7),s(fun(sK2,fun(sK2,bool)),sK4),s('type.sum.sum'(sK3,sK2),X2),s('type.sum.sum'(sK3,sK2),X3)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK3,sK2),X2))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK3,sK2),X3)))) )),
  inference(superposition,[],[f180,f308])).
fof(f8657,plain,(
  ( ! [X335,X334] : (s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK2),X334))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK2),X335))) | s(bool,'F_0') = s(bool,'const.quotient_sum.+++_4'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(sK2,fun(sK2,bool)),sK4),s('type.sum.sum'(sK2,sK2),X334),s('type.sum.sum'(sK2,sK2),X335)))) ) | $spl78),
  inference(cnf_transformation,[],[f8657_D])).
fof(f8657_D,plain,(
  ( ! [X335,X334] : (s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK2),X334))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK2),X335))) | s(bool,'F_0') = s(bool,'const.quotient_sum.+++_4'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(sK2,fun(sK2,bool)),sK4),s('type.sum.sum'(sK2,sK2),X334),s('type.sum.sum'(sK2,sK2),X335)))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f8694,plain,(
  ( ! [X335,X334,X336] : (s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK2),X334))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK2),X335))) | s(bool,'F_0') = s(bool,'const.quotient_sum.+++_4'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(sK2,fun(sK2,bool)),sK4),s('type.sum.sum'(sK2,sK2),X334),s('type.sum.sum'(sK2,sK2),X335))) | s(bool,'F_0') = s(bool,X336)) )),
  inference(subsumption_resolution,[],[f7936,f119])).
fof(f7936,plain,(
  ( ! [X335,X334,X336] : (~p(s(bool,X336)) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK2),X334))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK2),X335))) | s(bool,'F_0') = s(bool,'const.quotient_sum.+++_4'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(sK2,fun(sK2,bool)),sK4),s('type.sum.sum'(sK2,sK2),X334),s('type.sum.sum'(sK2,sK2),X335))) | s(bool,'F_0') = s(bool,X336)) )),
  inference(superposition,[],[f1864,f5187])).
fof(f1864,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.quotient_sum.+++_4'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(sK2,fun(sK2,bool)),sK4),s('type.sum.sum'(sK2,sK2),X0),s('type.sum.sum'(sK2,sK2),X1)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK2),X0))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK2),X1)))) )),
  inference(subsumption_resolution,[],[f1848,f141])).
fof(f1848,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(sK2,fun(sK2,bool)),sK4),s('type.sum.sum'(sK2,sK2),X0),s('type.sum.sum'(sK2,sK2),X1)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK2),X0))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK2),X1)))) )),
  inference(superposition,[],[f179,f187])).
fof(f8655,plain,(
  ( ! [X331,X332] : (s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),X331))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),X332))) | s(bool,'F_0') = s(bool,'const.quotient_sum.+++_4'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(sK3,fun(sK3,bool)),sK7),s('type.sum.sum'(sK2,sK3),X331),s('type.sum.sum'(sK2,sK3),X332)))) ) | $spl76),
  inference(cnf_transformation,[],[f8655_D])).
fof(f8655_D,plain,(
  ( ! [X331,X332] : (s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),X331))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),X332))) | s(bool,'F_0') = s(bool,'const.quotient_sum.+++_4'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(sK3,fun(sK3,bool)),sK7),s('type.sum.sum'(sK2,sK3),X331),s('type.sum.sum'(sK2,sK3),X332)))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f8693,plain,(
  ( ! [X333,X331,X332] : (s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),X331))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),X332))) | s(bool,'F_0') = s(bool,'const.quotient_sum.+++_4'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(sK3,fun(sK3,bool)),sK7),s('type.sum.sum'(sK2,sK3),X331),s('type.sum.sum'(sK2,sK3),X332))) | s(bool,'F_0') = s(bool,X333)) )),
  inference(subsumption_resolution,[],[f7935,f119])).
fof(f7935,plain,(
  ( ! [X333,X331,X332] : (~p(s(bool,X333)) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),X331))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),X332))) | s(bool,'F_0') = s(bool,'const.quotient_sum.+++_4'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(sK3,fun(sK3,bool)),sK7),s('type.sum.sum'(sK2,sK3),X331),s('type.sum.sum'(sK2,sK3),X332))) | s(bool,'F_0') = s(bool,X333)) )),
  inference(superposition,[],[f1865,f5187])).
fof(f1865,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.quotient_sum.+++_4'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(sK3,fun(sK3,bool)),sK7),s('type.sum.sum'(sK2,sK3),X2),s('type.sum.sum'(sK2,sK3),X3)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),X2))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),X3)))) )),
  inference(subsumption_resolution,[],[f1849,f141])).
fof(f1849,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(sK3,fun(sK3,bool)),sK7),s('type.sum.sum'(sK2,sK3),X2),s('type.sum.sum'(sK2,sK3),X3)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),X2))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),X3)))) )),
  inference(superposition,[],[f179,f308])).
fof(f8692,plain,(
  ( ! [X325,X327,X321,X323,X329,X318,X324,X326,X320,X322,X328,X317,X319] : (s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X317,X319),X321))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X317,X319),X322))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X319,fun(X319,bool)),X320),s(fun(X319,X324),X325),s(fun(X324,X319),X326)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X317,fun(X317,bool)),X318),s(fun(X317,X327),X328),s(fun(X327,X317),X329)))) | s(bool,'F_0') = s(bool,'const.quotient_sum.+++_4'(s(fun(X317,fun(X317,bool)),X318),s(fun(X319,fun(X319,bool)),X320),s('type.sum.sum'(X317,X319),X321),s('type.sum.sum'(X317,X319),X322))) | s(bool,'F_0') = s(bool,X323)) )),
  inference(subsumption_resolution,[],[f7933,f119])).
fof(f7933,plain,(
  ( ! [X325,X327,X321,X323,X329,X318,X324,X326,X320,X322,X328,X317,X319] : (~p(s(bool,X323)) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X317,X319),X321))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X317,X319),X322))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X319,fun(X319,bool)),X320),s(fun(X319,X324),X325),s(fun(X324,X319),X326)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X317,fun(X317,bool)),X318),s(fun(X317,X327),X328),s(fun(X327,X317),X329)))) | s(bool,'F_0') = s(bool,'const.quotient_sum.+++_4'(s(fun(X317,fun(X317,bool)),X318),s(fun(X319,fun(X319,bool)),X320),s('type.sum.sum'(X317,X319),X321),s('type.sum.sum'(X317,X319),X322))) | s(bool,'F_0') = s(bool,X323)) )),
  inference(superposition,[],[f130,f5187])).
fof(f8691,plain,(
  ( ! [X302,X304,X306,X301,X303,X305,X307] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X303,sK1),X307))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X301,sK3),'const.sum.++_3'(s(fun(X303,X301),X305),s(fun(sK1,sK3),sK9),s('type.sum.sum'(X303,sK1),X307))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X301,fun(X301,bool)),X302),s(fun(X301,X303),X304),s(fun(X303,X301),X305))) | s(bool,'F_0') = s(bool,X306)) )),
  inference(subsumption_resolution,[],[f7929,f119])).
fof(f7929,plain,(
  ( ! [X302,X304,X306,X301,X303,X305,X307] : (~p(s(bool,X306)) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X303,sK1),X307))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X301,sK3),'const.sum.++_3'(s(fun(X303,X301),X305),s(fun(sK1,sK3),sK9),s('type.sum.sum'(X303,sK1),X307))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X301,fun(X301,bool)),X302),s(fun(X301,X303),X304),s(fun(X303,X301),X305))) | s(bool,'F_0') = s(bool,X306)) )),
  inference(superposition,[],[f303,f5187])).
fof(f8690,plain,(
  ( ! [X294,X300,X296,X298,X295,X297,X299] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK1,X296),X300))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK3,X294),'const.sum.++_3'(s(fun(sK1,sK3),sK9),s(fun(X296,X294),X298),s('type.sum.sum'(sK1,X296),X300))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X294,fun(X294,bool)),X295),s(fun(X294,X296),X297),s(fun(X296,X294),X298))) | s(bool,'F_0') = s(bool,X299)) )),
  inference(subsumption_resolution,[],[f7928,f119])).
fof(f7928,plain,(
  ( ! [X294,X300,X296,X298,X295,X297,X299] : (~p(s(bool,X299)) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK1,X296),X300))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK3,X294),'const.sum.++_3'(s(fun(sK1,sK3),sK9),s(fun(X296,X294),X298),s('type.sum.sum'(sK1,X296),X300))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X294,fun(X294,bool)),X295),s(fun(X294,X296),X297),s(fun(X296,X294),X298))) | s(bool,'F_0') = s(bool,X299)) )),
  inference(superposition,[],[f302,f5187])).
fof(f8689,plain,(
  ( ! [X292,X288,X290,X287,X293,X289,X291,X286] : (~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X286,fun(X286,bool)),X287),s(fun(sK3,fun(sK3,bool)),sK7),s('type.sum.sum'(X286,sK3),X292),s('type.sum.sum'(X286,sK3),X293)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X286,sK3),X292))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X286,sK3),X293))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X286,fun(X286,bool)),X287),s(fun(X286,X288),X289),s(fun(X288,X286),X290))) | s(bool,'F_0') = s(bool,X291)) )),
  inference(subsumption_resolution,[],[f7927,f119])).
fof(f7927,plain,(
  ( ! [X292,X288,X290,X287,X293,X289,X291,X286] : (~p(s(bool,X291)) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X286,fun(X286,bool)),X287),s(fun(sK3,fun(sK3,bool)),sK7),s('type.sum.sum'(X286,sK3),X292),s('type.sum.sum'(X286,sK3),X293)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X286,sK3),X292))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X286,sK3),X293))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X286,fun(X286,bool)),X287),s(fun(X286,X288),X289),s(fun(X288,X286),X290))) | s(bool,'F_0') = s(bool,X291)) )),
  inference(superposition,[],[f301,f5187])).
fof(f8688,plain,(
  ( ! [X280,X282,X279,X285,X281,X283,X278,X284] : (~p(s(bool,'const.quotient_sum.+++_4'(s(fun(sK3,fun(sK3,bool)),sK7),s(fun(X278,fun(X278,bool)),X279),s('type.sum.sum'(sK3,X278),X284),s('type.sum.sum'(sK3,X278),X285)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK3,X278),X284))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK3,X278),X285))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X278,fun(X278,bool)),X279),s(fun(X278,X280),X281),s(fun(X280,X278),X282))) | s(bool,'F_0') = s(bool,X283)) )),
  inference(subsumption_resolution,[],[f7926,f119])).
fof(f7926,plain,(
  ( ! [X280,X282,X279,X285,X281,X283,X278,X284] : (~p(s(bool,X283)) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(sK3,fun(sK3,bool)),sK7),s(fun(X278,fun(X278,bool)),X279),s('type.sum.sum'(sK3,X278),X284),s('type.sum.sum'(sK3,X278),X285)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK3,X278),X284))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK3,X278),X285))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X278,fun(X278,bool)),X279),s(fun(X278,X280),X281),s(fun(X280,X278),X282))) | s(bool,'F_0') = s(bool,X283)) )),
  inference(superposition,[],[f300,f5187])).
fof(f8687,plain,(
  ( ! [X271,X277,X273,X275,X276,X272,X274] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X273,sK0),X277))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X271,sK2),'const.sum.++_3'(s(fun(X273,X271),X275),s(fun(sK0,sK2),sK6),s('type.sum.sum'(X273,sK0),X277))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X271,fun(X271,bool)),X272),s(fun(X271,X273),X274),s(fun(X273,X271),X275))) | s(bool,'F_0') = s(bool,X276)) )),
  inference(subsumption_resolution,[],[f7925,f119])).
fof(f7925,plain,(
  ( ! [X271,X277,X273,X275,X276,X272,X274] : (~p(s(bool,X276)) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X273,sK0),X277))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X271,sK2),'const.sum.++_3'(s(fun(X273,X271),X275),s(fun(sK0,sK2),sK6),s('type.sum.sum'(X273,sK0),X277))))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X271,fun(X271,bool)),X272),s(fun(X271,X273),X274),s(fun(X273,X271),X275))) | s(bool,'F_0') = s(bool,X276)) )),
  inference(superposition,[],[f182,f5187])).
fof(f8686,plain,(
  ( ! [X261,X263,X257,X259,X260,X262,X256,X258] : (~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X256,fun(X256,bool)),X257),s(fun(sK2,fun(sK2,bool)),sK4),s('type.sum.sum'(X256,sK2),X262),s('type.sum.sum'(X256,sK2),X263)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X256,sK2),X262))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X256,sK2),X263))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X256,fun(X256,bool)),X257),s(fun(X256,X258),X259),s(fun(X258,X256),X260))) | s(bool,'F_0') = s(bool,X261)) )),
  inference(subsumption_resolution,[],[f7923,f119])).
fof(f7923,plain,(
  ( ! [X261,X263,X257,X259,X260,X262,X256,X258] : (~p(s(bool,X261)) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X256,fun(X256,bool)),X257),s(fun(sK2,fun(sK2,bool)),sK4),s('type.sum.sum'(X256,sK2),X262),s('type.sum.sum'(X256,sK2),X263)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X256,sK2),X262))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X256,sK2),X263))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X256,fun(X256,bool)),X257),s(fun(X256,X258),X259),s(fun(X258,X256),X260))) | s(bool,'F_0') = s(bool,X261)) )),
  inference(superposition,[],[f180,f5187])).
fof(f8685,plain,(
  ( ! [X255,X253,X251,X249,X254,X252,X250,X248] : (~p(s(bool,'const.quotient_sum.+++_4'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(X248,fun(X248,bool)),X249),s('type.sum.sum'(sK2,X248),X254),s('type.sum.sum'(sK2,X248),X255)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,X248),X254))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,X248),X255))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X248,fun(X248,bool)),X249),s(fun(X248,X250),X251),s(fun(X250,X248),X252))) | s(bool,'F_0') = s(bool,X253)) )),
  inference(subsumption_resolution,[],[f7922,f119])).
fof(f7922,plain,(
  ( ! [X255,X253,X251,X249,X254,X252,X250,X248] : (~p(s(bool,X253)) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(X248,fun(X248,bool)),X249),s('type.sum.sum'(sK2,X248),X254),s('type.sum.sum'(sK2,X248),X255)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,X248),X254))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,X248),X255))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X248,fun(X248,bool)),X249),s(fun(X248,X250),X251),s(fun(X250,X248),X252))) | s(bool,'F_0') = s(bool,X253)) )),
  inference(superposition,[],[f179,f5187])).
fof(f8684,plain,(
  ( ! [X243,X241,X239,X237,X235,X246,X244,X242,X240,X238,X236,X247,X245] : (s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X241,X235),X242))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X241,X235),X243))) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X241,fun(X241,bool)),X244),s(fun(X235,fun(X235,bool)),X236),s('type.sum.sum'(X241,X235),X242),s('type.sum.sum'(X241,X235),X243)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X241,fun(X241,bool)),X244),s(fun(X241,X245),X246),s(fun(X245,X241),X247)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X235,fun(X235,bool)),X236),s(fun(X235,X237),X238),s(fun(X237,X235),X239))) | s(bool,'F_0') = s(bool,X240)) )),
  inference(subsumption_resolution,[],[f7921,f119])).
fof(f7921,plain,(
  ( ! [X243,X241,X239,X237,X235,X246,X244,X242,X240,X238,X236,X247,X245] : (~p(s(bool,X240)) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X241,X235),X242))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X241,X235),X243))) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X241,fun(X241,bool)),X244),s(fun(X235,fun(X235,bool)),X236),s('type.sum.sum'(X241,X235),X242),s('type.sum.sum'(X241,X235),X243)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X241,fun(X241,bool)),X244),s(fun(X241,X245),X246),s(fun(X245,X241),X247)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X235,fun(X235,bool)),X236),s(fun(X235,X237),X238),s(fun(X237,X235),X239))) | s(bool,'F_0') = s(bool,X240)) )),
  inference(superposition,[],[f130,f5187])).
fof(f8683,plain,(
  ( ! [X222,X231,X229,X227,X225,X233,X223,X230,X228,X226,X224,X234,X232] : (s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X222,X228),X229))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X222,X228),X230))) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X222,fun(X222,bool)),X223),s(fun(X228,fun(X228,bool)),X231),s('type.sum.sum'(X222,X228),X229),s('type.sum.sum'(X222,X228),X230)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X228,fun(X228,bool)),X231),s(fun(X228,X232),X233),s(fun(X232,X228),X234)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X222,fun(X222,bool)),X223),s(fun(X222,X224),X225),s(fun(X224,X222),X226))) | s(bool,'F_0') = s(bool,X227)) )),
  inference(subsumption_resolution,[],[f7920,f119])).
fof(f7920,plain,(
  ( ! [X222,X231,X229,X227,X225,X233,X223,X230,X228,X226,X224,X234,X232] : (~p(s(bool,X227)) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X222,X228),X229))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X222,X228),X230))) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X222,fun(X222,bool)),X223),s(fun(X228,fun(X228,bool)),X231),s('type.sum.sum'(X222,X228),X229),s('type.sum.sum'(X222,X228),X230)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X228,fun(X228,bool)),X231),s(fun(X228,X232),X233),s(fun(X232,X228),X234)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X222,fun(X222,bool)),X223),s(fun(X222,X224),X225),s(fun(X224,X222),X226))) | s(bool,'F_0') = s(bool,X227)) )),
  inference(superposition,[],[f130,f5187])).
fof(f8682,plain,(
  ( ! [X212,X218,X216,X210,X220,X215,X213,X219,X217,X211,X221,X214] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X216,X212),X217))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X218,X210),'const.sum.++_3'(s(fun(X216,X218),X219),s(fun(X212,X210),X214),s('type.sum.sum'(X216,X212),X217))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X218,fun(X218,bool)),X220),s(fun(X218,X216),X221),s(fun(X216,X218),X219)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X210,fun(X210,bool)),X211),s(fun(X210,X212),X213),s(fun(X212,X210),X214))) | s(bool,'F_0') = s(bool,X215)) )),
  inference(subsumption_resolution,[],[f7919,f119])).
fof(f7919,plain,(
  ( ! [X212,X218,X216,X210,X220,X215,X213,X219,X217,X211,X221,X214] : (~p(s(bool,X215)) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X216,X212),X217))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X218,X210),'const.sum.++_3'(s(fun(X216,X218),X219),s(fun(X212,X210),X214),s('type.sum.sum'(X216,X212),X217))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X218,fun(X218,bool)),X220),s(fun(X218,X216),X221),s(fun(X216,X218),X219)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X210,fun(X210,bool)),X211),s(fun(X210,X212),X213),s(fun(X212,X210),X214))) | s(bool,'F_0') = s(bool,X215)) )),
  inference(superposition,[],[f127,f5187])).
fof(f8681,plain,(
  ( ! [X208,X198,X206,X204,X202,X200,X209,X199,X207,X205,X203,X201] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X200,X204),X205))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X198,X206),'const.sum.++_3'(s(fun(X200,X198),X202),s(fun(X204,X206),X207),s('type.sum.sum'(X200,X204),X205))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X206,fun(X206,bool)),X208),s(fun(X206,X204),X209),s(fun(X204,X206),X207)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X198,fun(X198,bool)),X199),s(fun(X198,X200),X201),s(fun(X200,X198),X202))) | s(bool,'F_0') = s(bool,X203)) )),
  inference(subsumption_resolution,[],[f7918,f119])).
fof(f7918,plain,(
  ( ! [X208,X198,X206,X204,X202,X200,X209,X199,X207,X205,X203,X201] : (~p(s(bool,X203)) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X200,X204),X205))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X198,X206),'const.sum.++_3'(s(fun(X200,X198),X202),s(fun(X204,X206),X207),s('type.sum.sum'(X200,X204),X205))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X206,fun(X206,bool)),X208),s(fun(X206,X204),X209),s(fun(X204,X206),X207)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X198,fun(X198,bool)),X199),s(fun(X198,X200),X201),s(fun(X200,X198),X202))) | s(bool,'F_0') = s(bool,X203)) )),
  inference(superposition,[],[f127,f5187])).
fof(f8624,plain,(
  ( ! [X91,X89] : (s(bool,X89) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X91))) | s(bool,'F_0') = s(bool,X89) | s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X91)))) ) | $spl50),
  inference(cnf_transformation,[],[f8624_D])).
fof(f8624_D,plain,(
  ( ! [X91,X89] : (s(bool,X89) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X91))) | s(bool,'F_0') = s(bool,X89) | s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X91)))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f8675,plain,(
  ( ! [X90,X91,X89] : (s(bool,X89) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X91))) | s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X91))) | s(bool,'F_0') = s(bool,X89) | s(bool,'F_0') = s(bool,X90)) )),
  inference(subsumption_resolution,[],[f7878,f119])).
fof(f7878,plain,(
  ( ! [X90,X91,X89] : (~p(s(bool,X90)) | s(bool,X89) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X91))) | s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X91))) | s(bool,'F_0') = s(bool,X89) | s(bool,'F_0') = s(bool,X90)) )),
  inference(superposition,[],[f4569,f5187])).
fof(f4569,plain,(
  ( ! [X52,X53] : (~p(s(bool,X53)) | s(bool,X53) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X52))) | s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X52)))) )),
  inference(trivial_inequality_removal,[],[f4568])).
fof(f4568,plain,(
  ( ! [X52,X53] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,X53) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X52))) | ~p(s(bool,X53)) | s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X52)))) )),
  inference(superposition,[],[f4132,f126])).
fof(f8674,plain,(
  ( ! [X80,X81,X82] : (s(bool,X80) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X82))) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X82))) | s(bool,'F_0') = s(bool,X80) | s(bool,'F_0') = s(bool,X81)) )),
  inference(subsumption_resolution,[],[f7875,f119])).
fof(f7875,plain,(
  ( ! [X80,X81,X82] : (~p(s(bool,X81)) | s(bool,X80) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X82))) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X82))) | s(bool,'F_0') = s(bool,X80) | s(bool,'F_0') = s(bool,X81)) )),
  inference(superposition,[],[f3539,f5187])).
fof(f3539,plain,(
  ( ! [X6,X5] : (~p(s(bool,X6)) | s(bool,X6) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X5))) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X5)))) )),
  inference(resolution,[],[f3517,f115])).
fof(f8660,plain,(
  ( ! [X341,X340,X342] : (s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK3,sK3),X340))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK3,sK3),X341))) | s(bool,'F_0') = s(bool,X342) | s(bool,'F_0') = s(bool,'const.quotient_sum.+++_4'(s(fun(sK3,fun(sK3,bool)),sK7),s(fun(sK3,fun(sK3,bool)),sK7),s('type.sum.sum'(sK3,sK3),X340),s('type.sum.sum'(sK3,sK3),X341)))) )),
  inference(subsumption_resolution,[],[f7773,f119])).
fof(f7773,plain,(
  ( ! [X341,X340,X342] : (~p(s(bool,X342)) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK3,sK3),X340))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK3,sK3),X341))) | s(bool,'F_0') = s(bool,X342) | s(bool,'F_0') = s(bool,'const.quotient_sum.+++_4'(s(fun(sK3,fun(sK3,bool)),sK7),s(fun(sK3,fun(sK3,bool)),sK7),s('type.sum.sum'(sK3,sK3),X340),s('type.sum.sum'(sK3,sK3),X341)))) )),
  inference(superposition,[],[f2094,f5187])).
fof(f8658,plain,(
  ( ! [X337,X339,X338] : (s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK3,sK2),X337))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK3,sK2),X338))) | s(bool,'F_0') = s(bool,X339) | s(bool,'F_0') = s(bool,'const.quotient_sum.+++_4'(s(fun(sK3,fun(sK3,bool)),sK7),s(fun(sK2,fun(sK2,bool)),sK4),s('type.sum.sum'(sK3,sK2),X337),s('type.sum.sum'(sK3,sK2),X338)))) )),
  inference(subsumption_resolution,[],[f7772,f119])).
fof(f7772,plain,(
  ( ! [X337,X339,X338] : (~p(s(bool,X339)) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK3,sK2),X337))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK3,sK2),X338))) | s(bool,'F_0') = s(bool,X339) | s(bool,'F_0') = s(bool,'const.quotient_sum.+++_4'(s(fun(sK3,fun(sK3,bool)),sK7),s(fun(sK2,fun(sK2,bool)),sK4),s('type.sum.sum'(sK3,sK2),X337),s('type.sum.sum'(sK3,sK2),X338)))) )),
  inference(superposition,[],[f1927,f5187])).
fof(f8656,plain,(
  ( ! [X335,X334,X336] : (s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK2),X334))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK2),X335))) | s(bool,'F_0') = s(bool,X336) | s(bool,'F_0') = s(bool,'const.quotient_sum.+++_4'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(sK2,fun(sK2,bool)),sK4),s('type.sum.sum'(sK2,sK2),X334),s('type.sum.sum'(sK2,sK2),X335)))) )),
  inference(subsumption_resolution,[],[f7771,f119])).
fof(f7771,plain,(
  ( ! [X335,X334,X336] : (~p(s(bool,X336)) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK2),X334))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK2),X335))) | s(bool,'F_0') = s(bool,X336) | s(bool,'F_0') = s(bool,'const.quotient_sum.+++_4'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(sK2,fun(sK2,bool)),sK4),s('type.sum.sum'(sK2,sK2),X334),s('type.sum.sum'(sK2,sK2),X335)))) )),
  inference(superposition,[],[f1864,f5187])).
fof(f8654,plain,(
  ( ! [X333,X331,X332] : (s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),X331))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),X332))) | s(bool,'F_0') = s(bool,X333) | s(bool,'F_0') = s(bool,'const.quotient_sum.+++_4'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(sK3,fun(sK3,bool)),sK7),s('type.sum.sum'(sK2,sK3),X331),s('type.sum.sum'(sK2,sK3),X332)))) )),
  inference(subsumption_resolution,[],[f7770,f119])).
fof(f7770,plain,(
  ( ! [X333,X331,X332] : (~p(s(bool,X333)) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),X331))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),X332))) | s(bool,'F_0') = s(bool,X333) | s(bool,'F_0') = s(bool,'const.quotient_sum.+++_4'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(sK3,fun(sK3,bool)),sK7),s('type.sum.sum'(sK2,sK3),X331),s('type.sum.sum'(sK2,sK3),X332)))) )),
  inference(superposition,[],[f1865,f5187])).
fof(f8652,plain,(
  ( ! [X325,X327,X321,X323,X329,X318,X324,X326,X320,X322,X328,X317,X319] : (s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X317,X319),X321))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X317,X319),X322))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X319,fun(X319,bool)),X320),s(fun(X319,X324),X325),s(fun(X324,X319),X326)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X317,fun(X317,bool)),X318),s(fun(X317,X327),X328),s(fun(X327,X317),X329)))) | s(bool,'F_0') = s(bool,X323) | s(bool,'F_0') = s(bool,'const.quotient_sum.+++_4'(s(fun(X317,fun(X317,bool)),X318),s(fun(X319,fun(X319,bool)),X320),s('type.sum.sum'(X317,X319),X321),s('type.sum.sum'(X317,X319),X322)))) )),
  inference(subsumption_resolution,[],[f7768,f119])).
fof(f7768,plain,(
  ( ! [X325,X327,X321,X323,X329,X318,X324,X326,X320,X322,X328,X317,X319] : (~p(s(bool,X323)) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X317,X319),X321))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X317,X319),X322))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X319,fun(X319,bool)),X320),s(fun(X319,X324),X325),s(fun(X324,X319),X326)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X317,fun(X317,bool)),X318),s(fun(X317,X327),X328),s(fun(X327,X317),X329)))) | s(bool,'F_0') = s(bool,X323) | s(bool,'F_0') = s(bool,'const.quotient_sum.+++_4'(s(fun(X317,fun(X317,bool)),X318),s(fun(X319,fun(X319,bool)),X320),s('type.sum.sum'(X317,X319),X321),s('type.sum.sum'(X317,X319),X322)))) )),
  inference(superposition,[],[f130,f5187])).
fof(f8650,plain,(
  ( ! [X302,X304,X306,X301,X303,X305,X307] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X303,sK1),X307))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X301,sK3),'const.sum.++_3'(s(fun(X303,X301),X305),s(fun(sK1,sK3),sK9),s('type.sum.sum'(X303,sK1),X307))))) | s(bool,'F_0') = s(bool,X306) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X301,fun(X301,bool)),X302),s(fun(X301,X303),X304),s(fun(X303,X301),X305)))) )),
  inference(subsumption_resolution,[],[f7764,f119])).
fof(f7764,plain,(
  ( ! [X302,X304,X306,X301,X303,X305,X307] : (~p(s(bool,X306)) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X303,sK1),X307))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X301,sK3),'const.sum.++_3'(s(fun(X303,X301),X305),s(fun(sK1,sK3),sK9),s('type.sum.sum'(X303,sK1),X307))))) | s(bool,'F_0') = s(bool,X306) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X301,fun(X301,bool)),X302),s(fun(X301,X303),X304),s(fun(X303,X301),X305)))) )),
  inference(superposition,[],[f303,f5187])).
fof(f8648,plain,(
  ( ! [X294,X300,X296,X298,X295,X297,X299] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK1,X296),X300))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK3,X294),'const.sum.++_3'(s(fun(sK1,sK3),sK9),s(fun(X296,X294),X298),s('type.sum.sum'(sK1,X296),X300))))) | s(bool,'F_0') = s(bool,X299) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X294,fun(X294,bool)),X295),s(fun(X294,X296),X297),s(fun(X296,X294),X298)))) )),
  inference(subsumption_resolution,[],[f7763,f119])).
fof(f7763,plain,(
  ( ! [X294,X300,X296,X298,X295,X297,X299] : (~p(s(bool,X299)) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK1,X296),X300))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK3,X294),'const.sum.++_3'(s(fun(sK1,sK3),sK9),s(fun(X296,X294),X298),s('type.sum.sum'(sK1,X296),X300))))) | s(bool,'F_0') = s(bool,X299) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X294,fun(X294,bool)),X295),s(fun(X294,X296),X297),s(fun(X296,X294),X298)))) )),
  inference(superposition,[],[f302,f5187])).
fof(f8646,plain,(
  ( ! [X292,X288,X290,X287,X293,X289,X291,X286] : (~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X286,fun(X286,bool)),X287),s(fun(sK3,fun(sK3,bool)),sK7),s('type.sum.sum'(X286,sK3),X292),s('type.sum.sum'(X286,sK3),X293)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X286,sK3),X292))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X286,sK3),X293))) | s(bool,'F_0') = s(bool,X291) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X286,fun(X286,bool)),X287),s(fun(X286,X288),X289),s(fun(X288,X286),X290)))) )),
  inference(subsumption_resolution,[],[f7762,f119])).
fof(f7762,plain,(
  ( ! [X292,X288,X290,X287,X293,X289,X291,X286] : (~p(s(bool,X291)) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X286,fun(X286,bool)),X287),s(fun(sK3,fun(sK3,bool)),sK7),s('type.sum.sum'(X286,sK3),X292),s('type.sum.sum'(X286,sK3),X293)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X286,sK3),X292))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X286,sK3),X293))) | s(bool,'F_0') = s(bool,X291) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X286,fun(X286,bool)),X287),s(fun(X286,X288),X289),s(fun(X288,X286),X290)))) )),
  inference(superposition,[],[f301,f5187])).
fof(f8644,plain,(
  ( ! [X280,X282,X279,X285,X281,X283,X278,X284] : (~p(s(bool,'const.quotient_sum.+++_4'(s(fun(sK3,fun(sK3,bool)),sK7),s(fun(X278,fun(X278,bool)),X279),s('type.sum.sum'(sK3,X278),X284),s('type.sum.sum'(sK3,X278),X285)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK3,X278),X284))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK3,X278),X285))) | s(bool,'F_0') = s(bool,X283) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X278,fun(X278,bool)),X279),s(fun(X278,X280),X281),s(fun(X280,X278),X282)))) )),
  inference(subsumption_resolution,[],[f7761,f119])).
fof(f7761,plain,(
  ( ! [X280,X282,X279,X285,X281,X283,X278,X284] : (~p(s(bool,X283)) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(sK3,fun(sK3,bool)),sK7),s(fun(X278,fun(X278,bool)),X279),s('type.sum.sum'(sK3,X278),X284),s('type.sum.sum'(sK3,X278),X285)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK3,X278),X284))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK3,X278),X285))) | s(bool,'F_0') = s(bool,X283) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X278,fun(X278,bool)),X279),s(fun(X278,X280),X281),s(fun(X280,X278),X282)))) )),
  inference(superposition,[],[f300,f5187])).
fof(f8642,plain,(
  ( ! [X271,X277,X273,X275,X276,X272,X274] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X273,sK0),X277))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X271,sK2),'const.sum.++_3'(s(fun(X273,X271),X275),s(fun(sK0,sK2),sK6),s('type.sum.sum'(X273,sK0),X277))))) | s(bool,'F_0') = s(bool,X276) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X271,fun(X271,bool)),X272),s(fun(X271,X273),X274),s(fun(X273,X271),X275)))) )),
  inference(subsumption_resolution,[],[f7760,f119])).
fof(f7760,plain,(
  ( ! [X271,X277,X273,X275,X276,X272,X274] : (~p(s(bool,X276)) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X273,sK0),X277))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X271,sK2),'const.sum.++_3'(s(fun(X273,X271),X275),s(fun(sK0,sK2),sK6),s('type.sum.sum'(X273,sK0),X277))))) | s(bool,'F_0') = s(bool,X276) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X271,fun(X271,bool)),X272),s(fun(X271,X273),X274),s(fun(X273,X271),X275)))) )),
  inference(superposition,[],[f182,f5187])).
fof(f8640,plain,(
  ( ! [X261,X263,X257,X259,X260,X262,X256,X258] : (~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X256,fun(X256,bool)),X257),s(fun(sK2,fun(sK2,bool)),sK4),s('type.sum.sum'(X256,sK2),X262),s('type.sum.sum'(X256,sK2),X263)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X256,sK2),X262))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X256,sK2),X263))) | s(bool,'F_0') = s(bool,X261) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X256,fun(X256,bool)),X257),s(fun(X256,X258),X259),s(fun(X258,X256),X260)))) )),
  inference(subsumption_resolution,[],[f7758,f119])).
fof(f7758,plain,(
  ( ! [X261,X263,X257,X259,X260,X262,X256,X258] : (~p(s(bool,X261)) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X256,fun(X256,bool)),X257),s(fun(sK2,fun(sK2,bool)),sK4),s('type.sum.sum'(X256,sK2),X262),s('type.sum.sum'(X256,sK2),X263)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X256,sK2),X262))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X256,sK2),X263))) | s(bool,'F_0') = s(bool,X261) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X256,fun(X256,bool)),X257),s(fun(X256,X258),X259),s(fun(X258,X256),X260)))) )),
  inference(superposition,[],[f180,f5187])).
fof(f8638,plain,(
  ( ! [X255,X253,X251,X249,X254,X252,X250,X248] : (~p(s(bool,'const.quotient_sum.+++_4'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(X248,fun(X248,bool)),X249),s('type.sum.sum'(sK2,X248),X254),s('type.sum.sum'(sK2,X248),X255)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,X248),X254))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,X248),X255))) | s(bool,'F_0') = s(bool,X253) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X248,fun(X248,bool)),X249),s(fun(X248,X250),X251),s(fun(X250,X248),X252)))) )),
  inference(subsumption_resolution,[],[f7757,f119])).
fof(f7757,plain,(
  ( ! [X255,X253,X251,X249,X254,X252,X250,X248] : (~p(s(bool,X253)) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(X248,fun(X248,bool)),X249),s('type.sum.sum'(sK2,X248),X254),s('type.sum.sum'(sK2,X248),X255)))) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,X248),X254))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,X248),X255))) | s(bool,'F_0') = s(bool,X253) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X248,fun(X248,bool)),X249),s(fun(X248,X250),X251),s(fun(X250,X248),X252)))) )),
  inference(superposition,[],[f179,f5187])).
fof(f8636,plain,(
  ( ! [X243,X241,X239,X237,X235,X246,X244,X242,X240,X238,X236,X247,X245] : (s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X241,X235),X242))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X241,X235),X243))) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X241,fun(X241,bool)),X244),s(fun(X235,fun(X235,bool)),X236),s('type.sum.sum'(X241,X235),X242),s('type.sum.sum'(X241,X235),X243)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X241,fun(X241,bool)),X244),s(fun(X241,X245),X246),s(fun(X245,X241),X247)))) | s(bool,'F_0') = s(bool,X240) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X235,fun(X235,bool)),X236),s(fun(X235,X237),X238),s(fun(X237,X235),X239)))) )),
  inference(subsumption_resolution,[],[f7756,f119])).
fof(f7756,plain,(
  ( ! [X243,X241,X239,X237,X235,X246,X244,X242,X240,X238,X236,X247,X245] : (~p(s(bool,X240)) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X241,X235),X242))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X241,X235),X243))) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X241,fun(X241,bool)),X244),s(fun(X235,fun(X235,bool)),X236),s('type.sum.sum'(X241,X235),X242),s('type.sum.sum'(X241,X235),X243)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X241,fun(X241,bool)),X244),s(fun(X241,X245),X246),s(fun(X245,X241),X247)))) | s(bool,'F_0') = s(bool,X240) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X235,fun(X235,bool)),X236),s(fun(X235,X237),X238),s(fun(X237,X235),X239)))) )),
  inference(superposition,[],[f130,f5187])).
fof(f8634,plain,(
  ( ! [X222,X231,X229,X227,X225,X233,X223,X230,X228,X226,X224,X234,X232] : (s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X222,X228),X229))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X222,X228),X230))) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X222,fun(X222,bool)),X223),s(fun(X228,fun(X228,bool)),X231),s('type.sum.sum'(X222,X228),X229),s('type.sum.sum'(X222,X228),X230)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X228,fun(X228,bool)),X231),s(fun(X228,X232),X233),s(fun(X232,X228),X234)))) | s(bool,'F_0') = s(bool,X227) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X222,fun(X222,bool)),X223),s(fun(X222,X224),X225),s(fun(X224,X222),X226)))) )),
  inference(subsumption_resolution,[],[f7755,f119])).
fof(f7755,plain,(
  ( ! [X222,X231,X229,X227,X225,X233,X223,X230,X228,X226,X224,X234,X232] : (~p(s(bool,X227)) | s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X222,X228),X229))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(X222,X228),X230))) | ~p(s(bool,'const.quotient_sum.+++_4'(s(fun(X222,fun(X222,bool)),X223),s(fun(X228,fun(X228,bool)),X231),s('type.sum.sum'(X222,X228),X229),s('type.sum.sum'(X222,X228),X230)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X228,fun(X228,bool)),X231),s(fun(X228,X232),X233),s(fun(X232,X228),X234)))) | s(bool,'F_0') = s(bool,X227) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X222,fun(X222,bool)),X223),s(fun(X222,X224),X225),s(fun(X224,X222),X226)))) )),
  inference(superposition,[],[f130,f5187])).
fof(f8632,plain,(
  ( ! [X212,X218,X216,X210,X220,X215,X213,X219,X217,X211,X221,X214] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X216,X212),X217))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X218,X210),'const.sum.++_3'(s(fun(X216,X218),X219),s(fun(X212,X210),X214),s('type.sum.sum'(X216,X212),X217))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X218,fun(X218,bool)),X220),s(fun(X218,X216),X221),s(fun(X216,X218),X219)))) | s(bool,'F_0') = s(bool,X215) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X210,fun(X210,bool)),X211),s(fun(X210,X212),X213),s(fun(X212,X210),X214)))) )),
  inference(subsumption_resolution,[],[f7754,f119])).
fof(f7754,plain,(
  ( ! [X212,X218,X216,X210,X220,X215,X213,X219,X217,X211,X221,X214] : (~p(s(bool,X215)) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X216,X212),X217))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X218,X210),'const.sum.++_3'(s(fun(X216,X218),X219),s(fun(X212,X210),X214),s('type.sum.sum'(X216,X212),X217))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X218,fun(X218,bool)),X220),s(fun(X218,X216),X221),s(fun(X216,X218),X219)))) | s(bool,'F_0') = s(bool,X215) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X210,fun(X210,bool)),X211),s(fun(X210,X212),X213),s(fun(X212,X210),X214)))) )),
  inference(superposition,[],[f127,f5187])).
fof(f8630,plain,(
  ( ! [X208,X198,X206,X204,X202,X200,X209,X199,X207,X205,X203,X201] : (s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X200,X204),X205))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X198,X206),'const.sum.++_3'(s(fun(X200,X198),X202),s(fun(X204,X206),X207),s('type.sum.sum'(X200,X204),X205))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X206,fun(X206,bool)),X208),s(fun(X206,X204),X209),s(fun(X204,X206),X207)))) | s(bool,'F_0') = s(bool,X203) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X198,fun(X198,bool)),X199),s(fun(X198,X200),X201),s(fun(X200,X198),X202)))) )),
  inference(subsumption_resolution,[],[f7753,f119])).
fof(f7753,plain,(
  ( ! [X208,X198,X206,X204,X202,X200,X209,X199,X207,X205,X203,X201] : (~p(s(bool,X203)) | s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X200,X204),X205))) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(X198,X206),'const.sum.++_3'(s(fun(X200,X198),X202),s(fun(X204,X206),X207),s('type.sum.sum'(X200,X204),X205))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X206,fun(X206,bool)),X208),s(fun(X206,X204),X209),s(fun(X204,X206),X207)))) | s(bool,'F_0') = s(bool,X203) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X198,fun(X198,bool)),X199),s(fun(X198,X200),X201),s(fun(X200,X198),X202)))) )),
  inference(superposition,[],[f127,f5187])).
fof(f8623,plain,(
  ( ! [X90,X91,X89] : (s(bool,X89) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X91))) | s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X91))) | s(bool,'F_0') = s(bool,X90) | s(bool,'F_0') = s(bool,X89)) )),
  inference(subsumption_resolution,[],[f7713,f119])).
fof(f7713,plain,(
  ( ! [X90,X91,X89] : (~p(s(bool,X90)) | s(bool,X89) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X91))) | s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X91))) | s(bool,'F_0') = s(bool,X90) | s(bool,'F_0') = s(bool,X89)) )),
  inference(superposition,[],[f4569,f5187])).
fof(f8620,plain,(
  ( ! [X80,X81,X82] : (s(bool,X80) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X82))) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X82))) | s(bool,'F_0') = s(bool,X81) | s(bool,'F_0') = s(bool,X80)) )),
  inference(subsumption_resolution,[],[f7710,f119])).
fof(f7710,plain,(
  ( ! [X80,X81,X82] : (~p(s(bool,X81)) | s(bool,X80) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X82))) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X82))) | s(bool,'F_0') = s(bool,X81) | s(bool,'F_0') = s(bool,X80)) )),
  inference(superposition,[],[f3539,f5187])).
fof(f4655,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f4643,f252])).
fof(f4643,plain,(
  ( ! [X64] : (s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X64)))) ) | $spl16),
  inference(trivial_inequality_removal,[],[f4638])).
fof(f4638,plain,(
  ( ! [X64] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X64)))) ) | $spl16),
  inference(superposition,[],[f252,f1934])).
fof(f4646,plain,(
  $false | $spl16),
  inference(trivial_inequality_removal,[],[f4623])).
fof(f4623,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl16),
  inference(superposition,[],[f252,f1580])).
fof(f4645,plain,(
  $false | ($spl6 | $spl16)),
  inference(trivial_inequality_removal,[],[f4625])).
fof(f4625,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl6 | $spl16)),
  inference(superposition,[],[f252,f171])).
fof(f4644,plain,(
  $false | $spl16),
  inference(trivial_inequality_removal,[],[f4629])).
fof(f4629,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl16),
  inference(superposition,[],[f252,f1599])).
fof(f4642,plain,(
  $false | $spl16),
  inference(trivial_inequality_removal,[],[f4639])).
fof(f4639,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl16),
  inference(superposition,[],[f252,f2351])).
fof(f4641,plain,(
  $false | $spl16),
  inference(equality_resolution,[],[f252])).
fof(f4591,plain,(
  ( ! [X30,X32] : (s(bool,X32) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X30))) | p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X30)))) | ~p(s(bool,X32))) ) | $spl44),
  inference(cnf_transformation,[],[f4591_D])).
fof(f4591_D,plain,(
  ( ! [X30,X32] : (s(bool,X32) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X30))) | p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X30)))) | ~p(s(bool,X32))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f4592,plain,(
  ( ! [X35,X33,X34] : (s(bool,'F_0') != s(bool,X34) | s(bool,X35) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X33))) | ~p(s(bool,X35)) | p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X33))))) )),
  inference(subsumption_resolution,[],[f4560,f118])).
fof(f4560,plain,(
  ( ! [X35,X33,X34] : (s(bool,'F_0') != s(bool,X34) | s(bool,X35) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X33))) | ~p(s(bool,X35)) | p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X33)))) | p(s(bool,X34))) )),
  inference(superposition,[],[f4132,f112])).
fof(f4590,plain,(
  ( ! [X30,X31,X32] : (s(bool,'F_0') != s(bool,X31) | s(bool,X32) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X30))) | ~p(s(bool,X32)) | p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X30))))) )),
  inference(subsumption_resolution,[],[f4559,f118])).
fof(f4559,plain,(
  ( ! [X30,X31,X32] : (s(bool,'F_0') != s(bool,X31) | s(bool,X32) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X30))) | ~p(s(bool,X32)) | p(s(bool,X31)) | p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X30))))) )),
  inference(superposition,[],[f4132,f112])).
fof(f3029,plain,(
  ( ! [X34,X32] : (s(bool,X34) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X32))) | ~p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X32)))) | p(s(bool,X34))) ) | $spl42),
  inference(cnf_transformation,[],[f3029_D])).
fof(f3029_D,plain,(
  ( ! [X34,X32] : (s(bool,X34) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X32))) | ~p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X32)))) | p(s(bool,X34))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f3030,plain,(
  ( ! [X37,X35,X36] : (s(bool,'T_0') != s(bool,X36) | s(bool,X37) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X35))) | p(s(bool,X37)) | ~p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X35))))) )),
  inference(subsumption_resolution,[],[f3016,f122])).
fof(f3016,plain,(
  ( ! [X37,X35,X36] : (s(bool,'T_0') != s(bool,X36) | s(bool,X37) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X35))) | p(s(bool,X37)) | ~p(s(bool,X36)) | ~p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X35))))) )),
  inference(superposition,[],[f2606,f115])).
fof(f3028,plain,(
  ( ! [X33,X34,X32] : (s(bool,'T_0') != s(bool,X33) | s(bool,X34) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X32))) | p(s(bool,X34)) | ~p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X32))))) )),
  inference(subsumption_resolution,[],[f3015,f122])).
fof(f3015,plain,(
  ( ! [X33,X34,X32] : (s(bool,'T_0') != s(bool,X33) | s(bool,X34) = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X32))) | p(s(bool,X34)) | ~p(s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X32)))) | ~p(s(bool,X33))) )),
  inference(superposition,[],[f2606,f115])).
fof(f2764,plain,(
  ( ! [X20,X18] : (s(bool,X20) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X18))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X18)))) | p(s(bool,X20))) ) | $spl40),
  inference(cnf_transformation,[],[f2764_D])).
fof(f2764_D,plain,(
  ( ! [X20,X18] : (s(bool,X20) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X18))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X18)))) | p(s(bool,X20))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f2765,plain,(
  ( ! [X23,X21,X22] : (s(bool,'T_0') != s(bool,X22) | s(bool,X23) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X21))) | p(s(bool,X23)) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X21))))) )),
  inference(subsumption_resolution,[],[f2747,f122])).
fof(f2747,plain,(
  ( ! [X23,X21,X22] : (s(bool,'T_0') != s(bool,X22) | s(bool,X23) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X21))) | p(s(bool,X23)) | ~p(s(bool,X22)) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X21))))) )),
  inference(superposition,[],[f2230,f115])).
fof(f2763,plain,(
  ( ! [X19,X20,X18] : (s(bool,'T_0') != s(bool,X19) | s(bool,X20) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X18))) | p(s(bool,X20)) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X18))))) )),
  inference(subsumption_resolution,[],[f2746,f122])).
fof(f2746,plain,(
  ( ! [X19,X20,X18] : (s(bool,'T_0') != s(bool,X19) | s(bool,X20) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X18))) | p(s(bool,X20)) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X18)))) | ~p(s(bool,X19))) )),
  inference(superposition,[],[f2230,f115])).
fof(f2603,plain,(
  ( ! [X14] : (s(bool,'T_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X14))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X14))))) ) | $spl38),
  inference(cnf_transformation,[],[f2603_D])).
fof(f2603_D,plain,(
  ( ! [X14] : (s(bool,'T_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X14))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X14))))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f2604,plain,(
  ( ! [X17,X16] : (s(bool,'T_0') != s(bool,X17) | s(bool,'T_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X16))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X16))))) )),
  inference(subsumption_resolution,[],[f2590,f122])).
fof(f2590,plain,(
  ( ! [X17,X16] : (s(bool,'T_0') != s(bool,X17) | s(bool,'T_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X16))) | ~p(s(bool,X17)) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X16))))) )),
  inference(superposition,[],[f2234,f115])).
fof(f2602,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') != s(bool,X15) | s(bool,'T_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X14))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X14))))) )),
  inference(subsumption_resolution,[],[f2589,f122])).
fof(f2589,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') != s(bool,X15) | s(bool,'T_0') != s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X14))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X14)))) | ~p(s(bool,X15))) )),
  inference(superposition,[],[f2234,f115])).
fof(f2342,plain,(
  $false | $spl20),
  inference(trivial_inequality_removal,[],[f2322])).
fof(f2322,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl20),
  inference(superposition,[],[f350,f187])).
fof(f2341,plain,(
  $false | $spl20),
  inference(trivial_inequality_removal,[],[f2323])).
fof(f2323,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl20),
  inference(superposition,[],[f350,f308])).
fof(f2340,plain,(
  $false | $spl20),
  inference(trivial_inequality_removal,[],[f2324])).
fof(f2324,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl20),
  inference(superposition,[],[f350,f409])).
fof(f2339,plain,(
  $false | ($spl12 | $spl20)),
  inference(trivial_inequality_removal,[],[f2327])).
fof(f2327,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl12 | $spl20)),
  inference(superposition,[],[f350,f178])).
fof(f2338,plain,(
  $false | ($spl20 | $spl32)),
  inference(trivial_inequality_removal,[],[f2333])).
fof(f2333,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl20 | $spl32)),
  inference(superposition,[],[f350,f771])).
fof(f2337,plain,(
  $false | ($spl20 | $spl26 | $spl32)),
  inference(trivial_inequality_removal,[],[f2334])).
fof(f2334,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl20 | $spl26 | $spl32)),
  inference(superposition,[],[f350,f1035])).
fof(f2336,plain,(
  $false | $spl20),
  inference(equality_resolution,[],[f350])).
fof(f2294,plain,(
  ( ! [X14] : (s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X14))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X14))))) ) | $spl36),
  inference(cnf_transformation,[],[f2294_D])).
fof(f2294_D,plain,(
  ( ! [X14] : (s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X14))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X14))))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f2295,plain,(
  ( ! [X17,X16] : (s(bool,'T_0') != s(bool,X17) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X16))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X16))))) )),
  inference(subsumption_resolution,[],[f2280,f122])).
fof(f2280,plain,(
  ( ! [X17,X16] : (s(bool,'T_0') != s(bool,X17) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X16))) | ~p(s(bool,X17)) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X16))))) )),
  inference(superposition,[],[f2232,f115])).
fof(f2293,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') != s(bool,X15) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X14))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X14))))) )),
  inference(subsumption_resolution,[],[f2279,f122])).
fof(f2279,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') != s(bool,X15) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK0,sK0),X14))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK0,sK0),X14)))) | ~p(s(bool,X15))) )),
  inference(superposition,[],[f2232,f115])).
fof(f1007,plain,(
  $false | ($spl19 | $spl28 | $spl32)),
  inference(subsumption_resolution,[],[f955,f255])).
fof(f955,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl28 | $spl32)),
  inference(backward_demodulation,[],[f771,f767])).
fof(f744,plain,(
  s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) | $spl26),
  inference(superposition,[],[f126,f429])).
fof(f743,plain,(
  s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) | $spl26),
  inference(superposition,[],[f126,f429])).
fof(f716,plain,(
  s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) | $spl26),
  inference(superposition,[],[f429,f126])).
fof(f715,plain,(
  s(bool,'T_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'F_0') = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) | $spl26),
  inference(superposition,[],[f429,f126])).
fof(f677,plain,(
  $false | $spl24),
  inference(resolution,[],[f427,f108])).
fof(f649,plain,(
  $false | $spl22),
  inference(resolution,[],[f426,f107])).
fof(f635,plain,(
  $false | ($spl2 | $spl6)),
  inference(trivial_inequality_removal,[],[f630])).
fof(f630,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl2 | $spl6)),
  inference(superposition,[],[f602,f171])).
fof(f602,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) | ($spl2 | $spl6)),
  inference(subsumption_resolution,[],[f601,f120])).
fof(f601,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,X0))) ) | ($spl2 | $spl6)),
  inference(forward_demodulation,[],[f169,f171])).
fof(f634,plain,(
  $false | ($spl2 | $spl6)),
  inference(equality_resolution,[],[f602])).
fof(f404,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK10))) = s(bool,'const.sum.ISL_1'(s('type.sum.sum'(sK2,sK3),sK11))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK3,fun(sK3,bool)),sK7),s(fun(sK3,X0),X1),s(fun(X0,sK3),X2)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK2,fun(sK2,bool)),sK4),s(fun(sK2,X3),X4),s(fun(X3,sK2),X5))))) )),
  inference(resolution,[],[f109,f130])).
fof(f250,plain,(
  ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10)))) | $spl15),
  inference(cnf_transformation,[],[f250_D])).
fof(f250_D,plain,(
  ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10)))) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f343,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))))) ) | $spl9),
  inference(subsumption_resolution,[],[f332,f122])).
fof(f332,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10)))) | ~p(s(bool,X2))) ) | $spl9),
  inference(superposition,[],[f174,f115])).
fof(f345,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))))) ) | $spl9),
  inference(subsumption_resolution,[],[f333,f122])).
fof(f333,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,X3)) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))))) ) | $spl9),
  inference(superposition,[],[f174,f115])).
fof(f349,plain,(
  $false | ($spl7 | $spl9)),
  inference(subsumption_resolution,[],[f340,f172])).
fof(f340,plain,(
  s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | $spl9),
  inference(trivial_inequality_removal,[],[f338])).
fof(f338,plain,(
  s(bool,'T_0') != s(bool,'T_0') | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | $spl9),
  inference(superposition,[],[f174,f126])).
fof(f348,plain,(
  $false | ($spl7 | $spl9)),
  inference(subsumption_resolution,[],[f341,f247])).
fof(f247,plain,(
  p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10)))) | $spl7),
  inference(trivial_inequality_removal,[],[f241])).
fof(f241,plain,(
  s(bool,'F_0') != s(bool,'F_0') | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10)))) | $spl7),
  inference(superposition,[],[f172,f121])).
fof(f341,plain,(
  ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10)))) | $spl9),
  inference(trivial_inequality_removal,[],[f337])).
fof(f337,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10)))) | $spl9),
  inference(superposition,[],[f174,f125])).
fof(f347,plain,(
  $false | ($spl7 | $spl9)),
  inference(subsumption_resolution,[],[f342,f247])).
fof(f342,plain,(
  ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10)))) | $spl9),
  inference(trivial_inequality_removal,[],[f336])).
fof(f336,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10)))) | $spl9),
  inference(superposition,[],[f174,f123])).
fof(f243,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10)))) | $spl7),
  inference(superposition,[],[f172,f125])).
fof(f242,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10)))) | $spl7),
  inference(superposition,[],[f172,f123])).
fof(f253,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))))) ) | $spl7),
  inference(subsumption_resolution,[],[f237,f118])).
fof(f237,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10)))) | p(s(bool,X1))) ) | $spl7),
  inference(superposition,[],[f172,f112])).
fof(f249,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))))) ) | $spl7),
  inference(subsumption_resolution,[],[f236,f118])).
fof(f236,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,X0)) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))))) ) | $spl7),
  inference(superposition,[],[f172,f112])).
fof(f166,plain,(
  s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'T_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11)))),
  inference(superposition,[],[f110,f126])).
fof(f165,plain,(
  s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | s(bool,'F_0') = s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11)))),
  inference(superposition,[],[f110,f126])).
fof(f164,plain,(
  s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))))),
  inference(superposition,[],[f110,f125])).
fof(f163,plain,(
  s(bool,'T_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))))),
  inference(superposition,[],[f110,f123])).
fof(f162,plain,(
  s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))))),
  inference(superposition,[],[f110,f121])).
fof(f161,plain,(
  s(bool,'F_0') != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))))),
  inference(superposition,[],[f110,f119])).
fof(f160,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | ~p(s(bool,X3)) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))))) )),
  inference(superposition,[],[f110,f115])).
fof(f159,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | ~p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11)))) | ~p(s(bool,X2))) )),
  inference(superposition,[],[f110,f115])).
fof(f158,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11)))) | p(s(bool,X1))) )),
  inference(superposition,[],[f110,f112])).
fof(f157,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK10))) | p(s(bool,X0)) | p(s(bool,'const.sum.ISR_1'(s('type.sum.sum'(sK2,sK3),sK11))))) )),
  inference(superposition,[],[f110,f112])).
% SZS output end Proof for thm.quotient_sum
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 356
% Passive clauses: 5005
% Generated clauses: 35999
% Final active clauses: 145
% Final passive clauses: 3490
% Input formulas: 37
% Initial clauses: 50
% 
% Selected by SInE selection: 28
% SInE iterations: 1
% 
% Duplicate literals: 1542
% Trivial inequalities: 671
% Fw subsumption resolutions: 5488
% Fw demodulations: 4121
% Bw demodulations: 153
% 
% Simple tautologies: 1191
% Equational tautologies: 43
% Forward subsumptions: 17704
% Bw demodulations to eq. taut.: 7
% 
% Binary resolution: 1929
% Factoring: 4
% Forward superposition: 11760
% Backward superposition: 9160
% Self superposition: 125
% Equality factoring: 84
% Equality resolution: 46
% 
% Split clauses: 554
% Split components: 1111
% SAT solver clauses: 681
% SAT solver unit clauses: 41
% SAT solver binary clauses: 492
% 
% Sat splits: 554
% Sat splitting refutations: 124
% 
% Memory used [KB]: 13176
% Time elapsed: 0.729 s
% ------------------------------
% ------------------------------
% Success in time 12.45 s
