Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.rich_list
% (114136)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1119
% Passive clauses: 6416
% Generated clauses: 34763
% Final active clauses: 1063
% Final passive clauses: 4934
% Input formulas: 123
% Initial clauses: 301
% 
% Duplicate literals: 399
% Trivial inequalities: 42
% Fw subsumption resolutions: 2114
% Fw demodulations: 3206
% Bw demodulations: 15
% 
% Simple tautologies: 242
% Equational tautologies: 37
% Forward subsumptions: 21247
% Fw demodulations to eq. taut.: 2
% 
% Binary resolution: 10366
% Factoring: 1525
% Forward superposition: 7384
% Backward superposition: 9138
% Self superposition: 44
% Equality factoring: 2
% Equality resolution: 12
% 
% Split clauses: 848
% Split components: 1698
% SAT solver clauses: 1105
% SAT solver unit clauses: 11
% SAT solver binary clauses: 62
% 
% Sat splits: 848
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 1105
% TWLsolver calls for satisfiability: 411
% 
% Memory used [KB]: 17398
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.rich_list
% (114146)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit

% Active clauses: 657
% Passive clauses: 3435
% Generated clauses: 22964
% Final active clauses: 603
% Final passive clauses: 2482
% Input formulas: 123
% Initial clauses: 301
% 
% Duplicate literals: 127
% Trivial inequalities: 101
% Fw subsumption resolutions: 1161
% Fw demodulations: 2760
% Bw demodulations: 316
% Condensations: 86
% Global subsumptions: 24
% 
% Simple tautologies: 283
% Equational tautologies: 10
% Forward subsumptions: 11225
% Fw demodulations to eq. taut.: 2
% Bw demodulations to eq. taut.: 4
% 
% Binary resolution: 2879
% Factoring: 174
% Forward superposition: 5808
% Backward superposition: 5495
% Self superposition: 49
% Equality factoring: 2
% Equality resolution: 19
% 
% Split clauses: 22
% Split components: 44
% SAT solver clauses: 3868
% SAT solver unit clauses: 357
% SAT solver binary clauses: 99
% SAT solver learnt clauses: 48
% SAT solver learnt literals: 12
% 
% Sat splits: 22
% Sat splitting refutations: 41
% 
% TWLsolver clauses: 3772
% TWLsolver calls for satisfiability: 2065
% 
% Memory used [KB]: 20852
% Time elapsed: 0.726 s
% ------------------------------
% ------------------------------
% remaining time: 285 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.rich_list
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.rich_list
% SZS output start Proof for thm.rich_list
fof(f52764,plain,(
  $false),
  inference(subsumption_resolution,[],[f52763,f783])).
fof(f783,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f468,f462])).
fof(f462,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.rich_list.rich_list_58_prove_ltb',aHL_TRUTH)).
fof(f468,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f342])).
fof(f342,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f127])).
fof(f127,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f42])).
fof(f42,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.rich_list.rich_list_58_prove_ltb','thm.bool.NOT_CLAUSES|split|2')).
fof(f52763,plain,(
  p(s(bool,'F_0'))),
  inference(forward_demodulation,[],[f52762,f865])).
fof(f865,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK16)))),
  inference(unit_resulting_resolution,[],[f460,f518])).
fof(f518,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f346])).
fof(f346,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f174])).
fof(f174,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f173])).
fof(f173,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X16] : (s(bool,'F_0') = s(bool,X16) <=> ~p(s(bool,X16)))),
  file('thm.rich_list.rich_list_58_prove_ltb','thm.bool.EQ_CLAUSES|split|2')).
fof(f460,plain,(
  ~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK16))))),
  inference(cnf_transformation,[],[f341])).
fof(f341,plain,(
  ~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK16)))) & s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK16),s('type.list.list'(sK12),sK14))) = s('type.list.list'(sK12),'const.list.FILTER_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13,sK14,sK15,sK16])],[f281])).
fof(f281,plain,(
  ? [X0,X1,X2,X3,X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) & s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X2))) = s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X3),s('type.list.list'(X0),X1))))),
  inference(ennf_transformation,[],[f126])).
fof(f126,plain,(
  ~! [X0,X1,X2,X3,X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) => s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X2))) != s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X3),s('type.list.list'(X0),X1))))),
  inference(flattening,[],[f125])).
fof(f125,plain,(
  ~! [X0] : ! [X1,X2,X3,X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) => ~s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X2))) = s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X3),s('type.list.list'(X0),X1))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9,X10] : (~p(s(bool,happ(s(fun(X6,bool),X9),s(X6,X10)))) => ~s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X10),s('type.list.list'(X6),X8))) = s('type.list.list'(X6),'const.list.FILTER_2'(s(fun(X6,bool),X9),s('type.list.list'(X6),X7))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9,X10] : (~p(s(bool,happ(s(fun(X6,bool),X9),s(X6,X10)))) => ~s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X10),s('type.list.list'(X6),X8))) = s('type.list.list'(X6),'const.list.FILTER_2'(s(fun(X6,bool),X9),s('type.list.list'(X6),X7))))),
  file('thm.rich_list.rich_list_58_prove_ltb','thm.rich_list.rich_list_58_prove_ltb')).
fof(f52762,plain,(
  p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK16))))),
  inference(forward_demodulation,[],[f52761,f588])).
fof(f588,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.list.HD_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))))) )),
  inference(cnf_transformation,[],[f212])).
fof(f212,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.list.HD_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))))),
  inference(flattening,[],[f211])).
fof(f211,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.list.HD_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))))),
  inference(rectify,[],[f56])).
fof(f56,axiom,(
  ! [X6] : ! [X10,X16] : s(X6,X10) = s(X6,'const.list.HD_1'(s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X10),s('type.list.list'(X6),X16)))))),
  file('thm.rich_list.rich_list_58_prove_ltb','thm.list.HD')).
fof(f52761,plain,(
  p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,'const.list.HD_1'(s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK16),s('type.list.list'(sK12),sK14))))))))),
  inference(forward_demodulation,[],[f52760,f6166])).
fof(f6166,plain,(
  ( ! [X0,X1] : (s(X0,'const.list.HD_1'(s('type.list.list'(X0),X1))) = s(X0,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(X0),X1)))) )),
  inference(forward_demodulation,[],[f6018,f581])).
fof(f581,plain,(
  ( ! [X0,X1] : (s(X0,'const.list.HD_1'(s('type.list.list'(X0),X1))) = s(X0,happ(s(fun('type.list.list'(X0),X0),'const.list.HD_0'),s('type.list.list'(X0),X1)))) )),
  inference(cnf_transformation,[],[f206])).
fof(f206,plain,(
  ! [X0,X1] : s(X0,'const.list.HD_1'(s('type.list.list'(X0),X1))) = s(X0,happ(s(fun('type.list.list'(X0),X0),'const.list.HD_0'),s('type.list.list'(X0),X1)))),
  inference(flattening,[],[f205])).
fof(f205,plain,(
  ! [X0] : ! [X1] : s(X0,'const.list.HD_1'(s('type.list.list'(X0),X1))) = s(X0,happ(s(fun('type.list.list'(X0),X0),'const.list.HD_0'),s('type.list.list'(X0),X1)))),
  inference(rectify,[],[f124])).
fof(f124,axiom,(
  ! [X6] : ! [X28] : s(X6,'const.list.HD_1'(s('type.list.list'(X6),X28))) = s(X6,happ(s(fun('type.list.list'(X6),X6),'const.list.HD_0'),s('type.list.list'(X6),X28)))),
  file('thm.rich_list.rich_list_58_prove_ltb','ARITY_0_1_const.list.HD')).
fof(f6018,plain,(
  ( ! [X0,X1] : (s(X0,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(X0),X1))) = s(X0,happ(s(fun('type.list.list'(X0),X0),'const.list.HD_0'),s('type.list.list'(X0),X1)))) )),
  inference(superposition,[],[f590,f488])).
fof(f488,plain,(
  ( ! [X0] : (s(fun('type.list.list'(X0),X0),'const.list.HD_0') = s(fun('type.list.list'(X0),X0),'const.list.EL_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0] : s(fun('type.list.list'(X0),X0),'const.list.HD_0') = s(fun('type.list.list'(X0),X0),'const.list.EL_1'(s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f60])).
fof(f60,axiom,(
  ! [X6] : s(fun('type.list.list'(X6),X6),'const.list.HD_0') = s(fun('type.list.list'(X6),X6),'const.list.EL_1'(s('type.num.num','const.num.|48|_0')))),
  file('thm.rich_list.rich_list_58_prove_ltb','thm.list.EL_restricted|split|0')).
fof(f590,plain,(
  ( ! [X2,X0,X1] : (s(X0,'const.list.EL_2'(s('type.num.num',X1),s('type.list.list'(X0),X2))) = s(X0,happ(s(fun('type.list.list'(X0),X0),'const.list.EL_1'(s('type.num.num',X1))),s('type.list.list'(X0),X2)))) )),
  inference(cnf_transformation,[],[f216])).
fof(f216,plain,(
  ! [X0,X1,X2] : s(X0,'const.list.EL_2'(s('type.num.num',X1),s('type.list.list'(X0),X2))) = s(X0,happ(s(fun('type.list.list'(X0),X0),'const.list.EL_1'(s('type.num.num',X1))),s('type.list.list'(X0),X2)))),
  inference(flattening,[],[f215])).
fof(f215,plain,(
  ! [X0] : ! [X1,X2] : s(X0,'const.list.EL_2'(s('type.num.num',X1),s('type.list.list'(X0),X2))) = s(X0,happ(s(fun('type.list.list'(X0),X0),'const.list.EL_1'(s('type.num.num',X1))),s('type.list.list'(X0),X2)))),
  inference(rectify,[],[f123])).
fof(f123,axiom,(
  ! [X6] : ! [X28,X29] : s(X6,'const.list.EL_2'(s('type.num.num',X28),s('type.list.list'(X6),X29))) = s(X6,happ(s(fun('type.list.list'(X6),X6),'const.list.EL_1'(s('type.num.num',X28))),s('type.list.list'(X6),X29)))),
  file('thm.rich_list.rich_list_58_prove_ltb','ARITY_1_2_const.list.EL')).
fof(f52760,plain,(
  p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,'const.list.EL_2'(s('type.num.num','const.num.|48|_0'),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK16),s('type.list.list'(sK12),sK14))))))))),
  inference(forward_demodulation,[],[f52759,f49288])).
fof(f49288,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.ZERO_0')),
  inference(forward_demodulation,[],[f49247,f465])).
fof(f465,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0')))),
  inference(cnf_transformation,[],[f78])).
fof(f78,axiom,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0')))),
  file('thm.rich_list.rich_list_58_prove_ltb','thm.prim_rec.PRE|split|0')).
fof(f49247,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0')))),
  inference(unit_resulting_resolution,[],[f2672,f49106])).
fof(f49106,plain,(
  ( ! [X30] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X30))) | s('type.num.num',X30) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(backward_demodulation,[],[f49097,f3582])).
fof(f3582,plain,(
  ( ! [X30] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X30))) | s('type.num.num',X30) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f620,f465])).
fof(f620,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f397])).
fof(f397,plain,(
  ! [X0,X1] : ((s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) & (s('type.num.num',X0) != s('type.num.num',X1) | s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))),
  inference(nnf_transformation,[],[f246])).
fof(f246,plain,(
  ! [X0,X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))) <=> s('type.num.num',X0) = s('type.num.num',X1))),
  inference(rectify,[],[f82])).
fof(f82,axiom,(
  ! [X11,X12] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X11))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X12))) <=> s('type.num.num',X11) = s('type.num.num',X12))),
  file('thm.rich_list.rich_list_58_prove_ltb','thm.numeral.numeral_distrib|split|19')).
fof(f49097,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0')))),
  inference(forward_demodulation,[],[f49096,f465])).
fof(f49096,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(forward_demodulation,[],[f49095,f477])).
fof(f477,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f135])).
fof(f135,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X12] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X12),s('type.num.num','const.num.|48|_0')))),
  file('thm.rich_list.rich_list_58_prove_ltb','thm.arithmetic.MULT_|48|')).
fof(f49095,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))))))) )),
  inference(subsumption_resolution,[],[f49094,f462])).
fof(f49094,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))))))) )),
  inference(forward_demodulation,[],[f49093,f478])).
fof(f478,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f136])).
fof(f136,plain,(
  ! [X0] : s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f89])).
fof(f89,axiom,(
  ! [X11] : s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X11)))),
  file('thm.rich_list.rich_list_58_prove_ltb','thm.numeral.numeral_distrib|split|26')).
fof(f49093,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))))))) )),
  inference(forward_demodulation,[],[f49049,f465])).
fof(f49049,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.ZERO_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f21879,f477])).
fof(f21879,plain,(
  ( ! [X30] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',X30))),s('type.num.num','const.arithmetic.ZERO_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',X30)))))) )),
  inference(forward_demodulation,[],[f21878,f493])).
fof(f493,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',X0))))) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f151])).
fof(f151,plain,(
  ! [X0] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',X0))))) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f79])).
fof(f79,axiom,(
  ! [X11] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',X11))))) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X11)))))),
  file('thm.rich_list.rich_list_58_prove_ltb','thm.numeral.numeral_distrib|split|16')).
fof(f21878,plain,(
  ( ! [X30] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',X30))),s('type.num.num','const.arithmetic.ZERO_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X30)))))) )),
  inference(forward_demodulation,[],[f21837,f5422])).
fof(f5422,plain,(
  ( ! [X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.ZERO_0'))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(superposition,[],[f604,f2672])).
fof(f604,plain,(
  ( ! [X0,X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f230])).
fof(f230,plain,(
  ! [X0,X1] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f94])).
fof(f94,axiom,(
  ! [X11,X12] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X12),s('type.num.num',X11))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X11))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X12)))))),
  file('thm.rich_list.rich_list_58_prove_ltb','thm.numeral.numeral_distrib|split|31')).
fof(f21837,plain,(
  ( ! [X30] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',X30)))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X30)))))) )),
  inference(superposition,[],[f525,f493])).
fof(f525,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f350])).
fof(f350,plain,(
  ! [X0] : ((~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) & (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))))),
  inference(nnf_transformation,[],[f179])).
fof(f179,plain,(
  ! [X0] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X0))),
  inference(rectify,[],[f93])).
fof(f93,axiom,(
  ! [X11] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X11)))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X11))),
  file('thm.rich_list.rich_list_58_prove_ltb','thm.numeral.numeral_distrib|split|30')).
fof(f2672,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(forward_demodulation,[],[f2658,f787])).
fof(f787,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(forward_demodulation,[],[f498,f466])).
fof(f466,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(cnf_transformation,[],[f76])).
fof(f76,axiom,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  file('thm.rich_list.rich_list_58_prove_ltb','thm.numeral.numeral_distrib|split|13')).
fof(f498,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(cnf_transformation,[],[f156])).
fof(f156,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X12] : s('type.num.num',X12) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X12),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  file('thm.rich_list.rich_list_58_prove_ltb','thm.arithmetic.MULT_RIGHT_|49|')).
fof(f2658,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(unit_resulting_resolution,[],[f787,f522])).
fof(f522,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f348])).
fof(f348,plain,(
  ! [X0] : ((s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0)) & (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))),
  inference(nnf_transformation,[],[f177])).
fof(f177,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) <=> s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0))),
  inference(rectify,[],[f80])).
fof(f80,axiom,(
  ! [X11] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X11))) <=> s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X11))),
  file('thm.rich_list.rich_list_58_prove_ltb','thm.numeral.numeral_distrib|split|17')).
fof(f52759,plain,(
  p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,'const.list.EL_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK16),s('type.list.list'(sK12),sK14))))))))),
  inference(subsumption_resolution,[],[f52573,f49531])).
fof(f49531,plain,(
  ( ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))) )),
  inference(backward_demodulation,[],[f49288,f11331])).
fof(f11331,plain,(
  ( ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))) )),
  inference(unit_resulting_resolution,[],[f5467,f656])).
fof(f656,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f413])).
fof(f413,plain,(
  ! [X0,X1] : ((p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) & (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(nnf_transformation,[],[f260])).
fof(f260,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(flattening,[],[f259])).
fof(f259,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X12,X11] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X12),s('type.num.num',X11)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X11),s('type.num.num',X12)))))),
  file('thm.rich_list.rich_list_58_prove_ltb','thm.arithmetic.NOT_LESS')).
fof(f5467,plain,(
  ( ! [X21] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X21))),s('type.num.num','const.arithmetic.ZERO_0'))))) )),
  inference(backward_demodulation,[],[f5422,f3468])).
fof(f3468,plain,(
  ( ! [X21] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X21))))))))) )),
  inference(subsumption_resolution,[],[f3418,f579])).
fof(f579,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f202])).
fof(f202,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),
  inference(flattening,[],[f201])).
fof(f201,plain,(
  ! [X0] : ~s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),
  inference(rectify,[],[f62])).
fof(f62,axiom,(
  ! [X11] : ~s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num',X11)))),
  file('thm.rich_list.rich_list_58_prove_ltb','thm.num.NOT_SUC')).
fof(f3418,plain,(
  ( ! [X21] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X21)))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X21)))))) )),
  inference(superposition,[],[f525,f499])).
fof(f499,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f157])).
fof(f157,plain,(
  ! [X0] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f77])).
fof(f77,axiom,(
  ! [X11] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X11))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X11)))))),
  file('thm.rich_list.rich_list_58_prove_ltb','thm.numeral.numeral_distrib|split|14')).
fof(f52573,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14)))))))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,'const.list.EL_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK16),s('type.list.list'(sK12),sK14))))))))),
  inference(superposition,[],[f9434,f49288])).
fof(f9434,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14)))))))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK16),s('type.list.list'(sK12),sK14))))))))) )),
  inference(forward_demodulation,[],[f9433,f461])).
fof(f461,plain,(
  s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK16),s('type.list.list'(sK12),sK14))) = s('type.list.list'(sK12),'const.list.FILTER_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))),
  inference(cnf_transformation,[],[f341])).
fof(f9433,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),sK14)))))))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(sK12),'const.list.FILTER_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))))))))) )),
  inference(forward_demodulation,[],[f9390,f589])).
fof(f589,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2))))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))))) )),
  inference(cnf_transformation,[],[f214])).
fof(f214,plain,(
  ! [X0,X1,X2] : s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2))))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))))),
  inference(flattening,[],[f213])).
fof(f213,plain,(
  ! [X0] : ! [X1,X2] : s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X2))))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))))),
  inference(rectify,[],[f58])).
fof(f58,axiom,(
  ! [X6] : ! [X10,X16] : s('type.num.num','const.num.SUC_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),X16))))) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X10),s('type.list.list'(X6),X16)))))),
  file('thm.rich_list.rich_list_58_prove_ltb','thm.list.LENGTH|split|1')).
fof(f9390,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK16),s('type.list.list'(sK12),sK14)))))))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,'const.list.EL_2'(s('type.num.num',X0),s('type.list.list'(sK12),'const.list.FILTER_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))))))))) )),
  inference(superposition,[],[f598,f461])).
fof(f598,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X3),s('type.list.list'(X0),X2)))))))) | p(s(bool,happ(s(fun(X0,bool),X3),s(X0,'const.list.EL_2'(s('type.num.num',X1),s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X3),s('type.list.list'(X0),X2))))))))) )),
  inference(cnf_transformation,[],[f296])).
fof(f296,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X3),s('type.list.list'(X0),X2)))))))) | p(s(bool,happ(s(fun(X0,bool),X3),s(X0,'const.list.EL_2'(s('type.num.num',X1),s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X3),s('type.list.list'(X0),X2)))))))))),
  inference(ennf_transformation,[],[f224])).
fof(f224,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X3),s('type.list.list'(X0),X2)))))))) => p(s(bool,happ(s(fun(X0,bool),X3),s(X0,'const.list.EL_2'(s('type.num.num',X1),s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X3),s('type.list.list'(X0),X2)))))))))),
  inference(flattening,[],[f223])).
fof(f223,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X3),s('type.list.list'(X0),X2)))))))) => p(s(bool,happ(s(fun(X0,bool),X3),s(X0,'const.list.EL_2'(s('type.num.num',X1),s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X3),s('type.list.list'(X0),X2)))))))))),
  inference(rectify,[],[f107])).
fof(f107,axiom,(
  ! [X6] : ! [X25,X7,X9] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X25),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X6),'const.list.FILTER_2'(s(fun(X6,bool),X9),s('type.list.list'(X6),X7)))))))) => p(s(bool,happ(s(fun(X6,bool),X9),s(X6,'const.list.EL_2'(s('type.num.num',X25),s('type.list.list'(X6),'const.list.FILTER_2'(s(fun(X6,bool),X9),s('type.list.list'(X6),X7)))))))))),
  file('thm.rich_list.rich_list_58_prove_ltb','thm.rich_list.rich_list_57_prove_ltb')).
% SZS output end Proof for thm.rich_list
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 1161
% Passive clauses: 8620
% Generated clauses: 52643
% Final active clauses: 689
% Final passive clauses: 4241
% Input formulas: 123
% Initial clauses: 301
% 
% Duplicate literals: 125
% Trivial inequalities: 103
% Fw subsumption resolutions: 654
% Fw demodulations: 10840
% Bw demodulations: 3665
% Condensations: 186
% Global subsumptions: 284
% 
% Simple tautologies: 124
% Equational tautologies: 88
% Forward subsumptions: 30965
% Fw demodulations to eq. taut.: 217
% Bw demodulations to eq. taut.: 24
% 
% Unit resulting resolution: 10793
% Factoring: 3
% Forward superposition: 14154
% Backward superposition: 10902
% Self superposition: 72
% Equality factoring: 28
% Equality resolution: 22
% 
% Split clauses: 110
% Split components: 220
% SAT solver clauses: 9412
% SAT solver unit clauses: 3389
% SAT solver binary clauses: 4668
% 
% Sat splits: 110
% Sat splitting refutations: 286
% 
% Memory used [KB]: 39530
% Time elapsed: 2.199 s
% ------------------------------
% ------------------------------
% Success in time 3.658 s
