Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.extreal
% (22699)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1445
% Passive clauses: 8336
% Generated clauses: 40480
% Final active clauses: 1407
% Final passive clauses: 6578
% Input formulas: 261
% Initial clauses: 534
% 
% Duplicate literals: 240
% Trivial inequalities: 98
% Fw subsumption resolutions: 2611
% Fw demodulations: 3943
% Bw demodulations: 33
% 
% Simple tautologies: 358
% Equational tautologies: 34
% Forward subsumptions: 24436
% 
% Binary resolution: 7854
% Factoring: 755
% Forward superposition: 12804
% Backward superposition: 11333
% Self superposition: 38
% Equality factoring: 2
% Equality resolution: 27
% 
% Split clauses: 466
% Split components: 934
% SAT solver clauses: 650
% SAT solver unit clauses: 31
% SAT solver binary clauses: 94
% 
% Sat splits: 466
% 
% TWLsolver clauses: 650
% TWLsolver calls for satisfiability: 522
% 
% Memory used [KB]: 19573
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.extreal
% (22724)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 728
% Passive clauses: 3639
% Generated clauses: 24470
% Final active clauses: 704
% Final passive clauses: 2848
% Input formulas: 261
% Initial clauses: 534
% 
% Duplicate literals: 101
% Trivial inequalities: 69
% Fw subsumption resolutions: 909
% Fw demodulations: 3523
% Bw demodulations: 86
% Condensations: 107
% Global subsumptions: 35
% 
% Simple tautologies: 307
% Equational tautologies: 16
% Forward subsumptions: 11539
% Fw demodulations to eq. taut.: 2
% 
% Binary resolution: 2733
% Factoring: 110
% Forward superposition: 6493
% Backward superposition: 6011
% Self superposition: 53
% Equality factoring: 2
% Equality resolution: 25
% 
% Split clauses: 14
% Split components: 28
% SAT solver clauses: 4712
% SAT solver unit clauses: 690
% SAT solver binary clauses: 176
% SAT solver learnt clauses: 70
% SAT solver learnt literals: 18
% 
% Sat splits: 14
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 4572
% TWLsolver calls for satisfiability: 5862
% 
% Memory used [KB]: 23411
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.extreal
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.extreal
% SZS output start Proof for thm.extreal
fof(f78475,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f14461,f14476,f14475,f15404,f15407,f15409,f15411,f15413,f15417,f15419,f15423,f15425,f15447,f15449,f15471,f15473,f15495,f15517,f15537,f15654,f15656,f15658,f15748,f15759,f15770,f16785,f16799,f16926,f16939,f17065,f17078,f17711,f17713,f17715,f17805,f17816,f17827,f18842,f18856,f18983,f18996,f19122,f19135,f19687,f19689,f19691,f19719,f19721,f19723,f19781,f19783,f19809,f19811,f19869,f19871,f19897,f19899,f19957,f19959,f19985,f19987,f20003,f20005,f21822,f21823,f20025,f20029,f20051,f20056,f20060,f20064,f20086,f20091,f20107,f20123,f20149,f20151,f20153,f20168,f20170,f20185,f20187,f20202,f20204,f20219,f20221,f20240,f20264,f20288,f20502,f20504,f20506,f20511,f20513,f20535,f20537,f20539,f20541,f20543,f20545,f20552,f20554,f20556,f20558,f20631,f20638,f20640,f20642,f20644,f20646,f20664,f20666,f20668,f20670,f20672,f20674,f20676,f20678,f20680,f20682,f20684,f20686,f20688,f20690,f20692,f20694,f20696,f20698,f20700,f20731,f20753,f20783,f20805,f20807,f20838,f20840,f20857,f20859,f20890,f20892,f20909,f20911,f20944,f20946,f20963,f20965,f20990,f21014,f21061,f21063,f21084,f21086,f21107,f21109,f21130,f21204,f21206,f21208,f21213,f21215,f21241,f21243,f21245,f21274,f21276,f21292,f21307,f21322,f21346,f21374,f21376,f21378,f21408,f21410,f21412,f21440,f21442,f21444,f21476,f21499,f21531,f21554,f21556,f21559,f21561,f21563,f21565,f21567,f21569,f21585,f21587,f21589,f21591,f21593,f21595,f21597,f21599,f21601,f21603,f21605,f21654,f21656,f21658,f21660,f21662,f21664,f21669,f21671,f21673,f21675,f21711,f21713,f21749,f21751,f21753,f21758,f21760,f21762,f21764,f21766,f21768,f21770,f21772,f21774,f21776,f21778,f21780,f21782,f21784,f21786,f21788,f21824,f21825,f21826,f21827,f21828,f21829,f21844,f21845,f21846,f21847,f21848,f21849,f21850,f21851,f21852,f21853,f21854,f21855,f21856,f21857,f21858,f21859,f21860,f21861,f21862,f21863,f21864,f21865,f21866,f21867,f21868,f21869,f21870,f21871,f21872,f21873,f21874,f21875,f21876,f21877,f21878,f21879,f21880,f21881,f21882,f21883,f21884,f21885,f21886,f21887,f21888,f21889,f21890,f21891,f21892,f21893,f21894,f21895,f21896,f21897,f21898,f21899,f21900,f21901,f21902,f21903,f21904,f21905,f21906,f21907,f21908,f21909,f21910,f21911,f21912,f21913,f21914,f21915,f21916,f21917,f21918,f21919,f21920,f21921,f21922,f21923,f21924,f21925,f21926,f21927,f21928,f21929,f21930,f21931,f21932,f21933,f21934,f21935,f21936,f21937,f21938,f21939,f21940,f21941,f21942,f21943,f21944,f21945,f21946,f21947,f21948,f21949,f21950,f21951,f21952,f21953,f21954,f21955,f21956,f21957,f21958,f21959,f21960,f21961,f21962,f21963,f21964,f21965,f21966,f21967,f21968,f21969,f21970,f21971,f21972,f21973,f21974,f21975,f21976,f21977,f21978,f21979,f21980,f21981,f21982,f21983,f21984,f21985,f21986,f21987,f21988,f21989,f21990,f21991,f21992,f21993,f21994,f21995,f21996,f21997,f21998,f21999,f22000,f22001,f22002,f22003,f22004,f22005,f22006,f22007,f22008,f22009,f22010,f22011,f22012,f22013,f22014,f22015,f22016,f22017,f22018,f22019,f22020,f22021,f22022,f22023,f22024,f22025,f22026,f22027,f22028,f22029,f22030,f22031,f22032,f22033,f22034,f22035,f22036,f22037,f22038,f22039,f22040,f22041,f22042,f22043,f22044,f22045,f22046,f22047,f22048,f22049,f22050,f22051,f22052,f22053,f22054,f22055,f22056,f22057,f22058,f22059,f22060,f22061,f22062,f22063,f22064,f22065,f22066,f22067,f22068,f22069,f22070,f22071,f22072,f22073,f22074,f22075,f22076,f22077,f22078,f22079,f22080,f22081,f22082,f22083,f22084,f22085,f22086,f22087,f22088,f22089,f22090,f22091,f22092,f22093,f22094,f22095,f22096,f22097,f22098,f22099,f22100,f22101,f22102,f22103,f22104,f22105,f22106,f22107,f22108,f22109,f22110,f22111,f22112,f22113,f22114,f22115,f22116,f22117,f22118,f22119,f22120,f22121,f22122,f22123,f22124,f22125,f22126,f22127,f22128,f22129,f22130,f22131,f22132,f22133,f22134,f22135,f22136,f22137,f22138,f22139,f22140,f22141,f22142,f22143,f22144,f22145,f22146,f22147,f22148,f22149,f22150,f22151,f22152,f22153,f22154,f22155,f22156,f22157,f22158,f22159,f22160,f22161,f22162,f22163,f22164,f22165,f22166,f22167,f22168,f22169,f22170,f22171,f22172,f22173,f22174,f22175,f22176,f22177,f22178,f22179,f22180,f22181,f22182,f22183,f22184,f22185,f22186,f22187,f22188,f22189,f22190,f22191,f22192,f22193,f22194,f22195,f22196,f22197,f22198,f22199,f22200,f22201,f22202,f22203,f22204,f22205,f22206,f22207,f22208,f22209,f22210,f22211,f22212,f22213,f22214,f22215,f22216,f22217,f22218,f22219,f22220,f22221,f22222,f22223,f22224,f22225,f22226,f22227,f22228,f22229,f22230,f22231,f22232,f22233,f22234,f22235,f22236,f22237,f22238,f22239,f22240,f22241,f22242,f22243,f22244,f22245,f22246,f22247,f22248,f22249,f22250,f22251,f22252,f22253,f22254,f22255,f22256,f22257,f22258,f22259,f22260,f22261,f22262,f22263,f22264,f22265,f22266,f22267,f22268,f22269,f22270,f22271,f22272,f22273,f22274,f22275,f22276,f22277,f22278,f22279,f22280,f22281,f22282,f22283,f22284,f22285,f22286,f22287,f22288,f22289,f22290,f22291,f22292,f22293,f22294,f22295,f22296,f22297,f22298,f22299,f22300,f22301,f22302,f22303,f22304,f22305,f22306,f22307,f22308,f22309,f22310,f22311,f22312,f22313,f22314,f22315,f22316,f22317,f22318,f22319,f22320,f22321,f22322,f22323,f22324,f22325,f22326,f22327,f22328,f22329,f22330,f22331,f22332,f22333,f22334,f22335,f22336,f22337,f22338,f22339,f22340,f22341,f22342,f22343,f22344,f22345,f22346,f22347,f22348,f22349,f22350,f22351,f22352,f22353,f22354,f22355,f22356,f22357,f22358,f22359,f22360,f22361,f22362,f22363,f22364,f22365,f22366,f22367,f22368,f22369,f22370,f22371,f22372,f22373,f22374,f22375,f22376,f22377,f22378,f22379,f22380,f22381,f22382,f22383,f22384,f22385,f22386,f22387,f22388,f22389,f22390,f22391,f22392,f22393,f22394,f22395,f22396,f22397,f22398,f22399,f22400,f22401,f22402,f22403,f22404,f22405,f22406,f22407,f22408,f22409,f22410,f22411,f22412,f22413,f22414,f22415,f22416,f22417,f22418,f22419,f22420,f22421,f22422,f22423,f22424,f22425,f22426,f22427,f22428,f22429,f22430,f22431,f22432,f22433,f22434,f22435,f22436,f22437,f22438,f22439,f22440,f22441,f22442,f22443,f22444,f22445,f22446,f22447,f22448,f22449,f22450,f22451,f22452,f22453,f22454,f22455,f22456,f22457,f22458,f22459,f22460,f22461,f22462,f22463,f22464,f22465,f22466,f22467,f22468,f22469,f22470,f22471,f22472,f22473,f22474,f22475,f22476,f22477,f22478,f22479,f22480,f22481,f22482,f22483,f22484,f22485,f22486,f22487,f22488,f22489,f22490,f22491,f22492,f22493,f22494,f22495,f22496,f22497,f22498,f22499,f22500,f22501,f22502,f22503,f22504,f22505,f22506,f22507,f22508,f22509,f22510,f22511,f22512,f22513,f22514,f22515,f22516,f22517,f22518,f22519,f22520,f22521,f22522,f22523,f22524,f22525,f22526,f22527,f22528,f22529,f22530,f22531,f22532,f22533,f22534,f22535,f22536,f22537,f22538,f22539,f22540,f22541,f22542,f22543,f22544,f22545,f22546,f22547,f22548,f22549,f22550,f22551,f22552,f22553,f22554,f22555,f22556,f22557,f22558,f22559,f22560,f22561,f22562,f22563,f22564,f22565,f22566,f22567,f22568,f22569,f22570,f22571,f22572,f22573,f22574,f22575,f22576,f22577,f22578,f22579,f22580,f22581,f22582,f22583,f22584,f22585,f22586,f22587,f22588,f22589,f22590,f22591,f22592,f22593,f22594,f22595,f22596,f22597,f22598,f22599,f22600,f22601,f22602,f22603,f22604,f22605,f22606,f22607,f22608,f22609,f22610,f22611,f22612,f22613,f22614,f22615,f22616,f22617,f22618,f22619,f22620,f22621,f22622,f22623,f22624,f22625,f22626,f22627,f22628,f22629,f22630,f22631,f22632,f22633,f22634,f22635,f22636,f22637,f22638,f22639,f22640,f22641,f22642,f22643,f22644,f22645,f22646,f22647,f22648,f22649,f22841,f22843,f25689,f25691,f27491,f27492,f27649,f27650,f27783,f27784,f27954,f27955,f34923,f34926,f34924,f44955,f44956,f45110,f45111,f45261,f45262,f45416,f45417,f45656,f45660,f45658,f46153,f48234,f48238,f48236,f48751,f73764,f73766,f73765,f73860,f76746,f76747,f15405,f78470])).
fof(f78470,plain,(
  $false | $spl33),
  inference(subsumption_resolution,[],[f78469,f1890])).
fof(f1890,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f1508,f1067])).
fof(f1067,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f765])).
fof(f765,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f333])).
fof(f333,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X21] : (s(bool,'T_0') = s(bool,X21) <=> p(s(bool,X21)))),
  file('thm.extreal.add_not_infty|split|1','thm.bool.EQ_CLAUSES|split|1')).
fof(f1508,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f977,f952])).
fof(f952,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.extreal.add_not_infty|split|1',aHL_TRUTH)).
fof(f977,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f742])).
fof(f742,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f265])).
fof(f265,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f41])).
fof(f41,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.extreal.add_not_infty|split|1','thm.bool.NOT_CLAUSES|split|1')).
fof(f78469,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl33),
  inference(backward_demodulation,[],[f78468,f6866])).
fof(f6866,plain,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18)))),
  inference(unit_resulting_resolution,[],[f1508,f5894,f1265])).
fof(f1265,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f659])).
fof(f659,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f658])).
fof(f658,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f487])).
fof(f487,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X19,X20] : ((p(s(bool,X19)) => p(s(bool,X20))) => ((p(s(bool,X20)) => p(s(bool,X19))) => s(bool,X19) = s(bool,X20)))),
  file('thm.extreal.add_not_infty|split|1','thm.bool.IMP_ANTISYM_AX')).
fof(f5894,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18))))),
  inference(unit_resulting_resolution,[],[f950,f1087])).
fof(f1087,plain,(
  ( ! [X0] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X0)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)) )),
  inference(cnf_transformation,[],[f775])).
fof(f775,plain,(
  ! [X0] : ((~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X0)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)) & (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X0))))))),
  inference(nnf_transformation,[],[f345])).
fof(f345,plain,(
  ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X0)))) <=> s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0))),
  inference(rectify,[],[f78])).
fof(f78,axiom,(
  ! [X6] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X6)))) <=> s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X6))),
  file('thm.extreal.add_not_infty|split|1','thm.extreal.le_infty|split|3')).
fof(f950,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',sK18)),
  inference(cnf_transformation,[],[f741])).
fof(f741,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',sK17) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',sK18) & s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK17,sK18])],[f590])).
fof(f590,plain,(
  ? [X0,X1] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X1) & s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))),
  inference(flattening,[],[f589])).
fof(f589,plain,(
  ? [X0,X1] : ((s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X1)) & s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))),
  inference(ennf_transformation,[],[f264])).
fof(f264,plain,(
  ~! [X0,X1] : ((s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X1)) => s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))),
  inference(flattening,[],[f263])).
fof(f263,plain,(
  ~! [X0,X1] : ((~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) & ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X1)) => ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ((~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X6) & ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X7)) => ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X7))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ((~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X6) & ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X7)) => ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X7))))),
  file('thm.extreal.add_not_infty|split|1','thm.extreal.add_not_infty|split|1')).
fof(f78468,plain,(
  ( ! [X23] : (s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X23),s('type.extreal.extreal',sK18)))) ) | $spl33),
  inference(forward_demodulation,[],[f78467,f1911])).
fof(f1911,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.PosInf_0')))) )),
  inference(unit_resulting_resolution,[],[f983,f1068])).
fof(f1068,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f765])).
fof(f983,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.PosInf_0'))))) )),
  inference(cnf_transformation,[],[f270])).
fof(f270,plain,(
  ! [X0] : p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.PosInf_0'))))),
  inference(rectify,[],[f91])).
fof(f91,axiom,(
  ! [X6] : p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal','const.extreal.PosInf_0'))))),
  file('thm.extreal.add_not_infty|split|1','thm.extreal.le_infty|split|1')).
fof(f78467,plain,(
  ( ! [X23] : (s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X23),s('type.extreal.extreal',sK18))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',X23))),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl33),
  inference(subsumption_resolution,[],[f78466,f45657])).
fof(f45657,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',sK17) | $spl33),
  inference(cnf_transformation,[],[f45657_D])).
fof(f45657_D,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',sK17) <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f78466,plain,(
  ( ! [X23] : (s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X23),s('type.extreal.extreal',sK18))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',X23))),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK17)) )),
  inference(subsumption_resolution,[],[f78345,f949])).
fof(f949,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',sK17)),
  inference(cnf_transformation,[],[f741])).
fof(f78345,plain,(
  ( ! [X23] : (s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X23),s('type.extreal.extreal',sK18))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',X23))),s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK17) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK17)) )),
  inference(superposition,[],[f1351,f951])).
fof(f951,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))),
  inference(cnf_transformation,[],[f741])).
fof(f1351,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0)) )),
  inference(cnf_transformation,[],[f686])).
fof(f686,plain,(
  ! [X0,X1,X2] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2))))))),
  inference(flattening,[],[f685])).
fof(f685,plain,(
  ! [X0,X1,X2] : ((s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)) | s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2))))))),
  inference(ennf_transformation,[],[f547])).
fof(f547,plain,(
  ! [X0,X1,X2] : ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0)) => s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2))))))),
  inference(flattening,[],[f546])).
fof(f546,plain,(
  ! [X0,X1,X2] : ((~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) & ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)) => s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2))))))),
  inference(rectify,[],[f62])).
fof(f62,axiom,(
  ! [X6,X7,X38] : ((~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X6) & ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X6)) => s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X7),s('type.extreal.extreal',X38))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X7))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X38))))))),
  file('thm.extreal.add_not_infty|split|1','thm.extreal.le_ladd')).
fof(f15405,plain,(
  ( ! [X3] : (~p(s(bool,X3))) ) | $spl4),
  inference(cnf_transformation,[],[f15405_D])).
fof(f15405_D,plain,(
  ( ! [X3] : (~p(s(bool,X3))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f76747,plain,(
  ( ! [X8] : (p(s(bool,sK64(X8,'const.combin.I_0',bool))) | p(s(bool,happ(s(fun(bool,bool),X8),s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))))))) ) | $spl44),
  inference(cnf_transformation,[],[f76747_D])).
fof(f76747_D,plain,(
  ( ! [X8] : (p(s(bool,sK64(X8,'const.combin.I_0',bool))) | p(s(bool,happ(s(fun(bool,bool),X8),s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))))))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f76746,plain,(
  ( ! [X8,X7] : (p(s(bool,sK64(X8,'const.combin.I_0',bool))) | ~p(s(bool,X7)) | p(s(bool,happ(s(fun(bool,bool),X8),s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))))))) )),
  inference(forward_demodulation,[],[f76704,f1514])).
fof(f1514,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))) )),
  inference(forward_demodulation,[],[f1151,f1150])).
fof(f1150,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))) )),
  inference(cnf_transformation,[],[f389])).
fof(f389,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))),
  inference(flattening,[],[f388])).
fof(f388,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X8] : ! [X6] : s(X8,X6) = s(X8,'const.combin.I_1'(s(X8,X6)))),
  file('thm.extreal.add_not_infty|split|1','thm.combin.I_THM')).
fof(f1151,plain,(
  ( ! [X0,X1] : (s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f391])).
fof(f391,plain,(
  ! [X0,X1] : s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))),
  inference(flattening,[],[f390])).
fof(f390,plain,(
  ! [X0] : ! [X1] : s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X8] : ! [X13] : s(X8,'const.combin.I_1'(s(X8,X13))) = s(X8,happ(s(fun(X8,X8),'const.combin.I_0'),s(X8,X13)))),
  file('thm.extreal.add_not_infty|split|1','ARITY_0_1_const.combin.I')).
fof(f76704,plain,(
  ( ! [X8,X7] : (~p(s(bool,X7)) | p(s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,sK64(X8,'const.combin.I_0',bool))))) | p(s(bool,happ(s(fun(bool,bool),X8),s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))))))) )),
  inference(superposition,[],[f1168,f1514])).
fof(f1168,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK64(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,'const.min.@_1'(s(fun(X0,bool),X1))))))) )),
  inference(cnf_transformation,[],[f820])).
fof(f820,plain,(
  ! [X0,X1,X2] : (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK64(X2,X1,X0))))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK64(X2,X1,X0)))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK64])],[f629])).
fof(f629,plain,(
  ! [X0,X1,X2] : (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(flattening,[],[f628])).
fof(f628,plain,(
  ! [X0,X1,X2] : ((! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(ennf_transformation,[],[f413])).
fof(f413,plain,(
  ! [X0,X1,X2] : ((? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) => p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))) => p(s(bool,happ(s(fun(X0,bool),X2),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(flattening,[],[f412])).
fof(f412,plain,(
  ! [X0] : ! [X1,X2] : ((? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) => p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))) => p(s(bool,happ(s(fun(X0,bool),X2),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(rectify,[],[f221])).
fof(f221,axiom,(
  ! [X8] : ! [X14,X71] : ((? [X6] : p(s(bool,happ(s(fun(X8,bool),X14),s(X8,X6)))) & ! [X6] : (p(s(bool,happ(s(fun(X8,bool),X14),s(X8,X6)))) => p(s(bool,happ(s(fun(X8,bool),X71),s(X8,X6)))))) => p(s(bool,happ(s(fun(X8,bool),X71),s(X8,'const.min.@_1'(s(fun(X8,bool),X14)))))))),
  file('thm.extreal.add_not_infty|split|1','thm.bool.SELECT_ELIM_THM')).
fof(f73860,plain,(
  $false | ($spl33 | $spl40)),
  inference(subsumption_resolution,[],[f73856,f73765])).
fof(f73856,plain,(
  ~p(s(bool,sK83(sK17))) | $spl33),
  inference(unit_resulting_resolution,[],[f1890,f46393,f1384])).
fof(f1384,plain,(
  ( ! [X2,X0,X1] : (~sP6(X2,X1,X0) | ~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f909])).
fof(f909,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP6(X2,X1,X0)) & (~sP6(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f726])).
fof(f726,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP6(X2,X1,X0))),
  inference(definition_folding,[],[f557,f725,f724,f723,f722])).
fof(f722,plain,(
  ! [X0,X1,X2] : (sP3(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f723,plain,(
  ! [X1,X0,X2] : (sP4(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f724,plain,(
  ! [X2,X0,X1] : (sP5(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f725,plain,(
  ! [X2,X1,X0] : (sP6(X2,X1,X0) <=> (sP3(X0,X1,X2) & sP4(X1,X0,X2) & sP5(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f557,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f556])).
fof(f556,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f82])).
fof(f82,axiom,(
  ! [X35,X43,X44] : ((p(s(bool,X44)) <=> s(bool,X35) = s(bool,X43)) <=> ((p(s(bool,X44)) | p(s(bool,X43)) | p(s(bool,X35))) & (p(s(bool,X44)) | ~p(s(bool,X35)) | ~p(s(bool,X43))) & (p(s(bool,X43)) | ~p(s(bool,X35)) | ~p(s(bool,X44))) & (p(s(bool,X35)) | ~p(s(bool,X43)) | ~p(s(bool,X44)))))),
  file('thm.extreal.add_not_infty|split|1','thm.sat.sat_2|49|_prove_ltb')).
fof(f46393,plain,(
  sP6(sK83(sK17),'F_0','T_0') | $spl33),
  inference(unit_resulting_resolution,[],[f1890,f46160,f1382])).
fof(f1382,plain,(
  ( ! [X2,X0,X1] : (sP6(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f909])).
fof(f46160,plain,(
  ~p(s(bool,sK83(sK17))) | $spl33),
  inference(unit_resulting_resolution,[],[f949,f45657,f1230])).
fof(f1230,plain,(
  ( ! [X0] : (~p(s(bool,sK83(X0))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)) )),
  inference(cnf_transformation,[],[f861])).
fof(f861,plain,(
  ! [X0] : ((~p(s(bool,sK83(X0))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)) & ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0)) | p(s(bool,sK83(X0)))) & s('type.realax.real','const.extreal.real_1'(s('type.extreal.extreal',X0))) = s('type.realax.real','const.bool.COND_3'(s(bool,sK83(X0)),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.min.@_1'(s(fun('type.realax.real',bool),'_dst_x0x1_2'(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal',X0))),s(fun('type.realax.real','type.extreal.extreal'),'const.extreal.Normal_0'))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK83])],[f860])).
fof(f860,plain,(
  ! [X0] : ? [X1] : ((~p(s(bool,X1)) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)) & ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0)) | p(s(bool,X1))) & s('type.realax.real','const.extreal.real_1'(s('type.extreal.extreal',X0))) = s('type.realax.real','const.bool.COND_3'(s(bool,X1),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.min.@_1'(s(fun('type.realax.real',bool),'_dst_x0x1_2'(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal',X0))),s(fun('type.realax.real','type.extreal.extreal'),'const.extreal.Normal_0'))))))))),
  inference(flattening,[],[f859])).
fof(f859,plain,(
  ! [X0] : ? [X1] : (((~p(s(bool,X1)) | (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0))) & ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0)) | p(s(bool,X1)))) & s('type.realax.real','const.extreal.real_1'(s('type.extreal.extreal',X0))) = s('type.realax.real','const.bool.COND_3'(s(bool,X1),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.min.@_1'(s(fun('type.realax.real',bool),'_dst_x0x1_2'(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal',X0))),s(fun('type.realax.real','type.extreal.extreal'),'const.extreal.Normal_0'))))))))),
  inference(nnf_transformation,[],[f452])).
fof(f452,plain,(
  ! [X0] : ? [X1] : ((p(s(bool,X1)) <=> (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0))) & s('type.realax.real','const.extreal.real_1'(s('type.extreal.extreal',X0))) = s('type.realax.real','const.bool.COND_3'(s(bool,X1),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.min.@_1'(s(fun('type.realax.real',bool),'_dst_x0x1_2'(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal',X0))),s(fun('type.realax.real','type.extreal.extreal'),'const.extreal.Normal_0'))))))))),
  inference(rectify,[],[f96])).
fof(f96,axiom,(
  ! [X6] : ? [X45] : ((p(s(bool,X45)) <=> (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X6) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X6))) & s('type.realax.real','const.extreal.real_1'(s('type.extreal.extreal',X6))) = s('type.realax.real','const.bool.COND_3'(s(bool,X45),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.min.@_1'(s(fun('type.realax.real',bool),'_dst_x0x1_2'(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal',X6))),s(fun('type.realax.real','type.extreal.extreal'),'const.extreal.Normal_0'))))))))),
  file('thm.extreal.add_not_infty|split|1','thm.extreal.real_def')).
fof(f73765,plain,(
  ( ! [X7] : (p(s(bool,X7))) ) | $spl40),
  inference(cnf_transformation,[],[f73765_D])).
fof(f73765_D,plain,(
  ( ! [X7] : (p(s(bool,X7))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f73766,plain,(
  ( ! [X12,X10,X8,X11,X9] : (s(X8,X10) = s(X8,'const.bool.COND_3'(s(bool,X11),s(X8,X12),s(X8,X9))) | s(bool,'F_0') != s(bool,X11) | s(X8,X9) != s(X8,X10)) ) | $spl42),
  inference(cnf_transformation,[],[f73766_D])).
fof(f73766_D,plain,(
  ( ! [X12,X10,X8,X11,X9] : (s(X8,X10) = s(X8,'const.bool.COND_3'(s(bool,X11),s(X8,X12),s(X8,X9))) | s(bool,'F_0') != s(bool,X11) | s(X8,X9) != s(X8,X10)) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f73764,plain,(
  ( ! [X12,X10,X8,X7,X11,X9] : (s(X8,X10) = s(X8,'const.bool.COND_3'(s(bool,X11),s(X8,X12),s(X8,X9))) | p(s(bool,X7)) | s(X8,X9) != s(X8,X10) | s(bool,'F_0') != s(bool,X11)) )),
  inference(forward_demodulation,[],[f73763,f1164])).
fof(f1164,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f405])).
fof(f405,plain,(
  ! [X0,X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f404])).
fof(f404,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f170])).
fof(f170,axiom,(
  ! [X8] : ! [X19,X20] : s(X8,X20) = s(X8,'const.bool.COND_3'(s(bool,'F_0'),s(X8,X19),s(X8,X20)))),
  file('thm.extreal.add_not_infty|split|1','thm.bool.COND_CLAUSES|split|1')).
fof(f73763,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (p(s(bool,X7)) | s(X8,X9) != s(X8,X10) | s(X8,'const.bool.COND_3'(s(bool,'F_0'),s(X8,X13),s(X8,X10))) = s(X8,'const.bool.COND_3'(s(bool,X11),s(X8,X12),s(X8,X9))) | s(bool,'F_0') != s(bool,X11)) )),
  inference(subsumption_resolution,[],[f64851,f1067])).
fof(f64851,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (p(s(bool,X7)) | s(X8,X9) != s(X8,X10) | s(X8,'const.bool.COND_3'(s(bool,'F_0'),s(X8,X13),s(X8,X10))) = s(X8,'const.bool.COND_3'(s(bool,X11),s(X8,X12),s(X8,X9))) | s(bool,'F_0') != s(bool,X11) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f1202,f1021])).
fof(f1021,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.extreal.add_not_infty|split|1',aHL_BOOL_CASES)).
fof(f1202,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,X2)) | s(X0,X5) != s(X0,X6) | s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))) | s(bool,X1) != s(bool,X2)) )),
  inference(cnf_transformation,[],[f632])).
fof(f632,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : (s(bool,X1) != s(bool,X2) | (p(s(bool,X2)) & s(X0,X3) != s(X0,X4)) | (~p(s(bool,X2)) & s(X0,X5) != s(X0,X6)) | s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))))),
  inference(flattening,[],[f631])).
fof(f631,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ((s(bool,X1) != s(bool,X2) | (p(s(bool,X2)) & s(X0,X3) != s(X0,X4)) | (~p(s(bool,X2)) & s(X0,X5) != s(X0,X6))) | s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))))),
  inference(ennf_transformation,[],[f443])).
fof(f443,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ((s(bool,X1) = s(bool,X2) & (p(s(bool,X2)) => s(X0,X3) = s(X0,X4)) & (~p(s(bool,X2)) => s(X0,X5) = s(X0,X6))) => s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))))),
  inference(flattening,[],[f442])).
fof(f442,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6] : ((s(bool,X1) = s(bool,X2) & (p(s(bool,X2)) => s(X0,X3) = s(X0,X4)) & (~p(s(bool,X2)) => s(X0,X5) = s(X0,X6))) => s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))))),
  inference(rectify,[],[f181])).
fof(f181,axiom,(
  ! [X8] : ! [X14,X71,X6,X40,X7,X41] : ((s(bool,X14) = s(bool,X71) & (p(s(bool,X71)) => s(X8,X6) = s(X8,X40)) & (~p(s(bool,X71)) => s(X8,X7) = s(X8,X41))) => s(X8,'const.bool.COND_3'(s(bool,X14),s(X8,X6),s(X8,X7))) = s(X8,'const.bool.COND_3'(s(bool,X71),s(X8,X40),s(X8,X41))))),
  file('thm.extreal.add_not_infty|split|1','thm.bool.COND_CONG')).
fof(f48751,plain,(
  $false | ($spl34 | $spl36)),
  inference(subsumption_resolution,[],[f48750,f1890])).
fof(f48750,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl34 | $spl36)),
  inference(backward_demodulation,[],[f48723,f28562])).
fof(f28562,plain,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))),
  inference(forward_demodulation,[],[f28561,f966])).
fof(f966,plain,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))),
  inference(cnf_transformation,[],[f149])).
fof(f149,axiom,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))),
  file('thm.extreal.add_not_infty|split|1','thm.extreal.extreal_le_def|split|4')).
fof(f28561,plain,(
  s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))),
  inference(forward_demodulation,[],[f28492,f956])).
fof(f956,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.NegInf_0')))),
  inference(cnf_transformation,[],[f92])).
fof(f92,axiom,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.NegInf_0')))),
  file('thm.extreal.add_not_infty|split|1','thm.extreal.extreal_ainv_def|split|0')).
fof(f28492,plain,(
  s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0'))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))),
  inference(superposition,[],[f1250,f964])).
fof(f964,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))),
  inference(cnf_transformation,[],[f49])).
fof(f49,axiom,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))),
  file('thm.extreal.add_not_infty|split|1','thm.extreal.extreal_add_def|split|4')).
fof(f1250,plain,(
  ( ! [X0,X1] : (s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal',X0))),s('type.extreal.extreal',X1))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))))) )),
  inference(cnf_transformation,[],[f469])).
fof(f469,plain,(
  ! [X0,X1] : s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal',X0))),s('type.extreal.extreal',X1))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))))),
  inference(rectify,[],[f67])).
fof(f67,axiom,(
  ! [X6,X7] : s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal',X6))),s('type.extreal.extreal',X7))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X7)))))),
  file('thm.extreal.add_not_infty|split|1','thm.extreal.le_lneg')).
fof(f48723,plain,(
  s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0'))) | ($spl34 | $spl36)),
  inference(backward_demodulation,[],[f48236,f45660])).
fof(f48236,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK18) | $spl36),
  inference(cnf_transformation,[],[f48236_D])).
fof(f48236_D,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK18) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f48238,plain,(
  s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK17))) | $spl38),
  inference(cnf_transformation,[],[f48238_D])).
fof(f48238_D,plain,(
  s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK17))) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f48234,plain,(
  s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK17))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK18)),
  inference(forward_demodulation,[],[f48233,f8889])).
fof(f8889,plain,(
  s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK18),s('type.extreal.extreal','const.extreal.PosInf_0')))),
  inference(unit_resulting_resolution,[],[f952,f7411,f1268])).
fof(f1268,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f659])).
fof(f7411,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK18),s('type.extreal.extreal','const.extreal.PosInf_0'))))),
  inference(unit_resulting_resolution,[],[f950,f1136])).
fof(f1136,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)) )),
  inference(cnf_transformation,[],[f807])).
fof(f807,plain,(
  ! [X0] : ((s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.PosInf_0'))))) & (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0)))),
  inference(nnf_transformation,[],[f365])).
fof(f365,plain,(
  ! [X0] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0) <=> p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.PosInf_0')))))),
  inference(flattening,[],[f364])).
fof(f364,plain,(
  ! [X0] : (~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) <=> p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.PosInf_0')))))),
  inference(rectify,[],[f87])).
fof(f87,axiom,(
  ! [X6] : (~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X6) <=> p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal','const.extreal.PosInf_0')))))),
  file('thm.extreal.add_not_infty|split|1','thm.extreal.lt_infty|split|5')).
fof(f48233,plain,(
  s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK18),s('type.extreal.extreal','const.extreal.PosInf_0'))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK17))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK18)),
  inference(subsumption_resolution,[],[f48143,f950])).
fof(f48143,plain,(
  s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK18),s('type.extreal.extreal','const.extreal.PosInf_0'))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK17))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK18) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK18)),
  inference(superposition,[],[f1263,f951])).
fof(f1263,plain,(
  ( ! [X0,X1] : (s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X0))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X1) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X1)) )),
  inference(cnf_transformation,[],[f655])).
fof(f655,plain,(
  ! [X0,X1] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X1) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X1) | s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X0))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))))),
  inference(flattening,[],[f654])).
fof(f654,plain,(
  ! [X0,X1] : ((s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X1) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X1)) | s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X0))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))))),
  inference(ennf_transformation,[],[f484])).
fof(f484,plain,(
  ! [X0,X1] : ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X1) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X1)) => s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X0))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))))),
  inference(flattening,[],[f483])).
fof(f483,plain,(
  ! [X0,X1] : ((~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X1) & ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X1)) => s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X0))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X6,X7] : ((~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X7) & ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X7)) => s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X6))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X7),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X7))))))),
  file('thm.extreal.add_not_infty|split|1','thm.extreal.lt_addl')).
fof(f46153,plain,(
  $false | $spl32),
  inference(subsumption_resolution,[],[f46152,f7205])).
fof(f7205,plain,(
  s(bool,'T_0') != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18)))),
  inference(unit_resulting_resolution,[],[f952,f5894,f1494])).
fof(f1494,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s(bool,X0) != s(bool,X1) | p(s(bool,X1))) )),
  inference(condensation,[],[f1493])).
fof(f1493,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2)) )),
  inference(condensation,[],[f1448])).
fof(f1448,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f934])).
fof(f934,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f933])).
fof(f933,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f697])).
fof(f697,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f696])).
fof(f696,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f572])).
fof(f572,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f56])).
fof(f56,axiom,(
  ! [X6,X40,X7,X41] : ((s(bool,X6) = s(bool,X40) & (p(s(bool,X40)) => s(bool,X7) = s(bool,X41))) => ((p(s(bool,X6)) => p(s(bool,X7))) <=> (p(s(bool,X40)) => p(s(bool,X41)))))),
  file('thm.extreal.add_not_infty|split|1','thm.bool.IMP_CONG')).
fof(f46152,plain,(
  s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18))) | $spl32),
  inference(forward_demodulation,[],[f46149,f956])).
fof(f46149,plain,(
  s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),s('type.extreal.extreal',sK18))) | $spl32),
  inference(backward_demodulation,[],[f45658,f28570])).
fof(f28570,plain,(
  s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal',sK17))),s('type.extreal.extreal',sK18)))),
  inference(forward_demodulation,[],[f28498,f1911])).
fof(f28498,plain,(
  s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.PosInf_0'))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal',sK17))),s('type.extreal.extreal',sK18)))),
  inference(superposition,[],[f1250,f951])).
fof(f45658,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK17) | $spl32),
  inference(cnf_transformation,[],[f45658_D])).
fof(f45658_D,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK17) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f45660,plain,(
  s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK18))) | $spl34),
  inference(cnf_transformation,[],[f45660_D])).
fof(f45660_D,plain,(
  s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK18))) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f45656,plain,(
  s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK18))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK17)),
  inference(forward_demodulation,[],[f45655,f1911])).
fof(f45655,plain,(
  s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal','const.extreal.PosInf_0'))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK18))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK17)),
  inference(subsumption_resolution,[],[f45555,f949])).
fof(f45555,plain,(
  s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal','const.extreal.PosInf_0'))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK18))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK17) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK17)),
  inference(superposition,[],[f1262,f951])).
fof(f1262,plain,(
  ( ! [X0,X1] : (s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X1))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0)) )),
  inference(cnf_transformation,[],[f653])).
fof(f653,plain,(
  ! [X0,X1] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X1))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))))),
  inference(flattening,[],[f652])).
fof(f652,plain,(
  ! [X0,X1] : ((s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)) | s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X1))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))))),
  inference(ennf_transformation,[],[f482])).
fof(f482,plain,(
  ! [X0,X1] : ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0)) => s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X1))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))))),
  inference(flattening,[],[f481])).
fof(f481,plain,(
  ! [X0,X1] : ((~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) & ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)) => s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X1))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))))),
  inference(rectify,[],[f64])).
fof(f64,axiom,(
  ! [X6,X7] : ((~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X6) & ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X6)) => s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X7))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X7))))))),
  file('thm.extreal.add_not_infty|split|1','thm.extreal.le_addr')).
fof(f45417,plain,(
  ( ! [X8] : (p(s(bool,sK74('const.combin.I_0',X8,bool))) | p(s(bool,happ(s(fun(bool,bool),X8),s(bool,sK74('const.combin.I_0',X8,bool)))))) ) | $spl30),
  inference(cnf_transformation,[],[f45417_D])).
fof(f45417_D,plain,(
  ( ! [X8] : (p(s(bool,sK74('const.combin.I_0',X8,bool))) | p(s(bool,happ(s(fun(bool,bool),X8),s(bool,sK74('const.combin.I_0',X8,bool)))))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f45416,plain,(
  ( ! [X8,X7] : (p(s(bool,sK74('const.combin.I_0',X8,bool))) | ~p(s(bool,X7)) | p(s(bool,happ(s(fun(bool,bool),X8),s(bool,sK74('const.combin.I_0',X8,bool)))))) )),
  inference(forward_demodulation,[],[f45392,f1514])).
fof(f45392,plain,(
  ( ! [X8,X7] : (~p(s(bool,X7)) | p(s(bool,happ(s(fun(bool,bool),X8),s(bool,sK74('const.combin.I_0',X8,bool))))) | p(s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,sK74('const.combin.I_0',X8,bool)))))) )),
  inference(superposition,[],[f1189,f1514])).
fof(f1189,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK74(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK74(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f838])).
fof(f838,plain,(
  ! [X0,X1,X2] : ((! [X8] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X8)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X8))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK72(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK73(X2,X1,X0)))))) & ((! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK74(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK74(X2,X1,X0))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK72,sK73,sK74])],[f837])).
fof(f837,plain,(
  ! [X0,X1,X2] : ((! [X8] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X8)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X8))))) | ? [X7] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X7)))) | ? [X6] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) & ((! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(rectify,[],[f836])).
fof(f836,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) & ((! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(flattening,[],[f835])).
fof(f835,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & ((! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f425])).
fof(f425,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f424])).
fof(f424,plain,(
  ! [X0] : ! [X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f211])).
fof(f211,axiom,(
  ! [X8] : ! [X14,X71] : (? [X6] : (p(s(bool,happ(s(fun(X8,bool),X14),s(X8,X6)))) | p(s(bool,happ(s(fun(X8,bool),X71),s(X8,X6))))) <=> (? [X6] : p(s(bool,happ(s(fun(X8,bool),X14),s(X8,X6)))) | ? [X6] : p(s(bool,happ(s(fun(X8,bool),X71),s(X8,X6))))))),
  file('thm.extreal.add_not_infty|split|1','thm.bool.EXISTS_OR_THM')).
fof(f45262,plain,(
  ( ! [X8] : (p(s(bool,sK74(X8,'const.combin.I_0',bool))) | p(s(bool,happ(s(fun(bool,bool),X8),s(bool,sK74(X8,'const.combin.I_0',bool)))))) ) | $spl28),
  inference(cnf_transformation,[],[f45262_D])).
fof(f45262_D,plain,(
  ( ! [X8] : (p(s(bool,sK74(X8,'const.combin.I_0',bool))) | p(s(bool,happ(s(fun(bool,bool),X8),s(bool,sK74(X8,'const.combin.I_0',bool)))))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f45261,plain,(
  ( ! [X8,X7] : (p(s(bool,sK74(X8,'const.combin.I_0',bool))) | ~p(s(bool,X7)) | p(s(bool,happ(s(fun(bool,bool),X8),s(bool,sK74(X8,'const.combin.I_0',bool)))))) )),
  inference(forward_demodulation,[],[f45237,f1514])).
fof(f45237,plain,(
  ( ! [X8,X7] : (~p(s(bool,X7)) | p(s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,sK74(X8,'const.combin.I_0',bool))))) | p(s(bool,happ(s(fun(bool,bool),X8),s(bool,sK74(X8,'const.combin.I_0',bool)))))) )),
  inference(superposition,[],[f1188,f1514])).
fof(f1188,plain,(
  ( ! [X2,X0,X5,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK74(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK74(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f838])).
fof(f45111,plain,(
  ( ! [X8] : (p(s(bool,sK73('const.combin.I_0',X8,bool))) | p(s(bool,happ(s(fun(bool,bool),X8),s(bool,sK72('const.combin.I_0',X8,bool)))))) ) | $spl26),
  inference(cnf_transformation,[],[f45111_D])).
fof(f45111_D,plain,(
  ( ! [X8] : (p(s(bool,sK73('const.combin.I_0',X8,bool))) | p(s(bool,happ(s(fun(bool,bool),X8),s(bool,sK72('const.combin.I_0',X8,bool)))))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f45110,plain,(
  ( ! [X8,X7] : (p(s(bool,sK73('const.combin.I_0',X8,bool))) | ~p(s(bool,X7)) | p(s(bool,happ(s(fun(bool,bool),X8),s(bool,sK72('const.combin.I_0',X8,bool)))))) )),
  inference(forward_demodulation,[],[f45086,f1514])).
fof(f45086,plain,(
  ( ! [X8,X7] : (~p(s(bool,X7)) | p(s(bool,happ(s(fun(bool,bool),X8),s(bool,sK72('const.combin.I_0',X8,bool))))) | p(s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,sK73('const.combin.I_0',X8,bool)))))) )),
  inference(superposition,[],[f1187,f1514])).
fof(f1187,plain,(
  ( ! [X2,X0,X8,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X8)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK72(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK73(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f838])).
fof(f44956,plain,(
  ( ! [X8] : (p(s(bool,sK72(X8,'const.combin.I_0',bool))) | p(s(bool,happ(s(fun(bool,bool),X8),s(bool,sK73(X8,'const.combin.I_0',bool)))))) ) | $spl24),
  inference(cnf_transformation,[],[f44956_D])).
fof(f44956_D,plain,(
  ( ! [X8] : (p(s(bool,sK72(X8,'const.combin.I_0',bool))) | p(s(bool,happ(s(fun(bool,bool),X8),s(bool,sK73(X8,'const.combin.I_0',bool)))))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f44955,plain,(
  ( ! [X8,X7] : (p(s(bool,sK72(X8,'const.combin.I_0',bool))) | ~p(s(bool,X7)) | p(s(bool,happ(s(fun(bool,bool),X8),s(bool,sK73(X8,'const.combin.I_0',bool)))))) )),
  inference(forward_demodulation,[],[f44931,f1514])).
fof(f44931,plain,(
  ( ! [X8,X7] : (~p(s(bool,X7)) | p(s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,sK72(X8,'const.combin.I_0',bool))))) | p(s(bool,happ(s(fun(bool,bool),X8),s(bool,sK73(X8,'const.combin.I_0',bool)))))) )),
  inference(superposition,[],[f1186,f1514])).
fof(f1186,plain,(
  ( ! [X2,X0,X8,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X8)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK72(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK73(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f838])).
fof(f34924,plain,(
  ( ! [X72] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X72)))) ) | $spl20),
  inference(cnf_transformation,[],[f34924_D])).
fof(f34924_D,plain,(
  ( ! [X72] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X72)))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f34926,plain,(
  s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK41('const.extreal.extreal_inv_1'(s('type.extreal.extreal','const.extreal.PosInf_0')))))) | $spl22),
  inference(cnf_transformation,[],[f34926_D])).
fof(f34926_D,plain,(
  s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK41('const.extreal.extreal_inv_1'(s('type.extreal.extreal','const.extreal.PosInf_0')))))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f34923,plain,(
  ( ! [X72] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK41('const.extreal.extreal_inv_1'(s('type.extreal.extreal','const.extreal.PosInf_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X72)))) )),
  inference(forward_demodulation,[],[f34922,f1509])).
fof(f1509,plain,(
  s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.extreal_inv_1'(s('type.extreal.extreal','const.extreal.PosInf_0')))),
  inference(backward_demodulation,[],[f1004,f973])).
fof(f973,plain,(
  s('type.extreal.extreal','const.extreal.extreal_inv_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(cnf_transformation,[],[f136])).
fof(f136,axiom,(
  s('type.extreal.extreal','const.extreal.extreal_inv_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.extreal.add_not_infty|split|1','thm.extreal.extreal_inv_def|split|1')).
fof(f1004,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f291])).
fof(f291,plain,(
  ! [X0] : s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f85])).
fof(f85,axiom,(
  ! [X31] : s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X31))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X31)))))),
  file('thm.extreal.add_not_infty|split|1','thm.extreal.extreal_of_num_def')).
fof(f34922,plain,(
  ( ! [X72] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X72))) | s('type.extreal.extreal','const.extreal.extreal_inv_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK41('const.extreal.extreal_inv_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))))))) )),
  inference(subsumption_resolution,[],[f34921,f6642])).
fof(f6642,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.extreal_inv_1'(s('type.extreal.extreal','const.extreal.PosInf_0')))),
  inference(subsumption_resolution,[],[f6641,f1508])).
fof(f6641,plain,(
  p(s(bool,'F_0')) | s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.extreal_inv_1'(s('type.extreal.extreal','const.extreal.PosInf_0')))),
  inference(forward_demodulation,[],[f6542,f3894])).
fof(f3894,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0)))))) )),
  inference(superposition,[],[f1006,f1004])).
fof(f1006,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))))) )),
  inference(cnf_transformation,[],[f293])).
fof(f293,plain,(
  ! [X0] : s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))))),
  inference(rectify,[],[f148])).
fof(f148,axiom,(
  ! [X36] : s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X36)))))),
  file('thm.extreal.add_not_infty|split|1','thm.extreal.extreal_le_def|split|6')).
fof(f6542,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.extreal_inv_1'(s('type.extreal.extreal','const.extreal.PosInf_0')))),
  inference(superposition,[],[f1088,f1509])).
fof(f1088,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X0)))) | s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0)) )),
  inference(cnf_transformation,[],[f775])).
fof(f34921,plain,(
  ( ! [X72] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X72))) | s('type.extreal.extreal','const.extreal.extreal_inv_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK41('const.extreal.extreal_inv_1'(s('type.extreal.extreal','const.extreal.PosInf_0')))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_inv_1'(s('type.extreal.extreal','const.extreal.PosInf_0')))) )),
  inference(subsumption_resolution,[],[f34841,f5877])).
fof(f5877,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.extreal_inv_1'(s('type.extreal.extreal','const.extreal.PosInf_0')))),
  inference(subsumption_resolution,[],[f5778,f5659])).
fof(f5659,plain,(
  ( ! [X0] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0))),s('type.extreal.extreal','const.extreal.NegInf_0'))))) )),
  inference(unit_resulting_resolution,[],[f1143,f1085])).
fof(f1085,plain,(
  ( ! [X0] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0)) )),
  inference(cnf_transformation,[],[f774])).
fof(f774,plain,(
  ! [X0] : ((~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0)) & (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0'))))))),
  inference(nnf_transformation,[],[f344])).
fof(f344,plain,(
  ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0')))) <=> s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0))),
  inference(rectify,[],[f137])).
fof(f137,axiom,(
  ! [X6] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal','const.extreal.NegInf_0')))) <=> s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X6))),
  file('thm.extreal.add_not_infty|split|1','thm.extreal.le_infty|split|2')).
fof(f1143,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f375])).
fof(f375,plain,(
  ! [X0] : s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0)))),
  inference(flattening,[],[f374])).
fof(f374,plain,(
  ! [X0] : ~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0)))),
  inference(rectify,[],[f230])).
fof(f230,axiom,(
  ! [X31] : ~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X31)))),
  file('thm.extreal.add_not_infty|split|1','thm.extreal.num_not_infty|split|0')).
fof(f5778,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.extreal_inv_1'(s('type.extreal.extreal','const.extreal.PosInf_0')))),
  inference(superposition,[],[f1086,f1509])).
fof(f1086,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0)) )),
  inference(cnf_transformation,[],[f774])).
fof(f34841,plain,(
  ( ! [X72] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X72))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_inv_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))) | s('type.extreal.extreal','const.extreal.extreal_inv_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK41('const.extreal.extreal_inv_1'(s('type.extreal.extreal','const.extreal.PosInf_0')))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_inv_1'(s('type.extreal.extreal','const.extreal.PosInf_0')))) )),
  inference(superposition,[],[f1034,f1509])).
fof(f1034,plain,(
  ( ! [X2,X0] : (s('type.extreal.extreal',X0) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK41(X0)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)) )),
  inference(cnf_transformation,[],[f757])).
fof(f757,plain,(
  ! [X0] : ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0) & ! [X2] : s('type.extreal.extreal',X0) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) | s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK41(X0)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK41])],[f756])).
fof(f756,plain,(
  ! [X0] : ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0) & ! [X2] : s('type.extreal.extreal',X0) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) | ? [X1] : s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0))),
  inference(rectify,[],[f607])).
fof(f607,plain,(
  ! [X0] : ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0) & ! [X1] : s('type.extreal.extreal',X0) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) | ? [X2] : s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0))),
  inference(flattening,[],[f606])).
fof(f606,plain,(
  ! [X0] : ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0) & ! [X1] : s('type.extreal.extreal',X0) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) | (? [X2] : s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)))),
  inference(ennf_transformation,[],[f321])).
fof(f321,plain,(
  ! [X0] : ((s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | ? [X1] : s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) => (? [X2] : s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)))),
  inference(rectify,[],[f93])).
fof(f93,axiom,(
  ! [X12] : ((s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X12) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X12) | ? [X6] : s('type.extreal.extreal',X12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X6)))) => (? [X6] : s('type.extreal.extreal',X12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X6))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X12) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X12)))),
  file('thm.extreal.add_not_infty|split|1','thm.extreal.extreal_46_prove_ltb')).
fof(f27955,plain,(
  ( ! [X8] : (p(s(bool,sK69(X8,'const.combin.I_0',bool))) | p(s(bool,X8))) ) | $spl18),
  inference(cnf_transformation,[],[f27955_D])).
fof(f27955_D,plain,(
  ( ! [X8] : (p(s(bool,sK69(X8,'const.combin.I_0',bool))) | p(s(bool,X8))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f27954,plain,(
  ( ! [X8,X7] : (p(s(bool,sK69(X8,'const.combin.I_0',bool))) | ~p(s(bool,X7)) | p(s(bool,X8))) )),
  inference(forward_demodulation,[],[f27935,f1514])).
fof(f27935,plain,(
  ( ! [X8,X7] : (~p(s(bool,X7)) | p(s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,sK69(X8,'const.combin.I_0',bool))))) | p(s(bool,X8))) )),
  inference(superposition,[],[f1180,f1514])).
fof(f1180,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK69(X2,X1,X0))))) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f830])).
fof(f830,plain,(
  ! [X0,X1,X2] : (((! [X6] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) & ~p(s(bool,X2))) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK68(X2,X1,X0))))) | p(s(bool,X2)))) & (! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ~p(s(bool,X2))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK69(X2,X1,X0))))) | p(s(bool,X2))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK68,sK69])],[f829])).
fof(f829,plain,(
  ! [X0,X1,X2] : (((! [X6] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) & ~p(s(bool,X2))) | ? [X5] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | p(s(bool,X2)))) & (! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ~p(s(bool,X2))) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2))))),
  inference(rectify,[],[f828])).
fof(f828,plain,(
  ! [X0,X1,X2] : (((! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & ~p(s(bool,X2))) | ? [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,X2)))) & (! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ~p(s(bool,X2))) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2))))),
  inference(flattening,[],[f827])).
fof(f827,plain,(
  ! [X0,X1,X2] : (((! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & ~p(s(bool,X2))) | ? [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,X2)))) & (! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ~p(s(bool,X2))) | (? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2)))))),
  inference(nnf_transformation,[],[f421])).
fof(f421,plain,(
  ! [X0,X1,X2] : ((? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2))) <=> ? [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,X2))))),
  inference(flattening,[],[f420])).
fof(f420,plain,(
  ! [X0] : ! [X1,X2] : ((? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2))) <=> ? [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,X2))))),
  inference(rectify,[],[f219])).
fof(f219,axiom,(
  ! [X8] : ! [X14,X71] : ((? [X6] : p(s(bool,happ(s(fun(X8,bool),X14),s(X8,X6)))) | p(s(bool,X71))) <=> ? [X6] : (p(s(bool,happ(s(fun(X8,bool),X14),s(X8,X6)))) | p(s(bool,X71))))),
  file('thm.extreal.add_not_infty|split|1','thm.bool.LEFT_OR_EXISTS_THM')).
fof(f27784,plain,(
  ( ! [X8] : (p(s(bool,sK68(X8,'const.combin.I_0',bool))) | p(s(bool,X8))) ) | $spl16),
  inference(cnf_transformation,[],[f27784_D])).
fof(f27784_D,plain,(
  ( ! [X8] : (p(s(bool,sK68(X8,'const.combin.I_0',bool))) | p(s(bool,X8))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f27783,plain,(
  ( ! [X8,X7] : (p(s(bool,sK68(X8,'const.combin.I_0',bool))) | ~p(s(bool,X7)) | p(s(bool,X8))) )),
  inference(forward_demodulation,[],[f27764,f1514])).
fof(f27764,plain,(
  ( ! [X8,X7] : (~p(s(bool,X7)) | p(s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,sK68(X8,'const.combin.I_0',bool))))) | p(s(bool,X8))) )),
  inference(superposition,[],[f1178,f1514])).
fof(f1178,plain,(
  ( ! [X6,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK68(X2,X1,X0))))) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f830])).
fof(f27650,plain,(
  ( ! [X8] : (p(s(bool,sK67('const.combin.I_0',X8,bool))) | p(s(bool,X8))) ) | $spl14),
  inference(cnf_transformation,[],[f27650_D])).
fof(f27650_D,plain,(
  ( ! [X8] : (p(s(bool,sK67('const.combin.I_0',X8,bool))) | p(s(bool,X8))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f27649,plain,(
  ( ! [X8,X7] : (p(s(bool,sK67('const.combin.I_0',X8,bool))) | ~p(s(bool,X7)) | p(s(bool,X8))) )),
  inference(forward_demodulation,[],[f27630,f1514])).
fof(f27630,plain,(
  ( ! [X8,X7] : (~p(s(bool,X7)) | p(s(bool,X8)) | p(s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,sK67('const.combin.I_0',X8,bool)))))) )),
  inference(superposition,[],[f1177,f1514])).
fof(f1177,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK67(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f826])).
fof(f826,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X6] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK66(X2,X1,X0))))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK67(X2,X1,X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK66,sK67])],[f825])).
fof(f825,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X6] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | ? [X5] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(rectify,[],[f824])).
fof(f824,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(flattening,[],[f823])).
fof(f823,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f419])).
fof(f419,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f418])).
fof(f418,plain,(
  ! [X0] : ! [X1,X2] : ((p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f179])).
fof(f179,axiom,(
  ! [X8] : ! [X14,X71] : ((p(s(bool,X14)) | ? [X6] : p(s(bool,happ(s(fun(X8,bool),X71),s(X8,X6))))) <=> ? [X6] : (p(s(bool,X14)) | p(s(bool,happ(s(fun(X8,bool),X71),s(X8,X6))))))),
  file('thm.extreal.add_not_infty|split|1','thm.bool.RIGHT_OR_EXISTS_THM')).
fof(f27492,plain,(
  ( ! [X8] : (p(s(bool,sK66('const.combin.I_0',X8,bool))) | p(s(bool,X8))) ) | $spl12),
  inference(cnf_transformation,[],[f27492_D])).
fof(f27492_D,plain,(
  ( ! [X8] : (p(s(bool,sK66('const.combin.I_0',X8,bool))) | p(s(bool,X8))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f27491,plain,(
  ( ! [X8,X7] : (p(s(bool,sK66('const.combin.I_0',X8,bool))) | ~p(s(bool,X7)) | p(s(bool,X8))) )),
  inference(forward_demodulation,[],[f27472,f1514])).
fof(f27472,plain,(
  ( ! [X8,X7] : (~p(s(bool,X7)) | p(s(bool,X8)) | p(s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,sK66('const.combin.I_0',X8,bool)))))) )),
  inference(superposition,[],[f1175,f1514])).
fof(f1175,plain,(
  ( ! [X6,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK66(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f826])).
fof(f25691,plain,(
  p(s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0')))) | $spl10),
  inference(cnf_transformation,[],[f25691_D])).
fof(f25691_D,plain,(
  p(s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0')))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f25689,plain,(
  ( ! [X3] : (p(s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0')))) | ~p(s(bool,X3))) )),
  inference(forward_demodulation,[],[f25674,f1514])).
fof(f25674,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | p(s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))))))) )),
  inference(superposition,[],[f1167,f1514])).
fof(f1167,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1))))))) )),
  inference(cnf_transformation,[],[f627])).
fof(f627,plain,(
  ! [X0,X1,X2] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(ennf_transformation,[],[f411])).
fof(f411,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(flattening,[],[f410])).
fof(f410,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X8] : ! [X14,X6] : (p(s(bool,happ(s(fun(X8,bool),X14),s(X8,X6)))) => p(s(bool,happ(s(fun(X8,bool),X14),s(X8,'const.min.@_1'(s(fun(X8,bool),X14)))))))),
  file('thm.extreal.add_not_infty|split|1','thm.bool.SELECT_AX')).
fof(f22843,plain,(
  p(s(bool,sK62('const.combin.I_0',bool))) | $spl8),
  inference(cnf_transformation,[],[f22843_D])).
fof(f22843_D,plain,(
  p(s(bool,sK62('const.combin.I_0',bool))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f22841,plain,(
  ( ! [X3] : (p(s(bool,sK62('const.combin.I_0',bool))) | ~p(s(bool,X3))) )),
  inference(forward_demodulation,[],[f22826,f1514])).
fof(f22826,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | p(s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,sK62('const.combin.I_0',bool)))))) )),
  inference(superposition,[],[f1158,f1514])).
fof(f1158,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK62(X1,X0)))))) )),
  inference(cnf_transformation,[],[f814])).
fof(f814,plain,(
  ! [X0,X1] : ((p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK61(X1,X0))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) & (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK62(X1,X0))))) | ! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK61,sK62])],[f813])).
fof(f813,plain,(
  ! [X0,X1] : ((? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) & (? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(rectify,[],[f812])).
fof(f812,plain,(
  ! [X0,X1] : ((? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(nnf_transformation,[],[f626])).
fof(f626,plain,(
  ! [X0,X1] : (! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(ennf_transformation,[],[f399])).
fof(f399,plain,(
  ! [X0,X1] : (~? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(flattening,[],[f398])).
fof(f398,plain,(
  ! [X0] : ! [X1] : (~? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(rectify,[],[f178])).
fof(f178,axiom,(
  ! [X8] : ! [X14] : (~? [X6] : p(s(bool,happ(s(fun(X8,bool),X14),s(X8,X6)))) <=> ! [X6] : ~p(s(bool,happ(s(fun(X8,bool),X14),s(X8,X6)))))),
  file('thm.extreal.add_not_infty|split|1','thm.bool.NOT_EXISTS_THM')).
fof(f22649,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15412,f21752])).
fof(f21752,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) != s(X0,X2)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f15405,f1356])).
fof(f1356,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),'=_1'(s(X0,X1))),s(X0,X2)))) | s(X0,X1) != s(X0,X2)) )),
  inference(cnf_transformation,[],[f895])).
fof(f895,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X0,bool),'=_1'(s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),'=_1'(s(X0,X1))),s(X0,X2))))))),
  inference(nnf_transformation,[],[f554])).
fof(f554,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),'=_1'(s(X0,X1))),s(X0,X2)))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f117])).
fof(f117,axiom,(
  ! [X56,X57,X58] : (p(s(bool,happ(s(fun(X56,bool),'=_1'(s(X56,X57))),s(X56,X58)))) <=> s(X56,X57) = s(X56,X58))),
  file('thm.extreal.add_not_infty|split|1',arity_1_HL_EQ)).
fof(f15412,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f15405,f1072])).
fof(f1072,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f767])).
fof(f767,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f336])).
fof(f336,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f335])).
fof(f335,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X21] : (s(bool,'F_0') = s(bool,X21) <=> ~p(s(bool,X21)))),
  file('thm.extreal.add_not_infty|split|1','thm.bool.EQ_CLAUSES|split|3')).
fof(f22648,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15414,f21752])).
fof(f15414,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1508,f15405,f1265])).
fof(f22647,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15415,f21752])).
fof(f15415,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f15405,f15405,f1265])).
fof(f22646,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15418,f1508])).
fof(f15418,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f1890,f15405,f1265])).
fof(f22645,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15420,f21752])).
fof(f15420,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1508,f15405,f1265])).
fof(f22644,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15421,f21752])).
fof(f15421,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f15405,f15405,f1265])).
fof(f22643,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15424,f1508])).
fof(f15424,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f1890,f15405,f1265])).
fof(f22642,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15426,f1508])).
fof(f15426,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f1956,f15405,f1323])).
fof(f1323,plain,(
  ( ! [X0,X1] : (~sP2(X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f890])).
fof(f890,plain,(
  ! [X0,X1] : ((~sP2(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP2(X0,X1)))),
  inference(rectify,[],[f889])).
fof(f889,plain,(
  ! [X1,X0] : ((~sP2(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP2(X1,X0)))),
  inference(flattening,[],[f888])).
fof(f888,plain,(
  ! [X1,X0] : ((~sP2(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP2(X1,X0)))),
  inference(nnf_transformation,[],[f720])).
fof(f720,plain,(
  ! [X1,X0] : (sP2(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f1956,plain,(
  sP2('T_0','F_0')),
  inference(unit_resulting_resolution,[],[f1508,f952,f1325])).
fof(f1325,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP2(X0,X1)) )),
  inference(cnf_transformation,[],[f890])).
fof(f22641,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15427,f15405])).
fof(f15427,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK17)))) | $spl4),
  inference(unit_resulting_resolution,[],[f5960,f15405,f1323])).
fof(f5960,plain,(
  sP2('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK17)))),
  inference(unit_resulting_resolution,[],[f952,f5893,f1325])).
fof(f5893,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK17))))),
  inference(unit_resulting_resolution,[],[f949,f1087])).
fof(f22640,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15428,f15405])).
fof(f15428,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18)))) | $spl4),
  inference(unit_resulting_resolution,[],[f6872,f15405,f1323])).
fof(f6872,plain,(
  sP2('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18)))),
  inference(unit_resulting_resolution,[],[f952,f5894,f1325])).
fof(f22639,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15430,f15405])).
fof(f15430,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f7732,f15405,f1323])).
fof(f7732,plain,(
  sP2('F_0','const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal','const.extreal.PosInf_0')))),
  inference(unit_resulting_resolution,[],[f1508,f7410,f1328])).
fof(f1328,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP2(X0,X1)) )),
  inference(cnf_transformation,[],[f890])).
fof(f7410,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal','const.extreal.PosInf_0'))))),
  inference(unit_resulting_resolution,[],[f949,f1136])).
fof(f22638,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15431,f15405])).
fof(f15431,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK18),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f8898,f15405,f1323])).
fof(f8898,plain,(
  sP2('F_0','const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK18),s('type.extreal.extreal','const.extreal.PosInf_0')))),
  inference(unit_resulting_resolution,[],[f1508,f7411,f1328])).
fof(f22637,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15432,f15405])).
fof(f15432,plain,(
  p(s(bool,sK77(sK63('const.num.|48|_0','type.num.num')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2042,f15405,f1323])).
fof(f2042,plain,(
  sP2('F_0',sK77(sK63('const.num.|48|_0','type.num.num')))),
  inference(unit_resulting_resolution,[],[f1508,f2030,f1328])).
fof(f2030,plain,(
  p(s(bool,sK77(sK63('const.num.|48|_0','type.num.num'))))),
  inference(unit_resulting_resolution,[],[f1163,f1213])).
fof(f1213,plain,(
  ( ! [X0] : (p(s(bool,sK77(X0))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f843])).
fof(f843,plain,(
  ! [X0] : ((~p(s(bool,sK77(X0))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) & (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | p(s(bool,sK77(X0)))) & s('type.extreal.extreal','const.extreal.extreal_pow_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.num.num',X0))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,sK77(X0)),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.extreal.extreal','const.extreal.PosInf_0'))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK77])],[f842])).
fof(f842,plain,(
  ! [X0] : ? [X1] : ((~p(s(bool,X1)) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) & (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | p(s(bool,X1))) & s('type.extreal.extreal','const.extreal.extreal_pow_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.num.num',X0))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X1),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.extreal.extreal','const.extreal.PosInf_0'))))),
  inference(flattening,[],[f841])).
fof(f841,plain,(
  ! [X0] : ? [X1] : (((~p(s(bool,X1)) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) & (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | p(s(bool,X1)))) & s('type.extreal.extreal','const.extreal.extreal_pow_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.num.num',X0))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X1),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.extreal.extreal','const.extreal.PosInf_0'))))),
  inference(nnf_transformation,[],[f446])).
fof(f446,plain,(
  ! [X0] : ? [X1] : ((p(s(bool,X1)) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) & s('type.extreal.extreal','const.extreal.extreal_pow_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.num.num',X0))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X1),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.extreal.extreal','const.extreal.PosInf_0'))))),
  inference(rectify,[],[f133])).
fof(f133,axiom,(
  ! [X31] : ? [X45] : ((p(s(bool,X45)) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X31)) & s('type.extreal.extreal','const.extreal.extreal_pow_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.num.num',X31))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X45),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.extreal.extreal','const.extreal.PosInf_0'))))),
  file('thm.extreal.add_not_infty|split|1','thm.extreal.extreal_pow_def|split|1')).
fof(f1163,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,sK63(X1,X0))) )),
  inference(cnf_transformation,[],[f819])).
fof(f819,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,sK63(X1,X0))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK63])],[f403])).
fof(f403,plain,(
  ! [X0,X1] : ? [X2] : s(X0,X1) = s(X0,X2)),
  inference(flattening,[],[f402])).
fof(f402,plain,(
  ! [X0] : ! [X1] : ? [X2] : s(X0,X1) = s(X0,X2)),
  inference(rectify,[],[f217])).
fof(f217,axiom,(
  ! [X8] : ! [X12] : ? [X6] : s(X8,X6) = s(X8,X12)),
  file('thm.extreal.add_not_infty|split|1','thm.bool.EXISTS_REFL')).
fof(f22636,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15433,f15405])).
fof(f15433,plain,(
  p(s(bool,sK78(sK63('const.num.|48|_0','type.num.num')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2089,f15405,f1323])).
fof(f2089,plain,(
  sP2('F_0',sK78(sK63('const.num.|48|_0','type.num.num')))),
  inference(unit_resulting_resolution,[],[f1508,f2068,f1328])).
fof(f2068,plain,(
  p(s(bool,sK78(sK63('const.num.|48|_0','type.num.num'))))),
  inference(unit_resulting_resolution,[],[f1163,f1216])).
fof(f1216,plain,(
  ( ! [X0] : (p(s(bool,sK78(X0))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f846])).
fof(f846,plain,(
  ! [X0] : ((~p(s(bool,sK78(X0))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) & (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | p(s(bool,sK78(X0)))) & s('type.extreal.extreal','const.extreal.extreal_pow_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.num.num',X0))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,sK78(X0)),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num',X0))),s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK78])],[f845])).
fof(f845,plain,(
  ! [X0] : ? [X1] : ((~p(s(bool,X1)) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) & (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | p(s(bool,X1))) & s('type.extreal.extreal','const.extreal.extreal_pow_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.num.num',X0))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X1),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num',X0))),s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))))))),
  inference(flattening,[],[f844])).
fof(f844,plain,(
  ! [X0] : ? [X1] : (((~p(s(bool,X1)) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) & (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | p(s(bool,X1)))) & s('type.extreal.extreal','const.extreal.extreal_pow_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.num.num',X0))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X1),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num',X0))),s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))))))),
  inference(nnf_transformation,[],[f447])).
fof(f447,plain,(
  ! [X0] : ? [X1] : ((p(s(bool,X1)) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) & s('type.extreal.extreal','const.extreal.extreal_pow_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.num.num',X0))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X1),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num',X0))),s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))))))),
  inference(rectify,[],[f128])).
fof(f128,axiom,(
  ! [X31] : ? [X45] : ((p(s(bool,X45)) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X31)) & s('type.extreal.extreal','const.extreal.extreal_pow_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.num.num',X31))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X45),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num',X31))),s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))))))),
  file('thm.extreal.add_not_infty|split|1','thm.extreal.extreal_pow_def|split|2')).
fof(f22635,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15434,f15405])).
fof(f15434,plain,(
  p(s(bool,sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))))) | $spl4),
  inference(unit_resulting_resolution,[],[f2125,f15405,f1323])).
fof(f2125,plain,(
  sP2('F_0',sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))))),
  inference(unit_resulting_resolution,[],[f1508,f2108,f1328])).
fof(f2108,plain,(
  p(s(bool,sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))))),
  inference(unit_resulting_resolution,[],[f951,f1232])).
fof(f1232,plain,(
  ( ! [X0] : (p(s(bool,sK83(X0))) | s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0)) )),
  inference(cnf_transformation,[],[f861])).
fof(f22634,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15435,f15405])).
fof(f15435,plain,(
  p(s(bool,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2370,f15405,f1323])).
fof(f2370,plain,(
  sP2('F_0',sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))),
  inference(unit_resulting_resolution,[],[f1508,f2075,f1328])).
fof(f2075,plain,(
  p(s(bool,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal'))))),
  inference(unit_resulting_resolution,[],[f1163,f1231])).
fof(f1231,plain,(
  ( ! [X0] : (p(s(bool,sK83(X0))) | s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0)) )),
  inference(cnf_transformation,[],[f861])).
fof(f22633,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15436,f15405])).
fof(f15436,plain,(
  p(s(bool,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2489,f15405,f1323])).
fof(f2489,plain,(
  sP2('F_0',sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))),
  inference(unit_resulting_resolution,[],[f1508,f2114,f1328])).
fof(f2114,plain,(
  p(s(bool,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal'))))),
  inference(unit_resulting_resolution,[],[f1163,f1232])).
fof(f22632,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15437,f1508])).
fof(f15437,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f7731,f15405,f1323])).
fof(f7731,plain,(
  sP2('const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal','const.extreal.PosInf_0')),'F_0')),
  inference(unit_resulting_resolution,[],[f1508,f7410,f1325])).
fof(f22631,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15438,f1508])).
fof(f15438,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f8897,f15405,f1323])).
fof(f8897,plain,(
  sP2('const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK18),s('type.extreal.extreal','const.extreal.PosInf_0')),'F_0')),
  inference(unit_resulting_resolution,[],[f1508,f7411,f1325])).
fof(f22630,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15441,f1508])).
fof(f15441,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f2041,f15405,f1323])).
fof(f2041,plain,(
  sP2(sK77(sK63('const.num.|48|_0','type.num.num')),'F_0')),
  inference(unit_resulting_resolution,[],[f1508,f2030,f1325])).
fof(f22629,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15442,f1508])).
fof(f15442,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f2088,f15405,f1323])).
fof(f2088,plain,(
  sP2(sK78(sK63('const.num.|48|_0','type.num.num')),'F_0')),
  inference(unit_resulting_resolution,[],[f1508,f2068,f1325])).
fof(f22628,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15443,f1508])).
fof(f15443,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f2124,f15405,f1323])).
fof(f2124,plain,(
  sP2(sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))),'F_0')),
  inference(unit_resulting_resolution,[],[f1508,f2108,f1325])).
fof(f22627,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15444,f1508])).
fof(f15444,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f2369,f15405,f1323])).
fof(f2369,plain,(
  sP2(sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')),'F_0')),
  inference(unit_resulting_resolution,[],[f1508,f2075,f1325])).
fof(f22626,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15445,f1508])).
fof(f15445,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f2488,f15405,f1323])).
fof(f2488,plain,(
  sP2(sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')),'F_0')),
  inference(unit_resulting_resolution,[],[f1508,f2114,f1325])).
fof(f22625,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15453,f1508])).
fof(f15453,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f1965,f15405,f1323])).
fof(f1965,plain,(
  sP2('F_0','T_0')),
  inference(unit_resulting_resolution,[],[f1508,f952,f1328])).
fof(f22624,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15454,f1508])).
fof(f15454,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f7732,f15405,f1323])).
fof(f22623,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15455,f1508])).
fof(f15455,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f8898,f15405,f1323])).
fof(f22622,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15456,f1508])).
fof(f15456,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f2042,f15405,f1323])).
fof(f22621,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15457,f1508])).
fof(f15457,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f2089,f15405,f1323])).
fof(f22620,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15458,f1508])).
fof(f15458,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f2125,f15405,f1323])).
fof(f22619,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15459,f1508])).
fof(f15459,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f2370,f15405,f1323])).
fof(f22618,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15460,f1508])).
fof(f15460,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f2489,f15405,f1323])).
fof(f22617,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15461,f15405])).
fof(f15461,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f7731,f15405,f1323])).
fof(f22616,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15462,f15405])).
fof(f15462,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK18),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f8897,f15405,f1323])).
fof(f22615,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15463,f15405])).
fof(f15463,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK17)))) | $spl4),
  inference(unit_resulting_resolution,[],[f5961,f15405,f1323])).
fof(f5961,plain,(
  sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK17)),'T_0')),
  inference(unit_resulting_resolution,[],[f952,f5893,f1328])).
fof(f22614,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15464,f15405])).
fof(f15464,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18)))) | $spl4),
  inference(unit_resulting_resolution,[],[f6873,f15405,f1323])).
fof(f6873,plain,(
  sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18)),'T_0')),
  inference(unit_resulting_resolution,[],[f952,f5894,f1328])).
fof(f22613,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15465,f15405])).
fof(f15465,plain,(
  p(s(bool,sK77(sK63('const.num.|48|_0','type.num.num')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2041,f15405,f1323])).
fof(f22612,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15466,f15405])).
fof(f15466,plain,(
  p(s(bool,sK78(sK63('const.num.|48|_0','type.num.num')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2088,f15405,f1323])).
fof(f22611,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15467,f15405])).
fof(f15467,plain,(
  p(s(bool,sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))))) | $spl4),
  inference(unit_resulting_resolution,[],[f2124,f15405,f1323])).
fof(f22610,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15468,f15405])).
fof(f15468,plain,(
  p(s(bool,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2369,f15405,f1323])).
fof(f22609,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15469,f15405])).
fof(f15469,plain,(
  p(s(bool,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2488,f15405,f1323])).
fof(f22608,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15536,f17701])).
fof(f17701,plain,(
  ( ! [X0] : (sP6('T_0','T_0',X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1549,f1537,f1543,f15405,f1369])).
fof(f1369,plain,(
  ( ! [X2,X0,X1] : (~sP4(X1,X2,X0) | p(s(bool,X0)) | ~sP5(X0,X2,X1) | sP6(X0,X1,X2) | ~sP3(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f900])).
fof(f900,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (sP3(X2,X1,X0) & sP4(X1,X2,X0) & sP5(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP3(X2,X1,X0) | ~sP4(X1,X2,X0) | ~sP5(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(rectify,[],[f899])).
fof(f899,plain,(
  ! [X2,X1,X0] : ((~sP6(X2,X1,X0) | (sP3(X0,X1,X2) & sP4(X1,X0,X2) & sP5(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP3(X0,X1,X2) | ~sP4(X1,X0,X2) | ~sP5(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP6(X2,X1,X0)))),
  inference(flattening,[],[f898])).
fof(f898,plain,(
  ! [X2,X1,X0] : ((~sP6(X2,X1,X0) | (sP3(X0,X1,X2) & sP4(X1,X0,X2) & sP5(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP3(X0,X1,X2) | ~sP4(X1,X0,X2) | ~sP5(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP6(X2,X1,X0)))),
  inference(nnf_transformation,[],[f725])).
fof(f1543,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f952,f1375])).
fof(f1375,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f906])).
fof(f906,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f905])).
fof(f905,plain,(
  ! [X1,X0,X2] : ((~sP4(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP4(X1,X0,X2)))),
  inference(flattening,[],[f904])).
fof(f904,plain,(
  ! [X1,X0,X2] : ((~sP4(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP4(X1,X0,X2)))),
  inference(nnf_transformation,[],[f723])).
fof(f1537,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f952,f1371])).
fof(f1371,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f903])).
fof(f903,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f902])).
fof(f902,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP5(X2,X0,X1)))),
  inference(flattening,[],[f901])).
fof(f901,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP5(X2,X0,X1)))),
  inference(nnf_transformation,[],[f724])).
fof(f1549,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f952,f1379])).
fof(f1379,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f908])).
fof(f908,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(flattening,[],[f907])).
fof(f907,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(nnf_transformation,[],[f722])).
fof(f15536,plain,(
  ( ! [X0] : (~sP6('T_0','T_0',X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f952,f952,f15405,f1366])).
fof(f1366,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f900])).
fof(f22607,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15538,f19722])).
fof(f19722,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f15405,f1377])).
fof(f1377,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f906])).
fof(f15538,plain,(
  ~sP4('T_0','F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f1551,f3156,f1537,f15405,f1368])).
fof(f1368,plain,(
  ( ! [X2,X0,X1] : (~sP4(X1,X2,X0) | p(s(bool,X1)) | ~sP5(X0,X2,X1) | sP6(X0,X1,X2) | ~sP3(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f900])).
fof(f3156,plain,(
  ~sP6('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f952,f1890,f1384])).
fof(f1551,plain,(
  ( ! [X0,X1] : (sP3(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f952,f1380])).
fof(f1380,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f908])).
fof(f22606,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15539,f19722])).
fof(f15539,plain,(
  ~sP4('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK17)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f1551,f5962,f1537,f15405,f1368])).
fof(f5962,plain,(
  ~sP6('T_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK17)))),
  inference(unit_resulting_resolution,[],[f952,f952,f5893,f1366])).
fof(f22605,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15540,f19722])).
fof(f15540,plain,(
  ~sP4('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f1551,f6874,f1537,f15405,f1368])).
fof(f6874,plain,(
  ~sP6('T_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18)))),
  inference(unit_resulting_resolution,[],[f952,f952,f5894,f1366])).
fof(f22604,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15547,f19722])).
fof(f15547,plain,(
  ~sP4('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),'const.combin.I_1'(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f1608,f3184,f1605,f15405,f1368])).
fof(f1605,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) )),
  inference(unit_resulting_resolution,[],[f962,f1371])).
fof(f962,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))))),
  inference(cnf_transformation,[],[f119])).
fof(f119,axiom,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))))),
  file('thm.extreal.add_not_infty|split|1','thm.extreal.lt_infty|split|2')).
fof(f3184,plain,(
  ( ! [X0] : (~sP6('F_0',X0,'const.combin.I_1'(s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f1508,f1150,f1385])).
fof(f1385,plain,(
  ( ! [X2,X0,X1] : (~sP6(X2,X1,X0) | s(bool,X0) != s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f909])).
fof(f1608,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f962,f1380])).
fof(f22603,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15549,f19722])).
fof(f15549,plain,(
  ~sP4('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'const.combin.I_1'(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f1552,f3184,f1538,f15405,f1368])).
fof(f1538,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f954,f1371])).
fof(f954,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(cnf_transformation,[],[f245])).
fof(f245,axiom,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0'))))),
  file('thm.extreal.add_not_infty|split|1','thm.numeral.numeral_distrib|split|35')).
fof(f1552,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f954,f1380])).
fof(f22602,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15551,f19722])).
fof(f15551,plain,(
  ~sP4(sK77(sK63('const.num.|48|_0','type.num.num')),'const.combin.I_1'(s(bool,sK77(sK63('const.num.|48|_0','type.num.num')))),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f2046,f3184,f2043,f15405,f1368])).
fof(f2043,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,sK77(sK63('const.num.|48|_0','type.num.num')))) )),
  inference(unit_resulting_resolution,[],[f2030,f1371])).
fof(f2046,plain,(
  ( ! [X0,X1] : (sP3(X0,sK77(sK63('const.num.|48|_0','type.num.num')),X1)) )),
  inference(unit_resulting_resolution,[],[f2030,f1380])).
fof(f22601,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15553,f19722])).
fof(f15553,plain,(
  ~sP4(sK78(sK63('const.num.|48|_0','type.num.num')),'const.combin.I_1'(s(bool,sK78(sK63('const.num.|48|_0','type.num.num')))),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f2093,f3184,f2090,f15405,f1368])).
fof(f2090,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,sK78(sK63('const.num.|48|_0','type.num.num')))) )),
  inference(unit_resulting_resolution,[],[f2068,f1371])).
fof(f2093,plain,(
  ( ! [X0,X1] : (sP3(X0,sK78(sK63('const.num.|48|_0','type.num.num')),X1)) )),
  inference(unit_resulting_resolution,[],[f2068,f1380])).
fof(f22600,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15556,f19722])).
fof(f15556,plain,(
  ~sP4(sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))),'const.combin.I_1'(s(bool,sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))))),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f2129,f3184,f2126,f15405,f1368])).
fof(f2126,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))))) )),
  inference(unit_resulting_resolution,[],[f2108,f1371])).
fof(f2129,plain,(
  ( ! [X0,X1] : (sP3(X0,sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))),X1)) )),
  inference(unit_resulting_resolution,[],[f2108,f1380])).
fof(f22599,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15558,f19722])).
fof(f15558,plain,(
  ~sP4(sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')),'const.combin.I_1'(s(bool,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f2374,f3184,f2371,f15405,f1368])).
fof(f2371,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) )),
  inference(unit_resulting_resolution,[],[f2075,f1371])).
fof(f2374,plain,(
  ( ! [X0,X1] : (sP3(X0,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')),X1)) )),
  inference(unit_resulting_resolution,[],[f2075,f1380])).
fof(f22598,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15560,f19722])).
fof(f15560,plain,(
  ~sP4(sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')),'const.combin.I_1'(s(bool,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f2493,f3184,f2490,f15405,f1368])).
fof(f2490,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) )),
  inference(unit_resulting_resolution,[],[f2114,f1371])).
fof(f2493,plain,(
  ( ! [X0,X1] : (sP3(X0,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')),X1)) )),
  inference(unit_resulting_resolution,[],[f2114,f1380])).
fof(f22597,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15561,f19722])).
fof(f15561,plain,(
  ~sP4('T_0','F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f1551,f3156,f1539,f15405,f1368])).
fof(f1539,plain,(
  ( ! [X0,X1] : (sP5(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f1508,f1372])).
fof(f1372,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f903])).
fof(f22596,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15562,f19722])).
fof(f15562,plain,(
  ~sP4('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK17)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f1551,f5962,f5972,f15405,f1368])).
fof(f5972,plain,(
  ( ! [X0,X1] : (sP5(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK17)),X1)) )),
  inference(unit_resulting_resolution,[],[f5893,f1372])).
fof(f22595,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15563,f19722])).
fof(f15563,plain,(
  ~sP4('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f1551,f6874,f6884,f15405,f1368])).
fof(f6884,plain,(
  ( ! [X0,X1] : (sP5(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18)),X1)) )),
  inference(unit_resulting_resolution,[],[f5894,f1372])).
fof(f22594,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15573,f19722])).
fof(f15573,plain,(
  ~sP4('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),'const.combin.I_1'(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f1608,f3184,f1541,f15405,f1368])).
fof(f1541,plain,(
  ( ! [X0,X1] : (sP5('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f1508,f1373])).
fof(f1373,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f903])).
fof(f22593,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15574,f19722])).
fof(f15574,plain,(
  ~sP4('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'const.combin.I_1'(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f1552,f3184,f1541,f15405,f1368])).
fof(f22592,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15575,f19722])).
fof(f15575,plain,(
  ~sP4(sK77(sK63('const.num.|48|_0','type.num.num')),'const.combin.I_1'(s(bool,sK77(sK63('const.num.|48|_0','type.num.num')))),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f2046,f3184,f1541,f15405,f1368])).
fof(f22591,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15576,f19722])).
fof(f15576,plain,(
  ~sP4(sK78(sK63('const.num.|48|_0','type.num.num')),'const.combin.I_1'(s(bool,sK78(sK63('const.num.|48|_0','type.num.num')))),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f2093,f3184,f1541,f15405,f1368])).
fof(f22590,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15577,f19722])).
fof(f15577,plain,(
  ~sP4(sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))),'const.combin.I_1'(s(bool,sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))))),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f2129,f3184,f1541,f15405,f1368])).
fof(f22589,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15578,f19722])).
fof(f15578,plain,(
  ~sP4(sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')),'const.combin.I_1'(s(bool,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f2374,f3184,f1541,f15405,f1368])).
fof(f22588,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15579,f19722])).
fof(f15579,plain,(
  ~sP4(sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')),'const.combin.I_1'(s(bool,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f2493,f3184,f1541,f15405,f1368])).
fof(f22587,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15584,f19722])).
fof(f15584,plain,(
  ~sP4('const.combin.I_1'(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))),'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f1609,f3173,f1541,f15405,f1368])).
fof(f3173,plain,(
  ( ! [X0] : (~sP6('F_0','const.combin.I_1'(s(bool,X0)),X0)) )),
  inference(unit_resulting_resolution,[],[f1508,f1150,f1385])).
fof(f1609,plain,(
  ( ! [X0,X1] : (sP3('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f962,f1381])).
fof(f1381,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f908])).
fof(f22586,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15585,f19722])).
fof(f15585,plain,(
  ~sP4('const.combin.I_1'(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))),'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f1554,f3173,f1541,f15405,f1368])).
fof(f1554,plain,(
  ( ! [X0,X1] : (sP3('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f954,f1381])).
fof(f22585,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15586,f19722])).
fof(f15586,plain,(
  ~sP4('const.combin.I_1'(s(bool,sK77(sK63('const.num.|48|_0','type.num.num')))),sK77(sK63('const.num.|48|_0','type.num.num')),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f2047,f3173,f1541,f15405,f1368])).
fof(f2047,plain,(
  ( ! [X0,X1] : (sP3(sK77(sK63('const.num.|48|_0','type.num.num')),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f2030,f1381])).
fof(f22584,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15587,f19722])).
fof(f15587,plain,(
  ~sP4('const.combin.I_1'(s(bool,sK78(sK63('const.num.|48|_0','type.num.num')))),sK78(sK63('const.num.|48|_0','type.num.num')),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f2094,f3173,f1541,f15405,f1368])).
fof(f2094,plain,(
  ( ! [X0,X1] : (sP3(sK78(sK63('const.num.|48|_0','type.num.num')),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f2068,f1381])).
fof(f22583,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15588,f19722])).
fof(f15588,plain,(
  ~sP4('const.combin.I_1'(s(bool,sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))))),sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f2130,f3173,f1541,f15405,f1368])).
fof(f2130,plain,(
  ( ! [X0,X1] : (sP3(sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f2108,f1381])).
fof(f22582,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15589,f19722])).
fof(f15589,plain,(
  ~sP4('const.combin.I_1'(s(bool,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))),sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f2375,f3173,f1541,f15405,f1368])).
fof(f2375,plain,(
  ( ! [X0,X1] : (sP3(sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f2075,f1381])).
fof(f22581,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15590,f19722])).
fof(f15590,plain,(
  ~sP4('const.combin.I_1'(s(bool,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))),sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f2494,f3173,f1541,f15405,f1368])).
fof(f2494,plain,(
  ( ! [X0,X1] : (sP3(sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f2114,f1381])).
fof(f22580,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15602,f19690])).
fof(f19690,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f15405,f1373])).
fof(f15602,plain,(
  ~sP5('T_0','F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f1551,f3156,f1543,f15405,f1368])).
fof(f22579,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15603,f19690])).
fof(f15603,plain,(
  ~sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK17)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f1551,f5962,f1543,f15405,f1368])).
fof(f22578,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f15604,f19690])).
fof(f15604,plain,(
  ~sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f1551,f6874,f1543,f15405,f1368])).
fof(f22577,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f16767,f19690])).
fof(f16767,plain,(
  ~sP5('T_0','F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f1551,f3156,f1545,f15405,f1368])).
fof(f1545,plain,(
  ( ! [X0,X1] : (sP4(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f1508,f1376])).
fof(f1376,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f906])).
fof(f22576,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f16908,f19690])).
fof(f16908,plain,(
  ~sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK17)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f1551,f5962,f5979,f15405,f1368])).
fof(f5979,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK17)),X1)) )),
  inference(unit_resulting_resolution,[],[f5893,f1376])).
fof(f22575,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f17047,f19690])).
fof(f17047,plain,(
  ~sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f1551,f6874,f6891,f15405,f1368])).
fof(f6891,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18)),X1)) )),
  inference(unit_resulting_resolution,[],[f5894,f1376])).
fof(f22574,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f17595,f19722])).
fof(f17595,plain,(
  ~sP4('T_0','F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f1551,f3156,f1537,f15405,f1369])).
fof(f22573,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f17596,f19722])).
fof(f17596,plain,(
  ~sP4('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK17)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f1551,f5962,f1537,f15405,f1369])).
fof(f22572,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f17597,f19722])).
fof(f17597,plain,(
  ~sP4('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f1551,f6874,f1537,f15405,f1369])).
fof(f22571,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f17604,f19722])).
fof(f17604,plain,(
  ~sP4('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),'const.combin.I_1'(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f1608,f3184,f1605,f15405,f1369])).
fof(f22570,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f17606,f19722])).
fof(f17606,plain,(
  ~sP4('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'const.combin.I_1'(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f1552,f3184,f1538,f15405,f1369])).
fof(f22569,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f17608,f19722])).
fof(f17608,plain,(
  ~sP4(sK77(sK63('const.num.|48|_0','type.num.num')),'const.combin.I_1'(s(bool,sK77(sK63('const.num.|48|_0','type.num.num')))),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f2046,f3184,f2043,f15405,f1369])).
fof(f22568,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f17610,f19722])).
fof(f17610,plain,(
  ~sP4(sK78(sK63('const.num.|48|_0','type.num.num')),'const.combin.I_1'(s(bool,sK78(sK63('const.num.|48|_0','type.num.num')))),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f2093,f3184,f2090,f15405,f1369])).
fof(f22567,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f17613,f19722])).
fof(f17613,plain,(
  ~sP4(sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))),'const.combin.I_1'(s(bool,sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))))),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f2129,f3184,f2126,f15405,f1369])).
fof(f22566,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f17615,f19722])).
fof(f17615,plain,(
  ~sP4(sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')),'const.combin.I_1'(s(bool,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f2374,f3184,f2371,f15405,f1369])).
fof(f22565,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f17617,f19722])).
fof(f17617,plain,(
  ~sP4(sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')),'const.combin.I_1'(s(bool,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f2493,f3184,f2490,f15405,f1369])).
fof(f22564,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f17618,f19722])).
fof(f17618,plain,(
  ~sP4('T_0','F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f1551,f3156,f1539,f15405,f1369])).
fof(f22563,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f17619,f19722])).
fof(f17619,plain,(
  ~sP4('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK17)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f1551,f5962,f5972,f15405,f1369])).
fof(f22562,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f17620,f19722])).
fof(f17620,plain,(
  ~sP4('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f1551,f6874,f6884,f15405,f1369])).
fof(f22561,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f17630,f19722])).
fof(f17630,plain,(
  ~sP4('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),'const.combin.I_1'(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f1608,f3184,f1541,f15405,f1369])).
fof(f22560,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f17631,f19722])).
fof(f17631,plain,(
  ~sP4('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'const.combin.I_1'(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f1552,f3184,f1541,f15405,f1369])).
fof(f22559,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f17632,f19722])).
fof(f17632,plain,(
  ~sP4(sK77(sK63('const.num.|48|_0','type.num.num')),'const.combin.I_1'(s(bool,sK77(sK63('const.num.|48|_0','type.num.num')))),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f2046,f3184,f1541,f15405,f1369])).
fof(f22558,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f17633,f19722])).
fof(f17633,plain,(
  ~sP4(sK78(sK63('const.num.|48|_0','type.num.num')),'const.combin.I_1'(s(bool,sK78(sK63('const.num.|48|_0','type.num.num')))),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f2093,f3184,f1541,f15405,f1369])).
fof(f22557,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f17634,f19722])).
fof(f17634,plain,(
  ~sP4(sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))),'const.combin.I_1'(s(bool,sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))))),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f2129,f3184,f1541,f15405,f1369])).
fof(f22556,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f17635,f19722])).
fof(f17635,plain,(
  ~sP4(sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')),'const.combin.I_1'(s(bool,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f2374,f3184,f1541,f15405,f1369])).
fof(f22555,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f17636,f19722])).
fof(f17636,plain,(
  ~sP4(sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')),'const.combin.I_1'(s(bool,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f2493,f3184,f1541,f15405,f1369])).
fof(f22554,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f17641,f19722])).
fof(f17641,plain,(
  ~sP4('const.combin.I_1'(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))),'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f1609,f3173,f1541,f15405,f1369])).
fof(f22553,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f17642,f19722])).
fof(f17642,plain,(
  ~sP4('const.combin.I_1'(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))),'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f1554,f3173,f1541,f15405,f1369])).
fof(f22552,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f17643,f19722])).
fof(f17643,plain,(
  ~sP4('const.combin.I_1'(s(bool,sK77(sK63('const.num.|48|_0','type.num.num')))),sK77(sK63('const.num.|48|_0','type.num.num')),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f2047,f3173,f1541,f15405,f1369])).
fof(f22551,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f17644,f19722])).
fof(f17644,plain,(
  ~sP4('const.combin.I_1'(s(bool,sK78(sK63('const.num.|48|_0','type.num.num')))),sK78(sK63('const.num.|48|_0','type.num.num')),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f2094,f3173,f1541,f15405,f1369])).
fof(f22550,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f17645,f19722])).
fof(f17645,plain,(
  ~sP4('const.combin.I_1'(s(bool,sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))))),sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f2130,f3173,f1541,f15405,f1369])).
fof(f22549,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f17646,f19722])).
fof(f17646,plain,(
  ~sP4('const.combin.I_1'(s(bool,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))),sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f2375,f3173,f1541,f15405,f1369])).
fof(f22548,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f17647,f19722])).
fof(f17647,plain,(
  ~sP4('const.combin.I_1'(s(bool,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))),sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f2494,f3173,f1541,f15405,f1369])).
fof(f22547,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f17659,f19690])).
fof(f17659,plain,(
  ~sP5('T_0','F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f1551,f3156,f1543,f15405,f1369])).
fof(f22546,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f17660,f19690])).
fof(f17660,plain,(
  ~sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK17)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f1551,f5962,f1543,f15405,f1369])).
fof(f22545,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f17661,f19690])).
fof(f17661,plain,(
  ~sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f1551,f6874,f1543,f15405,f1369])).
fof(f22544,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f17844,f20063])).
fof(f20063,plain,(
  ( ! [X0] : (~sP6(X0,'T_0',sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2116,f15405,f1385])).
fof(f2116,plain,(
  s(bool,'T_0') = s(bool,sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))))),
  inference(unit_resulting_resolution,[],[f2108,f1068])).
fof(f17844,plain,(
  sP6('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),'T_0',sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))) | $spl4),
  inference(unit_resulting_resolution,[],[f2130,f1537,f1606,f15405,f1369])).
fof(f1606,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) )),
  inference(unit_resulting_resolution,[],[f962,f1375])).
fof(f22543,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f17968,f20063])).
fof(f17968,plain,(
  ( ! [X0] : (sP6('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.PosInf_0')),'T_0',sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2130,f1537,f1624,f15405,f1369])).
fof(f1624,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal','const.extreal.PosInf_0')))) )),
  inference(unit_resulting_resolution,[],[f983,f1375])).
fof(f22542,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f18081,f20063])).
fof(f18081,plain,(
  sP6('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'T_0',sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))) | $spl4),
  inference(unit_resulting_resolution,[],[f2130,f1537,f1544,f15405,f1369])).
fof(f1544,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f954,f1375])).
fof(f22541,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f18206,f20063])).
fof(f18206,plain,(
  sP6(sK77(sK63('const.num.|48|_0','type.num.num')),'T_0',sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))) | $spl4),
  inference(unit_resulting_resolution,[],[f2130,f1537,f2044,f15405,f1369])).
fof(f2044,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,sK77(sK63('const.num.|48|_0','type.num.num')))) )),
  inference(unit_resulting_resolution,[],[f2030,f1375])).
fof(f22540,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f18331,f20063])).
fof(f18331,plain,(
  sP6(sK78(sK63('const.num.|48|_0','type.num.num')),'T_0',sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))) | $spl4),
  inference(unit_resulting_resolution,[],[f2130,f1537,f2091,f15405,f1369])).
fof(f2091,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,sK78(sK63('const.num.|48|_0','type.num.num')))) )),
  inference(unit_resulting_resolution,[],[f2068,f1375])).
fof(f22539,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f18456,f20063])).
fof(f18456,plain,(
  sP6(sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))),'T_0',sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))) | $spl4),
  inference(unit_resulting_resolution,[],[f2130,f1537,f2127,f15405,f1369])).
fof(f2127,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))))) )),
  inference(unit_resulting_resolution,[],[f2108,f1375])).
fof(f22538,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f18581,f20063])).
fof(f18581,plain,(
  sP6(sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')),'T_0',sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))) | $spl4),
  inference(unit_resulting_resolution,[],[f2130,f1537,f2372,f15405,f1369])).
fof(f2372,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) )),
  inference(unit_resulting_resolution,[],[f2075,f1375])).
fof(f22537,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f18706,f20063])).
fof(f18706,plain,(
  sP6(sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')),'T_0',sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18)))) | $spl4),
  inference(unit_resulting_resolution,[],[f2130,f1537,f2491,f15405,f1369])).
fof(f2491,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) )),
  inference(unit_resulting_resolution,[],[f2114,f1375])).
fof(f22536,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f18824,f19690])).
fof(f18824,plain,(
  ~sP5('T_0','F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f1551,f3156,f1545,f15405,f1369])).
fof(f22535,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f18965,f19690])).
fof(f18965,plain,(
  ~sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK17)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f1551,f5962,f5979,f15405,f1369])).
fof(f22534,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f18987,f5962])).
fof(f18987,plain,(
  sP6('T_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK17))) | $spl4),
  inference(unit_resulting_resolution,[],[f1549,f1537,f5979,f15405,f1369])).
fof(f22533,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19104,f19690])).
fof(f19104,plain,(
  ~sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f1551,f6874,f6891,f15405,f1369])).
fof(f22532,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19126,f6874])).
fof(f19126,plain,(
  sP6('T_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18))) | $spl4),
  inference(unit_resulting_resolution,[],[f1549,f1537,f6891,f15405,f1369])).
fof(f22531,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19686,f19690])).
fof(f19686,plain,(
  ( ! [X0] : (~sP5('T_0','T_0',X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f952,f952,f15405,f1370])).
fof(f1370,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f903])).
fof(f22530,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19718,f19722])).
fof(f19718,plain,(
  ( ! [X0] : (~sP4('T_0','T_0',X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f952,f952,f15405,f1374])).
fof(f1374,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f906])).
fof(f22529,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19725,f15405])).
fof(f19725,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f1608,f1508,f15405,f1378])).
fof(f1378,plain,(
  ( ! [X2,X0,X1] : (~sP3(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f908])).
fof(f22528,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19726,f15405])).
fof(f19726,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f1552,f1508,f15405,f1378])).
fof(f22527,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19727,f15405])).
fof(f19727,plain,(
  p(s(bool,sK77(sK63('const.num.|48|_0','type.num.num')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2046,f1508,f15405,f1378])).
fof(f22526,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19728,f15405])).
fof(f19728,plain,(
  p(s(bool,sK78(sK63('const.num.|48|_0','type.num.num')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2093,f1508,f15405,f1378])).
fof(f22525,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19729,f15405])).
fof(f19729,plain,(
  p(s(bool,sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))))) | $spl4),
  inference(unit_resulting_resolution,[],[f2129,f1508,f15405,f1378])).
fof(f22524,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19730,f15405])).
fof(f19730,plain,(
  p(s(bool,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2374,f1508,f15405,f1378])).
fof(f22523,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19731,f15405])).
fof(f19731,plain,(
  p(s(bool,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2493,f1508,f15405,f1378])).
fof(f22522,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19732,f15405])).
fof(f19732,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1553,f1508,f15405,f1378])).
fof(f1553,plain,(
  ( ! [X0,X1] : (sP3('T_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f952,f1381])).
fof(f22521,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19733,f15405])).
fof(f19733,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1609,f1508,f15405,f1378])).
fof(f22520,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19734,f15405])).
fof(f19734,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1554,f1508,f15405,f1378])).
fof(f22519,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19735,f15405])).
fof(f19735,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2047,f1508,f15405,f1378])).
fof(f22518,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19736,f15405])).
fof(f19736,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2094,f1508,f15405,f1378])).
fof(f22517,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19737,f15405])).
fof(f19737,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2130,f1508,f15405,f1378])).
fof(f22516,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19738,f15405])).
fof(f19738,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2375,f1508,f15405,f1378])).
fof(f22515,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19739,f15405])).
fof(f19739,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2494,f1508,f15405,f1378])).
fof(f22514,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19740,f15405])).
fof(f19740,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1549,f15405,f15405,f1378])).
fof(f22513,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19741,f15405])).
fof(f19741,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1607,f15405,f15405,f1378])).
fof(f1607,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) )),
  inference(unit_resulting_resolution,[],[f962,f1379])).
fof(f22512,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19742,f15405])).
fof(f19742,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1550,f15405,f15405,f1378])).
fof(f1550,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f954,f1379])).
fof(f22511,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19743,f15405])).
fof(f19743,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2045,f15405,f15405,f1378])).
fof(f2045,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,sK77(sK63('const.num.|48|_0','type.num.num')))) )),
  inference(unit_resulting_resolution,[],[f2030,f1379])).
fof(f22510,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19744,f15405])).
fof(f19744,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2092,f15405,f15405,f1378])).
fof(f2092,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,sK78(sK63('const.num.|48|_0','type.num.num')))) )),
  inference(unit_resulting_resolution,[],[f2068,f1379])).
fof(f22509,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19745,f15405])).
fof(f19745,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2128,f15405,f15405,f1378])).
fof(f2128,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))))) )),
  inference(unit_resulting_resolution,[],[f2108,f1379])).
fof(f22508,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19746,f15405])).
fof(f19746,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2373,f15405,f15405,f1378])).
fof(f2373,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) )),
  inference(unit_resulting_resolution,[],[f2075,f1379])).
fof(f22507,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19747,f15405])).
fof(f19747,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2492,f15405,f15405,f1378])).
fof(f2492,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) )),
  inference(unit_resulting_resolution,[],[f2114,f1379])).
fof(f22506,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19749,f15405])).
fof(f19749,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f1608,f15405,f15405,f1378])).
fof(f22505,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19750,f15405])).
fof(f19750,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f1552,f15405,f15405,f1378])).
fof(f22504,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19751,f15405])).
fof(f19751,plain,(
  p(s(bool,sK77(sK63('const.num.|48|_0','type.num.num')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2046,f15405,f15405,f1378])).
fof(f22503,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19752,f15405])).
fof(f19752,plain,(
  p(s(bool,sK78(sK63('const.num.|48|_0','type.num.num')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2093,f15405,f15405,f1378])).
fof(f22502,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19753,f15405])).
fof(f19753,plain,(
  p(s(bool,sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))))) | $spl4),
  inference(unit_resulting_resolution,[],[f2129,f15405,f15405,f1378])).
fof(f22501,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19754,f15405])).
fof(f19754,plain,(
  p(s(bool,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2374,f15405,f15405,f1378])).
fof(f22500,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19755,f15405])).
fof(f19755,plain,(
  p(s(bool,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2493,f15405,f15405,f1378])).
fof(f22499,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19756,f15405])).
fof(f19756,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1553,f15405,f15405,f1378])).
fof(f22498,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19757,f15405])).
fof(f19757,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1609,f15405,f15405,f1378])).
fof(f22497,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19758,f15405])).
fof(f19758,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1554,f15405,f15405,f1378])).
fof(f22496,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19759,f15405])).
fof(f19759,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2047,f15405,f15405,f1378])).
fof(f22495,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19760,f15405])).
fof(f19760,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2094,f15405,f15405,f1378])).
fof(f22494,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19761,f15405])).
fof(f19761,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2130,f15405,f15405,f1378])).
fof(f22493,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19762,f15405])).
fof(f19762,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2375,f15405,f15405,f1378])).
fof(f22492,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19763,f15405])).
fof(f19763,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2494,f15405,f15405,f1378])).
fof(f22491,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19765,f15405])).
fof(f19765,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f1607,f1508,f15405,f1378])).
fof(f22490,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19766,f15405])).
fof(f19766,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f1550,f1508,f15405,f1378])).
fof(f22489,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19767,f15405])).
fof(f19767,plain,(
  p(s(bool,sK77(sK63('const.num.|48|_0','type.num.num')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2045,f1508,f15405,f1378])).
fof(f22488,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19768,f15405])).
fof(f19768,plain,(
  p(s(bool,sK78(sK63('const.num.|48|_0','type.num.num')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2092,f1508,f15405,f1378])).
fof(f22487,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19769,f15405])).
fof(f19769,plain,(
  p(s(bool,sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))))) | $spl4),
  inference(unit_resulting_resolution,[],[f2128,f1508,f15405,f1378])).
fof(f22486,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19770,f15405])).
fof(f19770,plain,(
  p(s(bool,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2373,f1508,f15405,f1378])).
fof(f22485,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19771,f15405])).
fof(f19771,plain,(
  p(s(bool,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2492,f1508,f15405,f1378])).
fof(f22484,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19772,f15405])).
fof(f19772,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1553,f1508,f15405,f1378])).
fof(f22483,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19773,f15405])).
fof(f19773,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1609,f1508,f15405,f1378])).
fof(f22482,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19774,f15405])).
fof(f19774,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1554,f1508,f15405,f1378])).
fof(f22481,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19775,f15405])).
fof(f19775,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2047,f1508,f15405,f1378])).
fof(f22480,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19776,f15405])).
fof(f19776,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2094,f1508,f15405,f1378])).
fof(f22479,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19777,f15405])).
fof(f19777,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2130,f1508,f15405,f1378])).
fof(f22478,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19778,f15405])).
fof(f19778,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2375,f1508,f15405,f1378])).
fof(f22477,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19779,f15405])).
fof(f19779,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2494,f1508,f15405,f1378])).
fof(f22476,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19785,f15405])).
fof(f19785,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f1607,f15405,f15405,f1378])).
fof(f22475,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19786,f15405])).
fof(f19786,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f1550,f15405,f15405,f1378])).
fof(f22474,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19787,f15405])).
fof(f19787,plain,(
  p(s(bool,sK77(sK63('const.num.|48|_0','type.num.num')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2045,f15405,f15405,f1378])).
fof(f22473,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19788,f15405])).
fof(f19788,plain,(
  p(s(bool,sK78(sK63('const.num.|48|_0','type.num.num')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2092,f15405,f15405,f1378])).
fof(f22472,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19789,f15405])).
fof(f19789,plain,(
  p(s(bool,sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))))) | $spl4),
  inference(unit_resulting_resolution,[],[f2128,f15405,f15405,f1378])).
fof(f22471,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19790,f15405])).
fof(f19790,plain,(
  p(s(bool,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2373,f15405,f15405,f1378])).
fof(f22470,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19791,f15405])).
fof(f19791,plain,(
  p(s(bool,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2492,f15405,f15405,f1378])).
fof(f22469,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19792,f15405])).
fof(f19792,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1551,f15405,f15405,f1378])).
fof(f22468,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19793,f15405])).
fof(f19793,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1608,f15405,f15405,f1378])).
fof(f22467,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19794,f15405])).
fof(f19794,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1552,f15405,f15405,f1378])).
fof(f22466,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19795,f15405])).
fof(f19795,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2046,f15405,f15405,f1378])).
fof(f22465,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19796,f15405])).
fof(f19796,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2093,f15405,f15405,f1378])).
fof(f22464,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19797,f15405])).
fof(f19797,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2129,f15405,f15405,f1378])).
fof(f22463,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19798,f15405])).
fof(f19798,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2374,f15405,f15405,f1378])).
fof(f22462,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19799,f15405])).
fof(f19799,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2493,f15405,f15405,f1378])).
fof(f22461,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19800,f15405])).
fof(f19800,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1553,f15405,f15405,f1378])).
fof(f22460,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19801,f15405])).
fof(f19801,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1609,f15405,f15405,f1378])).
fof(f22459,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19802,f15405])).
fof(f19802,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1554,f15405,f15405,f1378])).
fof(f22458,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19803,f15405])).
fof(f19803,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2047,f15405,f15405,f1378])).
fof(f22457,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19804,f15405])).
fof(f19804,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2094,f15405,f15405,f1378])).
fof(f22456,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19805,f15405])).
fof(f19805,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2130,f15405,f15405,f1378])).
fof(f22455,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19806,f15405])).
fof(f19806,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2375,f15405,f15405,f1378])).
fof(f22454,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19807,f15405])).
fof(f19807,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2494,f15405,f15405,f1378])).
fof(f22453,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19813,f15405])).
fof(f19813,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f1607,f1508,f15405,f1378])).
fof(f22452,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19814,f15405])).
fof(f19814,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f1550,f1508,f15405,f1378])).
fof(f22451,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19815,f15405])).
fof(f19815,plain,(
  p(s(bool,sK77(sK63('const.num.|48|_0','type.num.num')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2045,f1508,f15405,f1378])).
fof(f22450,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19816,f15405])).
fof(f19816,plain,(
  p(s(bool,sK78(sK63('const.num.|48|_0','type.num.num')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2092,f1508,f15405,f1378])).
fof(f22449,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19817,f15405])).
fof(f19817,plain,(
  p(s(bool,sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))))) | $spl4),
  inference(unit_resulting_resolution,[],[f2128,f1508,f15405,f1378])).
fof(f22448,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19818,f15405])).
fof(f19818,plain,(
  p(s(bool,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2373,f1508,f15405,f1378])).
fof(f22447,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19819,f15405])).
fof(f19819,plain,(
  p(s(bool,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2492,f1508,f15405,f1378])).
fof(f22446,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19820,f15405])).
fof(f19820,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1551,f1508,f15405,f1378])).
fof(f22445,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19821,f15405])).
fof(f19821,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1608,f1508,f15405,f1378])).
fof(f22444,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19822,f15405])).
fof(f19822,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1552,f1508,f15405,f1378])).
fof(f22443,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19823,f15405])).
fof(f19823,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2046,f1508,f15405,f1378])).
fof(f22442,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19824,f15405])).
fof(f19824,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2093,f1508,f15405,f1378])).
fof(f22441,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19825,f15405])).
fof(f19825,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2129,f1508,f15405,f1378])).
fof(f22440,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19826,f15405])).
fof(f19826,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2374,f1508,f15405,f1378])).
fof(f22439,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19827,f15405])).
fof(f19827,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2493,f1508,f15405,f1378])).
fof(f22438,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19829,f15405])).
fof(f19829,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f1607,f15405,f15405,f1378])).
fof(f22437,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19830,f15405])).
fof(f19830,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f1550,f15405,f15405,f1378])).
fof(f22436,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19831,f15405])).
fof(f19831,plain,(
  p(s(bool,sK77(sK63('const.num.|48|_0','type.num.num')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2045,f15405,f15405,f1378])).
fof(f22435,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19832,f15405])).
fof(f19832,plain,(
  p(s(bool,sK78(sK63('const.num.|48|_0','type.num.num')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2092,f15405,f15405,f1378])).
fof(f22434,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19833,f15405])).
fof(f19833,plain,(
  p(s(bool,sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))))) | $spl4),
  inference(unit_resulting_resolution,[],[f2128,f15405,f15405,f1378])).
fof(f22433,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19834,f15405])).
fof(f19834,plain,(
  p(s(bool,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2373,f15405,f15405,f1378])).
fof(f22432,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19835,f15405])).
fof(f19835,plain,(
  p(s(bool,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2492,f15405,f15405,f1378])).
fof(f22431,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19836,f15405])).
fof(f19836,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1551,f15405,f15405,f1378])).
fof(f22430,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19837,f15405])).
fof(f19837,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1608,f15405,f15405,f1378])).
fof(f22429,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19838,f15405])).
fof(f19838,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1552,f15405,f15405,f1378])).
fof(f22428,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19839,f15405])).
fof(f19839,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2046,f15405,f15405,f1378])).
fof(f22427,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19840,f15405])).
fof(f19840,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2093,f15405,f15405,f1378])).
fof(f22426,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19841,f15405])).
fof(f19841,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2129,f15405,f15405,f1378])).
fof(f22425,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19842,f15405])).
fof(f19842,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2374,f15405,f15405,f1378])).
fof(f22424,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19843,f15405])).
fof(f19843,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2493,f15405,f15405,f1378])).
fof(f22423,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19844,f15405])).
fof(f19844,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1553,f15405,f15405,f1378])).
fof(f22422,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19845,f15405])).
fof(f19845,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1609,f15405,f15405,f1378])).
fof(f22421,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19846,f15405])).
fof(f19846,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1554,f15405,f15405,f1378])).
fof(f22420,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19847,f15405])).
fof(f19847,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2047,f15405,f15405,f1378])).
fof(f22419,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19848,f15405])).
fof(f19848,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2094,f15405,f15405,f1378])).
fof(f22418,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19849,f15405])).
fof(f19849,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2130,f15405,f15405,f1378])).
fof(f22417,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19850,f15405])).
fof(f19850,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2375,f15405,f15405,f1378])).
fof(f22416,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19851,f15405])).
fof(f19851,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2494,f15405,f15405,f1378])).
fof(f22415,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19852,f15405])).
fof(f19852,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1551,f1508,f15405,f1378])).
fof(f22414,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19853,f15405])).
fof(f19853,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1608,f1508,f15405,f1378])).
fof(f22413,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19854,f15405])).
fof(f19854,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1552,f1508,f15405,f1378])).
fof(f22412,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19855,f15405])).
fof(f19855,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2046,f1508,f15405,f1378])).
fof(f22411,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19856,f15405])).
fof(f19856,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2093,f1508,f15405,f1378])).
fof(f22410,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19857,f15405])).
fof(f19857,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2129,f1508,f15405,f1378])).
fof(f22409,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19858,f15405])).
fof(f19858,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2374,f1508,f15405,f1378])).
fof(f22408,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19859,f15405])).
fof(f19859,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2493,f1508,f15405,f1378])).
fof(f22407,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19861,f15405])).
fof(f19861,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f1609,f1508,f15405,f1378])).
fof(f22406,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19862,f15405])).
fof(f19862,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f1554,f1508,f15405,f1378])).
fof(f22405,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19863,f15405])).
fof(f19863,plain,(
  p(s(bool,sK77(sK63('const.num.|48|_0','type.num.num')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2047,f1508,f15405,f1378])).
fof(f22404,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19864,f15405])).
fof(f19864,plain,(
  p(s(bool,sK78(sK63('const.num.|48|_0','type.num.num')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2094,f1508,f15405,f1378])).
fof(f22403,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19865,f15405])).
fof(f19865,plain,(
  p(s(bool,sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))))) | $spl4),
  inference(unit_resulting_resolution,[],[f2130,f1508,f15405,f1378])).
fof(f22402,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19866,f15405])).
fof(f19866,plain,(
  p(s(bool,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2375,f1508,f15405,f1378])).
fof(f22401,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19867,f15405])).
fof(f19867,plain,(
  p(s(bool,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2494,f1508,f15405,f1378])).
fof(f22400,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19872,f15405])).
fof(f19872,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1549,f15405,f15405,f1378])).
fof(f22399,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19873,f15405])).
fof(f19873,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1607,f15405,f15405,f1378])).
fof(f22398,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19874,f15405])).
fof(f19874,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1550,f15405,f15405,f1378])).
fof(f22397,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19875,f15405])).
fof(f19875,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2045,f15405,f15405,f1378])).
fof(f22396,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19876,f15405])).
fof(f19876,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2092,f15405,f15405,f1378])).
fof(f22395,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19877,f15405])).
fof(f19877,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2128,f15405,f15405,f1378])).
fof(f22394,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19878,f15405])).
fof(f19878,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2373,f15405,f15405,f1378])).
fof(f22393,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19879,f15405])).
fof(f19879,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2492,f15405,f15405,f1378])).
fof(f22392,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19880,f15405])).
fof(f19880,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1551,f15405,f15405,f1378])).
fof(f22391,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19881,f15405])).
fof(f19881,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1608,f15405,f15405,f1378])).
fof(f22390,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19882,f15405])).
fof(f19882,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1552,f15405,f15405,f1378])).
fof(f22389,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19883,f15405])).
fof(f19883,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2046,f15405,f15405,f1378])).
fof(f22388,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19884,f15405])).
fof(f19884,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2093,f15405,f15405,f1378])).
fof(f22387,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19885,f15405])).
fof(f19885,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2129,f15405,f15405,f1378])).
fof(f22386,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19886,f15405])).
fof(f19886,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2374,f15405,f15405,f1378])).
fof(f22385,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19887,f15405])).
fof(f19887,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2493,f15405,f15405,f1378])).
fof(f22384,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19889,f15405])).
fof(f19889,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f1609,f15405,f15405,f1378])).
fof(f22383,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19890,f15405])).
fof(f19890,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f1554,f15405,f15405,f1378])).
fof(f22382,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19891,f15405])).
fof(f19891,plain,(
  p(s(bool,sK77(sK63('const.num.|48|_0','type.num.num')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2047,f15405,f15405,f1378])).
fof(f22381,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19892,f15405])).
fof(f19892,plain,(
  p(s(bool,sK78(sK63('const.num.|48|_0','type.num.num')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2094,f15405,f15405,f1378])).
fof(f22380,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19893,f15405])).
fof(f19893,plain,(
  p(s(bool,sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))))) | $spl4),
  inference(unit_resulting_resolution,[],[f2130,f15405,f15405,f1378])).
fof(f22379,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19894,f15405])).
fof(f19894,plain,(
  p(s(bool,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2375,f15405,f15405,f1378])).
fof(f22378,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19895,f15405])).
fof(f19895,plain,(
  p(s(bool,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2494,f15405,f15405,f1378])).
fof(f22377,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19900,f15405])).
fof(f19900,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1549,f1508,f15405,f1378])).
fof(f22376,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19901,f15405])).
fof(f19901,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1607,f1508,f15405,f1378])).
fof(f22375,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19902,f15405])).
fof(f19902,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1550,f1508,f15405,f1378])).
fof(f22374,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19903,f15405])).
fof(f19903,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2045,f1508,f15405,f1378])).
fof(f22373,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19904,f15405])).
fof(f19904,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2092,f1508,f15405,f1378])).
fof(f22372,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19905,f15405])).
fof(f19905,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2128,f1508,f15405,f1378])).
fof(f22371,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19906,f15405])).
fof(f19906,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2373,f1508,f15405,f1378])).
fof(f22370,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19907,f15405])).
fof(f19907,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2492,f1508,f15405,f1378])).
fof(f22369,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19909,f15405])).
fof(f19909,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f1608,f1508,f15405,f1378])).
fof(f22368,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19910,f15405])).
fof(f19910,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f1552,f1508,f15405,f1378])).
fof(f22367,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19911,f15405])).
fof(f19911,plain,(
  p(s(bool,sK77(sK63('const.num.|48|_0','type.num.num')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2046,f1508,f15405,f1378])).
fof(f22366,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19912,f15405])).
fof(f19912,plain,(
  p(s(bool,sK78(sK63('const.num.|48|_0','type.num.num')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2093,f1508,f15405,f1378])).
fof(f22365,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19913,f15405])).
fof(f19913,plain,(
  p(s(bool,sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))))) | $spl4),
  inference(unit_resulting_resolution,[],[f2129,f1508,f15405,f1378])).
fof(f22364,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19914,f15405])).
fof(f19914,plain,(
  p(s(bool,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2374,f1508,f15405,f1378])).
fof(f22363,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19915,f15405])).
fof(f19915,plain,(
  p(s(bool,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2493,f1508,f15405,f1378])).
fof(f22362,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19916,f15405])).
fof(f19916,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1549,f15405,f15405,f1378])).
fof(f22361,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19917,f15405])).
fof(f19917,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1607,f15405,f15405,f1378])).
fof(f22360,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19918,f15405])).
fof(f19918,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1550,f15405,f15405,f1378])).
fof(f22359,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19919,f15405])).
fof(f19919,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2045,f15405,f15405,f1378])).
fof(f22358,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19920,f15405])).
fof(f19920,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2092,f15405,f15405,f1378])).
fof(f22357,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19921,f15405])).
fof(f19921,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2128,f15405,f15405,f1378])).
fof(f22356,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19922,f15405])).
fof(f19922,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2373,f15405,f15405,f1378])).
fof(f22355,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19923,f15405])).
fof(f19923,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2492,f15405,f15405,f1378])).
fof(f22354,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19925,f15405])).
fof(f19925,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f1608,f15405,f15405,f1378])).
fof(f22353,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19926,f15405])).
fof(f19926,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f1552,f15405,f15405,f1378])).
fof(f22352,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19927,f15405])).
fof(f19927,plain,(
  p(s(bool,sK77(sK63('const.num.|48|_0','type.num.num')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2046,f15405,f15405,f1378])).
fof(f22351,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19928,f15405])).
fof(f19928,plain,(
  p(s(bool,sK78(sK63('const.num.|48|_0','type.num.num')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2093,f15405,f15405,f1378])).
fof(f22350,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19929,f15405])).
fof(f19929,plain,(
  p(s(bool,sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))))) | $spl4),
  inference(unit_resulting_resolution,[],[f2129,f15405,f15405,f1378])).
fof(f22349,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19930,f15405])).
fof(f19930,plain,(
  p(s(bool,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2374,f15405,f15405,f1378])).
fof(f22348,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19931,f15405])).
fof(f19931,plain,(
  p(s(bool,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2493,f15405,f15405,f1378])).
fof(f22347,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19932,f15405])).
fof(f19932,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1553,f15405,f15405,f1378])).
fof(f22346,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19933,f15405])).
fof(f19933,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1609,f15405,f15405,f1378])).
fof(f22345,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19934,f15405])).
fof(f19934,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1554,f15405,f15405,f1378])).
fof(f22344,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19935,f15405])).
fof(f19935,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2047,f15405,f15405,f1378])).
fof(f22343,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19936,f15405])).
fof(f19936,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2094,f15405,f15405,f1378])).
fof(f22342,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19937,f15405])).
fof(f19937,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2130,f15405,f15405,f1378])).
fof(f22341,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19938,f15405])).
fof(f19938,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2375,f15405,f15405,f1378])).
fof(f22340,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19939,f15405])).
fof(f19939,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2494,f15405,f15405,f1378])).
fof(f22339,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19940,f15405])).
fof(f19940,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1549,f1508,f15405,f1378])).
fof(f22338,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19941,f15405])).
fof(f19941,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1607,f1508,f15405,f1378])).
fof(f22337,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19942,f15405])).
fof(f19942,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1550,f1508,f15405,f1378])).
fof(f22336,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19943,f15405])).
fof(f19943,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2045,f1508,f15405,f1378])).
fof(f22335,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19944,f15405])).
fof(f19944,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2092,f1508,f15405,f1378])).
fof(f22334,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19945,f15405])).
fof(f19945,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2128,f1508,f15405,f1378])).
fof(f22333,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19946,f15405])).
fof(f19946,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2373,f1508,f15405,f1378])).
fof(f22332,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19947,f15405])).
fof(f19947,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2492,f1508,f15405,f1378])).
fof(f22331,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19949,f15405])).
fof(f19949,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f1609,f1508,f15405,f1378])).
fof(f22330,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19950,f15405])).
fof(f19950,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f1554,f1508,f15405,f1378])).
fof(f22329,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19951,f15405])).
fof(f19951,plain,(
  p(s(bool,sK77(sK63('const.num.|48|_0','type.num.num')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2047,f1508,f15405,f1378])).
fof(f22328,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19952,f15405])).
fof(f19952,plain,(
  p(s(bool,sK78(sK63('const.num.|48|_0','type.num.num')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2094,f1508,f15405,f1378])).
fof(f22327,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19953,f15405])).
fof(f19953,plain,(
  p(s(bool,sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))))) | $spl4),
  inference(unit_resulting_resolution,[],[f2130,f1508,f15405,f1378])).
fof(f22326,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19954,f15405])).
fof(f19954,plain,(
  p(s(bool,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2375,f1508,f15405,f1378])).
fof(f22325,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19955,f15405])).
fof(f19955,plain,(
  p(s(bool,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2494,f1508,f15405,f1378])).
fof(f22324,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19960,f15405])).
fof(f19960,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1549,f15405,f15405,f1378])).
fof(f22323,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19961,f15405])).
fof(f19961,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1607,f15405,f15405,f1378])).
fof(f22322,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19962,f15405])).
fof(f19962,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1550,f15405,f15405,f1378])).
fof(f22321,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19963,f15405])).
fof(f19963,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2045,f15405,f15405,f1378])).
fof(f22320,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19964,f15405])).
fof(f19964,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2092,f15405,f15405,f1378])).
fof(f22319,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19965,f15405])).
fof(f19965,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2128,f15405,f15405,f1378])).
fof(f22318,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19966,f15405])).
fof(f19966,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2373,f15405,f15405,f1378])).
fof(f22317,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19967,f15405])).
fof(f19967,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2492,f15405,f15405,f1378])).
fof(f22316,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19968,f15405])).
fof(f19968,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1551,f15405,f15405,f1378])).
fof(f22315,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19969,f15405])).
fof(f19969,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1608,f15405,f15405,f1378])).
fof(f22314,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19970,f15405])).
fof(f19970,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1552,f15405,f15405,f1378])).
fof(f22313,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19971,f15405])).
fof(f19971,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2046,f15405,f15405,f1378])).
fof(f22312,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19972,f15405])).
fof(f19972,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2093,f15405,f15405,f1378])).
fof(f22311,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19973,f15405])).
fof(f19973,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2129,f15405,f15405,f1378])).
fof(f22310,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19974,f15405])).
fof(f19974,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2374,f15405,f15405,f1378])).
fof(f22309,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19975,f15405])).
fof(f19975,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2493,f15405,f15405,f1378])).
fof(f22308,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19977,f15405])).
fof(f19977,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f1609,f15405,f15405,f1378])).
fof(f22307,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19978,f15405])).
fof(f19978,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f1554,f15405,f15405,f1378])).
fof(f22306,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19979,f15405])).
fof(f19979,plain,(
  p(s(bool,sK77(sK63('const.num.|48|_0','type.num.num')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2047,f15405,f15405,f1378])).
fof(f22305,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19980,f15405])).
fof(f19980,plain,(
  p(s(bool,sK78(sK63('const.num.|48|_0','type.num.num')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2094,f15405,f15405,f1378])).
fof(f22304,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19981,f15405])).
fof(f19981,plain,(
  p(s(bool,sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))))) | $spl4),
  inference(unit_resulting_resolution,[],[f2130,f15405,f15405,f1378])).
fof(f22303,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19982,f15405])).
fof(f19982,plain,(
  p(s(bool,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2375,f15405,f15405,f1378])).
fof(f22302,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19983,f15405])).
fof(f19983,plain,(
  p(s(bool,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2494,f15405,f15405,f1378])).
fof(f22301,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19988,f1890])).
fof(f19988,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f3154,f15405,f1382])).
fof(f3154,plain,(
  ~sP6('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f952,f1890,f1384])).
fof(f22300,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19989,f1890])).
fof(f19989,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f3156,f15405,f1382])).
fof(f22299,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19990,f6293])).
fof(f6293,plain,(
  s(bool,'T_0') != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK17)))),
  inference(unit_resulting_resolution,[],[f952,f5893,f1494])).
fof(f19990,plain,(
  s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK17))) | $spl4),
  inference(unit_resulting_resolution,[],[f5962,f15405,f1382])).
fof(f22298,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f19991,f7205])).
fof(f19991,plain,(
  s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18))) | $spl4),
  inference(unit_resulting_resolution,[],[f6874,f15405,f1382])).
fof(f22297,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20106,f21637])).
fof(f21637,plain,(
  ( ! [X0,X1] : (sP8('T_0',X0,X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1555,f15405,f1497])).
fof(f1497,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0)) | ~sP7(X1,X2,X0)) )),
  inference(duplicate_literal_removal,[],[f1392])).
fof(f1392,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X0)) | ~sP7(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f912])).
fof(f912,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | (sP7(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP7(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f911])).
fof(f911,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (sP7(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP7(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f910])).
fof(f910,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (sP7(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP7(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f728])).
fof(f728,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (sP7(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f1555,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f952,f1394])).
fof(f1394,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f914])).
fof(f914,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(flattening,[],[f913])).
fof(f913,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(nnf_transformation,[],[f727])).
fof(f727,plain,(
  ! [X0,X1,X2] : (sP7(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f20106,plain,(
  ( ! [X0,X1] : (~sP8('T_0',X0,X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f952,f15405,f1387])).
fof(f1387,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f912])).
fof(f22296,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20122,f21637])).
fof(f20122,plain,(
  ( ! [X0,X1] : (~sP8('T_0',X0,X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f952,f15405,f1388])).
fof(f1388,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f912])).
fof(f22295,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20148,f20152])).
fof(f20152,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f15405,f1396])).
fof(f1396,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f914])).
fof(f20148,plain,(
  ( ! [X0] : (~sP7('T_0','T_0',X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f952,f952,f15405,f1393])).
fof(f1393,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f914])).
fof(f22294,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20160,f1508])).
fof(f20160,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f4461,f15405,f1397])).
fof(f1397,plain,(
  ( ! [X2,X0,X1] : (sP8(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f916])).
fof(f916,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP8(X2,X0,X1)) & (~sP8(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f915])).
fof(f915,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP8(X2,X0,X1)) & (~sP8(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f729])).
fof(f729,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP8(X2,X0,X1))),
  inference(definition_folding,[],[f559,f728,f727])).
fof(f559,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f558])).
fof(f558,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f90])).
fof(f90,axiom,(
  ! [X35,X43,X44] : ((p(s(bool,X44)) <=> (p(s(bool,X43)) & p(s(bool,X35)))) <=> ((p(s(bool,X44)) | ~p(s(bool,X43)) | ~p(s(bool,X35))) & (p(s(bool,X43)) | ~p(s(bool,X44))) & (p(s(bool,X35)) | ~p(s(bool,X44)))))),
  file('thm.extreal.add_not_infty|split|1','thm.sat.sat_23_prove_ltb')).
fof(f4461,plain,(
  ~sP8('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f4447,f1386])).
fof(f1386,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X1,X2) | sP7(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f912])).
fof(f4447,plain,(
  ~sP7('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f952,f1508,f952,f1393])).
fof(f22293,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20161,f15405])).
fof(f20161,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK17)))) | $spl4),
  inference(unit_resulting_resolution,[],[f6104,f15405,f1397])).
fof(f6104,plain,(
  ~sP8('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK17)),'T_0','T_0')),
  inference(unit_resulting_resolution,[],[f952,f952,f5893,f1402])).
fof(f1402,plain,(
  ( ! [X2,X0,X1] : (~sP8(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f916])).
fof(f22292,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20162,f15405])).
fof(f20162,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18)))) | $spl4),
  inference(unit_resulting_resolution,[],[f7016,f15405,f1397])).
fof(f7016,plain,(
  ~sP8('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18)),'T_0','T_0')),
  inference(unit_resulting_resolution,[],[f952,f952,f5894,f1402])).
fof(f22291,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20163,f15405])).
fof(f20163,plain,(
  p(s(bool,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2376,f15405,f1397])).
fof(f2376,plain,(
  ( ! [X0] : (~sP8(sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')),X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f1508,f2075,f1387])).
fof(f22290,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20164,f15405])).
fof(f20164,plain,(
  p(s(bool,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2377,f15405,f1397])).
fof(f2377,plain,(
  ( ! [X0] : (~sP8(sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')),'F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f1508,f2075,f1388])).
fof(f22289,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20165,f15405])).
fof(f20165,plain,(
  p(s(bool,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2495,f15405,f1397])).
fof(f2495,plain,(
  ( ! [X0] : (~sP8(sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')),X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f1508,f2114,f1387])).
fof(f22288,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20166,f15405])).
fof(f20166,plain,(
  p(s(bool,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2496,f15405,f1397])).
fof(f2496,plain,(
  ( ! [X0] : (~sP8(sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')),'F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f1508,f2114,f1388])).
fof(f22287,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20171,f1508])).
fof(f20171,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f2146,f15405,f1397])).
fof(f2146,plain,(
  ( ! [X0] : (~sP8('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f952,f1508,f1387])).
fof(f22286,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20172,f15405])).
fof(f20172,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK17)))) | $spl4),
  inference(unit_resulting_resolution,[],[f6077,f15405,f1397])).
fof(f6077,plain,(
  ( ! [X0] : (~sP8('T_0',X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK17)))) )),
  inference(unit_resulting_resolution,[],[f952,f5893,f1387])).
fof(f22285,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20173,f15405])).
fof(f20173,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18)))) | $spl4),
  inference(unit_resulting_resolution,[],[f6989,f15405,f1397])).
fof(f6989,plain,(
  ( ! [X0] : (~sP8('T_0',X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18)))) )),
  inference(unit_resulting_resolution,[],[f952,f5894,f1387])).
fof(f22284,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20174,f15405])).
fof(f20174,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2149,f15405,f1397])).
fof(f2149,plain,(
  ( ! [X0] : (~sP8('T_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f952,f1508,f1388])).
fof(f22283,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20175,f15405])).
fof(f20175,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f6080,f15405,f1397])).
fof(f6080,plain,(
  ( ! [X0] : (~sP8('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK17)),X0)) )),
  inference(unit_resulting_resolution,[],[f952,f5893,f1388])).
fof(f22282,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20176,f15405])).
fof(f20176,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f6992,f15405,f1397])).
fof(f6992,plain,(
  ( ! [X0] : (~sP8('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18)),X0)) )),
  inference(unit_resulting_resolution,[],[f952,f5894,f1388])).
fof(f22281,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20180,f1508])).
fof(f20180,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f2376,f15405,f1397])).
fof(f22280,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20181,f15405])).
fof(f20181,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2377,f15405,f1397])).
fof(f22279,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20182,f1508])).
fof(f20182,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f2495,f15405,f1397])).
fof(f22278,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20183,f15405])).
fof(f20183,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2496,f15405,f1397])).
fof(f22277,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20194,f1508])).
fof(f20194,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f4461,f15405,f1398])).
fof(f1398,plain,(
  ( ! [X2,X0,X1] : (sP8(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f916])).
fof(f22276,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20195,f15405])).
fof(f20195,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK17)))) | $spl4),
  inference(unit_resulting_resolution,[],[f6104,f15405,f1398])).
fof(f22275,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20196,f15405])).
fof(f20196,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18)))) | $spl4),
  inference(unit_resulting_resolution,[],[f7016,f15405,f1398])).
fof(f22274,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20197,f15405])).
fof(f20197,plain,(
  p(s(bool,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2376,f15405,f1398])).
fof(f22273,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20198,f15405])).
fof(f20198,plain,(
  p(s(bool,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2377,f15405,f1398])).
fof(f22272,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20199,f15405])).
fof(f20199,plain,(
  p(s(bool,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2495,f15405,f1398])).
fof(f22271,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20200,f15405])).
fof(f20200,plain,(
  p(s(bool,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2496,f15405,f1398])).
fof(f22270,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20205,f15405])).
fof(f20205,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2146,f15405,f1398])).
fof(f22269,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20206,f15405])).
fof(f20206,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f6077,f15405,f1398])).
fof(f22268,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20207,f15405])).
fof(f20207,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f6989,f15405,f1398])).
fof(f22267,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20208,f1508])).
fof(f20208,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f2149,f15405,f1398])).
fof(f22266,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20209,f15405])).
fof(f20209,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK17)))) | $spl4),
  inference(unit_resulting_resolution,[],[f6080,f15405,f1398])).
fof(f22265,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20210,f15405])).
fof(f20210,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18)))) | $spl4),
  inference(unit_resulting_resolution,[],[f6992,f15405,f1398])).
fof(f22264,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20214,f15405])).
fof(f20214,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2376,f15405,f1398])).
fof(f22263,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20215,f1508])).
fof(f20215,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f2377,f15405,f1398])).
fof(f22262,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20216,f15405])).
fof(f20216,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2495,f15405,f1398])).
fof(f22261,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20217,f1508])).
fof(f20217,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f2496,f15405,f1398])).
fof(f22260,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20263,f20964])).
fof(f20964,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f15405,f15405,f15405,f1414])).
fof(f1414,plain,(
  ( ! [X2,X0,X1] : (sP10(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f924])).
fof(f924,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP10(X2,X0,X1)) & (~sP10(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f923])).
fof(f923,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP10(X2,X0,X1)) & (~sP10(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f732])).
fof(f732,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP10(X2,X0,X1))),
  inference(definition_folding,[],[f561,f731,f730])).
fof(f730,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f731,plain,(
  ! [X2,X0,X1] : (sP10(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP9(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f561,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f560])).
fof(f560,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f88])).
fof(f88,axiom,(
  ! [X35,X43,X44] : ((p(s(bool,X44)) <=> (p(s(bool,X43)) | p(s(bool,X35)))) <=> ((p(s(bool,X44)) | ~p(s(bool,X43))) & (p(s(bool,X44)) | ~p(s(bool,X35))) & (p(s(bool,X43)) | p(s(bool,X35)) | ~p(s(bool,X44)))))),
  file('thm.extreal.add_not_infty|split|1','thm.sat.sat_25_prove_ltb')).
fof(f20263,plain,(
  ( ! [X0,X1] : (~sP10(X0,X1,'T_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f952,f15405,f1403])).
fof(f1403,plain,(
  ( ! [X2,X0,X1] : (~sP10(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f919])).
fof(f919,plain,(
  ! [X0,X1,X2] : ((~sP10(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP9(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP9(X0,X1,X2) | sP10(X0,X1,X2)))),
  inference(rectify,[],[f918])).
fof(f918,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP9(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP9(X2,X0,X1) | sP10(X2,X0,X1)))),
  inference(flattening,[],[f917])).
fof(f917,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP9(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP9(X2,X0,X1)) | sP10(X2,X0,X1)))),
  inference(nnf_transformation,[],[f731])).
fof(f22259,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20287,f20964])).
fof(f20287,plain,(
  ( ! [X0,X1] : (~sP10(X0,'T_0',X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f952,f15405,f1404])).
fof(f1404,plain,(
  ( ! [X2,X0,X1] : (~sP10(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f919])).
fof(f22258,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20435,f15405])).
fof(f20435,plain,(
  p(s(bool,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f1565,f2379,f15405,f1409])).
fof(f1409,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~sP9(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f919])).
fof(f2379,plain,(
  ( ! [X0] : (~sP10('F_0',X0,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) )),
  inference(unit_resulting_resolution,[],[f1508,f2075,f1403])).
fof(f1565,plain,(
  ( ! [X0,X1] : (sP9('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f1508,f1413])).
fof(f1413,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f922])).
fof(f922,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f921])).
fof(f921,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP9(X2,X0,X1)))),
  inference(flattening,[],[f920])).
fof(f920,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP9(X2,X0,X1)))),
  inference(nnf_transformation,[],[f730])).
fof(f22257,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20436,f15405])).
fof(f20436,plain,(
  p(s(bool,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f1565,f2498,f15405,f1409])).
fof(f2498,plain,(
  ( ! [X0] : (~sP10('F_0',X0,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) )),
  inference(unit_resulting_resolution,[],[f1508,f2114,f1403])).
fof(f22256,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20437,f15405])).
fof(f20437,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1565,f2171,f15405,f1409])).
fof(f2171,plain,(
  ( ! [X0] : (~sP10('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f952,f1508,f1404])).
fof(f22255,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20438,f15405])).
fof(f20438,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1563,f2171,f15405,f1409])).
fof(f1563,plain,(
  ( ! [X0,X1] : (sP9(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f952,f1412])).
fof(f1412,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f922])).
fof(f22254,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20440,f15405])).
fof(f20440,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f1611,f2171,f15405,f1409])).
fof(f1611,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) )),
  inference(unit_resulting_resolution,[],[f962,f1411])).
fof(f1411,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f922])).
fof(f22253,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20441,f15405])).
fof(f20441,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f1562,f2171,f15405,f1409])).
fof(f1562,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f954,f1411])).
fof(f22252,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20442,f15405])).
fof(f20442,plain,(
  p(s(bool,sK77(sK63('const.num.|48|_0','type.num.num')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2049,f2171,f15405,f1409])).
fof(f2049,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,sK77(sK63('const.num.|48|_0','type.num.num')))) )),
  inference(unit_resulting_resolution,[],[f2030,f1411])).
fof(f22251,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20443,f15405])).
fof(f20443,plain,(
  p(s(bool,sK78(sK63('const.num.|48|_0','type.num.num')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2096,f2171,f15405,f1409])).
fof(f2096,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,sK78(sK63('const.num.|48|_0','type.num.num')))) )),
  inference(unit_resulting_resolution,[],[f2068,f1411])).
fof(f22250,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20444,f15405])).
fof(f20444,plain,(
  p(s(bool,sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))))) | $spl4),
  inference(unit_resulting_resolution,[],[f2132,f2171,f15405,f1409])).
fof(f2132,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))))) )),
  inference(unit_resulting_resolution,[],[f2108,f1411])).
fof(f22249,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20445,f15405])).
fof(f20445,plain,(
  p(s(bool,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2385,f2171,f15405,f1409])).
fof(f2385,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) )),
  inference(unit_resulting_resolution,[],[f2075,f1411])).
fof(f22248,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20446,f15405])).
fof(f20446,plain,(
  p(s(bool,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2504,f2171,f15405,f1409])).
fof(f2504,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) )),
  inference(unit_resulting_resolution,[],[f2114,f1411])).
fof(f22247,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20447,f15405])).
fof(f20447,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1565,f2382,f15405,f1409])).
fof(f2382,plain,(
  ( ! [X0] : (~sP10('F_0',sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')),X0)) )),
  inference(unit_resulting_resolution,[],[f1508,f2075,f1404])).
fof(f22246,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20448,f15405])).
fof(f20448,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2386,f2382,f15405,f1409])).
fof(f2386,plain,(
  ( ! [X0,X1] : (sP9(X0,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')),X1)) )),
  inference(unit_resulting_resolution,[],[f2075,f1412])).
fof(f22245,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20450,f15405])).
fof(f20450,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f1611,f2382,f15405,f1409])).
fof(f22244,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20451,f15405])).
fof(f20451,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f1562,f2382,f15405,f1409])).
fof(f22243,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20452,f15405])).
fof(f20452,plain,(
  p(s(bool,sK77(sK63('const.num.|48|_0','type.num.num')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2049,f2382,f15405,f1409])).
fof(f22242,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20453,f15405])).
fof(f20453,plain,(
  p(s(bool,sK78(sK63('const.num.|48|_0','type.num.num')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2096,f2382,f15405,f1409])).
fof(f22241,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20454,f15405])).
fof(f20454,plain,(
  p(s(bool,sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))))) | $spl4),
  inference(unit_resulting_resolution,[],[f2132,f2382,f15405,f1409])).
fof(f22240,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20455,f15405])).
fof(f20455,plain,(
  p(s(bool,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2385,f2382,f15405,f1409])).
fof(f22239,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20456,f15405])).
fof(f20456,plain,(
  p(s(bool,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2504,f2382,f15405,f1409])).
fof(f22238,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20457,f15405])).
fof(f20457,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1565,f2501,f15405,f1409])).
fof(f2501,plain,(
  ( ! [X0] : (~sP10('F_0',sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')),X0)) )),
  inference(unit_resulting_resolution,[],[f1508,f2114,f1404])).
fof(f22237,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20458,f15405])).
fof(f20458,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2505,f2501,f15405,f1409])).
fof(f2505,plain,(
  ( ! [X0,X1] : (sP9(X0,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')),X1)) )),
  inference(unit_resulting_resolution,[],[f2114,f1412])).
fof(f22236,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20460,f15405])).
fof(f20460,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f1611,f2501,f15405,f1409])).
fof(f22235,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20461,f15405])).
fof(f20461,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f1562,f2501,f15405,f1409])).
fof(f22234,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20462,f15405])).
fof(f20462,plain,(
  p(s(bool,sK77(sK63('const.num.|48|_0','type.num.num')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2049,f2501,f15405,f1409])).
fof(f22233,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20463,f15405])).
fof(f20463,plain,(
  p(s(bool,sK78(sK63('const.num.|48|_0','type.num.num')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2096,f2501,f15405,f1409])).
fof(f22232,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20464,f15405])).
fof(f20464,plain,(
  p(s(bool,sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))))) | $spl4),
  inference(unit_resulting_resolution,[],[f2132,f2501,f15405,f1409])).
fof(f22231,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20465,f15405])).
fof(f20465,plain,(
  p(s(bool,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2385,f2501,f15405,f1409])).
fof(f22230,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20466,f15405])).
fof(f20466,plain,(
  p(s(bool,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2504,f2501,f15405,f1409])).
fof(f22229,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20468,f15405])).
fof(f20468,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f6175,f6106,f15405,f1409])).
fof(f6106,plain,(
  ( ! [X0] : (~sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK17)),'T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f952,f5893,f1404])).
fof(f6175,plain,(
  ( ! [X0,X1] : (sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK17)),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f5893,f1413])).
fof(f22228,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20469,f15405])).
fof(f20469,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1563,f6106,f15405,f1409])).
fof(f22227,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20471,f15405])).
fof(f20471,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f1611,f6106,f15405,f1409])).
fof(f22226,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20472,f15405])).
fof(f20472,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f1562,f6106,f15405,f1409])).
fof(f22225,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20473,f15405])).
fof(f20473,plain,(
  p(s(bool,sK77(sK63('const.num.|48|_0','type.num.num')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2049,f6106,f15405,f1409])).
fof(f22224,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20474,f15405])).
fof(f20474,plain,(
  p(s(bool,sK78(sK63('const.num.|48|_0','type.num.num')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2096,f6106,f15405,f1409])).
fof(f22223,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20475,f15405])).
fof(f20475,plain,(
  p(s(bool,sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))))) | $spl4),
  inference(unit_resulting_resolution,[],[f2132,f6106,f15405,f1409])).
fof(f22222,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20476,f15405])).
fof(f20476,plain,(
  p(s(bool,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2385,f6106,f15405,f1409])).
fof(f22221,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20477,f15405])).
fof(f20477,plain,(
  p(s(bool,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2504,f6106,f15405,f1409])).
fof(f22220,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20479,f15405])).
fof(f20479,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f7087,f7018,f15405,f1409])).
fof(f7018,plain,(
  ( ! [X0] : (~sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18)),'T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f952,f5894,f1404])).
fof(f7087,plain,(
  ( ! [X0,X1] : (sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18)),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f5894,f1413])).
fof(f22219,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20480,f15405])).
fof(f20480,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1563,f7018,f15405,f1409])).
fof(f22218,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20482,f15405])).
fof(f20482,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f1611,f7018,f15405,f1409])).
fof(f22217,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20483,f15405])).
fof(f20483,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f1562,f7018,f15405,f1409])).
fof(f22216,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20484,f15405])).
fof(f20484,plain,(
  p(s(bool,sK77(sK63('const.num.|48|_0','type.num.num')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2049,f7018,f15405,f1409])).
fof(f22215,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20485,f15405])).
fof(f20485,plain,(
  p(s(bool,sK78(sK63('const.num.|48|_0','type.num.num')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2096,f7018,f15405,f1409])).
fof(f22214,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20486,f15405])).
fof(f20486,plain,(
  p(s(bool,sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))))) | $spl4),
  inference(unit_resulting_resolution,[],[f2132,f7018,f15405,f1409])).
fof(f22213,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20487,f15405])).
fof(f20487,plain,(
  p(s(bool,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2385,f7018,f15405,f1409])).
fof(f22212,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20488,f15405])).
fof(f20488,plain,(
  p(s(bool,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2504,f7018,f15405,f1409])).
fof(f22211,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20501,f20806])).
fof(f20806,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f15405,f1413])).
fof(f20501,plain,(
  ~sP9('F_0','T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f2171,f1508,f15405,f1409])).
fof(f22210,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20503,f20806])).
fof(f20503,plain,(
  ~sP9('F_0',sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f2382,f1508,f15405,f1409])).
fof(f22209,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20505,f20806])).
fof(f20505,plain,(
  ~sP9('F_0',sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f2501,f1508,f15405,f1409])).
fof(f22208,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20510,f20806])).
fof(f20510,plain,(
  ~sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK17)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f6106,f1508,f15405,f1409])).
fof(f22207,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20512,f20806])).
fof(f20512,plain,(
  ~sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18)),'T_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f7018,f1508,f15405,f1409])).
fof(f22206,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20534,f20806])).
fof(f20534,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'T_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2168,f15405,f15405,f1409])).
fof(f2168,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f952,f1508,f1403])).
fof(f22205,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20536,f20806])).
fof(f20536,plain,(
  ( ! [X0] : (~sP9('F_0',X0,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2379,f15405,f15405,f1409])).
fof(f22204,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20538,f20806])).
fof(f20538,plain,(
  ( ! [X0] : (~sP9('F_0',X0,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2498,f15405,f15405,f1409])).
fof(f22203,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20540,f20806])).
fof(f20540,plain,(
  ( ! [X0] : (~sP9('F_0','T_0',X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2171,f15405,f15405,f1409])).
fof(f22202,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20542,f20806])).
fof(f20542,plain,(
  ( ! [X0] : (~sP9('F_0',sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')),X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2382,f15405,f15405,f1409])).
fof(f22201,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20544,f20806])).
fof(f20544,plain,(
  ( ! [X0] : (~sP9('F_0',sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')),X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2501,f15405,f15405,f1409])).
fof(f22200,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20551,f20806])).
fof(f20551,plain,(
  ( ! [X0] : (~sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK17)),X0,'T_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f6105,f15405,f15405,f1409])).
fof(f6105,plain,(
  ( ! [X0] : (~sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK17)),X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f952,f5893,f1403])).
fof(f22199,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20553,f20806])).
fof(f20553,plain,(
  ( ! [X0] : (~sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK17)),'T_0',X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f6106,f15405,f15405,f1409])).
fof(f22198,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20555,f20806])).
fof(f20555,plain,(
  ( ! [X0] : (~sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18)),X0,'T_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f7017,f15405,f15405,f1409])).
fof(f7017,plain,(
  ( ! [X0] : (~sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18)),X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f952,f5894,f1403])).
fof(f22197,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20557,f20806])).
fof(f20557,plain,(
  ( ! [X0] : (~sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18)),'T_0',X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f7018,f15405,f15405,f1409])).
fof(f22196,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20559,f15405])).
fof(f20559,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2168,f1561,f15405,f1409])).
fof(f1561,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f952,f1411])).
fof(f22195,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20561,f15405])).
fof(f20561,plain,(
  p(s(bool,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2382,f1561,f15405,f1409])).
fof(f22194,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20562,f15405])).
fof(f20562,plain,(
  p(s(bool,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2501,f1561,f15405,f1409])).
fof(f22193,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20563,f15405])).
fof(f20563,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f6105,f1561,f15405,f1409])).
fof(f22192,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20565,f15405])).
fof(f20565,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f7017,f1561,f15405,f1409])).
fof(f22191,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20568,f15405])).
fof(f20568,plain,(
  p(s(bool,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2382,f1611,f15405,f1409])).
fof(f22190,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20569,f15405])).
fof(f20569,plain,(
  p(s(bool,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2501,f1611,f15405,f1409])).
fof(f22189,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20573,f15405])).
fof(f20573,plain,(
  p(s(bool,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2382,f1562,f15405,f1409])).
fof(f22188,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20574,f15405])).
fof(f20574,plain,(
  p(s(bool,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2501,f1562,f15405,f1409])).
fof(f22187,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20578,f15405])).
fof(f20578,plain,(
  p(s(bool,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2382,f2049,f15405,f1409])).
fof(f22186,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20579,f15405])).
fof(f20579,plain,(
  p(s(bool,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2501,f2049,f15405,f1409])).
fof(f22185,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20583,f15405])).
fof(f20583,plain,(
  p(s(bool,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2382,f2096,f15405,f1409])).
fof(f22184,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20584,f15405])).
fof(f20584,plain,(
  p(s(bool,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2501,f2096,f15405,f1409])).
fof(f22183,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20588,f15405])).
fof(f20588,plain,(
  p(s(bool,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2382,f2132,f15405,f1409])).
fof(f22182,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20589,f15405])).
fof(f20589,plain,(
  p(s(bool,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2501,f2132,f15405,f1409])).
fof(f22181,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20592,f15405])).
fof(f20592,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2379,f2385,f15405,f1409])).
fof(f22180,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20594,f15405])).
fof(f20594,plain,(
  p(s(bool,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2382,f2385,f15405,f1409])).
fof(f22179,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20595,f15405])).
fof(f20595,plain,(
  p(s(bool,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2501,f2385,f15405,f1409])).
fof(f22178,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20598,f15405])).
fof(f20598,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2498,f2504,f15405,f1409])).
fof(f22177,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20600,f15405])).
fof(f20600,plain,(
  p(s(bool,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2382,f2504,f15405,f1409])).
fof(f22176,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20601,f15405])).
fof(f20601,plain,(
  p(s(bool,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2501,f2504,f15405,f1409])).
fof(f22175,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20605,f15405])).
fof(f20605,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2168,f1612,f15405,f1409])).
fof(f1612,plain,(
  ( ! [X0,X1] : (sP9(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f962,f1412])).
fof(f22174,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20606,f15405])).
fof(f20606,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2168,f1564,f15405,f1409])).
fof(f1564,plain,(
  ( ! [X0,X1] : (sP9(X0,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f954,f1412])).
fof(f22173,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20607,f15405])).
fof(f20607,plain,(
  p(s(bool,sK77(sK63('const.num.|48|_0','type.num.num')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2168,f2050,f15405,f1409])).
fof(f2050,plain,(
  ( ! [X0,X1] : (sP9(X0,sK77(sK63('const.num.|48|_0','type.num.num')),X1)) )),
  inference(unit_resulting_resolution,[],[f2030,f1412])).
fof(f22172,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20608,f15405])).
fof(f20608,plain,(
  p(s(bool,sK78(sK63('const.num.|48|_0','type.num.num')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2168,f2097,f15405,f1409])).
fof(f2097,plain,(
  ( ! [X0,X1] : (sP9(X0,sK78(sK63('const.num.|48|_0','type.num.num')),X1)) )),
  inference(unit_resulting_resolution,[],[f2068,f1412])).
fof(f22171,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20609,f15405])).
fof(f20609,plain,(
  p(s(bool,sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))))) | $spl4),
  inference(unit_resulting_resolution,[],[f2168,f2133,f15405,f1409])).
fof(f2133,plain,(
  ( ! [X0,X1] : (sP9(X0,sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))),X1)) )),
  inference(unit_resulting_resolution,[],[f2108,f1412])).
fof(f22170,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20610,f15405])).
fof(f20610,plain,(
  p(s(bool,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2382,f2386,f15405,f1409])).
fof(f22169,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20611,f15405])).
fof(f20611,plain,(
  p(s(bool,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2501,f2505,f15405,f1409])).
fof(f22168,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20612,f15405])).
fof(f20612,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2168,f1565,f15405,f1409])).
fof(f22167,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20613,f15405])).
fof(f20613,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2379,f1565,f15405,f1409])).
fof(f22166,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20614,f15405])).
fof(f20614,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2498,f1565,f15405,f1409])).
fof(f22165,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20616,f15405])).
fof(f20616,plain,(
  p(s(bool,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2382,f1565,f15405,f1409])).
fof(f22164,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20617,f15405])).
fof(f20617,plain,(
  p(s(bool,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2501,f1565,f15405,f1409])).
fof(f22163,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20618,f15405])).
fof(f20618,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f6105,f6175,f15405,f1409])).
fof(f22162,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20620,f15405])).
fof(f20620,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f7017,f7087,f15405,f1409])).
fof(f22161,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20622,f20806])).
fof(f20622,plain,(
  ~sP9('F_0','F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f2168,f1508,f15405,f1409])).
fof(f22160,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20623,f20806])).
fof(f20623,plain,(
  ~sP9('F_0','F_0',sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal'))) | $spl4),
  inference(unit_resulting_resolution,[],[f2379,f1508,f15405,f1409])).
fof(f22159,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20624,f20806])).
fof(f20624,plain,(
  ~sP9('F_0','F_0',sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal'))) | $spl4),
  inference(unit_resulting_resolution,[],[f2498,f1508,f15405,f1409])).
fof(f22158,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20628,f20806])).
fof(f20628,plain,(
  ~sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK17)),'F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f6105,f1508,f15405,f1409])).
fof(f22157,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20629,f20806])).
fof(f20629,plain,(
  ~sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18)),'F_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f7017,f1508,f15405,f1409])).
fof(f22156,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20648,f20806])).
fof(f20648,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'T_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2168,f15405,f15405,f1409])).
fof(f22155,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20649,f20806])).
fof(f20649,plain,(
  ( ! [X0] : (~sP9('F_0',X0,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2379,f15405,f15405,f1409])).
fof(f22154,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20650,f20806])).
fof(f20650,plain,(
  ( ! [X0] : (~sP9('F_0',X0,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2498,f15405,f15405,f1409])).
fof(f22153,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20651,f20806])).
fof(f20651,plain,(
  ( ! [X0] : (~sP9('F_0','T_0',X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2171,f15405,f15405,f1409])).
fof(f22152,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20652,f20806])).
fof(f20652,plain,(
  ( ! [X0] : (~sP9('F_0',sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')),X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2382,f15405,f15405,f1409])).
fof(f22151,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20653,f20806])).
fof(f20653,plain,(
  ( ! [X0] : (~sP9('F_0',sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')),X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2501,f15405,f15405,f1409])).
fof(f22150,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20659,f20806])).
fof(f20659,plain,(
  ( ! [X0] : (~sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK17)),X0,'T_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f6105,f15405,f15405,f1409])).
fof(f22149,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20660,f20806])).
fof(f20660,plain,(
  ( ! [X0] : (~sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK17)),'T_0',X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f6106,f15405,f15405,f1409])).
fof(f22148,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20661,f20806])).
fof(f20661,plain,(
  ( ! [X0] : (~sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18)),X0,'T_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f7017,f15405,f15405,f1409])).
fof(f22147,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20662,f20806])).
fof(f20662,plain,(
  ( ! [X0] : (~sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18)),'T_0',X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f7018,f15405,f15405,f1409])).
fof(f22146,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20703,f15405])).
fof(f20703,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f1611,f952,f15405,f1410])).
fof(f1410,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f922])).
fof(f22145,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20704,f15405])).
fof(f20704,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f1562,f952,f15405,f1410])).
fof(f22144,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20705,f15405])).
fof(f20705,plain,(
  p(s(bool,sK77(sK63('const.num.|48|_0','type.num.num')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2049,f952,f15405,f1410])).
fof(f22143,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20706,f15405])).
fof(f20706,plain,(
  p(s(bool,sK78(sK63('const.num.|48|_0','type.num.num')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2096,f952,f15405,f1410])).
fof(f22142,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20707,f15405])).
fof(f20707,plain,(
  p(s(bool,sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))))) | $spl4),
  inference(unit_resulting_resolution,[],[f2132,f952,f15405,f1410])).
fof(f22141,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20708,f15405])).
fof(f20708,plain,(
  p(s(bool,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2385,f952,f15405,f1410])).
fof(f22140,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20709,f15405])).
fof(f20709,plain,(
  p(s(bool,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2504,f952,f15405,f1410])).
fof(f22139,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20710,f15405])).
fof(f20710,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1563,f952,f15405,f1410])).
fof(f22138,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20711,f15405])).
fof(f20711,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1612,f952,f15405,f1410])).
fof(f22137,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20712,f15405])).
fof(f20712,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1564,f952,f15405,f1410])).
fof(f22136,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20713,f15405])).
fof(f20713,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2050,f952,f15405,f1410])).
fof(f22135,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20714,f15405])).
fof(f20714,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2097,f952,f15405,f1410])).
fof(f22134,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20715,f15405])).
fof(f20715,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2133,f952,f15405,f1410])).
fof(f22133,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20716,f15405])).
fof(f20716,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2386,f952,f15405,f1410])).
fof(f22132,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20717,f15405])).
fof(f20717,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2505,f952,f15405,f1410])).
fof(f22131,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20730,f20806])).
fof(f20730,plain,(
  ( ! [X0] : (~sP9('T_0',X0,'F_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f952,f1508,f15405,f1410])).
fof(f22130,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20752,f20806])).
fof(f20752,plain,(
  ( ! [X0,X1] : (~sP9('T_0',X0,X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f952,f15405,f15405,f1410])).
fof(f22129,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20754,f15405])).
fof(f20754,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1561,f952,f15405,f1410])).
fof(f22128,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20755,f15405])).
fof(f20755,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1611,f952,f15405,f1410])).
fof(f22127,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20756,f15405])).
fof(f20756,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1562,f952,f15405,f1410])).
fof(f22126,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20757,f15405])).
fof(f20757,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2049,f952,f15405,f1410])).
fof(f22125,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20758,f15405])).
fof(f20758,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2096,f952,f15405,f1410])).
fof(f22124,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20759,f15405])).
fof(f20759,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2132,f952,f15405,f1410])).
fof(f22123,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20760,f15405])).
fof(f20760,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2385,f952,f15405,f1410])).
fof(f22122,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20761,f15405])).
fof(f20761,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2504,f952,f15405,f1410])).
fof(f22121,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20763,f15405])).
fof(f20763,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f1612,f952,f15405,f1410])).
fof(f22120,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20764,f15405])).
fof(f20764,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f1564,f952,f15405,f1410])).
fof(f22119,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20765,f15405])).
fof(f20765,plain,(
  p(s(bool,sK77(sK63('const.num.|48|_0','type.num.num')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2050,f952,f15405,f1410])).
fof(f22118,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20766,f15405])).
fof(f20766,plain,(
  p(s(bool,sK78(sK63('const.num.|48|_0','type.num.num')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2097,f952,f15405,f1410])).
fof(f22117,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20767,f15405])).
fof(f20767,plain,(
  p(s(bool,sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))))) | $spl4),
  inference(unit_resulting_resolution,[],[f2133,f952,f15405,f1410])).
fof(f22116,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20768,f15405])).
fof(f20768,plain,(
  p(s(bool,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2386,f952,f15405,f1410])).
fof(f22115,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20769,f15405])).
fof(f20769,plain,(
  p(s(bool,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2505,f952,f15405,f1410])).
fof(f22114,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20782,f20806])).
fof(f20782,plain,(
  ( ! [X0] : (~sP9('T_0','F_0',X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f952,f1508,f15405,f1410])).
fof(f22113,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20804,f20806])).
fof(f20804,plain,(
  ( ! [X0,X1] : (~sP9('T_0',X0,X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f952,f15405,f15405,f1410])).
fof(f22112,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20809,f15405])).
fof(f20809,plain,(
  p(s(bool,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2379,f1508,f15405,f1414])).
fof(f22111,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20810,f15405])).
fof(f20810,plain,(
  p(s(bool,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2498,f1508,f15405,f1414])).
fof(f22110,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20811,f15405])).
fof(f20811,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2171,f1508,f15405,f1414])).
fof(f22109,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20812,f15405])).
fof(f20812,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2382,f1508,f15405,f1414])).
fof(f22108,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20813,f15405])).
fof(f20813,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2501,f1508,f15405,f1414])).
fof(f22107,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20815,f15405])).
fof(f20815,plain,(
  p(s(bool,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2379,f15405,f15405,f1414])).
fof(f22106,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20816,f15405])).
fof(f20816,plain,(
  p(s(bool,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2498,f15405,f15405,f1414])).
fof(f22105,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20817,f15405])).
fof(f20817,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2171,f15405,f15405,f1414])).
fof(f22104,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20818,f15405])).
fof(f20818,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2382,f15405,f15405,f1414])).
fof(f22103,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20819,f15405])).
fof(f20819,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2501,f15405,f15405,f1414])).
fof(f22102,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20820,f1508])).
fof(f20820,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f4572,f15405,f15405,f1414])).
fof(f4572,plain,(
  ~sP10('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f4553,f1405])).
fof(f1405,plain,(
  ( ! [X2,X0,X1] : (~sP10(X0,X1,X2) | sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f919])).
fof(f4553,plain,(
  ~sP9('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f952,f1508,f1508,f1410])).
fof(f22101,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20821,f1508])).
fof(f20821,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f6194,f15405,f15405,f1414])).
fof(f6194,plain,(
  ~sP10('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK17)),'F_0')),
  inference(unit_resulting_resolution,[],[f1508,f952,f5893,f1417])).
fof(f1417,plain,(
  ( ! [X2,X0,X1] : (~sP10(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f924])).
fof(f22100,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20822,f1508])).
fof(f20822,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f7106,f15405,f15405,f1414])).
fof(f7106,plain,(
  ~sP10('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18)),'F_0')),
  inference(unit_resulting_resolution,[],[f1508,f952,f5894,f1417])).
fof(f22099,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20823,f15405])).
fof(f20823,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK17)))) | $spl4),
  inference(unit_resulting_resolution,[],[f6201,f15405,f15405,f1414])).
fof(f6201,plain,(
  ~sP10('T_0','F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK17)))),
  inference(unit_resulting_resolution,[],[f1508,f952,f5893,f1417])).
fof(f22098,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20824,f15405])).
fof(f20824,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18)))) | $spl4),
  inference(unit_resulting_resolution,[],[f7113,f15405,f15405,f1414])).
fof(f7113,plain,(
  ~sP10('T_0','F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18)))),
  inference(unit_resulting_resolution,[],[f1508,f952,f5894,f1417])).
fof(f22097,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20826,f15405])).
fof(f20826,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f6106,f15405,f15405,f1414])).
fof(f22096,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20828,f15405])).
fof(f20828,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f7018,f15405,f15405,f1414])).
fof(f22095,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20829,f1508])).
fof(f20829,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f2171,f1508,f15405,f1414])).
fof(f22094,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20830,f1508])).
fof(f20830,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f2382,f1508,f15405,f1414])).
fof(f22093,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20831,f1508])).
fof(f20831,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f2501,f1508,f15405,f1414])).
fof(f22092,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20835,f15405])).
fof(f20835,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK17)))) | $spl4),
  inference(unit_resulting_resolution,[],[f6106,f1508,f15405,f1414])).
fof(f22091,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20836,f15405])).
fof(f20836,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18)))) | $spl4),
  inference(unit_resulting_resolution,[],[f7018,f1508,f15405,f1414])).
fof(f22090,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20841,f1508])).
fof(f20841,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f2168,f15405,f15405,f1414])).
fof(f22089,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20842,f1508])).
fof(f20842,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f2379,f15405,f15405,f1414])).
fof(f22088,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20843,f1508])).
fof(f20843,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f2498,f15405,f15405,f1414])).
fof(f22087,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20844,f1508])).
fof(f20844,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f2171,f15405,f15405,f1414])).
fof(f22086,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20845,f1508])).
fof(f20845,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f2382,f15405,f15405,f1414])).
fof(f22085,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20846,f1508])).
fof(f20846,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f2501,f15405,f15405,f1414])).
fof(f22084,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20852,f15405])).
fof(f20852,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK17)))) | $spl4),
  inference(unit_resulting_resolution,[],[f6105,f15405,f15405,f1414])).
fof(f22083,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20853,f15405])).
fof(f20853,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK17)))) | $spl4),
  inference(unit_resulting_resolution,[],[f6106,f15405,f15405,f1414])).
fof(f22082,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20854,f15405])).
fof(f20854,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18)))) | $spl4),
  inference(unit_resulting_resolution,[],[f7017,f15405,f15405,f1414])).
fof(f22081,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20855,f15405])).
fof(f20855,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18)))) | $spl4),
  inference(unit_resulting_resolution,[],[f7018,f15405,f15405,f1414])).
fof(f22080,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20860,f15405])).
fof(f20860,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2168,f1508,f15405,f1414])).
fof(f22079,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20861,f15405])).
fof(f20861,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2379,f1508,f15405,f1414])).
fof(f22078,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20862,f15405])).
fof(f20862,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2498,f1508,f15405,f1414])).
fof(f22077,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20864,f15405])).
fof(f20864,plain,(
  p(s(bool,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2382,f1508,f15405,f1414])).
fof(f22076,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20865,f15405])).
fof(f20865,plain,(
  p(s(bool,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2501,f1508,f15405,f1414])).
fof(f22075,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20866,f15405])).
fof(f20866,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2168,f15405,f15405,f1414])).
fof(f22074,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20867,f15405])).
fof(f20867,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2379,f15405,f15405,f1414])).
fof(f22073,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20868,f15405])).
fof(f20868,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2498,f15405,f15405,f1414])).
fof(f22072,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20870,f15405])).
fof(f20870,plain,(
  p(s(bool,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2382,f15405,f15405,f1414])).
fof(f22071,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20871,f15405])).
fof(f20871,plain,(
  p(s(bool,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2501,f15405,f15405,f1414])).
fof(f22070,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20872,f1508])).
fof(f20872,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f4572,f15405,f15405,f1414])).
fof(f22069,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20873,f15405])).
fof(f20873,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK17)))) | $spl4),
  inference(unit_resulting_resolution,[],[f6194,f15405,f15405,f1414])).
fof(f22068,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20874,f15405])).
fof(f20874,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18)))) | $spl4),
  inference(unit_resulting_resolution,[],[f7106,f15405,f15405,f1414])).
fof(f22067,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20875,f1508])).
fof(f20875,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f6201,f15405,f15405,f1414])).
fof(f22066,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20876,f1508])).
fof(f20876,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f7113,f15405,f15405,f1414])).
fof(f22065,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20877,f15405])).
fof(f20877,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f6105,f15405,f15405,f1414])).
fof(f22064,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20879,f15405])).
fof(f20879,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f7017,f15405,f15405,f1414])).
fof(f22063,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20881,f1508])).
fof(f20881,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f2168,f1508,f15405,f1414])).
fof(f22062,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20882,f1508])).
fof(f20882,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f2379,f1508,f15405,f1414])).
fof(f22061,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20883,f1508])).
fof(f20883,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f2498,f1508,f15405,f1414])).
fof(f22060,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20887,f15405])).
fof(f20887,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK17)))) | $spl4),
  inference(unit_resulting_resolution,[],[f6105,f1508,f15405,f1414])).
fof(f22059,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20888,f15405])).
fof(f20888,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18)))) | $spl4),
  inference(unit_resulting_resolution,[],[f7017,f1508,f15405,f1414])).
fof(f22058,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20893,f1508])).
fof(f20893,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f2168,f15405,f15405,f1414])).
fof(f22057,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20894,f1508])).
fof(f20894,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f2379,f15405,f15405,f1414])).
fof(f22056,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20895,f1508])).
fof(f20895,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f2498,f15405,f15405,f1414])).
fof(f22055,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20896,f1508])).
fof(f20896,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f2171,f15405,f15405,f1414])).
fof(f22054,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20897,f1508])).
fof(f20897,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f2382,f15405,f15405,f1414])).
fof(f22053,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20898,f1508])).
fof(f20898,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f2501,f15405,f15405,f1414])).
fof(f22052,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20904,f15405])).
fof(f20904,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK17)))) | $spl4),
  inference(unit_resulting_resolution,[],[f6105,f15405,f15405,f1414])).
fof(f22051,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20905,f15405])).
fof(f20905,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK17)))) | $spl4),
  inference(unit_resulting_resolution,[],[f6106,f15405,f15405,f1414])).
fof(f22050,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20906,f15405])).
fof(f20906,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18)))) | $spl4),
  inference(unit_resulting_resolution,[],[f7017,f15405,f15405,f1414])).
fof(f22049,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20907,f15405])).
fof(f20907,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18)))) | $spl4),
  inference(unit_resulting_resolution,[],[f7018,f15405,f15405,f1414])).
fof(f22048,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20913,f15405])).
fof(f20913,plain,(
  p(s(bool,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2382,f1508,f15405,f1414])).
fof(f22047,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20914,f15405])).
fof(f20914,plain,(
  p(s(bool,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2501,f1508,f15405,f1414])).
fof(f22046,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20915,f1508])).
fof(f20915,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f4572,f1508,f15405,f1414])).
fof(f22045,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20916,f15405])).
fof(f20916,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK17)))) | $spl4),
  inference(unit_resulting_resolution,[],[f6194,f1508,f15405,f1414])).
fof(f22044,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20917,f15405])).
fof(f20917,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18)))) | $spl4),
  inference(unit_resulting_resolution,[],[f7106,f1508,f15405,f1414])).
fof(f22043,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20920,f15405])).
fof(f20920,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2168,f15405,f15405,f1414])).
fof(f22042,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20921,f15405])).
fof(f20921,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2379,f15405,f15405,f1414])).
fof(f22041,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20922,f15405])).
fof(f20922,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2498,f15405,f15405,f1414])).
fof(f22040,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20924,f15405])).
fof(f20924,plain,(
  p(s(bool,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2382,f15405,f15405,f1414])).
fof(f22039,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20925,f15405])).
fof(f20925,plain,(
  p(s(bool,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2501,f15405,f15405,f1414])).
fof(f22038,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20926,f1508])).
fof(f20926,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f4572,f15405,f15405,f1414])).
fof(f22037,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20927,f15405])).
fof(f20927,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK17)))) | $spl4),
  inference(unit_resulting_resolution,[],[f6194,f15405,f15405,f1414])).
fof(f22036,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20928,f15405])).
fof(f20928,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18)))) | $spl4),
  inference(unit_resulting_resolution,[],[f7106,f15405,f15405,f1414])).
fof(f22035,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20929,f1508])).
fof(f20929,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f6201,f15405,f15405,f1414])).
fof(f22034,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20930,f1508])).
fof(f20930,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f7113,f15405,f15405,f1414])).
fof(f22033,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20931,f15405])).
fof(f20931,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f6105,f15405,f15405,f1414])).
fof(f22032,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20933,f15405])).
fof(f20933,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f7017,f15405,f15405,f1414])).
fof(f22031,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20936,f15405])).
fof(f20936,plain,(
  p(s(bool,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2379,f1508,f15405,f1414])).
fof(f22030,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20937,f15405])).
fof(f20937,plain,(
  p(s(bool,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2498,f1508,f15405,f1414])).
fof(f22029,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20938,f1508])).
fof(f20938,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f4572,f1508,f15405,f1414])).
fof(f22028,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20939,f15405])).
fof(f20939,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK17)))) | $spl4),
  inference(unit_resulting_resolution,[],[f6201,f1508,f15405,f1414])).
fof(f22027,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20940,f15405])).
fof(f20940,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18)))) | $spl4),
  inference(unit_resulting_resolution,[],[f7113,f1508,f15405,f1414])).
fof(f22026,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20948,f15405])).
fof(f20948,plain,(
  p(s(bool,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2379,f15405,f15405,f1414])).
fof(f22025,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20949,f15405])).
fof(f20949,plain,(
  p(s(bool,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2498,f15405,f15405,f1414])).
fof(f22024,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20950,f15405])).
fof(f20950,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2171,f15405,f15405,f1414])).
fof(f22023,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20951,f15405])).
fof(f20951,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2382,f15405,f15405,f1414])).
fof(f22022,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20952,f15405])).
fof(f20952,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2501,f15405,f15405,f1414])).
fof(f22021,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20953,f1508])).
fof(f20953,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f4572,f15405,f15405,f1414])).
fof(f22020,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20954,f1508])).
fof(f20954,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f6194,f15405,f15405,f1414])).
fof(f22019,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20955,f1508])).
fof(f20955,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f7106,f15405,f15405,f1414])).
fof(f22018,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20956,f15405])).
fof(f20956,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK17)))) | $spl4),
  inference(unit_resulting_resolution,[],[f6201,f15405,f15405,f1414])).
fof(f22017,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20957,f15405])).
fof(f20957,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18)))) | $spl4),
  inference(unit_resulting_resolution,[],[f7113,f15405,f15405,f1414])).
fof(f22016,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20959,f15405])).
fof(f20959,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f6106,f15405,f15405,f1414])).
fof(f22015,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20961,f15405])).
fof(f20961,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f7018,f15405,f15405,f1414])).
fof(f22014,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20967,f15405])).
fof(f20967,plain,(
  p(s(bool,sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))))) | $spl4),
  inference(unit_resulting_resolution,[],[f2227,f952,f15405,f1417])).
fof(f2227,plain,(
  ( ! [X0] : (sP10('T_0',X0,sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))))) )),
  inference(unit_resulting_resolution,[],[f2132,f952,f1496])).
fof(f1496,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~sP9(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(duplicate_literal_removal,[],[f1406])).
fof(f1406,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~sP9(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f919])).
fof(f22013,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20968,f15405])).
fof(f20968,plain,(
  p(s(bool,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2389,f952,f15405,f1417])).
fof(f2389,plain,(
  ( ! [X0] : (sP10('T_0',X0,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) )),
  inference(unit_resulting_resolution,[],[f952,f2075,f1415])).
fof(f1415,plain,(
  ( ! [X2,X0,X1] : (sP10(X2,X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f924])).
fof(f22012,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20969,f15405])).
fof(f20969,plain,(
  p(s(bool,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2508,f952,f15405,f1417])).
fof(f2508,plain,(
  ( ! [X0] : (sP10('T_0',X0,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) )),
  inference(unit_resulting_resolution,[],[f952,f2114,f1415])).
fof(f22011,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20970,f15405])).
fof(f20970,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2226,f952,f15405,f1417])).
fof(f2226,plain,(
  ( ! [X0] : (sP10('T_0',X0,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f1562,f952,f1496])).
fof(f22010,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20971,f15405])).
fof(f20971,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2180,f952,f15405,f1417])).
fof(f2180,plain,(
  ( ! [X0] : (sP10('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f952,f952,f1416])).
fof(f1416,plain,(
  ( ! [X2,X0,X1] : (sP10(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f924])).
fof(f22009,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20972,f15405])).
fof(f20972,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2230,f952,f15405,f1417])).
fof(f2230,plain,(
  ( ! [X0] : (sP10('T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X0)) )),
  inference(unit_resulting_resolution,[],[f1564,f952,f1496])).
fof(f22008,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20973,f15405])).
fof(f20973,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2231,f952,f15405,f1417])).
fof(f2231,plain,(
  ( ! [X0] : (sP10('T_0',sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))),X0)) )),
  inference(unit_resulting_resolution,[],[f2133,f952,f1496])).
fof(f22007,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20974,f15405])).
fof(f20974,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2392,f952,f15405,f1417])).
fof(f2392,plain,(
  ( ! [X0] : (sP10('T_0',sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')),X0)) )),
  inference(unit_resulting_resolution,[],[f952,f2075,f1416])).
fof(f22006,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f20975,f15405])).
fof(f20975,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2511,f952,f15405,f1417])).
fof(f2511,plain,(
  ( ! [X0] : (sP10('T_0',sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')),X0)) )),
  inference(unit_resulting_resolution,[],[f952,f2114,f1416])).
fof(f22005,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21050,f15405])).
fof(f21050,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2174,f952,f15405,f1417])).
fof(f2174,plain,(
  ( ! [X0] : (sP10('T_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f952,f952,f1415])).
fof(f22004,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21051,f15405])).
fof(f21051,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2227,f952,f15405,f1417])).
fof(f22003,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21052,f15405])).
fof(f21052,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2389,f952,f15405,f1417])).
fof(f22002,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21053,f15405])).
fof(f21053,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2508,f952,f15405,f1417])).
fof(f22001,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21054,f15405])).
fof(f21054,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2226,f952,f15405,f1417])).
fof(f22000,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21056,f15405])).
fof(f21056,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2230,f952,f15405,f1417])).
fof(f21999,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21057,f15405])).
fof(f21057,plain,(
  p(s(bool,sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))))) | $spl4),
  inference(unit_resulting_resolution,[],[f2231,f952,f15405,f1417])).
fof(f21998,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21058,f15405])).
fof(f21058,plain,(
  p(s(bool,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2392,f952,f15405,f1417])).
fof(f21997,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21059,f15405])).
fof(f21059,plain,(
  p(s(bool,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2511,f952,f15405,f1417])).
fof(f21996,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21064,f1508])).
fof(f21064,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f2188,f15405,f1420])).
fof(f1420,plain,(
  ( ! [X2,X0,X1] : (~sP12(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f927])).
fof(f927,plain,(
  ! [X0,X1,X2] : ((~sP12(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP11(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP11(X0,X1,X2) | sP12(X0,X1,X2)))),
  inference(rectify,[],[f926])).
fof(f926,plain,(
  ! [X2,X0,X1] : ((~sP12(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP11(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP11(X2,X0,X1) | sP12(X2,X0,X1)))),
  inference(flattening,[],[f925])).
fof(f925,plain,(
  ! [X2,X0,X1] : ((~sP12(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP11(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP11(X2,X0,X1)) | sP12(X2,X0,X1)))),
  inference(nnf_transformation,[],[f734])).
fof(f734,plain,(
  ! [X2,X0,X1] : (sP12(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP11(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP12])])).
fof(f2188,plain,(
  ( ! [X0] : (sP12('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f952,f1508,f1432])).
fof(f1432,plain,(
  ( ! [X2,X0,X1] : (sP12(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f932])).
fof(f932,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP12(X2,X0,X1)) & (~sP12(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f931])).
fof(f931,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP12(X2,X0,X1)) & (~sP12(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f735])).
fof(f735,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP12(X2,X0,X1))),
  inference(definition_folding,[],[f695,f734,f733])).
fof(f733,plain,(
  ! [X2,X0,X1] : (sP11(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f695,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f563])).
fof(f563,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f562])).
fof(f562,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f89])).
fof(f89,axiom,(
  ! [X35,X43,X44] : ((p(s(bool,X44)) <=> (p(s(bool,X43)) => p(s(bool,X35)))) <=> ((p(s(bool,X44)) | p(s(bool,X43))) & (p(s(bool,X44)) | ~p(s(bool,X35))) & (~p(s(bool,X43)) | p(s(bool,X35)) | ~p(s(bool,X44)))))),
  file('thm.extreal.add_not_infty|split|1','thm.sat.sat_27_prove_ltb')).
fof(f21995,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21065,f15405])).
fof(f21065,plain,(
  p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2207,f15405,f1420])).
fof(f2207,plain,(
  ( ! [X0] : (sP12('T_0',X0,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f1568,f952,f1495])).
fof(f1495,plain,(
  ( ! [X2,X0,X1] : (sP12(X0,X1,X2) | ~sP11(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(duplicate_literal_removal,[],[f1423])).
fof(f1423,plain,(
  ( ! [X2,X0,X1] : (sP12(X0,X1,X2) | ~sP11(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f927])).
fof(f1568,plain,(
  ( ! [X0,X1] : (sP11(X0,X1,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f982,f1428])).
fof(f1428,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f930])).
fof(f930,plain,(
  ! [X0,X1,X2] : ((~sP11(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP11(X0,X1,X2)))),
  inference(rectify,[],[f929])).
fof(f929,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP11(X2,X0,X1)))),
  inference(flattening,[],[f928])).
fof(f928,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP11(X2,X0,X1)))),
  inference(nnf_transformation,[],[f733])).
fof(f982,plain,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(cnf_transformation,[],[f269])).
fof(f269,plain,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(flattening,[],[f246])).
fof(f246,axiom,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0'))))),
  file('thm.extreal.add_not_infty|split|1','thm.numeral.numeral_distrib|split|34')).
fof(f21994,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21066,f15405])).
fof(f21066,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK17)))) | $spl4),
  inference(unit_resulting_resolution,[],[f6240,f15405,f1420])).
fof(f6240,plain,(
  ( ! [X0] : (sP12('T_0',X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK17)))) )),
  inference(unit_resulting_resolution,[],[f952,f5893,f1432])).
fof(f21993,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21067,f15405])).
fof(f21067,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18)))) | $spl4),
  inference(unit_resulting_resolution,[],[f7152,f15405,f1420])).
fof(f7152,plain,(
  ( ! [X0] : (sP12('T_0',X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18)))) )),
  inference(unit_resulting_resolution,[],[f952,f5894,f1432])).
fof(f21992,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21068,f15405])).
fof(f21068,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2194,f15405,f1420])).
fof(f2194,plain,(
  ( ! [X0] : (sP12('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f952,f952,f1433])).
fof(f1433,plain,(
  ( ! [X2,X0,X1] : (sP12(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f932])).
fof(f21991,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21069,f15405])).
fof(f21069,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2210,f15405,f1420])).
fof(f2210,plain,(
  ( ! [X0] : (sP12('T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X0)) )),
  inference(unit_resulting_resolution,[],[f1570,f952,f1495])).
fof(f1570,plain,(
  ( ! [X0,X1] : (sP11(X0,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f954,f1429])).
fof(f1429,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f930])).
fof(f21990,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21070,f15405])).
fof(f21070,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2211,f15405,f1420])).
fof(f2211,plain,(
  ( ! [X0] : (sP12('T_0',sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))),X0)) )),
  inference(unit_resulting_resolution,[],[f952,f2134,f1495])).
fof(f2134,plain,(
  ( ! [X0,X1] : (sP11(X0,sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))),X1)) )),
  inference(unit_resulting_resolution,[],[f2108,f1429])).
fof(f21989,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21071,f15405])).
fof(f21071,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2401,f15405,f1420])).
fof(f2401,plain,(
  ( ! [X0] : (sP12('T_0',sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')),X0)) )),
  inference(unit_resulting_resolution,[],[f952,f2075,f1433])).
fof(f21988,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21072,f15405])).
fof(f21072,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2520,f15405,f1420])).
fof(f2520,plain,(
  ( ! [X0] : (sP12('T_0',sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')),X0)) )),
  inference(unit_resulting_resolution,[],[f952,f2114,f1433])).
fof(f21987,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21079,f1508])).
fof(f21079,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f2411,f15405,f1420])).
fof(f2411,plain,(
  ( ! [X0] : (sP12(sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')),X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f1567,f2075,f1495])).
fof(f1567,plain,(
  ( ! [X0,X1] : (sP11(X0,X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f1508,f1428])).
fof(f21986,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21080,f15405])).
fof(f21080,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2413,f15405,f1420])).
fof(f2413,plain,(
  ( ! [X0] : (sP12(sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')),'T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f1569,f2075,f1495])).
fof(f1569,plain,(
  ( ! [X0,X1] : (sP11(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f952,f1429])).
fof(f21985,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21081,f1508])).
fof(f21081,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f2530,f15405,f1420])).
fof(f2530,plain,(
  ( ! [X0] : (sP12(sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')),X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f1567,f2114,f1495])).
fof(f21984,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21082,f15405])).
fof(f21082,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2532,f15405,f1420])).
fof(f2532,plain,(
  ( ! [X0] : (sP12(sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')),'T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f1569,f2114,f1495])).
fof(f21983,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21096,f1508])).
fof(f21096,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f4679,f15405,f1420])).
fof(f4679,plain,(
  sP12('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f1571,f952,f1508,f1426])).
fof(f1426,plain,(
  ( ! [X2,X0,X1] : (sP12(X0,X1,X2) | ~sP11(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f927])).
fof(f1571,plain,(
  ( ! [X0,X1] : (sP11('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f1508,f1430])).
fof(f1430,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f930])).
fof(f21982,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21097,f1508])).
fof(f21097,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f6235,f15405,f1420])).
fof(f6235,plain,(
  sP12('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK17)),'T_0')),
  inference(unit_resulting_resolution,[],[f1508,f952,f5893,f1431])).
fof(f1431,plain,(
  ( ! [X2,X0,X1] : (sP12(X2,X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f932])).
fof(f21981,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21098,f1508])).
fof(f21098,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f7147,f15405,f1420])).
fof(f7147,plain,(
  sP12('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18)),'T_0')),
  inference(unit_resulting_resolution,[],[f1508,f952,f5894,f1431])).
fof(f21980,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21099,f15405])).
fof(f21099,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK17)))) | $spl4),
  inference(unit_resulting_resolution,[],[f6238,f15405,f1420])).
fof(f6238,plain,(
  sP12('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK17)),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f1508,f952,f5893,f1431])).
fof(f21979,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21100,f15405])).
fof(f21100,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18)))) | $spl4),
  inference(unit_resulting_resolution,[],[f7150,f15405,f1420])).
fof(f7150,plain,(
  sP12('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18)),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f1508,f952,f5894,f1431])).
fof(f21978,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21101,f15405])).
fof(f21101,plain,(
  p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f4680,f15405,f1420])).
fof(f4680,plain,(
  sP12('const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f1572,f952,f1508,f1426])).
fof(f1572,plain,(
  ( ! [X0,X1] : (sP11('const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f982,f1430])).
fof(f21977,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21102,f15405])).
fof(f21102,plain,(
  p(s(bool,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2411,f15405,f1420])).
fof(f21976,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21103,f15405])).
fof(f21103,plain,(
  p(s(bool,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2413,f15405,f1420])).
fof(f21975,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21104,f15405])).
fof(f21104,plain,(
  p(s(bool,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2530,f15405,f1420])).
fof(f21974,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21105,f15405])).
fof(f21105,plain,(
  p(s(bool,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2532,f15405,f1420])).
fof(f21973,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21203,f21244])).
fof(f21244,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f15405,f1430])).
fof(f21203,plain,(
  ~sP11('F_0','T_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f2185,f952,f15405,f1426])).
fof(f2185,plain,(
  ( ! [X0] : (~sP12('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f952,f1508,f1421])).
fof(f1421,plain,(
  ( ! [X2,X0,X1] : (~sP12(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f927])).
fof(f21972,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21205,f21244])).
fof(f21205,plain,(
  ~sP11('F_0',sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f2397,f952,f15405,f1426])).
fof(f2397,plain,(
  ( ! [X0] : (~sP12('F_0',sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')),X0)) )),
  inference(unit_resulting_resolution,[],[f1508,f2075,f1421])).
fof(f21971,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21207,f21244])).
fof(f21207,plain,(
  ~sP11('F_0',sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f2516,f952,f15405,f1426])).
fof(f2516,plain,(
  ( ! [X0] : (~sP12('F_0',sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')),X0)) )),
  inference(unit_resulting_resolution,[],[f1508,f2114,f1421])).
fof(f21970,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21212,f21244])).
fof(f21212,plain,(
  ~sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK17)),'T_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f6219,f952,f15405,f1426])).
fof(f6219,plain,(
  ( ! [X0] : (~sP12('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK17)),'T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f952,f5893,f1421])).
fof(f21969,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21214,f21244])).
fof(f21214,plain,(
  ~sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18)),'T_0','T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f7131,f952,f15405,f1426])).
fof(f7131,plain,(
  ( ! [X0] : (~sP12('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18)),'T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f952,f5894,f1421])).
fof(f21968,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21240,f21244])).
fof(f21240,plain,(
  ( ! [X0] : (~sP11('T_0',X0,'T_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f952,f952,f15405,f1427])).
fof(f1427,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f930])).
fof(f21967,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21265,f1508])).
fof(f21265,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f2185,f952,f15405,f1431])).
fof(f21966,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21266,f1508])).
fof(f21266,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f2397,f952,f15405,f1431])).
fof(f21965,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21267,f1508])).
fof(f21267,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f2516,f952,f15405,f1431])).
fof(f21964,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21271,f15405])).
fof(f21271,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK17)))) | $spl4),
  inference(unit_resulting_resolution,[],[f6219,f952,f15405,f1431])).
fof(f21963,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21272,f15405])).
fof(f21272,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18)))) | $spl4),
  inference(unit_resulting_resolution,[],[f7131,f952,f15405,f1431])).
fof(f21962,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21278,f15405])).
fof(f21278,plain,(
  p(s(bool,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2397,f952,f15405,f1431])).
fof(f21961,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21279,f15405])).
fof(f21279,plain,(
  p(s(bool,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2516,f952,f15405,f1431])).
fof(f21960,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21280,f1508])).
fof(f21280,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f4712,f952,f15405,f1431])).
fof(f4712,plain,(
  ~sP12('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f4698,f1422])).
fof(f1422,plain,(
  ( ! [X2,X0,X1] : (~sP12(X0,X1,X2) | sP11(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f927])).
fof(f4698,plain,(
  ~sP11('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f952,f952,f1508,f1427])).
fof(f21959,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21281,f15405])).
fof(f21281,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK17)))) | $spl4),
  inference(unit_resulting_resolution,[],[f6244,f952,f15405,f1431])).
fof(f6244,plain,(
  ~sP12('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK17)),'T_0')),
  inference(unit_resulting_resolution,[],[f952,f952,f5893,f1434])).
fof(f1434,plain,(
  ( ! [X2,X0,X1] : (~sP12(X2,X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f932])).
fof(f21958,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21282,f15405])).
fof(f21282,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18)))) | $spl4),
  inference(unit_resulting_resolution,[],[f7156,f952,f15405,f1431])).
fof(f7156,plain,(
  ~sP12('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18)),'T_0')),
  inference(unit_resulting_resolution,[],[f952,f952,f5894,f1434])).
fof(f21957,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21373,f21377])).
fof(f21377,plain,(
  ( ! [X2,X0,X1] : (sP16(X0,X1,X2)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f15405,f1461])).
fof(f1461,plain,(
  ( ! [X2,X0,X1] : (sP16(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f937])).
fof(f937,plain,(
  ! [X0,X1,X2] : ((~sP16(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP16(X0,X1,X2)))),
  inference(rectify,[],[f936])).
fof(f936,plain,(
  ! [X3,X1,X2] : ((~sP16(X3,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X3))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X3))) | sP16(X3,X1,X2)))),
  inference(flattening,[],[f935])).
fof(f935,plain,(
  ! [X3,X1,X2] : ((~sP16(X3,X1,X2) | (~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X3)))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X3))) | sP16(X3,X1,X2)))),
  inference(nnf_transformation,[],[f739])).
fof(f739,plain,(
  ! [X3,X1,X2] : (sP16(X3,X1,X2) <=> (~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X3))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP16])])).
fof(f21373,plain,(
  ( ! [X0] : (~sP16('T_0',X0,'T_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f952,f952,f15405,f1458])).
fof(f1458,plain,(
  ( ! [X2,X0,X1] : (~sP16(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f937])).
fof(f21956,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21407,f21411])).
fof(f21411,plain,(
  ( ! [X2,X0,X1] : (sP15(X0,X1,X2)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f15405,f1465])).
fof(f1465,plain,(
  ( ! [X2,X0,X1] : (sP15(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f940])).
fof(f940,plain,(
  ! [X0,X1,X2] : ((~sP15(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP15(X0,X1,X2)))),
  inference(rectify,[],[f939])).
fof(f939,plain,(
  ! [X0,X1,X3] : ((~sP15(X0,X1,X3) | p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X3)) & p(s(bool,X1)) & p(s(bool,X0))) | sP15(X0,X1,X3)))),
  inference(flattening,[],[f938])).
fof(f938,plain,(
  ! [X0,X1,X3] : ((~sP15(X0,X1,X3) | (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X3)) & p(s(bool,X1)) & p(s(bool,X0))) | sP15(X0,X1,X3)))),
  inference(nnf_transformation,[],[f738])).
fof(f738,plain,(
  ! [X0,X1,X3] : (sP15(X0,X1,X3) <=> (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP15])])).
fof(f21407,plain,(
  ( ! [X0] : (~sP15('T_0','T_0',X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f952,f952,f15405,f1462])).
fof(f1462,plain,(
  ( ! [X2,X0,X1] : (~sP15(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f940])).
fof(f21955,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21439,f21443])).
fof(f21443,plain,(
  ( ! [X2,X0,X1] : (sP14(X0,X1,X2)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f15405,f1469])).
fof(f1469,plain,(
  ( ! [X2,X0,X1] : (sP14(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f943])).
fof(f943,plain,(
  ! [X0,X1,X2] : ((~sP14(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP14(X0,X1,X2)))),
  inference(rectify,[],[f942])).
fof(f942,plain,(
  ! [X2,X1,X3] : ((~sP14(X2,X1,X3) | p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2))) & ((~p(s(bool,X3)) & p(s(bool,X1)) & p(s(bool,X2))) | sP14(X2,X1,X3)))),
  inference(flattening,[],[f941])).
fof(f941,plain,(
  ! [X2,X1,X3] : ((~sP14(X2,X1,X3) | (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))) & ((~p(s(bool,X3)) & p(s(bool,X1)) & p(s(bool,X2))) | sP14(X2,X1,X3)))),
  inference(nnf_transformation,[],[f737])).
fof(f737,plain,(
  ! [X2,X1,X3] : (sP14(X2,X1,X3) <=> (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP14])])).
fof(f21439,plain,(
  ( ! [X0] : (~sP14('T_0','T_0',X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f952,f952,f15405,f1466])).
fof(f1466,plain,(
  ( ! [X2,X0,X1] : (~sP14(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f943])).
fof(f21954,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21446,f15405])).
fof(f21446,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f1617,f952,f15405,f1470])).
fof(f1470,plain,(
  ( ! [X2,X0,X1] : (~sP13(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f946])).
fof(f946,plain,(
  ! [X0,X1,X2] : ((~sP13(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP13(X0,X1,X2)))),
  inference(rectify,[],[f945])).
fof(f945,plain,(
  ! [X0,X2,X3] : ((~sP13(X0,X2,X3) | p(s(bool,X3)) | p(s(bool,X2)) | ~p(s(bool,X0))) & ((~p(s(bool,X3)) & ~p(s(bool,X2)) & p(s(bool,X0))) | sP13(X0,X2,X3)))),
  inference(flattening,[],[f944])).
fof(f944,plain,(
  ! [X0,X2,X3] : ((~sP13(X0,X2,X3) | (p(s(bool,X3)) | p(s(bool,X2)) | ~p(s(bool,X0)))) & ((~p(s(bool,X3)) & ~p(s(bool,X2)) & p(s(bool,X0))) | sP13(X0,X2,X3)))),
  inference(nnf_transformation,[],[f736])).
fof(f736,plain,(
  ! [X0,X2,X3] : (sP13(X0,X2,X3) <=> (p(s(bool,X3)) | p(s(bool,X2)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP13])])).
fof(f1617,plain,(
  ( ! [X0,X1] : (sP13(X0,X1,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) )),
  inference(unit_resulting_resolution,[],[f962,f1471])).
fof(f1471,plain,(
  ( ! [X2,X0,X1] : (sP13(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f946])).
fof(f21953,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21447,f15405])).
fof(f21447,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f1592,f952,f15405,f1470])).
fof(f1592,plain,(
  ( ! [X0,X1] : (sP13(X0,X1,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f954,f1471])).
fof(f21952,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21448,f15405])).
fof(f21448,plain,(
  p(s(bool,sK77(sK63('const.num.|48|_0','type.num.num')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2055,f952,f15405,f1470])).
fof(f2055,plain,(
  ( ! [X0,X1] : (sP13(X0,X1,sK77(sK63('const.num.|48|_0','type.num.num')))) )),
  inference(unit_resulting_resolution,[],[f2030,f1471])).
fof(f21951,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21449,f15405])).
fof(f21449,plain,(
  p(s(bool,sK78(sK63('const.num.|48|_0','type.num.num')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2102,f952,f15405,f1470])).
fof(f2102,plain,(
  ( ! [X0,X1] : (sP13(X0,X1,sK78(sK63('const.num.|48|_0','type.num.num')))) )),
  inference(unit_resulting_resolution,[],[f2068,f1471])).
fof(f21950,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21450,f15405])).
fof(f21450,plain,(
  p(s(bool,sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))))) | $spl4),
  inference(unit_resulting_resolution,[],[f2138,f952,f15405,f1470])).
fof(f2138,plain,(
  ( ! [X0,X1] : (sP13(X0,X1,sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))))) )),
  inference(unit_resulting_resolution,[],[f2108,f1471])).
fof(f21949,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21451,f15405])).
fof(f21451,plain,(
  p(s(bool,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2409,f952,f15405,f1470])).
fof(f2409,plain,(
  ( ! [X0,X1] : (sP13(X0,X1,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) )),
  inference(unit_resulting_resolution,[],[f2075,f1471])).
fof(f21948,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21452,f15405])).
fof(f21452,plain,(
  p(s(bool,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2528,f952,f15405,f1470])).
fof(f2528,plain,(
  ( ! [X0,X1] : (sP13(X0,X1,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) )),
  inference(unit_resulting_resolution,[],[f2114,f1471])).
fof(f21947,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21453,f15405])).
fof(f21453,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1593,f952,f15405,f1470])).
fof(f1593,plain,(
  ( ! [X0,X1] : (sP13(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f952,f1472])).
fof(f1472,plain,(
  ( ! [X2,X0,X1] : (sP13(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f946])).
fof(f21946,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21455,f15405])).
fof(f21455,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1618,f952,f15405,f1470])).
fof(f1618,plain,(
  ( ! [X0,X1] : (sP13(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f962,f1472])).
fof(f21945,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21456,f15405])).
fof(f21456,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1594,f952,f15405,f1470])).
fof(f1594,plain,(
  ( ! [X0,X1] : (sP13(X0,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f954,f1472])).
fof(f21944,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21457,f15405])).
fof(f21457,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2056,f952,f15405,f1470])).
fof(f2056,plain,(
  ( ! [X0,X1] : (sP13(X0,sK77(sK63('const.num.|48|_0','type.num.num')),X1)) )),
  inference(unit_resulting_resolution,[],[f2030,f1472])).
fof(f21943,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21458,f15405])).
fof(f21458,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2103,f952,f15405,f1470])).
fof(f2103,plain,(
  ( ! [X0,X1] : (sP13(X0,sK78(sK63('const.num.|48|_0','type.num.num')),X1)) )),
  inference(unit_resulting_resolution,[],[f2068,f1472])).
fof(f21942,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21459,f15405])).
fof(f21459,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2139,f952,f15405,f1470])).
fof(f2139,plain,(
  ( ! [X0,X1] : (sP13(X0,sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))),X1)) )),
  inference(unit_resulting_resolution,[],[f2108,f1472])).
fof(f21941,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21460,f15405])).
fof(f21460,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2410,f952,f15405,f1470])).
fof(f2410,plain,(
  ( ! [X0,X1] : (sP13(X0,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')),X1)) )),
  inference(unit_resulting_resolution,[],[f2075,f1472])).
fof(f21940,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21461,f15405])).
fof(f21461,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2529,f952,f15405,f1470])).
fof(f2529,plain,(
  ( ! [X0,X1] : (sP13(X0,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')),X1)) )),
  inference(unit_resulting_resolution,[],[f2114,f1472])).
fof(f21939,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21475,f21555])).
fof(f21555,plain,(
  ( ! [X2,X0,X1] : (sP13(X0,X1,X2)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f15405,f1473])).
fof(f1473,plain,(
  ( ! [X2,X0,X1] : (sP13(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f946])).
fof(f21475,plain,(
  ( ! [X0] : (~sP13('T_0',X0,'F_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f952,f1508,f15405,f1470])).
fof(f21938,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21498,f21555])).
fof(f21498,plain,(
  ( ! [X0,X1] : (~sP13('T_0',X0,X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f952,f15405,f15405,f1470])).
fof(f21937,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21500,f15405])).
fof(f21500,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1591,f952,f15405,f1470])).
fof(f1591,plain,(
  ( ! [X0,X1] : (sP13(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f952,f1471])).
fof(f21936,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21501,f15405])).
fof(f21501,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1617,f952,f15405,f1470])).
fof(f21935,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21502,f15405])).
fof(f21502,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1592,f952,f15405,f1470])).
fof(f21934,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21503,f15405])).
fof(f21503,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2055,f952,f15405,f1470])).
fof(f21933,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21504,f15405])).
fof(f21504,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2102,f952,f15405,f1470])).
fof(f21932,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21505,f15405])).
fof(f21505,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2138,f952,f15405,f1470])).
fof(f21931,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21506,f15405])).
fof(f21506,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2409,f952,f15405,f1470])).
fof(f21930,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21507,f15405])).
fof(f21507,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2528,f952,f15405,f1470])).
fof(f21929,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21509,f1508])).
fof(f21509,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(unit_resulting_resolution,[],[f11963,f952,f15405,f1470])).
fof(f11963,plain,(
  ( ! [X0] : (sP13(X0,'F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f1164,f1474])).
fof(f1474,plain,(
  ( ! [X2,X0,X3,X1] : (sP13(X0,X2,X3) | s(bool,X3) != s(bool,'const.bool.COND_3'(s(bool,X2),s(bool,X1),s(bool,X0)))) )),
  inference(cnf_transformation,[],[f948])).
fof(f948,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X3) != s(bool,'const.bool.COND_3'(s(bool,X2),s(bool,X1),s(bool,X0))) | (sP13(X0,X2,X3) & sP14(X2,X1,X3) & sP15(X0,X1,X3) & sP16(X3,X1,X2) & (p(s(bool,X2)) | p(s(bool,X0)) | ~p(s(bool,X3))))) & (~sP13(X0,X2,X3) | ~sP14(X2,X1,X3) | ~sP15(X0,X1,X3) | ~sP16(X3,X1,X2) | (~p(s(bool,X2)) & ~p(s(bool,X0)) & p(s(bool,X3))) | s(bool,X3) = s(bool,'const.bool.COND_3'(s(bool,X2),s(bool,X1),s(bool,X0)))))),
  inference(flattening,[],[f947])).
fof(f947,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X3) != s(bool,'const.bool.COND_3'(s(bool,X2),s(bool,X1),s(bool,X0))) | (sP13(X0,X2,X3) & sP14(X2,X1,X3) & sP15(X0,X1,X3) & sP16(X3,X1,X2) & (p(s(bool,X2)) | p(s(bool,X0)) | ~p(s(bool,X3))))) & ((~sP13(X0,X2,X3) | ~sP14(X2,X1,X3) | ~sP15(X0,X1,X3) | ~sP16(X3,X1,X2) | (~p(s(bool,X2)) & ~p(s(bool,X0)) & p(s(bool,X3)))) | s(bool,X3) = s(bool,'const.bool.COND_3'(s(bool,X2),s(bool,X1),s(bool,X0)))))),
  inference(nnf_transformation,[],[f740])).
fof(f740,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X3) = s(bool,'const.bool.COND_3'(s(bool,X2),s(bool,X1),s(bool,X0))) <=> (sP13(X0,X2,X3) & sP14(X2,X1,X3) & sP15(X0,X1,X3) & sP16(X3,X1,X2) & (p(s(bool,X2)) | p(s(bool,X0)) | ~p(s(bool,X3)))))),
  inference(definition_folding,[],[f586,f739,f738,f737,f736])).
fof(f586,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X3) = s(bool,'const.bool.COND_3'(s(bool,X2),s(bool,X1),s(bool,X0))) <=> ((p(s(bool,X3)) | p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X3))) & (p(s(bool,X2)) | p(s(bool,X0)) | ~p(s(bool,X3)))))),
  inference(flattening,[],[f585])).
fof(f585,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X3) = s(bool,'const.bool.COND_3'(s(bool,X2),s(bool,X1),s(bool,X0))) <=> ((p(s(bool,X3)) | p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X3))) & (p(s(bool,X2)) | p(s(bool,X0)) | ~p(s(bool,X3)))))),
  inference(rectify,[],[f228])).
fof(f228,axiom,(
  ! [X73,X35,X43,X44] : (s(bool,X44) = s(bool,'const.bool.COND_3'(s(bool,X43),s(bool,X35),s(bool,X73))) <=> ((p(s(bool,X44)) | p(s(bool,X43)) | ~p(s(bool,X73))) & (p(s(bool,X44)) | ~p(s(bool,X35)) | ~p(s(bool,X43))) & (p(s(bool,X44)) | ~p(s(bool,X35)) | ~p(s(bool,X73))) & (~p(s(bool,X43)) | p(s(bool,X35)) | ~p(s(bool,X44))) & (p(s(bool,X43)) | p(s(bool,X73)) | ~p(s(bool,X44)))))),
  file('thm.extreal.add_not_infty|split|1','thm.sat.sat_3|49|_prove_ltb')).
fof(f21928,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21510,f15405])).
fof(f21510,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f1618,f952,f15405,f1470])).
fof(f21927,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21511,f15405])).
fof(f21511,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f1594,f952,f15405,f1470])).
fof(f21926,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21512,f15405])).
fof(f21512,plain,(
  p(s(bool,sK77(sK63('const.num.|48|_0','type.num.num')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2056,f952,f15405,f1470])).
fof(f21925,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21513,f15405])).
fof(f21513,plain,(
  p(s(bool,sK78(sK63('const.num.|48|_0','type.num.num')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2103,f952,f15405,f1470])).
fof(f21924,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21514,f15405])).
fof(f21514,plain,(
  p(s(bool,sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))))) | $spl4),
  inference(unit_resulting_resolution,[],[f2139,f952,f15405,f1470])).
fof(f21923,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21515,f15405])).
fof(f21515,plain,(
  p(s(bool,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2410,f952,f15405,f1470])).
fof(f21922,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21516,f15405])).
fof(f21516,plain,(
  p(s(bool,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2529,f952,f15405,f1470])).
fof(f21921,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21530,f21555])).
fof(f21530,plain,(
  ( ! [X0] : (~sP13('T_0','F_0',X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f952,f1508,f15405,f1470])).
fof(f21920,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21553,f21555])).
fof(f21553,plain,(
  ( ! [X0,X1] : (~sP13('T_0',X0,X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f952,f15405,f15405,f1470])).
fof(f21919,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21653,f1555])).
fof(f21653,plain,(
  ( ! [X0] : (~sP7(X0,'F_0','T_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2146,f15405,f1497])).
fof(f21918,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21655,f1555])).
fof(f21655,plain,(
  ( ! [X0] : (~sP7(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK17)),'T_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f6077,f15405,f1497])).
fof(f21917,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21657,f1555])).
fof(f21657,plain,(
  ( ! [X0] : (~sP7(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18)),'T_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f6989,f15405,f1497])).
fof(f21916,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21659,f1555])).
fof(f21659,plain,(
  ( ! [X0] : (~sP7('F_0',X0,'T_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2149,f15405,f1497])).
fof(f21915,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21661,f1555])).
fof(f21661,plain,(
  ( ! [X0] : (~sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK17)),X0,'T_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f6080,f15405,f1497])).
fof(f21914,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21663,f1555])).
fof(f21663,plain,(
  ( ! [X0] : (~sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18)),X0,'T_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f6992,f15405,f1497])).
fof(f21913,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21668,f2378])).
fof(f2378,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) )),
  inference(unit_resulting_resolution,[],[f2075,f1394])).
fof(f21668,plain,(
  ( ! [X0] : (~sP7(X0,'F_0',sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2376,f15405,f1497])).
fof(f21912,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21670,f2378])).
fof(f21670,plain,(
  ( ! [X0] : (~sP7('F_0',X0,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2377,f15405,f1497])).
fof(f21911,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21672,f2497])).
fof(f2497,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) )),
  inference(unit_resulting_resolution,[],[f2114,f1394])).
fof(f21672,plain,(
  ( ! [X0] : (~sP7(X0,'F_0',sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2495,f15405,f1497])).
fof(f21910,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21674,f2497])).
fof(f21674,plain,(
  ( ! [X0] : (~sP7('F_0',X0,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2496,f15405,f1497])).
fof(f21909,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21677,f15405])).
fof(f21677,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f1603,f15405,f1526])).
fof(f1526,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f1308,f1508])).
fof(f1308,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f883])).
fof(f883,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP1(X0,X1)))),
  inference(flattening,[],[f882])).
fof(f882,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP1(X0,X1)))),
  inference(nnf_transformation,[],[f718])).
fof(f718,plain,(
  ! [X0,X1] : (sP1(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f1603,plain,(
  ( ! [X0] : (sP1(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) )),
  inference(unit_resulting_resolution,[],[f962,f1309])).
fof(f1309,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f883])).
fof(f21908,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21678,f15405])).
fof(f21678,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f7725,f15405,f1526])).
fof(f7725,plain,(
  ( ! [X0] : (sP1(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal','const.extreal.PosInf_0')))) )),
  inference(unit_resulting_resolution,[],[f7410,f1309])).
fof(f21907,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21679,f15405])).
fof(f21679,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK18),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f8891,f15405,f1526])).
fof(f8891,plain,(
  ( ! [X0] : (sP1(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK18),s('type.extreal.extreal','const.extreal.PosInf_0')))) )),
  inference(unit_resulting_resolution,[],[f7411,f1309])).
fof(f21906,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21680,f15405])).
fof(f21680,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.PosInf_0'))))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1621,f15405,f1526])).
fof(f1621,plain,(
  ( ! [X0,X1] : (sP1(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal','const.extreal.PosInf_0')))) )),
  inference(unit_resulting_resolution,[],[f983,f1309])).
fof(f21905,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21681,f15405])).
fof(f21681,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0))))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1663,f15405,f1526])).
fof(f1663,plain,(
  ( ! [X0,X1] : (sP1(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X1)))) )),
  inference(unit_resulting_resolution,[],[f984,f1309])).
fof(f984,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0))))) )),
  inference(cnf_transformation,[],[f271])).
fof(f271,plain,(
  ! [X0] : p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0))))),
  inference(rectify,[],[f143])).
fof(f143,axiom,(
  ! [X6] : p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X6))))),
  file('thm.extreal.add_not_infty|split|1','thm.extreal.le_refl')).
fof(f21904,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21682,f15405])).
fof(f21682,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0))))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1705,f15405,f1526])).
fof(f1705,plain,(
  ( ! [X0,X1] : (sP1(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X1)))) )),
  inference(unit_resulting_resolution,[],[f985,f1309])).
fof(f985,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0))))) )),
  inference(cnf_transformation,[],[f272])).
fof(f272,plain,(
  ! [X0] : p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0))))),
  inference(rectify,[],[f138])).
fof(f138,axiom,(
  ! [X6] : p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X6))))),
  file('thm.extreal.add_not_infty|split|1','thm.extreal.le_infty|split|0')).
fof(f21903,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21683,f15405])).
fof(f21683,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f1534,f15405,f1526])).
fof(f1534,plain,(
  ( ! [X0] : (sP1(X0,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f954,f1309])).
fof(f21902,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21684,f15405])).
fof(f21684,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.DATATYPE_1'(s(X0,X1))))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1814,f15405,f1526])).
fof(f1814,plain,(
  ( ! [X4,X5,X3] : (sP1(X5,'const.bool.DATATYPE_1'(s(X3,X4)))) )),
  inference(subsumption_resolution,[],[f1812,f952])).
fof(f1812,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'T_0')) | sP1(X5,'const.bool.DATATYPE_1'(s(X3,X4)))) )),
  inference(superposition,[],[f1309,f1149])).
fof(f1149,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.DATATYPE_1'(s(X0,X1)))) )),
  inference(cnf_transformation,[],[f387])).
fof(f387,plain,(
  ! [X0,X1] : s(bool,'T_0') = s(bool,'const.bool.DATATYPE_1'(s(X0,X1)))),
  inference(flattening,[],[f386])).
fof(f386,plain,(
  ! [X0] : ! [X1] : s(bool,'T_0') = s(bool,'const.bool.DATATYPE_1'(s(X0,X1)))),
  inference(rectify,[],[f164])).
fof(f164,axiom,(
  ! [X8] : ! [X6] : s(bool,'T_0') = s(bool,'const.bool.DATATYPE_1'(s(X8,X6)))),
  file('thm.extreal.add_not_infty|split|1','thm.bool.DATATYPE_TAG_THM')).
fof(f21901,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21685,f15405])).
fof(f21685,plain,(
  p(s(bool,sK63('T_0',bool))) | $spl4),
  inference(unit_resulting_resolution,[],[f1601,f15405,f1526])).
fof(f1601,plain,(
  ( ! [X0] : (sP1(X0,sK63('T_0',bool))) )),
  inference(unit_resulting_resolution,[],[f952,f1598])).
fof(f1598,plain,(
  ( ! [X2,X3] : (sP1(X3,sK63(X2,bool)) | ~p(s(bool,X2))) )),
  inference(superposition,[],[f1309,f1163])).
fof(f21900,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21686,f15405])).
fof(f21686,plain,(
  p(s(bool,sK63('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),bool))) | $spl4),
  inference(unit_resulting_resolution,[],[f1602,f15405,f1526])).
fof(f1602,plain,(
  ( ! [X0] : (sP1(X0,sK63('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),bool))) )),
  inference(unit_resulting_resolution,[],[f954,f1598])).
fof(f21899,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21687,f15405])).
fof(f21687,plain,(
  p(s(bool,sK63(sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))),bool))) | $spl4),
  inference(unit_resulting_resolution,[],[f2141,f15405,f1526])).
fof(f2141,plain,(
  ( ! [X0] : (sP1(X0,sK63(sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))),bool))) )),
  inference(unit_resulting_resolution,[],[f2108,f1598])).
fof(f21898,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21688,f15405])).
fof(f21688,plain,(
  p(s(bool,sK77(sK63('const.num.|48|_0','type.num.num')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2035,f15405,f1526])).
fof(f2035,plain,(
  ( ! [X0] : (sP1(X0,sK77(sK63('const.num.|48|_0','type.num.num')))) )),
  inference(unit_resulting_resolution,[],[f2030,f1309])).
fof(f21897,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21689,f15405])).
fof(f21689,plain,(
  p(s(bool,sK78(sK63('const.num.|48|_0','type.num.num')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2082,f15405,f1526])).
fof(f2082,plain,(
  ( ! [X0] : (sP1(X0,sK78(sK63('const.num.|48|_0','type.num.num')))) )),
  inference(unit_resulting_resolution,[],[f2068,f1309])).
fof(f21896,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21690,f15405])).
fof(f21690,plain,(
  p(s(bool,sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))))) | $spl4),
  inference(unit_resulting_resolution,[],[f2118,f15405,f1526])).
fof(f2118,plain,(
  ( ! [X0] : (sP1(X0,sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))))) )),
  inference(unit_resulting_resolution,[],[f2108,f1309])).
fof(f21895,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21691,f15405])).
fof(f21691,plain,(
  p(s(bool,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2363,f15405,f1526])).
fof(f2363,plain,(
  ( ! [X0] : (sP1(X0,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) )),
  inference(unit_resulting_resolution,[],[f2075,f1309])).
fof(f21894,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21692,f15405])).
fof(f21692,plain,(
  p(s(bool,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2482,f15405,f1526])).
fof(f2482,plain,(
  ( ! [X0] : (sP1(X0,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) )),
  inference(unit_resulting_resolution,[],[f2114,f1309])).
fof(f21893,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21693,f15405])).
fof(f21693,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1535,f15405,f1526])).
fof(f1535,plain,(
  ( ! [X0] : (sP1('T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f952,f1310])).
fof(f1310,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f883])).
fof(f21892,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21694,f15405])).
fof(f21694,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1604,f15405,f1526])).
fof(f1604,plain,(
  ( ! [X0] : (sP1('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0)) )),
  inference(unit_resulting_resolution,[],[f962,f1310])).
fof(f21891,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21695,f15405])).
fof(f21695,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f7726,f15405,f1526])).
fof(f7726,plain,(
  ( ! [X0] : (sP1('const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal','const.extreal.PosInf_0')),X0)) )),
  inference(unit_resulting_resolution,[],[f7410,f1310])).
fof(f21890,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21696,f15405])).
fof(f21696,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f8892,f15405,f1526])).
fof(f8892,plain,(
  ( ! [X0] : (sP1('const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK18),s('type.extreal.extreal','const.extreal.PosInf_0')),X0)) )),
  inference(unit_resulting_resolution,[],[f7411,f1310])).
fof(f21889,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21697,f15405])).
fof(f21697,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1622,f15405,f1526])).
fof(f1622,plain,(
  ( ! [X0,X1] : (sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.PosInf_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f983,f1310])).
fof(f21888,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21698,f15405])).
fof(f21698,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1664,f15405,f1526])).
fof(f1664,plain,(
  ( ! [X0,X1] : (sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),X1)) )),
  inference(unit_resulting_resolution,[],[f984,f1310])).
fof(f21887,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21699,f15405])).
fof(f21699,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1706,f15405,f1526])).
fof(f1706,plain,(
  ( ! [X0,X1] : (sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)),X1)) )),
  inference(unit_resulting_resolution,[],[f985,f1310])).
fof(f21886,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21700,f15405])).
fof(f21700,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1536,f15405,f1526])).
fof(f1536,plain,(
  ( ! [X0] : (sP1('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X0)) )),
  inference(unit_resulting_resolution,[],[f954,f1310])).
fof(f21885,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21701,f15405])).
fof(f21701,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1813,f15405,f1526])).
fof(f1813,plain,(
  ( ! [X2,X0,X1] : (sP1('const.bool.DATATYPE_1'(s(X0,X1)),X2)) )),
  inference(subsumption_resolution,[],[f1811,f952])).
fof(f1811,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | sP1('const.bool.DATATYPE_1'(s(X0,X1)),X2)) )),
  inference(superposition,[],[f1310,f1149])).
fof(f21884,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21702,f15405])).
fof(f21702,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1599,f15405,f1526])).
fof(f1599,plain,(
  ( ! [X0] : (sP1(sK63('T_0',bool),X0)) )),
  inference(unit_resulting_resolution,[],[f952,f1597])).
fof(f1597,plain,(
  ( ! [X0,X1] : (sP1(sK63(X0,bool),X1) | ~p(s(bool,X0))) )),
  inference(superposition,[],[f1310,f1163])).
fof(f21883,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21703,f15405])).
fof(f21703,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1600,f15405,f1526])).
fof(f1600,plain,(
  ( ! [X0] : (sP1(sK63('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),bool),X0)) )),
  inference(unit_resulting_resolution,[],[f954,f1597])).
fof(f21882,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21704,f15405])).
fof(f21704,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2140,f15405,f1526])).
fof(f2140,plain,(
  ( ! [X0] : (sP1(sK63(sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))),bool),X0)) )),
  inference(unit_resulting_resolution,[],[f2108,f1597])).
fof(f21881,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21705,f15405])).
fof(f21705,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2036,f15405,f1526])).
fof(f2036,plain,(
  ( ! [X0] : (sP1(sK77(sK63('const.num.|48|_0','type.num.num')),X0)) )),
  inference(unit_resulting_resolution,[],[f2030,f1310])).
fof(f21880,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21706,f15405])).
fof(f21706,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2083,f15405,f1526])).
fof(f2083,plain,(
  ( ! [X0] : (sP1(sK78(sK63('const.num.|48|_0','type.num.num')),X0)) )),
  inference(unit_resulting_resolution,[],[f2068,f1310])).
fof(f21879,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21707,f15405])).
fof(f21707,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2119,f15405,f1526])).
fof(f2119,plain,(
  ( ! [X0] : (sP1(sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))),X0)) )),
  inference(unit_resulting_resolution,[],[f2108,f1310])).
fof(f21878,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21708,f15405])).
fof(f21708,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2364,f15405,f1526])).
fof(f2364,plain,(
  ( ! [X0] : (sP1(sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')),X0)) )),
  inference(unit_resulting_resolution,[],[f2075,f1310])).
fof(f21877,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21709,f15405])).
fof(f21709,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2483,f15405,f1526])).
fof(f2483,plain,(
  ( ! [X0] : (sP1(sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')),X0)) )),
  inference(unit_resulting_resolution,[],[f2114,f1310])).
fof(f21876,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21714,f15405])).
fof(f21714,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1533,f15405,f1526])).
fof(f1533,plain,(
  ( ! [X0] : (sP1(X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f952,f1309])).
fof(f21875,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21715,f15405])).
fof(f21715,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1603,f15405,f1526])).
fof(f21874,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21716,f15405])).
fof(f21716,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f7725,f15405,f1526])).
fof(f21873,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21717,f15405])).
fof(f21717,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f8891,f15405,f1526])).
fof(f21872,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21718,f15405])).
fof(f21718,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1621,f15405,f1526])).
fof(f21871,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21719,f15405])).
fof(f21719,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1663,f15405,f1526])).
fof(f21870,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21720,f15405])).
fof(f21720,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1705,f15405,f1526])).
fof(f21869,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21721,f15405])).
fof(f21721,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1534,f15405,f1526])).
fof(f21868,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21722,f15405])).
fof(f21722,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1814,f15405,f1526])).
fof(f21867,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21723,f15405])).
fof(f21723,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1601,f15405,f1526])).
fof(f21866,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21724,f15405])).
fof(f21724,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1602,f15405,f1526])).
fof(f21865,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21725,f15405])).
fof(f21725,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2141,f15405,f1526])).
fof(f21864,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21726,f15405])).
fof(f21726,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2035,f15405,f1526])).
fof(f21863,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21727,f15405])).
fof(f21727,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2082,f15405,f1526])).
fof(f21862,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21728,f15405])).
fof(f21728,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2118,f15405,f1526])).
fof(f21861,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21729,f15405])).
fof(f21729,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2363,f15405,f1526])).
fof(f21860,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21730,f15405])).
fof(f21730,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f2482,f15405,f1526])).
fof(f21859,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21732,f15405])).
fof(f21732,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f1604,f15405,f1526])).
fof(f21858,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21733,f15405])).
fof(f21733,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f7726,f15405,f1526])).
fof(f21857,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21734,f15405])).
fof(f21734,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK18),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f8892,f15405,f1526])).
fof(f21856,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21735,f15405])).
fof(f21735,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.PosInf_0'))))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1622,f15405,f1526])).
fof(f21855,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21736,f15405])).
fof(f21736,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0))))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1664,f15405,f1526])).
fof(f21854,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21737,f15405])).
fof(f21737,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0))))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1706,f15405,f1526])).
fof(f21853,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21738,f15405])).
fof(f21738,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl4),
  inference(unit_resulting_resolution,[],[f1536,f15405,f1526])).
fof(f21852,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21739,f15405])).
fof(f21739,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.DATATYPE_1'(s(X0,X1))))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f1813,f15405,f1526])).
fof(f21851,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21740,f15405])).
fof(f21740,plain,(
  p(s(bool,sK63('T_0',bool))) | $spl4),
  inference(unit_resulting_resolution,[],[f1599,f15405,f1526])).
fof(f21850,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21741,f15405])).
fof(f21741,plain,(
  p(s(bool,sK63('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),bool))) | $spl4),
  inference(unit_resulting_resolution,[],[f1600,f15405,f1526])).
fof(f21849,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21742,f15405])).
fof(f21742,plain,(
  p(s(bool,sK63(sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))),bool))) | $spl4),
  inference(unit_resulting_resolution,[],[f2140,f15405,f1526])).
fof(f21848,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21743,f15405])).
fof(f21743,plain,(
  p(s(bool,sK77(sK63('const.num.|48|_0','type.num.num')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2036,f15405,f1526])).
fof(f21847,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21744,f15405])).
fof(f21744,plain,(
  p(s(bool,sK78(sK63('const.num.|48|_0','type.num.num')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2083,f15405,f1526])).
fof(f21846,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21745,f15405])).
fof(f21745,plain,(
  p(s(bool,sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))))) | $spl4),
  inference(unit_resulting_resolution,[],[f2119,f15405,f1526])).
fof(f21845,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21746,f15405])).
fof(f21746,plain,(
  p(s(bool,sK83(sK63('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2364,f15405,f1526])).
fof(f21844,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21747,f15405])).
fof(f21747,plain,(
  p(s(bool,sK83(sK63('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl4),
  inference(unit_resulting_resolution,[],[f2483,f15405,f1526])).
fof(f21829,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21757,f21787])).
fof(f21787,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f15405,f1231])).
fof(f21757,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f15405,f1134])).
fof(f1134,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0)) )),
  inference(cnf_transformation,[],[f806])).
fof(f806,plain,(
  ! [X0] : ((s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0))))) & (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)))) | s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0)))),
  inference(nnf_transformation,[],[f363])).
fof(f363,plain,(
  ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) <=> p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)))))),
  inference(flattening,[],[f362])).
fof(f362,plain,(
  ! [X0] : (~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) <=> p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)))))),
  inference(rectify,[],[f145])).
fof(f145,axiom,(
  ! [X6] : (~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X6) <=> p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X6)))))),
  file('thm.extreal.add_not_infty|split|1','thm.extreal.lt_infty|split|6')).
fof(f21828,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21759,f21785])).
fof(f21785,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f15405,f1232])).
fof(f21759,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f15405,f1136])).
fof(f21827,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21761,f15405])).
fof(f21761,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f15405,f1276])).
fof(f1276,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))) )),
  inference(cnf_transformation,[],[f863])).
fof(f863,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0))))) & (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))))),
  inference(nnf_transformation,[],[f498])).
fof(f498,plain,(
  ! [X0,X1] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) <=> ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))))),
  inference(flattening,[],[f497])).
fof(f497,plain,(
  ! [X0,X1] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) <=> ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))))),
  inference(rectify,[],[f86])).
fof(f86,axiom,(
  ! [X6,X7] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X7)))) <=> ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X7),s('type.extreal.extreal',X6)))))),
  file('thm.extreal.add_not_infty|split|1','thm.extreal.extreal_lt_def')).
fof(f21826,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21765,f15405])).
fof(f21765,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f15405,f1276])).
fof(f21825,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21769,f15405])).
fof(f21769,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1))))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f15405,f1274])).
fof(f1274,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X0))))) )),
  inference(cnf_transformation,[],[f862])).
fof(f862,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X0))))) & (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X0)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1))))))),
  inference(nnf_transformation,[],[f496])).
fof(f496,plain,(
  ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) <=> ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X0)))))),
  inference(flattening,[],[f495])).
fof(f495,plain,(
  ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) <=> ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X0)))))),
  inference(rectify,[],[f227])).
fof(f227,axiom,(
  ! [X6,X7] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X7)))) <=> ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X7),s('type.realax.real',X6)))))),
  file('thm.extreal.add_not_infty|split|1','thm.real.real_lte')).
fof(f21824,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f21771,f15405])).
fof(f21771,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1))))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f15405,f1274])).
fof(f21788,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f964,f15405,f1231])).
fof(f21786,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f989,f15405,f1232])).
fof(f989,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X0)))) )),
  inference(cnf_transformation,[],[f276])).
fof(f276,plain,(
  ! [X0] : s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X0)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X12] : s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X12)))),
  file('thm.extreal.add_not_infty|split|1','thm.extreal.extreal_add_def|split|1')).
fof(f21784,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1015,f15405,f1228])).
fof(f1228,plain,(
  ( ! [X0] : (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,sK82(X0)))) )),
  inference(cnf_transformation,[],[f858])).
fof(f858,plain,(
  ! [X0] : ((~p(s(bool,sK82(X0))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,sK82(X0)))) & s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.NegInf_0'))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,sK82(X0)),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK82])],[f857])).
fof(f857,plain,(
  ! [X0] : ? [X1] : ((~p(s(bool,X1)) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,X1))) & s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.NegInf_0'))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X1),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))))))),
  inference(flattening,[],[f856])).
fof(f856,plain,(
  ! [X0] : ? [X1] : (((~p(s(bool,X1)) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,X1)))) & s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.NegInf_0'))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X1),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))))))),
  inference(nnf_transformation,[],[f451])).
fof(f451,plain,(
  ! [X0] : ? [X1] : ((p(s(bool,X1)) <=> s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.NegInf_0'))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X1),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))))))),
  inference(rectify,[],[f134])).
fof(f134,axiom,(
  ! [X6] : ? [X45] : ((p(s(bool,X45)) <=> s('type.realax.real',X6) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X6))),s('type.extreal.extreal','const.extreal.NegInf_0'))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X45),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X6))),s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))))))),
  file('thm.extreal.add_not_infty|split|1','thm.extreal.extreal_mul_def|split|4')).
fof(f1015,plain,(
  ( ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) )),
  inference(cnf_transformation,[],[f302])).
fof(f302,plain,(
  ! [X0] : s('type.realax.real',X0) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))),
  inference(rectify,[],[f238])).
fof(f238,axiom,(
  ! [X6] : s('type.realax.real',X6) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X6),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))),
  file('thm.extreal.add_not_infty|split|1','thm.real.REAL_MUL_RID')).
fof(f21782,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1015,f15405,f1225])).
fof(f1225,plain,(
  ( ! [X0] : (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,sK81(X0)))) )),
  inference(cnf_transformation,[],[f855])).
fof(f855,plain,(
  ! [X0] : ((~p(s(bool,sK81(X0))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,sK81(X0)))) & s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,sK81(X0)),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK81])],[f854])).
fof(f854,plain,(
  ! [X0] : ? [X1] : ((~p(s(bool,X1)) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,X1))) & s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X1),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))))))),
  inference(flattening,[],[f853])).
fof(f853,plain,(
  ! [X0] : ? [X1] : (((~p(s(bool,X1)) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,X1)))) & s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X1),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))))))),
  inference(nnf_transformation,[],[f450])).
fof(f450,plain,(
  ! [X0] : ? [X1] : ((p(s(bool,X1)) <=> s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X1),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))))))),
  inference(rectify,[],[f132])).
fof(f132,axiom,(
  ! [X7] : ? [X45] : ((p(s(bool,X45)) <=> s('type.realax.real',X7) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X7))))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X45),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X7))),s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))))))),
  file('thm.extreal.add_not_infty|split|1','thm.extreal.extreal_mul_def|split|7')).
fof(f21780,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1015,f15405,f1222])).
fof(f1222,plain,(
  ( ! [X0] : (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,sK80(X0)))) )),
  inference(cnf_transformation,[],[f852])).
fof(f852,plain,(
  ! [X0] : ((~p(s(bool,sK80(X0))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,sK80(X0)))) & s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0'))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,sK80(X0)),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK80])],[f851])).
fof(f851,plain,(
  ! [X0] : ? [X1] : ((~p(s(bool,X1)) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,X1))) & s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0'))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X1),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))))))),
  inference(flattening,[],[f850])).
fof(f850,plain,(
  ! [X0] : ? [X1] : (((~p(s(bool,X1)) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,X1)))) & s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0'))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X1),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))))))),
  inference(nnf_transformation,[],[f449])).
fof(f449,plain,(
  ! [X0] : ? [X1] : ((p(s(bool,X1)) <=> s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0'))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X1),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))))))),
  inference(rectify,[],[f130])).
fof(f130,axiom,(
  ! [X6] : ? [X45] : ((p(s(bool,X45)) <=> s('type.realax.real',X6) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X6))),s('type.extreal.extreal','const.extreal.PosInf_0'))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X45),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X6))),s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))))))),
  file('thm.extreal.add_not_infty|split|1','thm.extreal.extreal_mul_def|split|6')).
fof(f21778,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1015,f15405,f1219])).
fof(f1219,plain,(
  ( ! [X0] : (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,sK79(X0)))) )),
  inference(cnf_transformation,[],[f849])).
fof(f849,plain,(
  ! [X0] : ((~p(s(bool,sK79(X0))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,sK79(X0)))) & s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,sK79(X0)),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK79])],[f848])).
fof(f848,plain,(
  ! [X0] : ? [X1] : ((~p(s(bool,X1)) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,X1))) & s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X1),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))))))),
  inference(flattening,[],[f847])).
fof(f847,plain,(
  ! [X0] : ? [X1] : (((~p(s(bool,X1)) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,X1)))) & s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X1),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))))))),
  inference(nnf_transformation,[],[f448])).
fof(f448,plain,(
  ! [X0] : ? [X1] : ((p(s(bool,X1)) <=> s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X1),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))))))),
  inference(rectify,[],[f131])).
fof(f131,axiom,(
  ! [X7] : ? [X45] : ((p(s(bool,X45)) <=> s('type.realax.real',X7) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X7))))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X45),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X7))),s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))))))),
  file('thm.extreal.add_not_infty|split|1','thm.extreal.extreal_mul_def|split|5')).
fof(f21776,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f958,f15405,f1216])).
fof(f958,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.PosInf_0')))),
  inference(cnf_transformation,[],[f150])).
fof(f150,axiom,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.PosInf_0')))),
  file('thm.extreal.add_not_infty|split|1','thm.extreal.extreal_size_def|split|1')).
fof(f21774,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f958,f15405,f1213])).
fof(f21772,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f15405,f15405,f1274])).
fof(f21770,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f15405,f15405,f1274])).
fof(f21768,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f989,f15405,f1088])).
fof(f21766,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f15405,f15405,f1276])).
fof(f21764,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f964,f15405,f1086])).
fof(f21762,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f15405,f15405,f1276])).
fof(f21760,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f981,f15405,f1136])).
fof(f981,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.NegInf_0')),
  inference(cnf_transformation,[],[f268])).
fof(f268,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.NegInf_0')),
  inference(flattening,[],[f111])).
fof(f111,axiom,(
  ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.NegInf_0')),
  file('thm.extreal.add_not_infty|split|1','thm.extreal.extreal_distinct|split|0')).
fof(f21758,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1141,f15405,f1134])).
fof(f1141,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))) )),
  inference(cnf_transformation,[],[f371])).
fof(f371,plain,(
  ! [X0] : s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))),
  inference(flattening,[],[f370])).
fof(f370,plain,(
  ! [X0] : ~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))),
  inference(rectify,[],[f225])).
fof(f225,axiom,(
  ! [X6] : ~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X6)))),
  file('thm.extreal.add_not_infty|split|1','thm.extreal.extreal_not_infty|split|0')).
fof(f21753,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1190,f15405,f1356])).
fof(f1190,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X1) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(X0,X1),s(X0,X2),s(fun('type.realax.real',X0),X3)))) )),
  inference(cnf_transformation,[],[f427])).
fof(f427,plain,(
  ! [X0,X1,X2,X3] : s(X0,X1) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(X0,X1),s(X0,X2),s(fun('type.realax.real',X0),X3)))),
  inference(flattening,[],[f426])).
fof(f426,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,X1) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(X0,X1),s(X0,X2),s(fun('type.realax.real',X0),X3)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X8] : ! [X9,X10,X11] : s(X8,X9) = s(X8,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(X8,X9),s(X8,X10),s(fun('type.realax.real',X8),X11)))),
  file('thm.extreal.add_not_infty|split|1','thm.extreal.extreal_case_def|split|0')).
fof(f21751,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1533,f15405,f15405,f1526])).
fof(f21749,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1533,f1508,f15405,f1526])).
fof(f21713,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1533,f15405,f15405,f1526])).
fof(f21711,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1535,f1508,f15405,f1526])).
fof(f21675,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1559,f2496,f15405,f1497])).
fof(f1559,plain,(
  ( ! [X0,X1] : (sP7('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f1508,f1396])).
fof(f21673,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1557,f2495,f15405,f1497])).
fof(f1557,plain,(
  ( ! [X0,X1] : (sP7(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f1508,f1395])).
fof(f1395,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f914])).
fof(f21671,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1559,f2377,f15405,f1497])).
fof(f21669,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1557,f2376,f15405,f1497])).
fof(f21664,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f7001,f6992,f15405,f1497])).
fof(f7001,plain,(
  ( ! [X0,X1] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18)),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f5894,f1396])).
fof(f21662,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f6089,f6080,f15405,f1497])).
fof(f6089,plain,(
  ( ! [X0,X1] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK17)),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f5893,f1396])).
fof(f21660,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1559,f2149,f15405,f1497])).
fof(f21658,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f7000,f6989,f15405,f1497])).
fof(f7000,plain,(
  ( ! [X0,X1] : (sP7(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18)),X1)) )),
  inference(unit_resulting_resolution,[],[f5894,f1395])).
fof(f21656,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f6088,f6077,f15405,f1497])).
fof(f6088,plain,(
  ( ! [X0,X1] : (sP7(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK17)),X1)) )),
  inference(unit_resulting_resolution,[],[f5893,f1395])).
fof(f21654,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1557,f2146,f15405,f1497])).
fof(f21605,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f952,f1163,f15405,f1494])).
fof(f21603,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f952,f1165,f15405,f1494])).
fof(f1165,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f407])).
fof(f407,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f406])).
fof(f406,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f169])).
fof(f169,axiom,(
  ! [X8] : ! [X19,X20] : s(X8,X19) = s(X8,'const.bool.COND_3'(s(bool,'T_0'),s(X8,X19),s(X8,X20)))),
  file('thm.extreal.add_not_infty|split|1','thm.bool.COND_CLAUSES|split|0')).
fof(f21601,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f952,f1164,f15405,f1494])).
fof(f21599,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f952,f1166,f15405,f1494])).
fof(f1166,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f409])).
fof(f409,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(flattening,[],[f408])).
fof(f408,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X8] : ! [X12,X11] : s(X8,X12) = s(X8,'const.ind_type.FCONS_3'(s(X8,X12),s(fun('type.num.num',X8),X11),s('type.num.num','const.num.|48|_0')))),
  file('thm.extreal.add_not_infty|split|1','thm.ind_type.FCONS|split|0')).
fof(f21597,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f952,f1150,f15405,f1494])).
fof(f21595,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f952,f1196,f15405,f1494])).
fof(f1196,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X2) = s(X0,happ(s(fun('type.extreal.extreal',X0),sK75(X3,X2,X1,X0)),s('type.extreal.extreal','const.extreal.PosInf_0')))) )),
  inference(cnf_transformation,[],[f839])).
fof(f839,plain,(
  ! [X0,X1,X2,X3] : (s(X0,X1) = s(X0,happ(s(fun('type.extreal.extreal',X0),sK75(X3,X2,X1,X0)),s('type.extreal.extreal','const.extreal.NegInf_0'))) & s(X0,X2) = s(X0,happ(s(fun('type.extreal.extreal',X0),sK75(X3,X2,X1,X0)),s('type.extreal.extreal','const.extreal.PosInf_0'))) & ! [X5] : s(X0,happ(s(fun('type.realax.real',X0),X3),s('type.realax.real',X5))) = s(X0,happ(s(fun('type.extreal.extreal',X0),sK75(X3,X2,X1,X0)),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK75])],[f437])).
fof(f437,plain,(
  ! [X0,X1,X2,X3] : ? [X4] : (s(X0,X1) = s(X0,happ(s(fun('type.extreal.extreal',X0),X4),s('type.extreal.extreal','const.extreal.NegInf_0'))) & s(X0,X2) = s(X0,happ(s(fun('type.extreal.extreal',X0),X4),s('type.extreal.extreal','const.extreal.PosInf_0'))) & ! [X5] : s(X0,happ(s(fun('type.realax.real',X0),X3),s('type.realax.real',X5))) = s(X0,happ(s(fun('type.extreal.extreal',X0),X4),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5))))))),
  inference(flattening,[],[f436])).
fof(f436,plain,(
  ! [X0] : ! [X1,X2,X3] : ? [X4] : (s(X0,X1) = s(X0,happ(s(fun('type.extreal.extreal',X0),X4),s('type.extreal.extreal','const.extreal.NegInf_0'))) & s(X0,X2) = s(X0,happ(s(fun('type.extreal.extreal',X0),X4),s('type.extreal.extreal','const.extreal.PosInf_0'))) & ! [X5] : s(X0,happ(s(fun('type.realax.real',X0),X3),s('type.realax.real',X5))) = s(X0,happ(s(fun('type.extreal.extreal',X0),X4),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5))))))),
  inference(rectify,[],[f146])).
fof(f146,axiom,(
  ! [X8] : ! [X60,X61,X62] : ? [X63] : (s(X8,X60) = s(X8,happ(s(fun('type.extreal.extreal',X8),X63),s('type.extreal.extreal','const.extreal.NegInf_0'))) & s(X8,X61) = s(X8,happ(s(fun('type.extreal.extreal',X8),X63),s('type.extreal.extreal','const.extreal.PosInf_0'))) & ! [X12] : s(X8,happ(s(fun('type.realax.real',X8),X62),s('type.realax.real',X12))) = s(X8,happ(s(fun('type.extreal.extreal',X8),X63),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X12))))))),
  file('thm.extreal.add_not_infty|split|1','thm.extreal.extreal_Axiom')).
fof(f21593,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f952,f1195,f15405,f1494])).
fof(f1195,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X1) = s(X0,happ(s(fun('type.extreal.extreal',X0),sK75(X3,X2,X1,X0)),s('type.extreal.extreal','const.extreal.NegInf_0')))) )),
  inference(cnf_transformation,[],[f839])).
fof(f21591,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f952,f1335,f15405,f1494])).
fof(f1335,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f525])).
fof(f525,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))),
  inference(flattening,[],[f524])).
fof(f524,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))),
  inference(rectify,[],[f160])).
fof(f160,axiom,(
  ! [X24,X22] : ! [X25,X27] : s(X22,X25) = s(X22,happ(s(fun(X24,X22),'_dst_x0_1'(s(X22,X25))),s(X24,X27)))),
  file('thm.extreal.add_not_infty|split|1','_dst_x0_ax')).
fof(f21589,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f952,f1514,f15405,f1494])).
fof(f21587,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f952,f1191,f15405,f1494])).
fof(f1191,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X2) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.PosInf_0'),s(X0,X1),s(X0,X2),s(fun('type.realax.real',X0),X3)))) )),
  inference(cnf_transformation,[],[f429])).
fof(f429,plain,(
  ! [X0,X1,X2,X3] : s(X0,X2) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.PosInf_0'),s(X0,X1),s(X0,X2),s(fun('type.realax.real',X0),X3)))),
  inference(flattening,[],[f428])).
fof(f428,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,X2) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.PosInf_0'),s(X0,X1),s(X0,X2),s(fun('type.realax.real',X0),X3)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X8] : ! [X9,X10,X11] : s(X8,X10) = s(X8,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.PosInf_0'),s(X8,X9),s(X8,X10),s(fun('type.realax.real',X8),X11)))),
  file('thm.extreal.add_not_infty|split|1','thm.extreal.extreal_case_def|split|1')).
fof(f21585,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f952,f1190,f15405,f1494])).
fof(f21569,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f952,f2116,f15405,f1494])).
fof(f21567,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f952,f988,f15405,f1494])).
fof(f988,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f275])).
fof(f275,plain,(
  ! [X0] : s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f262])).
fof(f262,axiom,(
  ! [X31] : s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X31),s('type.num.num','const.num.|48|_0')))),
  file('thm.extreal.add_not_infty|split|1','thm.numeral.numeral_distrib|split|29')).
fof(f21565,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f952,f991,f15405,f1494])).
fof(f991,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f278])).
fof(f278,plain,(
  ! [X0] : s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f251])).
fof(f251,axiom,(
  ! [X31] : s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X31)))),
  file('thm.extreal.add_not_infty|split|1','thm.numeral.numeral_distrib|split|26')).
fof(f21563,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f952,f1149,f15405,f1494])).
fof(f21561,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f952,f1005,f15405,f1494])).
fof(f1005,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0')))) )),
  inference(cnf_transformation,[],[f292])).
fof(f292,plain,(
  ! [X0] : s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0')))),
  inference(rectify,[],[f147])).
fof(f147,axiom,(
  ! [X28] : s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X28))),s('type.extreal.extreal','const.extreal.PosInf_0')))),
  file('thm.extreal.add_not_infty|split|1','thm.extreal.extreal_le_def|split|3')).
fof(f21559,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f952,f967,f15405,f1494])).
fof(f967,plain,(
  s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))),
  inference(cnf_transformation,[],[f151])).
fof(f151,axiom,(
  s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))),
  file('thm.extreal.add_not_infty|split|1','thm.extreal.extreal_le_def|split|2')).
fof(f21556,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f4847,f15405,f1473])).
fof(f4847,plain,(
  ~sP13('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f952,f1508,f1508,f1470])).
fof(f21554,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1591,f952,f15405,f15405,f1470])).
fof(f21531,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f11963,f952,f1508,f15405,f1470])).
fof(f21499,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1591,f952,f15405,f15405,f1470])).
fof(f21476,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1593,f952,f1508,f15405,f1470])).
fof(f21444,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f4817,f15405,f1469])).
fof(f4817,plain,(
  ~sP14('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f952,f1508,f952,f1466])).
fof(f21442,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f4817,f15405,f1468])).
fof(f1468,plain,(
  ( ! [X2,X0,X1] : (sP14(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f943])).
fof(f21440,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f11988,f952,f952,f15405,f1466])).
fof(f11988,plain,(
  ( ! [X0] : (sP14('T_0',X0,X0)) )),
  inference(unit_resulting_resolution,[],[f1165,f1475])).
fof(f1475,plain,(
  ( ! [X2,X0,X3,X1] : (sP14(X2,X1,X3) | s(bool,X3) != s(bool,'const.bool.COND_3'(s(bool,X2),s(bool,X1),s(bool,X0)))) )),
  inference(cnf_transformation,[],[f948])).
fof(f21412,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f4792,f15405,f1465])).
fof(f4792,plain,(
  ~sP15('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f952,f1508,f952,f1462])).
fof(f21410,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f4792,f15405,f1464])).
fof(f1464,plain,(
  ( ! [X2,X0,X1] : (sP15(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f940])).
fof(f21408,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f12006,f952,f952,f15405,f1462])).
fof(f12006,plain,(
  ( ! [X0,X1] : (sP15(X0,X1,X0)) )),
  inference(unit_resulting_resolution,[],[f1164,f1476])).
fof(f1476,plain,(
  ( ! [X2,X0,X3,X1] : (sP15(X0,X1,X3) | s(bool,X3) != s(bool,'const.bool.COND_3'(s(bool,X2),s(bool,X1),s(bool,X0)))) )),
  inference(cnf_transformation,[],[f948])).
fof(f21378,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f4767,f15405,f1461])).
fof(f4767,plain,(
  ~sP16('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f952,f952,f1508,f1458])).
fof(f21376,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f4767,f15405,f1459])).
fof(f1459,plain,(
  ( ! [X2,X0,X1] : (sP16(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f937])).
fof(f21374,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f12025,f952,f952,f15405,f1458])).
fof(f12025,plain,(
  ( ! [X0] : (sP16(X0,X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f1165,f1477])).
fof(f1477,plain,(
  ( ! [X2,X0,X3,X1] : (sP16(X3,X1,X2) | s(bool,X3) != s(bool,'const.bool.COND_3'(s(bool,X2),s(bool,X1),s(bool,X0)))) )),
  inference(cnf_transformation,[],[f948])).
fof(f21346,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f2194,f952,f952,f15405,f1434])).
fof(f21322,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f4712,f952,f15405,f1432])).
fof(f21307,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f2185,f952,f15405,f15405,f1431])).
fof(f21292,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f4712,f952,f1508,f15405,f1431])).
fof(f21276,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f2185,f15405,f952,f15405,f1431])).
fof(f21274,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f2185,f1508,f952,f15405,f1431])).
fof(f21245,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f4698,f15405,f1430])).
fof(f21243,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f4698,f15405,f1428])).
fof(f21241,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1569,f952,f952,f15405,f1427])).
fof(f21215,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f7146,f7131,f952,f15405,f1426])).
fof(f7146,plain,(
  ( ! [X0,X1] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK18)),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f5894,f1430])).
fof(f21213,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f6234,f6219,f952,f15405,f1426])).
fof(f6234,plain,(
  ( ! [X0,X1] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK17)),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f5893,f1430])).
fof(f21208,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1571,f2516,f952,f15405,f1426])).
fof(f21206,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1571,f2397,f952,f15405,f1426])).
fof(f21204,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1571,f2185,f952,f15405,f1426])).
fof(f21130,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f2194,f952,f15405,f1421])).
fof(f21109,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f2188,f15405,f15405,f1420])).
fof(f21107,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f4679,f1508,f15405,f1420])).
fof(f21086,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f2188,f15405,f15405,f1420])).
fof(f21084,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f2188,f1508,f15405,f1420])).
fof(f21063,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f2174,f15405,f952,f15405,f1417])).
fof(f21061,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f2174,f1508,f952,f15405,f1417])).
fof(f21014,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f2174,f952,f15405,f15405,f1417])).
fof(f20990,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f2180,f952,f1508,f15405,f1417])).
fof(f20965,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f2168,f15405,f15405,f15405,f1414])).
fof(f20963,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f2171,f1508,f15405,f15405,f1414])).
fof(f20946,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f2168,f15405,f1508,f15405,f1414])).
fof(f20944,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f4572,f1508,f1508,f15405,f1414])).
fof(f20911,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f2168,f15405,f15405,f15405,f1414])).
fof(f20909,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f2168,f1508,f15405,f15405,f1414])).
fof(f20892,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f2168,f15405,f1508,f15405,f1414])).
fof(f20890,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f2168,f1508,f1508,f15405,f1414])).
fof(f20859,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f2168,f15405,f15405,f15405,f1414])).
fof(f20857,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f2168,f1508,f15405,f15405,f1414])).
fof(f20840,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f2171,f15405,f1508,f15405,f1414])).
fof(f20838,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f2171,f1508,f1508,f15405,f1414])).
fof(f20807,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f4553,f15405,f1413])).
fof(f20805,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1561,f952,f15405,f15405,f1410])).
fof(f20783,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1561,f952,f1508,f15405,f1410])).
fof(f20753,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1561,f952,f15405,f15405,f1410])).
fof(f20731,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1563,f952,f1508,f15405,f1410])).
fof(f20700,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f7017,f7087,f15405,f15405,f1409])).
fof(f20698,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f6105,f6175,f15405,f15405,f1409])).
fof(f20696,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f2168,f1565,f15405,f15405,f1409])).
fof(f20694,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f2501,f2505,f15405,f15405,f1409])).
fof(f20692,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f2382,f2386,f15405,f15405,f1409])).
fof(f20690,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f2168,f2133,f15405,f15405,f1409])).
fof(f20688,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f2168,f2097,f15405,f15405,f1409])).
fof(f20686,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f2168,f2050,f15405,f15405,f1409])).
fof(f20684,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f2168,f1564,f15405,f15405,f1409])).
fof(f20682,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f2168,f1612,f15405,f15405,f1409])).
fof(f20680,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f2171,f1563,f15405,f15405,f1409])).
fof(f20678,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f2498,f2504,f15405,f15405,f1409])).
fof(f20676,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f2379,f2385,f15405,f15405,f1409])).
fof(f20674,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f2171,f2132,f15405,f15405,f1409])).
fof(f20672,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f2171,f2096,f15405,f15405,f1409])).
fof(f20670,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f2171,f2049,f15405,f15405,f1409])).
fof(f20668,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f2171,f1562,f15405,f15405,f1409])).
fof(f20666,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f2171,f1611,f15405,f15405,f1409])).
fof(f20664,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f2168,f1561,f15405,f15405,f1409])).
fof(f20646,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f7017,f7087,f1508,f15405,f1409])).
fof(f20644,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f6105,f6175,f1508,f15405,f1409])).
fof(f20642,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f2168,f1565,f1508,f15405,f1409])).
fof(f20640,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f2498,f2504,f1508,f15405,f1409])).
fof(f20638,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f2379,f2385,f1508,f15405,f1409])).
fof(f20631,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f2168,f1561,f1508,f15405,f1409])).
fof(f20558,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f7087,f7018,f15405,f15405,f1409])).
fof(f20556,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f7087,f7017,f15405,f15405,f1409])).
fof(f20554,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f6175,f6106,f15405,f15405,f1409])).
fof(f20552,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f6175,f6105,f15405,f15405,f1409])).
fof(f20545,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1565,f2501,f15405,f15405,f1409])).
fof(f20543,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1565,f2382,f15405,f15405,f1409])).
fof(f20541,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1565,f2171,f15405,f15405,f1409])).
fof(f20539,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1565,f2498,f15405,f15405,f1409])).
fof(f20537,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1565,f2379,f15405,f15405,f1409])).
fof(f20535,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1565,f2168,f15405,f15405,f1409])).
fof(f20513,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f7087,f7018,f1508,f15405,f1409])).
fof(f20511,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f6175,f6106,f1508,f15405,f1409])).
fof(f20506,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1565,f2501,f1508,f15405,f1409])).
fof(f20504,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1565,f2382,f1508,f15405,f1409])).
fof(f20502,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1565,f2171,f1508,f15405,f1409])).
fof(f20288,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f2180,f952,f15405,f1404])).
fof(f20264,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f2174,f952,f15405,f1403])).
fof(f20240,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f4427,f952,f952,f15405,f1402])).
fof(f4427,plain,(
  sP8('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f1555,f952,f952,f1389])).
fof(f1389,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP7(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f912])).
fof(f20221,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f2146,f15405,f15405,f1398])).
fof(f20219,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f2149,f1508,f15405,f1398])).
fof(f20204,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f2146,f15405,f15405,f1398])).
fof(f20202,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f4461,f1508,f15405,f1398])).
fof(f20187,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f2146,f15405,f15405,f1397])).
fof(f20185,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f2146,f1508,f15405,f1397])).
fof(f20170,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f2146,f15405,f15405,f1397])).
fof(f20168,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f4461,f1508,f15405,f1397])).
fof(f20153,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f4447,f15405,f1396])).
fof(f20151,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f4447,f15405,f1395])).
fof(f20149,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1555,f952,f952,f15405,f1393])).
fof(f20123,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f4427,f952,f15405,f1388])).
fof(f20107,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f4427,f952,f15405,f1387])).
fof(f20091,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f3143,f1163,f15405,f1385])).
fof(f3143,plain,(
  ( ! [X0] : (sP6('T_0',X0,sK63(X0,bool))) )),
  inference(unit_resulting_resolution,[],[f952,f1163,f1383])).
fof(f1383,plain,(
  ( ! [X2,X0,X1] : (sP6(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f909])).
fof(f20086,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f3144,f1150,f15405,f1385])).
fof(f3144,plain,(
  ( ! [X0] : (sP6('T_0',X0,'const.combin.I_1'(s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f952,f1150,f1383])).
fof(f20064,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f3139,f2116,f15405,f1385])).
fof(f3139,plain,(
  sP6('T_0','T_0',sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))))),
  inference(unit_resulting_resolution,[],[f952,f2116,f1383])).
fof(f20060,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f3136,f1149,f15405,f1385])).
fof(f3136,plain,(
  ( ! [X0,X1] : (sP6('T_0','T_0','const.bool.DATATYPE_1'(s(X0,X1)))) )),
  inference(unit_resulting_resolution,[],[f952,f1149,f1383])).
fof(f20056,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f3133,f1163,f15405,f1385])).
fof(f3133,plain,(
  ( ! [X0] : (sP6('T_0',sK63(X0,bool),X0)) )),
  inference(unit_resulting_resolution,[],[f952,f1163,f1383])).
fof(f20051,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f3134,f1150,f15405,f1385])).
fof(f3134,plain,(
  ( ! [X0] : (sP6('T_0','const.combin.I_1'(s(bool,X0)),X0)) )),
  inference(unit_resulting_resolution,[],[f952,f1150,f1383])).
fof(f20029,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f3129,f2116,f15405,f1385])).
fof(f3129,plain,(
  sP6('T_0',sK83('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK17),s('type.extreal.extreal',sK18))),'T_0')),
  inference(unit_resulting_resolution,[],[f952,f2116,f1383])).
fof(f20025,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f3126,f1149,f15405,f1385])).
fof(f3126,plain,(
  ( ! [X0,X1] : (sP6('T_0','const.bool.DATATYPE_1'(s(X0,X1)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f952,f1149,f1383])).
fof(f21823,plain,(
  $false | $spl4),
  inference(trivial_inequality_removal,[],[f20013])).
fof(f20013,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f11712,f15405,f1385])).
fof(f11712,plain,(
  sP6('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1549,f1539,f1545,f1508,f1368])).
fof(f21822,plain,(
  $false | $spl4),
  inference(trivial_inequality_removal,[],[f20012])).
fof(f20012,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f11502,f15405,f1385])).
fof(f11502,plain,(
  sP6('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f1553,f1537,f1543,f952,f1367])).
fof(f1367,plain,(
  ( ! [X2,X0,X1] : (~sP4(X1,X2,X0) | ~p(s(bool,X2)) | ~sP5(X0,X2,X1) | sP6(X0,X1,X2) | ~sP3(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f900])).
fof(f20005,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f3156,f1890,f15405,f1382])).
fof(f20003,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f3154,f1890,f15405,f1382])).
fof(f19987,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1549,f15405,f15405,f15405,f1378])).
fof(f19985,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1549,f1508,f15405,f15405,f1378])).
fof(f19959,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1549,f15405,f1508,f15405,f1378])).
fof(f19957,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1549,f1508,f1508,f15405,f1378])).
fof(f19899,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1549,f15405,f15405,f15405,f1378])).
fof(f19897,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1549,f1508,f15405,f15405,f1378])).
fof(f19871,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1551,f15405,f1508,f15405,f1378])).
fof(f19869,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1551,f1508,f1508,f15405,f1378])).
fof(f19811,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1549,f15405,f15405,f15405,f1378])).
fof(f19809,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1551,f1508,f15405,f15405,f1378])).
fof(f19783,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1549,f15405,f1508,f15405,f1378])).
fof(f19781,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1553,f1508,f1508,f15405,f1378])).
fof(f19723,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f4361,f15405,f1377])).
fof(f4361,plain,(
  ~sP4('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f952,f1508,f952,f1374])).
fof(f19721,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f4361,f15405,f1376])).
fof(f19719,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1543,f952,f952,f15405,f1374])).
fof(f19691,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f4333,f15405,f1373])).
fof(f4333,plain,(
  ~sP5('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f952,f1508,f952,f1370])).
fof(f19689,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f4333,f15405,f1372])).
fof(f19687,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f3417,f952,f952,f15405,f1370])).
fof(f3417,plain,(
  ( ! [X0,X1] : (sP5('T_0','T_0','const.bool.DATATYPE_1'(s(X0,X1)))) )),
  inference(unit_resulting_resolution,[],[f3126,f1365])).
fof(f1365,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | sP5(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f900])).
fof(f19135,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1551,f6874,f1537,f6891,f15405,f1369])).
fof(f19122,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1551,f6874,f6884,f6891,f15405,f1369])).
fof(f18996,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1551,f5962,f1537,f5979,f15405,f1369])).
fof(f18983,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1551,f5962,f5972,f5979,f15405,f1369])).
fof(f18856,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1551,f3156,f1537,f1545,f15405,f1369])).
fof(f18842,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1551,f3156,f1539,f1545,f15405,f1369])).
fof(f17827,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1551,f6874,f6884,f1543,f15405,f1369])).
fof(f17816,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1551,f5962,f5972,f1543,f15405,f1369])).
fof(f17805,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1551,f3156,f1539,f1543,f15405,f1369])).
fof(f17715,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1551,f6874,f1537,f1543,f15405,f1369])).
fof(f17713,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1551,f5962,f1537,f1543,f15405,f1369])).
fof(f17711,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1551,f3156,f1537,f1543,f15405,f1369])).
fof(f17078,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1551,f6874,f1537,f6891,f15405,f1368])).
fof(f17065,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1551,f6874,f6884,f6891,f15405,f1368])).
fof(f16939,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1551,f5962,f1537,f5979,f15405,f1368])).
fof(f16926,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1551,f5962,f5972,f5979,f15405,f1368])).
fof(f16799,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1551,f3156,f1537,f1545,f15405,f1368])).
fof(f16785,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1551,f3156,f1539,f1545,f15405,f1368])).
fof(f15770,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1551,f6874,f6884,f1543,f15405,f1368])).
fof(f15759,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1551,f5962,f5972,f1543,f15405,f1368])).
fof(f15748,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1551,f3156,f1539,f1543,f15405,f1368])).
fof(f15658,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1551,f6874,f1537,f1543,f15405,f1368])).
fof(f15656,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1551,f5962,f1537,f1543,f15405,f1368])).
fof(f15654,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1551,f3156,f1537,f1543,f15405,f1368])).
fof(f15537,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f3139,f952,f952,f15405,f1366])).
fof(f15517,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1953,f952,f15405,f1328])).
fof(f1953,plain,(
  ~sP2('T_0','T_0')),
  inference(unit_resulting_resolution,[],[f952,f952,f1324])).
fof(f1324,plain,(
  ( ! [X0,X1] : (~sP2(X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f890])).
fof(f15495,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1953,f952,f15405,f1325])).
fof(f15473,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1956,f15405,f15405,f1323])).
fof(f15471,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1965,f1508,f15405,f1323])).
fof(f15449,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1956,f15405,f15405,f1323])).
fof(f15447,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1956,f1508,f15405,f1323])).
fof(f15425,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1508,f1890,f15405,f1265])).
fof(f15423,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f15405,f1890,f15405,f1265])).
fof(f15419,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1508,f1890,f15405,f1265])).
fof(f15417,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f15405,f1890,f15405,f1265])).
fof(f15413,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f1890,f15405,f1072])).
fof(f15411,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f967,f15405,f1067])).
fof(f15409,plain,(
  $false | $spl4),
  inference(unit_resulting_resolution,[],[f952,f15405])).
fof(f15407,plain,(
  p(s(bool,sK61('const.combin.I_0',bool))) | $spl6),
  inference(cnf_transformation,[],[f15407_D])).
fof(f15407_D,plain,(
  p(s(bool,sK61('const.combin.I_0',bool))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f15404,plain,(
  ( ! [X3] : (p(s(bool,sK61('const.combin.I_0',bool))) | ~p(s(bool,X3))) )),
  inference(forward_demodulation,[],[f15389,f1514])).
fof(f15389,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | p(s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,sK61('const.combin.I_0',bool)))))) )),
  inference(superposition,[],[f1157,f1514])).
fof(f1157,plain,(
  ( ! [X4,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK61(X1,X0)))))) )),
  inference(cnf_transformation,[],[f814])).
fof(f14475,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'const.combin.I_0')))) | $spl0),
  inference(cnf_transformation,[],[f14475_D])).
fof(f14475_D,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'const.combin.I_0')))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f14476,plain,(
  ( ! [X3] : (~p(s(bool,sK65(X3,'const.combin.I_0',bool)))) ) | $spl2),
  inference(cnf_transformation,[],[f14476_D])).
fof(f14476_D,plain,(
  ( ! [X3] : (~p(s(bool,sK65(X3,'const.combin.I_0',bool)))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f14461,plain,(
  ( ! [X3] : (~p(s(bool,sK65(X3,'const.combin.I_0',bool))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'const.combin.I_0'))))) )),
  inference(superposition,[],[f1171,f1514])).
fof(f1171,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK65(X2,X1,X0))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f821])).
fof(f821,plain,(
  ! [X0,X1,X2] : ((s(X0,X2) = s(X0,sK65(X2,X1,X0)) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK65(X2,X1,X0)))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK65])],[f630])).
fof(f630,plain,(
  ! [X0,X1,X2] : (? [X3] : (s(X0,X2) = s(X0,X3) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(ennf_transformation,[],[f415])).
fof(f415,plain,(
  ! [X0,X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f414])).
fof(f414,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X8] : ! [X14,X21] : (! [X6] : (s(X8,X6) = s(X8,X21) => p(s(bool,happ(s(fun(X8,bool),X14),s(X8,X6))))) => p(s(bool,'|dollar|exists_1'(s(fun(X8,bool),X14)))))),
  file('thm.extreal.add_not_infty|split|1','thm.bool.JRH_INDUCT_UTIL')).
% SZS output end Proof for thm.extreal
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 965
% Passive clauses: 8457
% Generated clauses: 78510
% Final active clauses: 646
% Final passive clauses: 4488
% Input formulas: 261
% Initial clauses: 534
% 
% Duplicate literals: 10
% Trivial inequalities: 15
% Fw subsumption resolutions: 1415
% Fw demodulations: 34066
% Bw demodulations: 1141
% Condensations: 53
% 
% Simple tautologies: 113
% Equational tautologies: 48
% Forward subsumptions: 33037
% Fw demodulations to eq. taut.: 237
% Bw demodulations to eq. taut.: 12
% 
% Unit resulting resolution: 24685
% Factoring: 3
% Forward superposition: 9026
% Backward superposition: 6414
% Self superposition: 81
% Equality factoring: 2
% Equality resolution: 23
% 
% Split clauses: 17
% Split components: 34
% SAT solver clauses: 9535
% SAT solver unit clauses: 4883
% SAT solver binary clauses: 3362
% 
% Sat splits: 17
% Sat splitting refutations: 1042
% 
% Memory used [KB]: 45031
% Time elapsed: 1.813 s
% ------------------------------
% ------------------------------
% Success in time 3.15 s
