Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.real_sigma
% (54026)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 728
% Passive clauses: 8961
% Generated clauses: 19473
% Final active clauses: 718
% Final passive clauses: 7889
% Input formulas: 43
% Initial clauses: 192
% 
% Duplicate literals: 504
% Trivial inequalities: 22
% Fw subsumption resolutions: 1841
% Fw demodulations: 305
% Bw demodulations: 9
% 
% Simple tautologies: 235
% Equational tautologies: 56
% Forward subsumptions: 6880
% 
% Binary resolution: 10488
% Factoring: 1336
% Forward superposition: 2573
% Backward superposition: 1895
% Self superposition: 24
% Equality factoring: 28
% Equality resolution: 7
% 
% Split clauses: 593
% Split components: 1193
% SAT solver clauses: 1021
% SAT solver unit clauses: 19
% SAT solver binary clauses: 191
% 
% Sat splits: 593
% 
% TWLsolver clauses: 1020
% TWLsolver calls for satisfiability: 432
% 
% Memory used [KB]: 9594
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.real_sigma
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.real_sigma
% SZS output start Proof for thm.real_sigma
fof(f8144,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f7922,f7952,f274,f475,f7560,f7562,f2837,f2838,f602,f607,f612,f616,f600,f604,f613,f618,f2810,f2811,f7554,f7557,f7556,f8036,f258,f458,f456,f601,f605,f608,f610,f257,f457,f455,f7264,f7265,f8135,f256,f460,f7477,f7507,f255,f462,f259,f467,f305,f484,f7411,f7440,f253,f465,f254,f463,f304,f485,f7402,f302,f487,f303,f486])).
fof(f486,plain,(
  ( ! [X2,X0,X3,X1] : (s(fun(X0,bool),'const.pred_set.EMPTY_0') != s(fun(X0,bool),X1) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,sK18(X2,X1,X0))))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1))))) ) | $spl32),
  inference(cnf_transformation,[],[f486_D])).
fof(f486_D,plain,(
  ( ! [X2,X0,X3,X1] : (s(fun(X0,bool),'const.pred_set.EMPTY_0') != s(fun(X0,bool),X1) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,sK18(X2,X1,X0))))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1))))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f303,plain,(
  ( ! [X2,X0,X3,X1] : (s(fun(X0,bool),'const.pred_set.EMPTY_0') != s(fun(X0,bool),X1) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,sK18(X2,X1,X0))))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f187])).
fof(f187,plain,(
  ! [X0,X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) | ! [X2] : ((p(s(bool,'const.bool.IN_2'(s(X0,sK18(X2,X1,X0)),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,sK18(X2,X1,X0)))))))) | ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3)))) | ((s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') != s(fun(X0,bool),X1)) & (s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK18])],[f186])).
fof(f186,plain,(
  ! [X0,X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) | ! [X2] : (? [X4] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X4))))))) | ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3)))) | ((s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') != s(fun(X0,bool),X1)) & (s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1)))))))),
  inference(nnf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0,X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) | ! [X2] : (? [X4] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X4))))))) | ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3)))) | (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) <=> s(fun(X0,bool),'const.pred_set.EMPTY_0') != s(fun(X0,bool),X1))))),
  inference(flattening,[],[f124])).
fof(f124,plain,(
  ! [X0,X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) | ! [X2] : ((? [X4] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X4))))))) | ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3))))) | (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) <=> s(fun(X0,bool),'const.pred_set.EMPTY_0') != s(fun(X0,bool),X1))))),
  inference(ennf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0,X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) => ! [X2] : ((! [X4] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X4))))))) & ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3))))) => (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) <=> s(fun(X0,bool),'const.pred_set.EMPTY_0') != s(fun(X0,bool),X1))))),
  inference(flattening,[],[f76])).
fof(f76,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) => ! [X2] : ((! [X4] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X4))))))) & ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3))))) => (~s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) <=> ~s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1))))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X6] : ! [X7] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X7)))) => ! [X8] : ((! [X9] : (p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X7)))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun(X6,'type.realax.real'),X8),s(X6,X9))))))) & ? [X9] : (p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X7)))) & ~s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(X6,'type.realax.real'),X8),s(X6,X9))))) => (~s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X6,'type.realax.real'),X8),s(fun(X6,bool),X7))) <=> ~s(fun(X6,bool),'const.pred_set.EMPTY_0') = s(fun(X6,bool),X7))))),
  file('thm.real_sigma.REAL_SUM_IMAGE_NONZERO','thm.real_sigma.real_sigma_5_prove_ltb')).
fof(f487,plain,(
  ( ! [X2,X0,X3,X1] : (s(fun(X0,bool),'const.pred_set.EMPTY_0') != s(fun(X0,bool),X1) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK18(X2,X1,X0)),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1))))) ) | $spl34),
  inference(cnf_transformation,[],[f487_D])).
fof(f487_D,plain,(
  ( ! [X2,X0,X3,X1] : (s(fun(X0,bool),'const.pred_set.EMPTY_0') != s(fun(X0,bool),X1) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK18(X2,X1,X0)),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1))))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f302,plain,(
  ( ! [X2,X0,X3,X1] : (s(fun(X0,bool),'const.pred_set.EMPTY_0') != s(fun(X0,bool),X1) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK18(X2,X1,X0)),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f187])).
fof(f7402,plain,(
  $false | ($spl5 | $spl12 | $spl242 | $spl2632)),
  inference(subsumption_resolution,[],[f7401,f460])).
fof(f7401,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),sK16),s(sK14,sK17))) | ($spl12 | $spl242 | $spl2632)),
  inference(subsumption_resolution,[],[f7356,f467])).
fof(f7356,plain,(
  ~p(s(bool,'T_0')) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),sK16),s(sK14,sK17))) | ($spl242 | $spl2632)),
  inference(superposition,[],[f7265,f618])).
fof(f485,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK18(X2,X1,X0)),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1))))) ) | $spl30),
  inference(cnf_transformation,[],[f485_D])).
fof(f485_D,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK18(X2,X1,X0)),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1))))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f304,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK18(X2,X1,X0)),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f187])).
fof(f463,plain,(
  ( ! [X4] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),sK16),s(sK14,X4)))))) | ~p(s(bool,'const.bool.IN_2'(s(sK14,X4),s(fun(sK14,bool),sK15))))) ) | $spl8),
  inference(cnf_transformation,[],[f463_D])).
fof(f463_D,plain,(
  ( ! [X4] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),sK16),s(sK14,X4)))))) | ~p(s(bool,'const.bool.IN_2'(s(sK14,X4),s(fun(sK14,bool),sK15))))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f254,plain,(
  ( ! [X4] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),sK16),s(sK14,X4)))))) | ~p(s(bool,'const.bool.IN_2'(s(sK14,X4),s(fun(sK14,bool),sK15))))) )),
  inference(cnf_transformation,[],[f167])).
fof(f167,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK14,bool),sK15)))) & (! [X4] : (~p(s(bool,'const.bool.IN_2'(s(sK14,X4),s(fun(sK14,bool),sK15)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),sK16),s(sK14,X4))))))) & (p(s(bool,'const.bool.IN_2'(s(sK14,sK17),s(fun(sK14,bool),sK15)))) & s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),sK16),s(sK14,sK17)))) & (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK14,'type.realax.real'),sK16),s(fun(sK14,bool),sK15))) | s(fun(sK14,bool),'const.pred_set.EMPTY_0') != s(fun(sK14,bool),sK15)) & (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK14,'type.realax.real'),sK16),s(fun(sK14,bool),sK15))) | s(fun(sK14,bool),'const.pred_set.EMPTY_0') = s(fun(sK14,bool),sK15)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK14,sK15,sK16,sK17])],[f166])).
fof(f166,plain,(
  ? [X0,X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) & ? [X2] : (! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X4))))))) & ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3)))) & (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') != s(fun(X0,bool),X1)) & (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1))))),
  inference(flattening,[],[f165])).
fof(f165,plain,(
  ? [X0,X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) & ? [X2] : (! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X4))))))) & ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3)))) & ((s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') != s(fun(X0,bool),X1)) & (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1)))))),
  inference(nnf_transformation,[],[f113])).
fof(f113,plain,(
  ? [X0,X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) & ? [X2] : (! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X4))))))) & ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3)))) & (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) <~> s(fun(X0,bool),'const.pred_set.EMPTY_0') != s(fun(X0,bool),X1))))),
  inference(flattening,[],[f112])).
fof(f112,plain,(
  ? [X0,X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) & ? [X2] : ((! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X4))))))) & ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3))))) & (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) <~> s(fun(X0,bool),'const.pred_set.EMPTY_0') != s(fun(X0,bool),X1))))),
  inference(ennf_transformation,[],[f46])).
fof(f46,plain,(
  ~! [X0,X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) => ! [X2] : ((! [X4] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X4))))))) & ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3))))) => (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) <=> s(fun(X0,bool),'const.pred_set.EMPTY_0') != s(fun(X0,bool),X1))))),
  inference(flattening,[],[f45])).
fof(f45,plain,(
  ~! [X0] : ! [X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) => ! [X2] : ((! [X4] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X4))))))) & ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3))))) => (~s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) <=> ~s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X7)))) => ! [X8] : ((! [X9] : (p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X7)))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun(X6,'type.realax.real'),X8),s(X6,X9))))))) & ? [X9] : (p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X7)))) & ~s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(X6,'type.realax.real'),X8),s(X6,X9))))) => (~s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X6,'type.realax.real'),X8),s(fun(X6,bool),X7))) <=> ~s(fun(X6,bool),'const.pred_set.EMPTY_0') = s(fun(X6,bool),X7))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X7)))) => ! [X8] : ((! [X9] : (p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X7)))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun(X6,'type.realax.real'),X8),s(X6,X9))))))) & ? [X9] : (p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X7)))) & ~s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(X6,'type.realax.real'),X8),s(X6,X9))))) => (~s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X6,'type.realax.real'),X8),s(fun(X6,bool),X7))) <=> ~s(fun(X6,bool),'const.pred_set.EMPTY_0') = s(fun(X6,bool),X7))))),
  file('thm.real_sigma.REAL_SUM_IMAGE_NONZERO','thm.real_sigma.REAL_SUM_IMAGE_NONZERO')).
fof(f465,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK14,bool),sK15)))) | $spl10),
  inference(cnf_transformation,[],[f465_D])).
fof(f465_D,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK14,bool),sK15)))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f253,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK14,bool),sK15))))),
  inference(cnf_transformation,[],[f167])).
fof(f7440,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK14,bool),sK15))) | $spl2770),
  inference(cnf_transformation,[],[f7440_D])).
fof(f7440_D,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK14,bool),sK15))) <=> ~$spl2770),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2770])])).
fof(f7411,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK14,bool),sK15))) | ($spl10 | $spl18)),
  inference(resolution,[],[f465,f475])).
fof(f484,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,sK18(X2,X1,X0))))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1))))) ) | $spl28),
  inference(cnf_transformation,[],[f484_D])).
fof(f484_D,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,sK18(X2,X1,X0))))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1))))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f305,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X3))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,sK18(X2,X1,X0))))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f187])).
fof(f467,plain,(
  p(s(bool,'T_0')) | $spl12),
  inference(cnf_transformation,[],[f467_D])).
fof(f467_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f259,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.real_sigma.REAL_SUM_IMAGE_NONZERO',aHL_TRUTH)).
fof(f462,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK14,sK17),s(fun(sK14,bool),sK15)))) | $spl6),
  inference(cnf_transformation,[],[f462_D])).
fof(f462_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK14,sK17),s(fun(sK14,bool),sK15)))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f255,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK14,sK17),s(fun(sK14,bool),sK15))))),
  inference(cnf_transformation,[],[f167])).
fof(f7507,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK14,sK17),s(fun(sK14,bool),sK15))) | $spl2810),
  inference(cnf_transformation,[],[f7507_D])).
fof(f7507_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK14,sK17),s(fun(sK14,bool),sK15))) <=> ~$spl2810),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2810])])).
fof(f7477,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK14,sK17),s(fun(sK14,bool),sK15))) | ($spl6 | $spl18)),
  inference(resolution,[],[f462,f475])).
fof(f460,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),sK16),s(sK14,sK17))) | $spl5),
  inference(cnf_transformation,[],[f460_D])).
fof(f460_D,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),sK16),s(sK14,sK17))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f256,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),sK16),s(sK14,sK17)))),
  inference(cnf_transformation,[],[f167])).
fof(f8135,plain,(
  $false | ($spl5 | $spl12 | $spl2810 | $spl2852)),
  inference(subsumption_resolution,[],[f8134,f460])).
fof(f8134,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),sK16),s(sK14,sK17))) | ($spl12 | $spl2810 | $spl2852)),
  inference(subsumption_resolution,[],[f8092,f467])).
fof(f8092,plain,(
  ~p(s(bool,'T_0')) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),sK16),s(sK14,sK17))) | ($spl2810 | $spl2852)),
  inference(superposition,[],[f7557,f7507])).
fof(f7265,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK14,X0),s(fun(sK14,bool),'const.pred_set.EMPTY_0')))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),sK16),s(sK14,X0)))) ) | $spl2632),
  inference(cnf_transformation,[],[f7265_D])).
fof(f7265_D,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK14,X0),s(fun(sK14,bool),'const.pred_set.EMPTY_0')))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),sK16),s(sK14,X0)))) ) <=> ~$spl2632),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2632])])).
fof(f7264,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK14,X0),s(fun(sK14,bool),'const.pred_set.EMPTY_0')))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),sK16),s(sK14,X0)))) ) | ($spl234 | $spl239 | $spl1128 | $spl1132)),
  inference(subsumption_resolution,[],[f7263,f610])).
fof(f7263,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK14,X0),s(fun(sK14,bool),'const.pred_set.EMPTY_0')))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),sK16),s(sK14,X0))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK14,'type.realax.real'),sK16),s(fun(sK14,bool),'const.pred_set.EMPTY_0')))) ) | ($spl234 | $spl1128 | $spl1132)),
  inference(subsumption_resolution,[],[f7175,f2811])).
fof(f7175,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK14,X0),s(fun(sK14,bool),'const.pred_set.EMPTY_0')))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),sK16),s(sK14,X0))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK14,'type.realax.real'),sK16),s(fun(sK14,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.bool.IN_2'(s(sK14,sK18(sK16,sK15,sK14)),s(fun(sK14,bool),'const.pred_set.EMPTY_0'))))) ) | ($spl234 | $spl1132)),
  inference(resolution,[],[f2838,f605])).
fof(f455,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK14,'type.realax.real'),sK16),s(fun(sK14,bool),sK15))) | $spl1),
  inference(cnf_transformation,[],[f455_D])).
fof(f455_D,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK14,'type.realax.real'),sK16),s(fun(sK14,bool),sK15))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f457,plain,(
  s(fun(sK14,bool),'const.pred_set.EMPTY_0') != s(fun(sK14,bool),sK15) | $spl3),
  inference(cnf_transformation,[],[f457_D])).
fof(f457_D,plain,(
  s(fun(sK14,bool),'const.pred_set.EMPTY_0') != s(fun(sK14,bool),sK15) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f257,plain,(
  s(fun(sK14,bool),'const.pred_set.EMPTY_0') != s(fun(sK14,bool),sK15) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK14,'type.realax.real'),sK16),s(fun(sK14,bool),sK15)))),
  inference(cnf_transformation,[],[f167])).
fof(f610,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK14,'type.realax.real'),sK16),s(fun(sK14,bool),'const.pred_set.EMPTY_0'))) | $spl239),
  inference(cnf_transformation,[],[f610_D])).
fof(f610_D,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK14,'type.realax.real'),sK16),s(fun(sK14,bool),'const.pred_set.EMPTY_0'))) <=> ~$spl239),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl239])])).
fof(f608,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK14,'type.realax.real'),sK16),s(fun(sK14,bool),'const.pred_set.EMPTY_0'))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f455,f458])).
fof(f605,plain,(
  ( ! [X4] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),sK16),s(sK14,X4)))))) | ~p(s(bool,'const.bool.IN_2'(s(sK14,X4),s(fun(sK14,bool),'const.pred_set.EMPTY_0'))))) ) | $spl234),
  inference(cnf_transformation,[],[f605_D])).
fof(f605_D,plain,(
  ( ! [X4] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),sK16),s(sK14,X4)))))) | ~p(s(bool,'const.bool.IN_2'(s(sK14,X4),s(fun(sK14,bool),'const.pred_set.EMPTY_0'))))) ) <=> ~$spl234),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl234])])).
fof(f601,plain,(
  ( ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(sK14,X4),s(fun(sK14,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),sK16),s(sK14,X4))))))) ) | ($spl2 | $spl8)),
  inference(backward_demodulation,[],[f458,f463])).
fof(f456,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK14,'type.realax.real'),sK16),s(fun(sK14,bool),sK15))) | $spl0),
  inference(cnf_transformation,[],[f456_D])).
fof(f456_D,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK14,'type.realax.real'),sK16),s(fun(sK14,bool),sK15))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f458,plain,(
  s(fun(sK14,bool),'const.pred_set.EMPTY_0') = s(fun(sK14,bool),sK15) | $spl2),
  inference(cnf_transformation,[],[f458_D])).
fof(f458_D,plain,(
  s(fun(sK14,bool),'const.pred_set.EMPTY_0') = s(fun(sK14,bool),sK15) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f258,plain,(
  s(fun(sK14,bool),'const.pred_set.EMPTY_0') = s(fun(sK14,bool),sK15) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK14,'type.realax.real'),sK16),s(fun(sK14,bool),sK15)))),
  inference(cnf_transformation,[],[f167])).
fof(f8036,plain,(
  $false | ($spl8 | $spl12 | $spl2851 | $spl2896)),
  inference(subsumption_resolution,[],[f8035,f467])).
fof(f8035,plain,(
  ~p(s(bool,'T_0')) | ($spl8 | $spl2851 | $spl2896)),
  inference(forward_demodulation,[],[f7990,f7952])).
fof(f7990,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK14,sK18(sK16,sK15,sK14)),s(fun(sK14,bool),sK15)))) | ($spl8 | $spl2851)),
  inference(resolution,[],[f7556,f463])).
fof(f7556,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),sK16),s(sK14,sK18(sK16,sK15,sK14))))))) | $spl2851),
  inference(cnf_transformation,[],[f7556_D])).
fof(f7556_D,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),sK16),s(sK14,sK18(sK16,sK15,sK14))))))) <=> ~$spl2851),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2851])])).
fof(f7557,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK14,X0),s(fun(sK14,bool),sK15)))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),sK16),s(sK14,X0)))) ) | $spl2852),
  inference(cnf_transformation,[],[f7557_D])).
fof(f7557_D,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK14,X0),s(fun(sK14,bool),sK15)))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),sK16),s(sK14,X0)))) ) <=> ~$spl2852),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2852])])).
fof(f7554,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),sK16),s(sK14,X0))) | ~p(s(bool,'const.bool.IN_2'(s(sK14,X0),s(fun(sK14,bool),sK15)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),sK16),s(sK14,sK18(sK16,sK15,sK14)))))))) ) | ($spl0 | $spl3 | $spl12 | $spl28 | $spl2770)),
  inference(subsumption_resolution,[],[f7553,f467])).
fof(f7553,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),sK16),s(sK14,X0))) | ~p(s(bool,'const.bool.IN_2'(s(sK14,X0),s(fun(sK14,bool),sK15)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),sK16),s(sK14,sK18(sK16,sK15,sK14)))))))) ) | ($spl0 | $spl3 | $spl28 | $spl2770)),
  inference(forward_demodulation,[],[f7552,f7440])).
fof(f7552,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),sK16),s(sK14,X0))) | ~p(s(bool,'const.bool.IN_2'(s(sK14,X0),s(fun(sK14,bool),sK15)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),sK16),s(sK14,sK18(sK16,sK15,sK14))))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK14,bool),sK15))))) ) | ($spl0 | $spl3 | $spl28)),
  inference(subsumption_resolution,[],[f7551,f457])).
fof(f7551,plain,(
  ( ! [X0] : (s(fun(sK14,bool),'const.pred_set.EMPTY_0') = s(fun(sK14,bool),sK15) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),sK16),s(sK14,X0))) | ~p(s(bool,'const.bool.IN_2'(s(sK14,X0),s(fun(sK14,bool),sK15)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),sK16),s(sK14,sK18(sK16,sK15,sK14))))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK14,bool),sK15))))) ) | ($spl0 | $spl28)),
  inference(trivial_inequality_removal,[],[f7545])).
fof(f7545,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(fun(sK14,bool),'const.pred_set.EMPTY_0') = s(fun(sK14,bool),sK15) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),sK16),s(sK14,X0))) | ~p(s(bool,'const.bool.IN_2'(s(sK14,X0),s(fun(sK14,bool),sK15)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),sK16),s(sK14,sK18(sK16,sK15,sK14))))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK14,bool),sK15))))) ) | ($spl0 | $spl28)),
  inference(superposition,[],[f484,f456])).
fof(f2811,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(sK14,sK18(X0,sK15,sK14)),s(fun(sK14,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.bool.IN_2'(s(sK14,X1),s(fun(sK14,bool),'const.pred_set.EMPTY_0')))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK14,'type.realax.real'),X0),s(fun(sK14,bool),'const.pred_set.EMPTY_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),X0),s(sK14,X1)))) ) | $spl1128),
  inference(cnf_transformation,[],[f2811_D])).
fof(f2811_D,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(sK14,sK18(X0,sK15,sK14)),s(fun(sK14,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.bool.IN_2'(s(sK14,X1),s(fun(sK14,bool),'const.pred_set.EMPTY_0')))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK14,'type.realax.real'),X0),s(fun(sK14,bool),'const.pred_set.EMPTY_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),X0),s(sK14,X1)))) ) <=> ~$spl1128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1128])])).
fof(f2810,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(sK14,sK18(X0,sK15,sK14)),s(fun(sK14,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.bool.IN_2'(s(sK14,X1),s(fun(sK14,bool),'const.pred_set.EMPTY_0')))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK14,'type.realax.real'),X0),s(fun(sK14,bool),'const.pred_set.EMPTY_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),X0),s(sK14,X1)))) ) | ($spl2 | $spl12 | $spl34 | $spl240)),
  inference(subsumption_resolution,[],[f2809,f467])).
fof(f2809,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK14,sK18(X0,sK15,sK14)),s(fun(sK14,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.bool.IN_2'(s(sK14,X1),s(fun(sK14,bool),'const.pred_set.EMPTY_0')))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK14,'type.realax.real'),X0),s(fun(sK14,bool),'const.pred_set.EMPTY_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),X0),s(sK14,X1)))) ) | ($spl2 | $spl34 | $spl240)),
  inference(forward_demodulation,[],[f2808,f616])).
fof(f2808,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK14,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,'const.bool.IN_2'(s(sK14,sK18(X0,sK15,sK14)),s(fun(sK14,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.bool.IN_2'(s(sK14,X1),s(fun(sK14,bool),'const.pred_set.EMPTY_0')))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK14,'type.realax.real'),X0),s(fun(sK14,bool),'const.pred_set.EMPTY_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),X0),s(sK14,X1)))) ) | ($spl2 | $spl34)),
  inference(forward_demodulation,[],[f2807,f458])).
fof(f2807,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(sK14,sK18(X0,sK15,sK14)),s(fun(sK14,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.bool.IN_2'(s(sK14,X1),s(fun(sK14,bool),'const.pred_set.EMPTY_0')))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK14,'type.realax.real'),X0),s(fun(sK14,bool),'const.pred_set.EMPTY_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),X0),s(sK14,X1))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK14,bool),sK15))))) ) | ($spl2 | $spl34)),
  inference(forward_demodulation,[],[f2806,f458])).
fof(f2806,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(sK14,X1),s(fun(sK14,bool),'const.pred_set.EMPTY_0')))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK14,'type.realax.real'),X0),s(fun(sK14,bool),'const.pred_set.EMPTY_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),X0),s(sK14,X1))) | p(s(bool,'const.bool.IN_2'(s(sK14,sK18(X0,sK15,sK14)),s(fun(sK14,bool),sK15)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK14,bool),sK15))))) ) | ($spl2 | $spl34)),
  inference(forward_demodulation,[],[f2805,f458])).
fof(f2805,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK14,'type.realax.real'),X0),s(fun(sK14,bool),'const.pred_set.EMPTY_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),X0),s(sK14,X1))) | ~p(s(bool,'const.bool.IN_2'(s(sK14,X1),s(fun(sK14,bool),sK15)))) | p(s(bool,'const.bool.IN_2'(s(sK14,sK18(X0,sK15,sK14)),s(fun(sK14,bool),sK15)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK14,bool),sK15))))) ) | ($spl2 | $spl34)),
  inference(forward_demodulation,[],[f2804,f458])).
fof(f2804,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK14,'type.realax.real'),X0),s(fun(sK14,bool),sK15))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),X0),s(sK14,X1))) | ~p(s(bool,'const.bool.IN_2'(s(sK14,X1),s(fun(sK14,bool),sK15)))) | p(s(bool,'const.bool.IN_2'(s(sK14,sK18(X0,sK15,sK14)),s(fun(sK14,bool),sK15)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK14,bool),sK15))))) ) | ($spl2 | $spl34)),
  inference(trivial_inequality_removal,[],[f2802])).
fof(f2802,plain,(
  ( ! [X0,X1] : (s(fun(sK14,bool),'const.pred_set.EMPTY_0') != s(fun(sK14,bool),'const.pred_set.EMPTY_0') | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK14,'type.realax.real'),X0),s(fun(sK14,bool),sK15))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),X0),s(sK14,X1))) | ~p(s(bool,'const.bool.IN_2'(s(sK14,X1),s(fun(sK14,bool),sK15)))) | p(s(bool,'const.bool.IN_2'(s(sK14,sK18(X0,sK15,sK14)),s(fun(sK14,bool),sK15)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK14,bool),sK15))))) ) | ($spl2 | $spl34)),
  inference(superposition,[],[f487,f458])).
fof(f618,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK14,sK17),s(fun(sK14,bool),'const.pred_set.EMPTY_0'))) | $spl242),
  inference(cnf_transformation,[],[f618_D])).
fof(f618_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK14,sK17),s(fun(sK14,bool),'const.pred_set.EMPTY_0'))) <=> ~$spl242),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl242])])).
fof(f613,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK14,sK17),s(fun(sK14,bool),'const.pred_set.EMPTY_0'))) | ($spl18 | $spl232)),
  inference(resolution,[],[f475,f604])).
fof(f604,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK14,sK17),s(fun(sK14,bool),'const.pred_set.EMPTY_0')))) | $spl232),
  inference(cnf_transformation,[],[f604_D])).
fof(f604_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK14,sK17),s(fun(sK14,bool),'const.pred_set.EMPTY_0')))) <=> ~$spl232),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl232])])).
fof(f600,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK14,sK17),s(fun(sK14,bool),'const.pred_set.EMPTY_0')))) | ($spl2 | $spl6)),
  inference(backward_demodulation,[],[f458,f462])).
fof(f616,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK14,bool),'const.pred_set.EMPTY_0'))) | $spl240),
  inference(cnf_transformation,[],[f616_D])).
fof(f616_D,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK14,bool),'const.pred_set.EMPTY_0'))) <=> ~$spl240),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl240])])).
fof(f612,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK14,bool),'const.pred_set.EMPTY_0'))) | ($spl18 | $spl236)),
  inference(resolution,[],[f475,f607])).
fof(f607,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK14,bool),'const.pred_set.EMPTY_0')))) | $spl236),
  inference(cnf_transformation,[],[f607_D])).
fof(f607_D,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK14,bool),'const.pred_set.EMPTY_0')))) <=> ~$spl236),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl236])])).
fof(f602,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK14,bool),'const.pred_set.EMPTY_0')))) | ($spl2 | $spl10)),
  inference(backward_demodulation,[],[f458,f465])).
fof(f2838,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(sK14,X1),s(fun(sK14,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),X0),s(sK14,sK18(X0,sK15,sK14))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),X0),s(sK14,X1))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK14,'type.realax.real'),X0),s(fun(sK14,bool),'const.pred_set.EMPTY_0')))) ) | $spl1132),
  inference(cnf_transformation,[],[f2838_D])).
fof(f2838_D,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(sK14,X1),s(fun(sK14,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),X0),s(sK14,sK18(X0,sK15,sK14))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),X0),s(sK14,X1))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK14,'type.realax.real'),X0),s(fun(sK14,bool),'const.pred_set.EMPTY_0')))) ) <=> ~$spl1132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1132])])).
fof(f2837,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(sK14,X1),s(fun(sK14,bool),'const.pred_set.EMPTY_0')))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK14,'type.realax.real'),X0),s(fun(sK14,bool),'const.pred_set.EMPTY_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),X0),s(sK14,X1))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),X0),s(sK14,sK18(X0,sK15,sK14)))))))) ) | ($spl2 | $spl12 | $spl32 | $spl240)),
  inference(subsumption_resolution,[],[f2836,f467])).
fof(f2836,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK14,X1),s(fun(sK14,bool),'const.pred_set.EMPTY_0')))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK14,'type.realax.real'),X0),s(fun(sK14,bool),'const.pred_set.EMPTY_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),X0),s(sK14,X1))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),X0),s(sK14,sK18(X0,sK15,sK14)))))))) ) | ($spl2 | $spl32 | $spl240)),
  inference(forward_demodulation,[],[f2835,f616])).
fof(f2835,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK14,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.bool.IN_2'(s(sK14,X1),s(fun(sK14,bool),'const.pred_set.EMPTY_0')))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK14,'type.realax.real'),X0),s(fun(sK14,bool),'const.pred_set.EMPTY_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),X0),s(sK14,X1))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),X0),s(sK14,sK18(X0,sK15,sK14)))))))) ) | ($spl2 | $spl32)),
  inference(forward_demodulation,[],[f2834,f458])).
fof(f2834,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(sK14,X1),s(fun(sK14,bool),'const.pred_set.EMPTY_0')))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK14,'type.realax.real'),X0),s(fun(sK14,bool),'const.pred_set.EMPTY_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),X0),s(sK14,X1))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),X0),s(sK14,sK18(X0,sK15,sK14))))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK14,bool),sK15))))) ) | ($spl2 | $spl32)),
  inference(forward_demodulation,[],[f2833,f458])).
fof(f2833,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK14,'type.realax.real'),X0),s(fun(sK14,bool),'const.pred_set.EMPTY_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),X0),s(sK14,X1))) | ~p(s(bool,'const.bool.IN_2'(s(sK14,X1),s(fun(sK14,bool),sK15)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),X0),s(sK14,sK18(X0,sK15,sK14))))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK14,bool),sK15))))) ) | ($spl2 | $spl32)),
  inference(forward_demodulation,[],[f2832,f458])).
fof(f2832,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK14,'type.realax.real'),X0),s(fun(sK14,bool),sK15))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),X0),s(sK14,X1))) | ~p(s(bool,'const.bool.IN_2'(s(sK14,X1),s(fun(sK14,bool),sK15)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),X0),s(sK14,sK18(X0,sK15,sK14))))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK14,bool),sK15))))) ) | ($spl2 | $spl32)),
  inference(trivial_inequality_removal,[],[f2830])).
fof(f2830,plain,(
  ( ! [X0,X1] : (s(fun(sK14,bool),'const.pred_set.EMPTY_0') != s(fun(sK14,bool),'const.pred_set.EMPTY_0') | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK14,'type.realax.real'),X0),s(fun(sK14,bool),sK15))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),X0),s(sK14,X1))) | ~p(s(bool,'const.bool.IN_2'(s(sK14,X1),s(fun(sK14,bool),sK15)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),X0),s(sK14,sK18(X0,sK15,sK14))))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK14,bool),sK15))))) ) | ($spl2 | $spl32)),
  inference(superposition,[],[f486,f458])).
fof(f7562,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK14,sK18(sK16,sK15,sK14)),s(fun(sK14,bool),sK15)))) | $spl2854),
  inference(cnf_transformation,[],[f7562_D])).
fof(f7562_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK14,sK18(sK16,sK15,sK14)),s(fun(sK14,bool),sK15)))) <=> ~$spl2854),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2854])])).
fof(f7560,plain,(
  ( ! [X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),sK16),s(sK14,X1))) | ~p(s(bool,'const.bool.IN_2'(s(sK14,X1),s(fun(sK14,bool),sK15)))) | p(s(bool,'const.bool.IN_2'(s(sK14,sK18(sK16,sK15,sK14)),s(fun(sK14,bool),sK15))))) ) | ($spl0 | $spl3 | $spl12 | $spl30 | $spl2770)),
  inference(subsumption_resolution,[],[f7559,f467])).
fof(f7559,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),sK16),s(sK14,X1))) | ~p(s(bool,'const.bool.IN_2'(s(sK14,X1),s(fun(sK14,bool),sK15)))) | p(s(bool,'const.bool.IN_2'(s(sK14,sK18(sK16,sK15,sK14)),s(fun(sK14,bool),sK15))))) ) | ($spl0 | $spl3 | $spl30 | $spl2770)),
  inference(forward_demodulation,[],[f7558,f7440])).
fof(f7558,plain,(
  ( ! [X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),sK16),s(sK14,X1))) | ~p(s(bool,'const.bool.IN_2'(s(sK14,X1),s(fun(sK14,bool),sK15)))) | p(s(bool,'const.bool.IN_2'(s(sK14,sK18(sK16,sK15,sK14)),s(fun(sK14,bool),sK15)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK14,bool),sK15))))) ) | ($spl0 | $spl3 | $spl30)),
  inference(subsumption_resolution,[],[f7550,f457])).
fof(f7550,plain,(
  ( ! [X1] : (s(fun(sK14,bool),'const.pred_set.EMPTY_0') = s(fun(sK14,bool),sK15) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),sK16),s(sK14,X1))) | ~p(s(bool,'const.bool.IN_2'(s(sK14,X1),s(fun(sK14,bool),sK15)))) | p(s(bool,'const.bool.IN_2'(s(sK14,sK18(sK16,sK15,sK14)),s(fun(sK14,bool),sK15)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK14,bool),sK15))))) ) | ($spl0 | $spl30)),
  inference(trivial_inequality_removal,[],[f7546])).
fof(f7546,plain,(
  ( ! [X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(fun(sK14,bool),'const.pred_set.EMPTY_0') = s(fun(sK14,bool),sK15) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun(sK14,'type.realax.real'),sK16),s(sK14,X1))) | ~p(s(bool,'const.bool.IN_2'(s(sK14,X1),s(fun(sK14,bool),sK15)))) | p(s(bool,'const.bool.IN_2'(s(sK14,sK18(sK16,sK15,sK14)),s(fun(sK14,bool),sK15)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK14,bool),sK15))))) ) | ($spl0 | $spl30)),
  inference(superposition,[],[f485,f456])).
fof(f475,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl18),
  inference(cnf_transformation,[],[f475_D])).
fof(f475_D,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f274,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f170])).
fof(f170,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f56])).
fof(f56,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X12] : (s(bool,'T_0') = s(bool,X12) <=> p(s(bool,X12)))),
  file('thm.real_sigma.REAL_SUM_IMAGE_NONZERO','thm.bool.EQ_CLAUSES|split|1')).
fof(f7952,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK14,sK18(sK16,sK15,sK14)),s(fun(sK14,bool),sK15))) | $spl2896),
  inference(cnf_transformation,[],[f7952_D])).
fof(f7952_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK14,sK18(sK16,sK15,sK14)),s(fun(sK14,bool),sK15))) <=> ~$spl2896),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2896])])).
fof(f7922,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK14,sK18(sK16,sK15,sK14)),s(fun(sK14,bool),sK15))) | ($spl18 | $spl2854)),
  inference(resolution,[],[f7562,f475])).
% SZS output end Proof for thm.real_sigma
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 377
% Passive clauses: 1474
% Generated clauses: 7812
% Final active clauses: 362
% Final passive clauses: 1040
% Input formulas: 43
% Initial clauses: 192
% 
% Duplicate literals: 216
% Trivial inequalities: 31
% Fw subsumption resolutions: 721
% Fw demodulations: 444
% Bw demodulations: 64
% Condensations: 222
% Global subsumptions: 14
% 
% Simple tautologies: 247
% Equational tautologies: 21
% Forward subsumptions: 2608
% 
% Binary resolution: 2078
% Factoring: 68
% Forward superposition: 784
% Backward superposition: 1418
% Self superposition: 28
% Equality factoring: 32
% Equality resolution: 12
% 
% Split clauses: 26
% Split components: 52
% SAT solver clauses: 1988
% SAT solver unit clauses: 206
% SAT solver binary clauses: 80
% SAT solver learnt clauses: 38
% SAT solver learnt literals: 35
% 
% Sat splits: 26
% Sat splitting refutations: 6
% 
% TWLsolver clauses: 1922
% TWLsolver calls for satisfiability: 2141
% 
% Memory used [KB]: 7931
% Time elapsed: 0.320 s
% ------------------------------
% ------------------------------
% Success in time 0.939 s
